Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% (110388)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 764
% Passive clauses: 8726
% Generated clauses: 21058
% Final active clauses: 733
% Final passive clauses: 6976
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 458
% Trivial inequalities: 85
% Fw subsumption resolutions: 1627
% Fw demodulations: 623
% Bw demodulations: 11
% 
% Simple tautologies: 229
% Equational tautologies: 73
% Forward subsumptions: 9259
% Fw demodulations to eq. taut.: 11
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 10802
% Factoring: 2537
% Forward superposition: 1678
% Backward superposition: 2708
% Self superposition: 45
% Equality factoring: 26
% Equality resolution: 32
% 
% Split clauses: 311
% Split components: 623
% SAT solver clauses: 636
% SAT solver unit clauses: 24
% SAT solver binary clauses: 109
% 
% Sat splits: 311
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 626
% TWLsolver calls for satisfiability: 432
% 
% Memory used [KB]: 9338
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (110403)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 575
% Passive clauses: 2713
% Generated clauses: 18355
% Final active clauses: 388
% Final passive clauses: 1250
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 265
% Trivial inequalities: 208
% Fw subsumption resolutions: 1236
% Fw demodulations: 2238
% Bw demodulations: 858
% Condensations: 427
% Global subsumptions: 22
% 
% Simple tautologies: 335
% Equational tautologies: 35
% Forward subsumptions: 7691
% Fw demodulations to eq. taut.: 25
% Bw demodulations to eq. taut.: 92
% 
% Binary resolution: 3112
% Factoring: 256
% Forward superposition: 3390
% Backward superposition: 2527
% Self superposition: 63
% Equality factoring: 28
% Equality resolution: 50
% 
% Split clauses: 280
% Split components: 576
% SAT solver clauses: 3050
% SAT solver unit clauses: 255
% SAT solver binary clauses: 119
% SAT solver learnt clauses: 48
% SAT solver learnt literals: 15
% 
% Sat splits: 280
% Sat splitting refutations: 16
% 
% TWLsolver clauses: 2958
% TWLsolver calls for satisfiability: 1968
% 
% Memory used [KB]: 12665
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f21641,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f306,f545,f543,f305,f544,f542,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f5004,f4920,f5006,f4924,f5008,f4927,f5010,f4930,f5012,f4933,f5014,f4936,f5016,f4939,f5018,f4942,f5020,f4945,f5022,f4948,f5024,f4951,f5026,f4954,f5028,f4957,f5030,f4960,f5032,f4963,f5034,f4966,f5036,f4921,f4969,f5259,f5260,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342,f5344,f5346,f5348,f5350,f5352,f5354,f5356,f5358,f5359,f5360,f5361,f5362,f5364,f5366,f5689,f5691,f5724,f5726,f5727,f5729,f5734,f5736,f5738,f5740,f5744,f5746,f5748,f5750,f5752,f5754,f5760,f5762,f5768,f5770,f5776,f5782,f5795,f5819,f5843,f5867,f5925,f5927,f5929,f5932,f5934,f5939,f5941,f5946,f5948,f5961,f5963,f5965,f5968,f5970,f5974,f5976,f6007,f6018,f6034,f6045,f6047,f6056,f6063,f6072,f6079,f6081,f6083,f6085,f6087,f6612,f6613,f6132,f6134,f6136,f6138,f6145,f6152,f6162,f6170,f6172,f6183,f6191,f6193,f6203,f6210,f6219,f6226,f6228,f6237,f6239,f6244,f6246,f6255,f6257,f6262,f6264,f6273,f6275,f6280,f6282,f6291,f6298,f6307,f6314,f6320,f6326,f6335,f6345,f6347,f6356,f6358,f6360,f6365,f6367,f6372,f6374,f6379,f6381,f6386,f6388,f6396,f6405,f6415,f6423,f6425,f6441,f6443,f6445,f6447,f6458,f6461,f6472,f6475,f6490,f6492,f6506,f6508,f6510,f6516,f6521,f6526,f6528,f6533,f6535,f6546,f6555,f6566,f6575,f6580,f6582,f6584,f6586,f6592,f6594,f6596,f6614,f6615,f6616,f6617,f6618,f6619,f6620,f6621,f6622,f6623,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6633,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6644,f6645,f6646,f6647,f6648,f6649,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f6670,f6671,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6680,f6681,f6682,f6683,f6684,f6685,f6686,f6687,f6688,f6689,f6690,f6691,f6692,f6693,f6694,f6695,f6696,f6697,f6698,f6699,f6700,f6701,f6702,f6703,f6704,f6705,f6706,f6707,f6708,f6709,f6710,f6711,f6712,f6713,f6714,f6715,f6716,f6717,f6718,f6719,f6720,f6721,f6722,f6723,f6724,f6725,f6726,f6727,f6729,f6730,f6731,f6732,f6733,f6734,f6735,f6736,f6737,f6738,f6739,f6740,f6741,f6742,f6743,f6744,f6745,f6746,f6747,f6748,f6749,f6750,f6751,f6752,f6753,f6754,f6755,f6756,f6757,f6758,f6759,f6760,f6761,f6762,f6763,f6764,f6765,f6766,f6767,f6768,f6769,f6770,f6771,f6772,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6781,f6782,f6783,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6791,f6792,f6793,f6794,f6795,f6796,f6797,f6798,f6799,f6800,f6801,f6802,f6803,f6804,f6805,f6806,f6807,f6808,f6809,f6810,f6811,f6812,f6813,f6814,f6815,f6816,f6817,f6818,f6819,f6820,f6821,f6822,f8670,f8672,f8671,f12789,f12791,f12790,f12975,f12982,f12983,f13126,f13130,f13133,f13135,f13128,f13644,f13645,f13652,f13649,f13646,f13647,f13655,f13648,f13657,f13650,f13654,f13664,f13665,f13677,f13675,f20740,f20742,f20741,f20744,f20745,f20778,f20780,f20779,f20782,f5692,f20783,f21325,f21346,f21359,f21363,f21365,f21369,f21373,f21377,f21380,f21384,f21388,f21392,f21396,f21400,f21404,f21408,f21412,f21416,f21420,f21424,f21428,f21432,f21436,f21440,f21444,f21448,f21452,f21455,f21457,f21458,f21598])).
fof(f21598,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21597,f546])).
fof(f546,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f310,f307])).
fof(f307,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11',aHL_TRUTH)).
fof(f310,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f38])).
fof(f38,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.bool.NOT_CLAUSES|split|2')).
fof(f21597,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21596,f15754])).
fof(f15754,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f600,f13011,f782])).
fof(f782,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f776,f546])).
fof(f776,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f406,f321])).
fof(f321,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11',aHL_BOOL_CASES)).
fof(f406,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.bool.IMP_ANTISYM_AX')).
fof(f13011,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f12987,f2227])).
fof(f2227,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2222,f547])).
fof(f547,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f318,f314])).
fof(f314,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.arithmetic.NUMERAL_DEF')).
fof(f318,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_MOD_REDUCE|split|0')).
fof(f2222,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(trivial_inequality_removal,[],[f2216])).
fof(f2216,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(superposition,[],[f413,f547])).
fof(f413,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) & (((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & (((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & ~s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X16,X17] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int',X17)))) <=> ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X17),s('type.integer.int',X16))) & ~s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X17) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_DIVIDES_MOD|48|')).
fof(f12987,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f12977,f416])).
fof(f416,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X7,X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_INJ')).
fof(f12977,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))) | $spl3),
  inference(forward_demodulation,[],[f12976,f314])).
fof(f12976,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))) | $spl3),
  inference(forward_demodulation,[],[f545,f314])).
fof(f600,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f546,f332])).
fof(f332,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.bool.EQ_CLAUSES|split|0')).
fof(f21596,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21317,f557])).
fof(f557,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f556,f314])).
fof(f556,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f403,f314])).
fof(f403,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X7,X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_MOD_REDUCE|split|3')).
fof(f21317,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(trivial_inequality_removal,[],[f21280])).
fof(f21280,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f2307,f13137])).
fof(f13137,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f13136,f2414])).
fof(f2414,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f2396,f416])).
fof(f2396,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(condensation,[],[f2395])).
fof(f2395,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9)) )),
  inference(subsumption_resolution,[],[f2394,f1642])).
fof(f1642,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f389,f420])).
fof(f420,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.integer_244_prove_ltb|split|1')).
fof(f389,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.arithmetic.SUC_NOT')).
fof(f2394,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9)) )),
  inference(forward_demodulation,[],[f2343,f400])).
fof(f400,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f2343,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))))))) )),
  inference(superposition,[],[f422,f553])).
fof(f553,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f552,f315])).
fof(f315,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.arithmetic.ADD_|48|')).
fof(f552,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(forward_demodulation,[],[f550,f400])).
fof(f550,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(backward_demodulation,[],[f400,f319])).
fof(f319,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.arithmetic.BIT2')).
fof(f422,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f13136,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f13083,f307])).
fof(f13083,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f410,f12988])).
fof(f12988,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f5690,f5687])).
fof(f5687,plain,(
  ( ! [X4] : (sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl4),
  inference(global_subsumption,[],[f306,f305,f307,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f389,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1439,f1438,f1437,f1436,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2300,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2848,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4123,f4122,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4442,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4677,f4676,f4535,f4534,f392,f4997,f393,f4998,f2296,f2220,f2307,f2227,f2879,f4092,f2224,f4110,f4174,f4156,f2880,f2872,f2869,f5369,f4996,f4995,f4455,f4419,f4416,f2792,f2789,f2718,f2715,f1470,f1469,f1468,f1467,f1466,f1465,f1464,f1463,f1462,f1461,f1460,f1459,f1458,f1457,f1456,f1455,f1454,f1453,f1452,f1413,f1412,f1411,f1410,f1409,f1408,f1407,f1406,f1405,f1404,f1403,f1402,f1401,f1400,f1399,f1398,f1397,f1396,f1395,f415,f414,f413,f412,f410,f399,f398,f4533,f4675,f546,f337,f336,f4920,f5686])).
fof(f5686,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl4),
  inference(subsumption_resolution,[],[f5677,f332])).
fof(f5677,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X5)) ) | $spl4),
  inference(superposition,[],[f4920,f731])).
fof(f336,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.bool.EQ_CLAUSES|split|3')).
fof(f337,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f208])).
fof(f4675,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,'T_0') != s(bool,f(s(X159,sK16(X160,X161,bool,X159)))) | s(fun(X159,bool),X160) = s(fun(X159,bool),X161) | s(bool,X162) = s(bool,g(s(X159,sK16(X160,X161,bool,X159)))) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f455,f731])).
fof(f4533,plain,(
  ( ! [X167,X169,X170,X168] : (s(bool,X170) != s(bool,f(s(X167,sK16(X168,X169,bool,X167)))) | s(fun(X167,bool),X168) = s(fun(X167,bool),X169) | s(bool,'T_0') = s(bool,X170) | s(bool,'T_0') = s(bool,g(s(X167,sK16(X168,X169,bool,X167))))) )),
  inference(superposition,[],[f455,f731])).
fof(f398,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X16,X17] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int',X17))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X17)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_DIVIDES_NEG|split|0')).
fof(f399,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X16,X17] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int',X17))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X16))),s('type.integer.int',X17)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_DIVIDES_NEG|split|1')).
fof(f412,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f250])).
fof(f414,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f415,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f1395,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int',X20))))) )),
  inference(superposition,[],[f431,f398])).
fof(f1396,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X23)))),X24) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int',X23))))) )),
  inference(superposition,[],[f432,f398])).
fof(f1397,plain,(
  ( ! [X26,X27,X25] : (sP1('const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X26)))),X27) | p(s(bool,X27)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int',X26))))) )),
  inference(superposition,[],[f447,f398])).
fof(f1398,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29))))) | p(s(bool,X30)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int',X29))))) )),
  inference(superposition,[],[f450,f398])).
fof(f1399,plain,(
  ( ! [X33,X31,X34,X32] : (~sP5(X33,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))),X34) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int',X32))))) )),
  inference(superposition,[],[f459,f398])).
fof(f1400,plain,(
  ( ! [X37,X35,X38,X36] : (~sP4(X37,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X36)))),X38) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int',X36))))) )),
  inference(superposition,[],[f463,f398])).
fof(f1401,plain,(
  ( ! [X39,X41,X42,X40] : (~sP3(X41,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40)))),X42) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int',X40))))) )),
  inference(superposition,[],[f467,f398])).
fof(f1402,plain,(
  ( ! [X45,X43,X46,X44] : (~sP6(X45,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X44)))),X46) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int',X44))))) )),
  inference(superposition,[],[f486,f398])).
fof(f1403,plain,(
  ( ! [X47,X50,X48,X49] : (~sP7(X50,X49,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X48))))) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int',X48))))) )),
  inference(superposition,[],[f495,f398])).
fof(f1404,plain,(
  ( ! [X54,X52,X53,X51] : (~sP8('const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X52)))),X53,X54) | p(s(bool,X53)) | p(s(bool,X54)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int',X52))))) )),
  inference(superposition,[],[f503,f398])).
fof(f1405,plain,(
  ( ! [X57,X58,X56,X55] : (sP9('const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X56)))),X58,X57) | ~p(s(bool,X57)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int',X56))))) )),
  inference(superposition,[],[f508,f398])).
fof(f1406,plain,(
  ( ! [X61,X59,X62,X60] : (sP9('const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X60)))),X61,X62) | ~p(s(bool,X61)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))))) )),
  inference(superposition,[],[f509,f398])).
fof(f1407,plain,(
  ( ! [X66,X64,X65,X63] : (~sP9('const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X64)))),X66,X65) | p(s(bool,X65)) | p(s(bool,X66)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64))))) )),
  inference(superposition,[],[f510,f398])).
fof(f1408,plain,(
  ( ! [X70,X68,X69,X67] : (~sP10(X70,X69,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X68))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int',X68))))) )),
  inference(superposition,[],[f520,f398])).
fof(f1409,plain,(
  ( ! [X74,X72,X71,X73] : (sP11('const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X72)))),X74,X73) | p(s(bool,X73)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int',X72))))) )),
  inference(superposition,[],[f525,f398])).
fof(f1410,plain,(
  ( ! [X78,X76,X77,X75] : (sP11('const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X76)))),X77,X78) | ~p(s(bool,X77)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int',X76))))) )),
  inference(superposition,[],[f526,f398])).
fof(f1411,plain,(
  ( ! [X80,X81,X79,X82] : (~sP11('const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X80)))),X82,X81) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int',X80))))) )),
  inference(superposition,[],[f527,f398])).
fof(f1412,plain,(
  ( ! [X85,X83,X86,X84] : (sP11('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | ~sP10('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int',X84))))) )),
  inference(superposition,[],[f530,f398])).
fof(f1413,plain,(
  ( ! [X90,X88,X87,X89] : (sP9('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int',X88))))) )),
  inference(superposition,[],[f531,f398])).
fof(f1452,plain,(
  ( ! [X23,X21,X22] : (sP0(X23,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X21))),s('type.integer.int',X22))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X21),s('type.integer.int',X22))))) )),
  inference(superposition,[],[f431,f399])).
fof(f1453,plain,(
  ( ! [X26,X24,X25] : (sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))),s('type.integer.int',X25)),X26) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X24),s('type.integer.int',X25))))) )),
  inference(superposition,[],[f432,f399])).
fof(f1454,plain,(
  ( ! [X28,X29,X27] : (sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27))),s('type.integer.int',X28)),X29) | p(s(bool,X29)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X27),s('type.integer.int',X28))))) )),
  inference(superposition,[],[f447,f399])).
fof(f1455,plain,(
  ( ! [X30,X31,X32] : (sP1(X32,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30))),s('type.integer.int',X31))) | p(s(bool,X32)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X30),s('type.integer.int',X31))))) )),
  inference(superposition,[],[f450,f399])).
fof(f1456,plain,(
  ( ! [X35,X33,X36,X34] : (~sP5(X35,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33))),s('type.integer.int',X34)),X36) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X33),s('type.integer.int',X34))))) )),
  inference(superposition,[],[f459,f399])).
fof(f1457,plain,(
  ( ! [X39,X37,X38,X40] : (~sP4(X39,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)),X40) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X37),s('type.integer.int',X38))))) )),
  inference(superposition,[],[f463,f399])).
fof(f1458,plain,(
  ( ! [X43,X41,X44,X42] : (~sP3(X43,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X41))),s('type.integer.int',X42)),X44) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X41),s('type.integer.int',X42))))) )),
  inference(superposition,[],[f467,f399])).
fof(f1459,plain,(
  ( ! [X47,X45,X48,X46] : (~sP6(X47,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X45))),s('type.integer.int',X46)),X48) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X45),s('type.integer.int',X46))))) )),
  inference(superposition,[],[f486,f399])).
fof(f1460,plain,(
  ( ! [X52,X50,X51,X49] : (~sP7(X52,X51,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X49))),s('type.integer.int',X50))) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X49),s('type.integer.int',X50))))) )),
  inference(superposition,[],[f495,f399])).
fof(f1461,plain,(
  ( ! [X54,X56,X55,X53] : (~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X53))),s('type.integer.int',X54)),X55,X56) | p(s(bool,X55)) | p(s(bool,X56)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X53),s('type.integer.int',X54))))) )),
  inference(superposition,[],[f503,f399])).
fof(f1462,plain,(
  ( ! [X59,X57,X60,X58] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X58)),X60,X59) | ~p(s(bool,X59)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X57),s('type.integer.int',X58))))) )),
  inference(superposition,[],[f508,f399])).
fof(f1463,plain,(
  ( ! [X61,X64,X62,X63] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X61))),s('type.integer.int',X62)),X63,X64) | ~p(s(bool,X63)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X61),s('type.integer.int',X62))))) )),
  inference(superposition,[],[f509,f399])).
fof(f1464,plain,(
  ( ! [X68,X66,X67,X65] : (~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X65))),s('type.integer.int',X66)),X68,X67) | p(s(bool,X67)) | p(s(bool,X68)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int',X66))))) )),
  inference(superposition,[],[f510,f399])).
fof(f1465,plain,(
  ( ! [X70,X72,X71,X69] : (~sP10(X72,X71,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X69))),s('type.integer.int',X70))) | p(s(bool,X71)) | ~p(s(bool,X72)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X69),s('type.integer.int',X70))))) )),
  inference(superposition,[],[f520,f399])).
fof(f1466,plain,(
  ( ! [X76,X74,X75,X73] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X73))),s('type.integer.int',X74)),X76,X75) | p(s(bool,X75)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X73),s('type.integer.int',X74))))) )),
  inference(superposition,[],[f525,f399])).
fof(f1467,plain,(
  ( ! [X80,X78,X79,X77] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X77))),s('type.integer.int',X78)),X79,X80) | ~p(s(bool,X79)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X77),s('type.integer.int',X78))))) )),
  inference(superposition,[],[f526,f399])).
fof(f1468,plain,(
  ( ! [X83,X81,X84,X82] : (~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X81))),s('type.integer.int',X82)),X84,X83) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X81),s('type.integer.int',X82))))) )),
  inference(superposition,[],[f527,f399])).
fof(f1469,plain,(
  ( ! [X88,X87,X85,X86] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X85),s('type.integer.int',X86))))) )),
  inference(superposition,[],[f530,f399])).
fof(f1470,plain,(
  ( ! [X92,X90,X91,X89] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X89),s('type.integer.int',X90))))) )),
  inference(superposition,[],[f531,f399])).
fof(f2715,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X3),s('type.integer.int',X4)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f410,f398])).
fof(f2718,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))) )),
  inference(superposition,[],[f410,f399])).
fof(f2789,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X3),s('type.integer.int',X4)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4)))) )),
  inference(superposition,[],[f412,f398])).
fof(f2792,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f412,f399])).
fof(f4416,plain,(
  ( ! [X30,X28,X29] : (s(bool,X30) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int',X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29)))))) )),
  inference(superposition,[],[f731,f398])).
fof(f4419,plain,(
  ( ! [X39,X37,X38] : (s(bool,X39) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X37),s('type.integer.int',X38))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)))) )),
  inference(superposition,[],[f731,f399])).
fof(f4455,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X21),s('type.integer.int',X22))) | s(bool,X23) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X21),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f731,f398])).
fof(f4995,plain,(
  ( ! [X127,X125,X126] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X126),s('type.integer.int',X125))) | s('type.integer.int',X126) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X127) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X125),s('type.integer.int',X126))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(subsumption_resolution,[],[f4664,f307])).
fof(f4664,plain,(
  ( ! [X127,X125,X126] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X126),s('type.integer.int',X125))) | s('type.integer.int',X126) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X127) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X125),s('type.integer.int',X126))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f412,f731])).
fof(f4996,plain,(
  ( ! [X130,X128,X129] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X129),s('type.integer.int',X128))) | s('type.integer.int',X128) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X130) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X128),s('type.integer.int',X129))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(subsumption_resolution,[],[f4665,f307])).
fof(f4665,plain,(
  ( ! [X130,X128,X129] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X129),s('type.integer.int',X128))) | s('type.integer.int',X128) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X130) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X128),s('type.integer.int',X129))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f410,f731])).
fof(f5369,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2709,f2869])).
fof(f2869,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2868,f547])).
fof(f2868,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))))) )),
  inference(subsumption_resolution,[],[f2849,f331])).
fof(f2849,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))))) )),
  inference(superposition,[],[f414,f547])).
fof(f2872,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4)))))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2851,f331])).
fof(f2851,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4)))))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f414,f331])).
fof(f2880,plain,(
  ( ! [X6,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) )),
  inference(subsumption_resolution,[],[f2863,f2414])).
fof(f2863,plain,(
  ( ! [X6,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) )),
  inference(superposition,[],[f414,f557])).
fof(f4156,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f548,f2414,f414])).
fof(f4174,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) )),
  inference(forward_demodulation,[],[f4159,f399])).
fof(f4159,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) )),
  inference(unit_resulting_resolution,[],[f2349,f2414,f413])).
fof(f4110,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) )),
  inference(forward_demodulation,[],[f4095,f399])).
fof(f4095,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) )),
  inference(unit_resulting_resolution,[],[f2349,f2401,f413])).
fof(f2224,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(forward_demodulation,[],[f2214,f548])).
fof(f2214,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(unit_resulting_resolution,[],[f1061,f548,f413])).
fof(f4092,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f547,f2401,f414])).
fof(f2879,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))))) )),
  inference(subsumption_resolution,[],[f2862,f2401])).
fof(f2862,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))))) )),
  inference(superposition,[],[f414,f555])).
fof(f2220,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(trivial_inequality_removal,[],[f2218])).
fof(f2218,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f413,f331])).
fof(f2296,plain,(
  ( ! [X4,X5] : (s('type.integer.int',X5) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(trivial_inequality_removal,[],[f2294])).
fof(f2294,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X5) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f415,f331])).
fof(f4998,plain,(
  ( ! [X154,X152,X155,X153] : (p(s(bool,happ(s(fun(X152,bool),X153),s(X152,X154)))) | s(bool,X155) = s(bool,happ(s(fun(X152,bool),X153),s(X152,sK15(X153,X152)))) | s(bool,'T_0') = s(bool,X155)) )),
  inference(subsumption_resolution,[],[f4674,f307])).
fof(f4674,plain,(
  ( ! [X154,X152,X155,X153] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X152,bool),X153),s(X152,X154)))) | s(bool,X155) = s(bool,happ(s(fun(X152,bool),X153),s(X152,sK15(X153,X152)))) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f393,f731])).
fof(f393,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f244])).
fof(f244,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f243])).
fof(f243,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X12] : ! [X14] : (~! [X8] : p(s(bool,happ(s(fun(X12,bool),X14),s(X12,X8)))) <=> ? [X8] : ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,X8)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.bool.NOT_FORALL_THM')).
fof(f4997,plain,(
  ( ! [X151,X149,X150,X148] : (p(s(bool,happ(s(fun(X148,bool),X149),s(X148,X150)))) | s(bool,X151) = s(bool,happ(s(fun(X148,bool),X149),s(X148,sK14(X149,X148)))) | s(bool,'T_0') = s(bool,X151)) )),
  inference(subsumption_resolution,[],[f4673,f307])).
fof(f4673,plain,(
  ( ! [X151,X149,X150,X148] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X148,bool),X149),s(X148,X150)))) | s(bool,X151) = s(bool,happ(s(fun(X148,bool),X149),s(X148,sK14(X149,X148)))) | s(bool,'T_0') = s(bool,X151)) )),
  inference(superposition,[],[f392,f731])).
fof(f392,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f4534,plain,(
  ( ! [X173,X171,X172] : (~p(s(bool,happ(s(fun(bool,bool),X171),s(bool,X172)))) | p(s(bool,happ(s(fun(bool,bool),X171),s(bool,X173)))) | s(bool,'T_0') = s(bool,X172) | s(bool,'T_0') = s(bool,sK14(X171,bool))) )),
  inference(superposition,[],[f392,f731])).
fof(f4535,plain,(
  ( ! [X175,X176,X174] : (~p(s(bool,happ(s(fun(bool,bool),X174),s(bool,X175)))) | p(s(bool,happ(s(fun(bool,bool),X174),s(bool,X176)))) | s(bool,'T_0') = s(bool,X175) | s(bool,'T_0') = s(bool,sK15(X174,bool))) )),
  inference(superposition,[],[f393,f731])).
fof(f4676,plain,(
  ( ! [X165,X163,X164] : (~p(s(bool,happ(s(fun(bool,bool),X163),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X163),s(bool,X164)))) | s(bool,X165) = s(bool,sK14(X163,bool)) | s(bool,'T_0') = s(bool,X165)) )),
  inference(superposition,[],[f392,f731])).
fof(f4677,plain,(
  ( ! [X167,X166,X168] : (~p(s(bool,happ(s(fun(bool,bool),X166),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X166),s(bool,X167)))) | s(bool,X168) = s(bool,sK15(X166,bool)) | s(bool,'T_0') = s(bool,X168)) )),
  inference(superposition,[],[f393,f731])).
fof(f4678,plain,(
  ( ! [X171,X169,X172,X170] : (s(X171,f(s(bool,'T_0'))) != s(X171,g(s(bool,'T_0'))) | s(fun(bool,X171),X169) = s(fun(bool,X171),X170) | s(bool,X172) = s(bool,sK16(X169,X170,X171,bool)) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f455,f731])).
fof(f4993,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f4650,f307])).
fof(f4650,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f531,f731])).
fof(f4992,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f4649,f307])).
fof(f4649,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f530,f731])).
fof(f4991,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f4648,f307])).
fof(f4648,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f527,f731])).
fof(f4990,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f4647,f307])).
fof(f4647,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f526,f731])).
fof(f4989,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f4646,f307])).
fof(f4646,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f525,f731])).
fof(f4988,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f4645,f307])).
fof(f4645,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f520,f731])).
fof(f4987,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f4644,f307])).
fof(f4644,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f510,f731])).
fof(f4986,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f4643,f307])).
fof(f4643,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f509,f731])).
fof(f4985,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f4642,f307])).
fof(f4642,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f508,f731])).
fof(f4984,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f4641,f307])).
fof(f4641,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f503,f731])).
fof(f4983,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f4640,f307])).
fof(f4640,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f495,f731])).
fof(f4982,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f4639,f307])).
fof(f4639,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f486,f731])).
fof(f4981,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f4638,f307])).
fof(f4638,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f467,f731])).
fof(f4980,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f4637,f307])).
fof(f4637,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f463,f731])).
fof(f4979,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4636,f307])).
fof(f4636,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f459,f731])).
fof(f4978,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f4635,f307])).
fof(f4635,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f450,f731])).
fof(f4977,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f4634,f307])).
fof(f4634,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f447,f731])).
fof(f4976,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f4633,f307])).
fof(f4633,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f432,f731])).
fof(f4975,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f4632,f307])).
fof(f4632,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f431,f731])).
fof(f4974,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f4631,f307])).
fof(f4631,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f409,f731])).
fof(f4973,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4627,f307])).
fof(f4627,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f336,f731])).
fof(f4594,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f332,f731])).
fof(f4536,plain,(
  ( ! [X177,X180,X178,X179] : (s(X179,f(s(bool,X180))) != s(X179,g(s(bool,X180))) | s(fun(bool,X179),X177) = s(fun(bool,X179),X178) | s(bool,'T_0') = s(bool,X180) | s(bool,'T_0') = s(bool,sK16(X177,X178,X179,bool))) )),
  inference(superposition,[],[f455,f731])).
fof(f4442,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f731,f731])).
fof(f4394,plain,(
  ~sP5('F_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f3485,f457])).
fof(f4389,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f3480,f457])).
fof(f4381,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) )),
  inference(forward_demodulation,[],[f4337,f1815])).
fof(f4337,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))))) )),
  inference(superposition,[],[f553,f1885])).
fof(f4336,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))))) )),
  inference(superposition,[],[f551,f1885])).
fof(f4335,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3))),s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s('type.num.num',X4)))) )),
  inference(superposition,[],[f401,f1885])).
fof(f4334,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f400,f1885])).
fof(f4313,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f572,f577,f3479,f460])).
fof(f4304,plain,(
  ~sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f3474,f458])).
fof(f4261,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X25))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X24))))) | s('type.num.num',X25) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24)))))) )),
  inference(superposition,[],[f416,f1788])).
fof(f4298,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))) )),
  inference(forward_demodulation,[],[f4249,f1815])).
fof(f4249,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))) )),
  inference(superposition,[],[f553,f1788])).
fof(f4248,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5)))))) )),
  inference(superposition,[],[f551,f1788])).
fof(f4247,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3))),s('type.num.num',X4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3))),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f401,f1788])).
fof(f4246,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f400,f1788])).
fof(f4190,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1658,f418])).
fof(f4122,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1642,f331])).
fof(f4123,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))))) )),
  inference(unit_resulting_resolution,[],[f1642,f418])).
fof(f4063,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f1439,f458])).
fof(f4010,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f1437,f2349])).
fof(f4007,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f1437,f458])).
fof(f3978,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f1436,f2349])).
fof(f3950,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f1384,f458])).
fof(f3903,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1382,f2349])).
fof(f3900,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f1382,f458])).
fof(f3857,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1381,f2349])).
fof(f3772,plain,(
  ( ! [X63] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X63))) | s('type.integer.int',X63) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f3694,f2349])).
fof(f3694,plain,(
  ( ! [X63] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X63))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X63) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f418,f2349])).
fof(f3658,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f418,f2349])).
fof(f3453,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2300,f409])).
fof(f3455,plain,(
  ( ! [X0] : (sP0(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2300,f431])).
fof(f3456,plain,(
  ( ! [X0] : (sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f2300,f432])).
fof(f3457,plain,(
  ~sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f2300,f446])).
fof(f3459,plain,(
  ~sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2300,f446])).
fof(f3460,plain,(
  ~sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2300,f2300,f446])).
fof(f3461,plain,(
  sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f546,f2300,f447])).
fof(f3462,plain,(
  sP1('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f2300,f450])).
fof(f3465,plain,(
  ~sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2300,f459])).
fof(f3466,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2300,f546,f2300,f459])).
fof(f3468,plain,(
  sP5('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f572,f569,f576,f2300,f460])).
fof(f3469,plain,(
  ~sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2300,f463])).
fof(f3474,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2300,f463])).
fof(f3475,plain,(
  ~sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2300,f546,f2300,f463])).
fof(f3479,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2300,f464])).
fof(f3480,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2300,f467])).
fof(f3485,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2300,f467])).
fof(f3486,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2300,f546,f2300,f467])).
fof(f3490,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2300,f468])).
fof(f3491,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2300,f472])).
fof(f3492,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) )),
  inference(unit_resulting_resolution,[],[f2300,f473])).
fof(f3493,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f2300,f474])).
fof(f3498,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f398,f2300,f476])).
fof(f3499,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f399,f2300,f476])).
fof(f3504,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f398,f2300,f476])).
fof(f3505,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f399,f2300,f476])).
fof(f3506,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f600,f2300,f477])).
fof(f3508,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f600,f2300,f477])).
fof(f3510,plain,(
  ( ! [X0] : (~sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f2300,f480])).
fof(f3511,plain,(
  ( ! [X0] : (~sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f546,f2300,f481])).
fof(f3518,plain,(
  ~sP6('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2300,f486])).
fof(f3523,plain,(
  ~sP6('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2300,f486])).
fof(f3524,plain,(
  ~sP6('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2300,f546,f2300,f486])).
fof(f3528,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2300,f487])).
fof(f3529,plain,(
  sP7('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f307,f2300,f492])).
fof(f3534,plain,(
  sP7('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f307,f2300,f492])).
fof(f3536,plain,(
  sP7('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2300,f2300,f492])).
fof(f3539,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f307,f2300,f492])).
fof(f3540,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f2300,f307,f2300,f492])).
fof(f3541,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2300,f2300,f492])).
fof(f3542,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2300,f2300,f2300,f492])).
fof(f3543,plain,(
  ~sP7('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2300,f495])).
fof(f3547,plain,(
  ~sP7('F_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f546,f2300,f495])).
fof(f3548,plain,(
  ~sP7('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2300,f546,f2300,f495])).
fof(f3551,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f546,f2300,f496])).
fof(f3553,plain,(
  ( ! [X0] : (~sP9('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f546,f2300,f497])).
fof(f3557,plain,(
  ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f2300,f503])).
fof(f3560,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2300,f504])).
fof(f3561,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) )),
  inference(unit_resulting_resolution,[],[f2300,f505])).
fof(f3562,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f307,f2300,f508])).
fof(f3566,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2300,f2300,f508])).
fof(f3567,plain,(
  ( ! [X0] : (sP9('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f307,f2300,f509])).
fof(f3571,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f2300,f2300,f509])).
fof(f3572,plain,(
  ~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f2300,f510])).
fof(f3573,plain,(
  ( ! [X0] : (~sP11('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f546,f2300,f514])).
fof(f3588,plain,(
  ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f2300,f520])).
fof(f3593,plain,(
  ~sP10('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f307,f2300,f520])).
fof(f3594,plain,(
  ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f2300,f2300,f520])).
fof(f3598,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) )),
  inference(unit_resulting_resolution,[],[f2300,f522])).
fof(f3599,plain,(
  sP11('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f546,f2300,f524])).
fof(f3602,plain,(
  ( ! [X0] : (sP11('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f307,f2300,f526])).
fof(f3606,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f2300,f2300,f526])).
fof(f3607,plain,(
  ~sP11('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f546,f2300,f527])).
fof(f3610,plain,(
  ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2300,f527])).
fof(f3611,plain,(
  ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2300,f546,f2300,f527])).
fof(f3612,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f584,f2300,f530])).
fof(f3614,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f585,f2300,f530])).
fof(f3615,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f581,f2300,f531])).
fof(f3616,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f582,f2300,f531])).
fof(f3409,plain,(
  ( ! [X21,X20] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X20))))) | s('type.num.num',X21) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))) )),
  inference(superposition,[],[f416,f1815])).
fof(f3399,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f551,f1815])).
fof(f3433,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) )),
  inference(forward_demodulation,[],[f3397,f1885])).
fof(f3397,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4)))))))) )),
  inference(superposition,[],[f553,f1815])).
fof(f3396,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f400,f1815])).
fof(f3395,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num',X1)))))) )),
  inference(superposition,[],[f401,f1815])).
fof(f3338,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f416,f1573])).
fof(f3328,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f551,f1573])).
fof(f3365,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f3364,f1046])).
fof(f3364,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f3326,f1885])).
fof(f3326,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f553,f1573])).
fof(f3363,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f3325,f1810])).
fof(f3325,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f400,f1573])).
fof(f3362,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f3324,f1881])).
fof(f3324,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f401,f1573])).
fof(f3141,plain,(
  ( ! [X39,X38,X40] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X39))))))))) | s('type.integer.int',X40) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X39))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X39)))))))))) )),
  inference(superposition,[],[f418,f559])).
fof(f3128,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))) )),
  inference(superposition,[],[f331,f559])).
fof(f2953,plain,(
  ( ! [X39,X38,X40] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39))))))))) | s('type.integer.int',X40) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39)))))))))) )),
  inference(superposition,[],[f418,f558])).
fof(f2940,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(superposition,[],[f331,f558])).
fof(f2848,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f1062,f547,f414])).
fof(f2709,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f410,f331])).
fof(f2606,plain,(
  ( ! [X24,X23,X22] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X23))))))))) | s('type.integer.int',X24) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X23)))))))) )),
  inference(superposition,[],[f418,f557])).
fof(f2595,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f557,f548])).
fof(f2489,plain,(
  ( ! [X24,X23,X22] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X23))))))))) | s('type.integer.int',X24) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X23)))))))) )),
  inference(superposition,[],[f418,f555])).
fof(f2478,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f555,f547])).
fof(f2411,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f2396,f421])).
fof(f2412,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f2396,f420])).
fof(f2398,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f2393,f421])).
fof(f2399,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f2393,f420])).
fof(f2401,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f2393,f416])).
fof(f2393,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(condensation,[],[f2392])).
fof(f2392,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(subsumption_resolution,[],[f2391,f1642])).
fof(f2391,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(forward_demodulation,[],[f2342,f400])).
fof(f2342,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))))) )),
  inference(superposition,[],[f422,f551])).
fof(f2354,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))) | s('type.integer.int',X11) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X10))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f2349,f2138])).
fof(f2353,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))))))) )),
  inference(backward_demodulation,[],[f2349,f1999])).
fof(f2352,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8)))))))) )),
  inference(backward_demodulation,[],[f2349,f1932])).
fof(f2349,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f2337,f315])).
fof(f2337,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f316,f316,f422])).
fof(f2300,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f2299,f548])).
fof(f2299,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f2291,f548])).
fof(f2291,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))))) )),
  inference(unit_resulting_resolution,[],[f548,f548,f415])).
fof(f2223,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f2215,f548])).
fof(f2215,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f1062,f548,f413])).
fof(f2140,plain,(
  ( ! [X14,X15] : (s('type.integer.int',X15) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14))))) | s('type.integer.int',X14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f419,f331])).
fof(f2138,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X11) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X10))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f418,f331])).
fof(f2129,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1061,f331])).
fof(f1999,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f418,f548])).
fof(f1932,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f418,f547])).
fof(f1885,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f1884,f1815])).
fof(f1884,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3)))))))) )),
  inference(forward_demodulation,[],[f1883,f1788])).
fof(f1883,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3)))))))))) )),
  inference(forward_demodulation,[],[f1882,f400])).
fof(f1882,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))))) )),
  inference(forward_demodulation,[],[f1856,f400])).
fof(f1856,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))))) )),
  inference(superposition,[],[f401,f551])).
fof(f1881,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f1845,f316])).
fof(f1845,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(superposition,[],[f401,f1046])).
fof(f1815,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f1814,f1788])).
fof(f1814,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f1789,f400])).
fof(f1789,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f400,f553])).
fof(f1812,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f1788,f1570])).
fof(f1811,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))))))) | s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))) )),
  inference(backward_demodulation,[],[f1788,f1577])).
fof(f1788,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f400,f551])).
fof(f1810,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f1787,f315])).
fof(f1787,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f400,f1046])).
fof(f1665,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))) )),
  inference(superposition,[],[f421,f316])).
fof(f1664,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f421,f553])).
fof(f1663,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f421,f551])).
fof(f1662,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f421,f315])).
fof(f1660,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f421,f314])).
fof(f1658,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f389,f421])).
fof(f1659,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1053,f421])).
fof(f1656,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))) )),
  inference(superposition,[],[f420,f316])).
fof(f1655,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f420,f553])).
fof(f1654,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f420,f551])).
fof(f1653,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f420,f315])).
fof(f1643,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(unit_resulting_resolution,[],[f1053,f420])).
fof(f1577,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X8))))))))) | s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))) )),
  inference(forward_demodulation,[],[f1576,f400])).
fof(f1576,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))))) | s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))) )),
  inference(forward_demodulation,[],[f1564,f400])).
fof(f1564,plain,(
  ( ! [X8,X9] : (s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))))) )),
  inference(superposition,[],[f417,f553])).
fof(f1562,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) )),
  inference(superposition,[],[f416,f553])).
fof(f1573,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1556,f1046])).
fof(f1556,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f553,f316])).
fof(f1555,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f553,f1046])).
fof(f1570,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f1569,f400])).
fof(f1569,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f1547,f400])).
fof(f1547,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f553,f417])).
fof(f1436,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f546,f399,f478])).
fof(f1437,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f307,f399,f476])).
fof(f1438,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f546,f399,f478])).
fof(f1439,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f307,f399,f476])).
fof(f1381,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f546,f398,f478])).
fof(f1382,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f307,f398,f476])).
fof(f1383,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f546,f398,f478])).
fof(f1384,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f307,f398,f476])).
fof(f1317,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f574,f570,f575,f546,f461])).
fof(f1283,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f572,f569,f576,f307,f460])).
fof(f1246,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1061,f418])).
fof(f1247,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1062,f418])).
fof(f1225,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))))) | s('type.num.num',X7) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) )),
  inference(forward_demodulation,[],[f1210,f400])).
fof(f1210,plain,(
  ( ! [X6,X7] : (s('type.num.num',X7) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) )),
  inference(superposition,[],[f417,f551])).
fof(f1219,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f1204,f400])).
fof(f1204,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f551,f417])).
fof(f1068,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) )),
  inference(superposition,[],[f416,f316])).
fof(f1067,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) | s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f416,f551])).
fof(f1066,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f416,f315])).
fof(f1063,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f416,f314])).
fof(f1061,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f389,f416])).
fof(f1062,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f1053,f416])).
fof(f1053,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f389,f1046])).
fof(f1046,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f551,f316])).
fof(f1009,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1000,f515])).
fof(f1000,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f307,f520])).
fof(f989,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f586,f307,f546,f519])).
fof(f958,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f950,f498])).
fof(f950,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f307,f503])).
fof(f944,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f583,f546,f546,f502])).
fof(f916,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f907,f479])).
fof(f907,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f307,f486])).
fof(f896,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f578,f307,f307,f482])).
fof(f892,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f887,f456])).
fof(f887,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f546,f471])).
fof(f883,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f874,f457])).
fof(f874,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f307,f467])).
fof(f857,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f307,f463])).
fof(f813,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f600,f477])).
fof(f817,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f600,f477])).
fof(f797,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f546,f600,f475])).
fof(f799,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f600,f475])).
fof(f731,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f321,f321])).
fof(f681,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f307,f526])).
fof(f673,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f307,f525])).
fof(f671,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f546,f514])).
fof(f670,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f546,f513])).
fof(f662,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f307,f509])).
fof(f658,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f307,f508])).
fof(f655,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f546,f497])).
fof(f651,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f546,f496])).
fof(f650,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f546,f546,f491])).
fof(f644,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f546,f490])).
fof(f630,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f546,f481])).
fof(f627,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f307,f546,f480])).
fof(f622,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f560])).
fof(f619,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f450])).
fof(f614,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f447])).
fof(f611,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f307,f446])).
fof(f610,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f445])).
fof(f586,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f523])).
fof(f585,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f307,f522])).
fof(f584,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f521])).
fof(f583,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f506])).
fof(f582,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f307,f505])).
fof(f581,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f504])).
fof(f580,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f489])).
fof(f579,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f546,f488])).
fof(f578,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f487])).
fof(f577,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f307,f474])).
fof(f576,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f307,f473])).
fof(f575,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f472])).
fof(f574,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f470])).
fof(f573,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f546,f469])).
fof(f572,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f468])).
fof(f571,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f466])).
fof(f570,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f546,f465])).
fof(f569,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f464])).
fof(f568,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f432])).
fof(f567,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f431])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f183,f200,f199])).
fof(f199,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f200,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X17)) => p(s(bool,X18)))) <=> ((p(s(bool,X16)) | p(s(bool,X17))) & (p(s(bool,X16)) | ~p(s(bool,X18))) & (~p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.sat.sat_27_prove_ltb')).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f301])).
fof(f301,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f199])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f302])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f302])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f302])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f298])).
fof(f298,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f200])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f299])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f299])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f516])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f299])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f299])).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f157,f197,f196])).
fof(f196,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f157,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X17)) | p(s(bool,X18)))) <=> ((p(s(bool,X16)) | ~p(s(bool,X17))) & (p(s(bool,X16)) | ~p(s(bool,X18))) & (p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.sat.sat_25_prove_ltb')).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f293])).
fof(f293,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f196])).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f294])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f294])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f294])).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f290])).
fof(f290,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f291])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f291])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f499])).
fof(f499,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f155,f194,f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f194,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X17)) & p(s(bool,X18)))) <=> ((p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18))) & (p(s(bool,X17)) | ~p(s(bool,X16))) & (p(s(bool,X18)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.sat.sat_23_prove_ltb')).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f288])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f288])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f288])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f193])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f286])).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f286])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f283])).
fof(f283,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f194])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f284])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f284])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f153,f191,f190,f189,f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f189,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f190,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f191,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> s(bool,X17) = s(bool,X18)) <=> ((p(s(bool,X16)) | p(s(bool,X17)) | p(s(bool,X18))) & (p(s(bool,X16)) | ~p(s(bool,X18)) | ~p(s(bool,X17))) & (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X16))) & (p(s(bool,X18)) | ~p(s(bool,X17)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.sat.sat_2|49|_prove_ltb')).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f281])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f281])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f281])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f188])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f280])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f280])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f277])).
fof(f277,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f189])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f278])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f278])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f278])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f274])).
fof(f274,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f190])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f275])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f275])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f275])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f271])).
fof(f271,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f191])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f272])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f272])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f455,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11',aHL_EXT)).
fof(f445,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f266])).
fof(f266,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f446,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f267])).
fof(f447,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f450,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f560,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f430,f546])).
fof(f430,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f431,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f432,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f421,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f254])).
fof(f418,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X8,X13] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X8))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X13))) <=> s('type.integer.int',X8) = s('type.integer.int',X13))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_EQ_NEG')).
fof(f419,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f252])).
fof(f417,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f251])).
fof(f409,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f559,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))) )),
  inference(forward_demodulation,[],[f405,f314])).
fof(f405,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X8,X6] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_MOD_REDUCE|split|4')).
fof(f558,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(forward_demodulation,[],[f404,f314])).
fof(f404,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X8,X6] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_MOD_REDUCE|split|5')).
fof(f555,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f554,f314])).
fof(f554,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f402,f314])).
fof(f402,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X7,X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_MOD_REDUCE|split|2')).
fof(f401,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X6)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.arithmetic.ADD_CLAUSES|split|2')).
fof(f551,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f549,f315])).
fof(f549,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(backward_demodulation,[],[f400,f317])).
fof(f333,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f331,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : (~s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X16] : (~s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_MOD|48|')).
fof(f548,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f320,f314])).
fof(f320,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_MOD_REDUCE|split|1')).
fof(f317,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.arithmetic.BIT|49|')).
fof(f316,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.arithmetic.ADD_CLAUSES|split|0')).
fof(f5690,plain,(
  ~sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl41),
  inference(cnf_transformation,[],[f5690_D])).
fof(f5690_D,plain,(
  ~sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f410,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f250])).
fof(f2307,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) )),
  inference(forward_demodulation,[],[f2298,f547])).
fof(f2298,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(trivial_inequality_removal,[],[f2292])).
fof(f2292,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(superposition,[],[f415,f547])).
fof(f21458,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21161,f12987])).
fof(f21161,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f557,f13137])).
fof(f21457,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21160,f12987])).
fof(f21160,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f13137,f557])).
fof(f21455,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21454,f12985])).
fof(f12985,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f12977,f420])).
fof(f21454,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21453,f2349])).
fof(f21453,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21118,f557])).
fof(f21118,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13137,f419])).
fof(f21452,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21451,f12985])).
fof(f21451,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21450,f2349])).
fof(f21450,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21449,f557])).
fof(f21449,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21119,f548])).
fof(f21119,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2414,f13137,f2140])).
fof(f21448,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21447,f12985])).
fof(f21447,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21446,f2349])).
fof(f21446,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21445,f557])).
fof(f21445,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21120,f13015])).
fof(f13015,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f2349,f12987,f2354])).
fof(f21120,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12987,f13137,f2140])).
fof(f21444,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21443,f12985])).
fof(f21443,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21442,f2349])).
fof(f21442,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21441,f547])).
fof(f21441,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21121,f557])).
fof(f21121,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1062,f13137,f2140])).
fof(f21440,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21439,f12985])).
fof(f21439,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21438,f2349])).
fof(f21438,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21437,f557])).
fof(f21437,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21122,f547])).
fof(f21122,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2401,f13137,f2140])).
fof(f21436,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21435,f12985])).
fof(f21435,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21434,f2349])).
fof(f21434,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21433,f557])).
fof(f21433,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21123,f2129])).
fof(f21123,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1061,f13137,f2140])).
fof(f21432,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21431,f12985])).
fof(f21431,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21430,f2349])).
fof(f21430,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21429,f557])).
fof(f21429,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21124,f19550])).
fof(f19550,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(forward_demodulation,[],[f19518,f4122])).
fof(f19518,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(unit_resulting_resolution,[],[f1642,f2412,f5369])).
fof(f21124,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2412,f13137,f2140])).
fof(f21428,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21427,f12985])).
fof(f21427,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21426,f2349])).
fof(f21426,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21425,f557])).
fof(f21425,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21125,f13204])).
fof(f13204,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f2349,f12985,f2354])).
fof(f21125,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12985,f13137,f2140])).
fof(f21424,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21423,f12985])).
fof(f21423,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21422,f2349])).
fof(f21422,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21421,f557])).
fof(f21421,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21126,f19269])).
fof(f19269,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) ) | $spl3),
  inference(forward_demodulation,[],[f19236,f13204])).
fof(f19236,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f12985,f2399,f5369])).
fof(f21126,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2399,f13137,f2140])).
fof(f21420,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21419,f12985])).
fof(f21419,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21418,f2349])).
fof(f21418,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21417,f557])).
fof(f21417,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21127,f4122])).
fof(f21127,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1642,f13137,f2140])).
fof(f21416,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21415,f546])).
fof(f21415,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21414,f15754])).
fof(f21414,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21413,f548])).
fof(f21413,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21128,f557])).
fof(f21128,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2414,f13137,f2296])).
fof(f21412,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21411,f546])).
fof(f21411,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21410,f15754])).
fof(f21410,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21409,f13015])).
fof(f21409,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21129,f557])).
fof(f21129,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12987,f13137,f2296])).
fof(f21408,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21407,f546])).
fof(f21407,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21406,f15754])).
fof(f21406,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21405,f547])).
fof(f21405,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21130,f557])).
fof(f21130,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1062,f13137,f2296])).
fof(f21404,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21403,f546])).
fof(f21403,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21402,f15754])).
fof(f21402,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21401,f547])).
fof(f21401,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21131,f557])).
fof(f21131,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2401,f13137,f2296])).
fof(f21400,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21399,f546])).
fof(f21399,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21398,f15754])).
fof(f21398,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21397,f2129])).
fof(f21397,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21132,f557])).
fof(f21132,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1061,f13137,f2296])).
fof(f21396,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21395,f546])).
fof(f21395,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21394,f15754])).
fof(f21394,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21393,f19550])).
fof(f21393,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21133,f557])).
fof(f21133,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2412,f13137,f2296])).
fof(f21392,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21391,f546])).
fof(f21391,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21390,f15754])).
fof(f21390,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21389,f13204])).
fof(f21389,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21134,f557])).
fof(f21134,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12985,f13137,f2296])).
fof(f21388,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21387,f546])).
fof(f21387,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21386,f15754])).
fof(f21386,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21385,f19269])).
fof(f21385,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21135,f557])).
fof(f21135,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2399,f13137,f2296])).
fof(f21384,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21383,f546])).
fof(f21383,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21382,f15754])).
fof(f21382,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21381,f4122])).
fof(f21381,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21136,f557])).
fof(f21136,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1642,f13137,f2296])).
fof(f21380,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21379,f546])).
fof(f21379,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21378,f15754])).
fof(f21378,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21137,f557])).
fof(f21137,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13137,f2307])).
fof(f21377,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21376,f546])).
fof(f21376,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21375,f15754])).
fof(f21375,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21374,f547])).
fof(f21374,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21138,f557])).
fof(f21138,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f547,f13137,f415])).
fof(f21373,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21372,f546])).
fof(f21372,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21371,f15754])).
fof(f21371,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21370,f548])).
fof(f21370,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21139,f557])).
fof(f21139,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f548,f13137,f415])).
fof(f21369,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21368,f546])).
fof(f21368,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21367,f15754])).
fof(f21367,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21366,f13015])).
fof(f21366,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21140,f557])).
fof(f21140,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13015,f13137,f415])).
fof(f21365,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21364,f21331])).
fof(f21331,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21330,f557])).
fof(f21330,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21156,f398])).
fof(f21156,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12985,f13137,f413])).
fof(f21364,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21141,f557])).
fof(f21141,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13137,f13137,f415])).
fof(f21363,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21362,f546])).
fof(f21362,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21361,f15754])).
fof(f21361,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21360,f13204])).
fof(f21360,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21142,f557])).
fof(f21142,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13204,f13137,f415])).
fof(f21359,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21358,f546])).
fof(f21358,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21357,f15754])).
fof(f21357,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21356,f557])).
fof(f21356,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21143,f399])).
fof(f21143,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2349,f13137,f415])).
fof(f21346,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21345,f21331])).
fof(f21345,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21147,f557])).
fof(f21147,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13137,f13137,f415])).
fof(f21325,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21324,f12985])).
fof(f21324,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21323,f2349])).
fof(f21323,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21159,f557])).
fof(f21159,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13137,f419])).
fof(f20783,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl78),
  inference(cnf_transformation,[],[f20783_D])).
fof(f20783_D,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5692,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl42),
  inference(cnf_transformation,[],[f5692_D])).
fof(f5692_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f20782,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f20767,f332])).
fof(f20767,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f4978,f321])).
fof(f20779,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl76),
  inference(cnf_transformation,[],[f20779_D])).
fof(f20779_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f20780,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f20763,f332])).
fof(f20763,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f4978,f731])).
fof(f20778,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f20762,f332])).
fof(f20762,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f4978,f731])).
fof(f20745,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl74),
  inference(cnf_transformation,[],[f20745_D])).
fof(f20745_D,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f20744,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f20729,f332])).
fof(f20729,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f4977,f321])).
fof(f20741,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl72),
  inference(cnf_transformation,[],[f20741_D])).
fof(f20741_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f20742,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f20725,f332])).
fof(f20725,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f4977,f731])).
fof(f20740,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f20724,f332])).
fof(f20724,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f4977,f731])).
fof(f13675,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) | $spl68),
  inference(cnf_transformation,[],[f13675_D])).
fof(f13675_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f13677,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl71),
  inference(cnf_transformation,[],[f13677_D])).
fof(f13677_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f13665,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK16(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2427,f731])).
fof(f2427,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(X3,sK16(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,g(s(X3,sK16(X4,X5,bool,X3))))) )),
  inference(superposition,[],[f455,f321])).
fof(f13664,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2427,f731])).
fof(f13654,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl62),
  inference(cnf_transformation,[],[f13654_D])).
fof(f13654_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f13650,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool))) ) | $spl58),
  inference(cnf_transformation,[],[f13650_D])).
fof(f13650_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f13657,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl67),
  inference(cnf_transformation,[],[f13657_D])).
fof(f13657_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f13648,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK16(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2425,f321])).
fof(f2425,plain,(
  ( ! [X4,X5,X3] : (s(X5,f(s(bool,'F_0'))) != s(X5,g(s(bool,'F_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,'T_0') = s(bool,sK16(X3,X4,X5,bool))) )),
  inference(superposition,[],[f455,f321])).
fof(f13655,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl64),
  inference(cnf_transformation,[],[f13655_D])).
fof(f13655_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13647,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK16(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2425,f731])).
fof(f13646,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK16(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2425,f731])).
fof(f13649,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl56),
  inference(cnf_transformation,[],[f13649_D])).
fof(f13649_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f13652,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl61),
  inference(cnf_transformation,[],[f13652_D])).
fof(f13652_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f13645,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK16(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2425,f731])).
fof(f13644,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2425,f731])).
fof(f13128,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | $spl50),
  inference(cnf_transformation,[],[f13128_D])).
fof(f13128_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f13135,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | $spl54),
  inference(cnf_transformation,[],[f13135_D])).
fof(f13135_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f13133,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f13082,f307])).
fof(f13082,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(superposition,[],[f412,f12988])).
fof(f13130,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | $spl52),
  inference(cnf_transformation,[],[f13130_D])).
fof(f13130_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f13126,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f13080,f307])).
fof(f13080,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(superposition,[],[f2792,f12988])).
fof(f12983,plain,(
  $false | ($spl0 | $spl40)),
  inference(subsumption_resolution,[],[f7022,f12981])).
fof(f12981,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f12980,f314])).
fof(f12980,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f12979,f314])).
fof(f12979,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))) | $spl0),
  inference(forward_demodulation,[],[f12978,f398])).
fof(f12978,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) | $spl0),
  inference(forward_demodulation,[],[f542,f399])).
fof(f7022,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl40),
  inference(unit_resulting_resolution,[],[f307,f5691,f446])).
fof(f12982,plain,(
  $false | ($spl0 | $spl44)),
  inference(subsumption_resolution,[],[f12981,f7024])).
fof(f7024,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl44),
  inference(unit_resulting_resolution,[],[f307,f5726,f446])).
fof(f12975,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f12974,f546])).
fof(f12974,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f12936,f607])).
fof(f607,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f566,f337])).
fof(f566,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f565,f314])).
fof(f565,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f564,f314])).
fof(f564,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))) | $spl1),
  inference(forward_demodulation,[],[f563,f399])).
fof(f563,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))) | $spl1),
  inference(forward_demodulation,[],[f543,f398])).
fof(f12936,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl2),
  inference(trivial_inequality_removal,[],[f12935])).
fof(f12935,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl2),
  inference(superposition,[],[f2880,f562])).
fof(f562,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))) | $spl2),
  inference(forward_demodulation,[],[f561,f314])).
fof(f561,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))) | $spl2),
  inference(forward_demodulation,[],[f544,f314])).
fof(f12790,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13))))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl48),
  inference(cnf_transformation,[],[f12790_D])).
fof(f12790_D,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13))))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f12791,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | s('type.integer.int',X15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))))) )),
  inference(subsumption_resolution,[],[f12716,f332])).
fof(f12716,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | s('type.integer.int',X15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))))) )),
  inference(superposition,[],[f2872,f731])).
fof(f12789,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X14)) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13)))))) )),
  inference(subsumption_resolution,[],[f12715,f332])).
fof(f12715,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X14)) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13))))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f2872,f731])).
fof(f8671,plain,(
  ( ! [X7] : (s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl46),
  inference(cnf_transformation,[],[f8671_D])).
fof(f8671_D,plain,(
  ( ! [X7] : (s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f8672,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f8661,f332])).
fof(f8661,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f2869,f731])).
fof(f8670,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f8660,f332])).
fof(f8660,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2869,f731])).
fof(f6822,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5730,f5742])).
fof(f5742,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f5692,f409])).
fof(f5730,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f336])).
fof(f6821,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5733,f307])).
fof(f5733,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f600,f5692,f409])).
fof(f6820,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5735,f307])).
fof(f5735,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f601,f5692,f409])).
fof(f601,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f566,f332])).
fof(f6819,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5739,f5692])).
fof(f5739,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f601,f5692,f409])).
fof(f6818,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5743,f307])).
fof(f5743,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f600,f5692,f409])).
fof(f6817,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5745,f307])).
fof(f5745,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f601,f5692,f409])).
fof(f6816,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5749,f5692])).
fof(f5749,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f601,f5692,f409])).
fof(f6815,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5756,f307])).
fof(f5756,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f619,f5692,f446])).
fof(f6814,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5757,f5692])).
fof(f5757,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3462,f5692,f446])).
fof(f6813,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5763,f307])).
fof(f5763,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f614,f5692,f446])).
fof(f6812,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5766,f5692])).
fof(f5766,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3461,f5692,f446])).
fof(f6811,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5783,f307])).
fof(f5783,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f799,f546,f5692,f459])).
fof(f6810,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5784,f5692])).
fof(f5784,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f804,f546,f5692,f459])).
fof(f804,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f607,f476])).
fof(f6809,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5785,f5692])).
fof(f5785,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1724,f546,f5692,f459])).
fof(f1724,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f1673,f476])).
fof(f1673,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f601,f1521])).
fof(f1521,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X5))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X4),s('type.integer.int',X5)))) ) | $spl1),
  inference(forward_demodulation,[],[f1486,f398])).
fof(f1486,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X5)))) ) | $spl1),
  inference(superposition,[],[f1416,f399])).
fof(f1416,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) ) | $spl1),
  inference(forward_demodulation,[],[f1385,f607])).
fof(f1385,plain,(
  ( ! [X2,X3] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int',X3))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3)))))) ) | $spl1),
  inference(superposition,[],[f398,f1039])).
fof(f1039,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(subsumption_resolution,[],[f1036,f546])).
fof(f1036,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f782,f607])).
fof(f6808,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5786,f5692])).
fof(f5786,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2233,f546,f5692,f459])).
fof(f2233,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2187,f476])).
fof(f2187,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f1673,f2066])).
fof(f2066,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X8),s('type.integer.int',X9))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X8))),s('type.integer.int',X9)))) ) | $spl1),
  inference(forward_demodulation,[],[f2051,f398])).
fof(f2051,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X8))),s('type.integer.int',X9)))) ) | $spl1),
  inference(superposition,[],[f1524,f399])).
fof(f1524,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int',X11)))) ) | $spl1),
  inference(subsumption_resolution,[],[f1494,f307])).
fof(f1494,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int',X11)))) ) | $spl1),
  inference(superposition,[],[f336,f1416])).
fof(f6807,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5787,f5692])).
fof(f5787,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3261,f546,f5692,f459])).
fof(f3261,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2230,f476])).
fof(f2230,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2187,f2066])).
fof(f6806,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5789,f5692])).
fof(f5789,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3494,f546,f5692,f459])).
fof(f3494,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f607,f2300,f476])).
fof(f6805,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5821,f307])).
fof(f5821,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f804,f546,f5692,f459])).
fof(f6804,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5822,f307])).
fof(f5822,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1724,f546,f5692,f459])).
fof(f6803,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5823,f307])).
fof(f5823,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2233,f546,f5692,f459])).
fof(f6802,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5824,f307])).
fof(f5824,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3261,f546,f5692,f459])).
fof(f6801,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5825,f307])).
fof(f5825,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f1317,f546,f5692,f459])).
fof(f6800,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5826,f5692])).
fof(f5826,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3494,f546,f5692,f459])).
fof(f6799,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5868,f6086])).
fof(f6086,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f474])).
fof(f5868,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f572,f569,f5692,f460])).
fof(f6798,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5869,f6086])).
fof(f5869,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f572,f569,f5692,f460])).
fof(f818,plain,(
  ~sP5('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f601,f477])).
fof(f6797,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5870,f6086])).
fof(f5870,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f573,f569,f5692,f460])).
fof(f6796,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5871,f6086])).
fof(f5871,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f589,f569,f5692,f460])).
fof(f589,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f566,f469])).
fof(f6795,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5872,f6086])).
fof(f5872,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3507,f589,f569,f5692,f460])).
fof(f3507,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f601,f2300,f477])).
fof(f6794,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5873,f6086])).
fof(f5873,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3465,f572,f3479,f5692,f460])).
fof(f6793,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5874,f6086])).
fof(f5874,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3465,f573,f3479,f5692,f460])).
fof(f6792,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5875,f6086])).
fof(f5875,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f573,f570,f5692,f460])).
fof(f6791,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5876,f6086])).
fof(f5876,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3465,f573,f570,f5692,f460])).
fof(f6790,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5882,f6086])).
fof(f5882,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f572,f570,f5692,f460])).
fof(f6789,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5883,f6086])).
fof(f5883,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3465,f572,f570,f5692,f460])).
fof(f6788,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5886,f6086])).
fof(f5886,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f589,f587,f5692,f460])).
fof(f587,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f566,f465])).
fof(f6787,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5890,f6086])).
fof(f5890,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3507,f589,f587,f5692,f460])).
fof(f6786,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5891,f6086])).
fof(f5891,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f572,f587,f5692,f460])).
fof(f6785,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5910,f6046])).
fof(f6046,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f464])).
fof(f5910,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f572,f575,f5692,f460])).
fof(f6784,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5911,f6046])).
fof(f5911,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f572,f575,f5692,f460])).
fof(f6783,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5913,f6046])).
fof(f5913,plain,(
  ~sP4('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3465,f572,f575,f5692,f460])).
fof(f6782,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5914,f6046])).
fof(f5914,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f573,f575,f5692,f460])).
fof(f6781,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5915,f6046])).
fof(f5915,plain,(
  ~sP4('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3465,f573,f575,f5692,f460])).
fof(f6780,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5916,f6046])).
fof(f5916,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f589,f575,f5692,f460])).
fof(f6779,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5922,f6080])).
fof(f6080,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f468])).
fof(f5922,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f569,f575,f5692,f460])).
fof(f6778,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5923,f6080])).
fof(f5923,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f569,f575,f5692,f460])).
fof(f6777,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5930,f6080])).
fof(f5930,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3465,f3479,f575,f5692,f460])).
fof(f6776,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5936,f6080])).
fof(f5936,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f570,f575,f5692,f460])).
fof(f6775,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5937,f6080])).
fof(f5937,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3465,f570,f575,f5692,f460])).
fof(f6774,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5944,f6080])).
fof(f5944,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f587,f575,f5692,f460])).
fof(f6773,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5951,f6046])).
fof(f5951,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f572,f576,f5692,f460])).
fof(f6772,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5952,f6046])).
fof(f5952,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f572,f576,f5692,f460])).
fof(f6771,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5953,f6046])).
fof(f5953,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f573,f576,f5692,f460])).
fof(f6770,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5954,f6046])).
fof(f5954,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f589,f576,f5692,f460])).
fof(f6769,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5955,f6046])).
fof(f5955,plain,(
  ~sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3507,f589,f576,f5692,f460])).
fof(f6768,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5956,f6080])).
fof(f5956,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f569,f576,f5692,f460])).
fof(f6767,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5957,f6080])).
fof(f5957,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f569,f576,f5692,f460])).
fof(f6766,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5959,f6080])).
fof(f5959,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3507,f569,f576,f5692,f460])).
fof(f6765,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5962,f6131])).
fof(f6131,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f600,f5692,f477])).
fof(f5962,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f573,f569,f576,f5692,f460])).
fof(f6764,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5964,f6133])).
fof(f6133,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f601,f5692,f477])).
fof(f5964,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f589,f569,f576,f5692,f460])).
fof(f6763,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5966,f6080])).
fof(f5966,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f817,f570,f576,f5692,f460])).
fof(f6762,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5967,f6131])).
fof(f5967,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f573,f570,f576,f5692,f460])).
fof(f6761,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5969,f817])).
fof(f5969,plain,(
  sP5('T_0','T_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f572,f570,f576,f5692,f460])).
fof(f6760,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5971,f6080])).
fof(f5971,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f587,f576,f5692,f460])).
fof(f6759,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5972,f6080])).
fof(f5972,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3507,f587,f576,f5692,f460])).
fof(f6758,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5973,f6133])).
fof(f5973,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f589,f587,f576,f5692,f460])).
fof(f6757,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5975,f818])).
fof(f5975,plain,(
  sP5('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f572,f587,f576,f5692,f460])).
fof(f6756,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5980,f6131])).
fof(f5980,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','F_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f573,f588,f576,f5692,f460])).
fof(f588,plain,(
  ( ! [X0,X1] : (sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f566,f466])).
fof(f6755,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5981,f6133])).
fof(f5981,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f589,f588,f576,f5692,f460])).
fof(f6754,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5996,f5692])).
fof(f5996,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f587,f546,f5692,f463])).
fof(f6753,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5997,f5692])).
fof(f5997,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1914,f546,f5692,f463])).
fof(f1914,plain,(
  sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1726,f458])).
fof(f1726,plain,(
  sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f1673,f476])).
fof(f6752,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5998,f5692])).
fof(f5998,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2455,f546,f5692,f463])).
fof(f2455,plain,(
  sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f2235,f458])).
fof(f2235,plain,(
  sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2187,f476])).
fof(f6751,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5999,f5692])).
fof(f5999,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3438,f546,f5692,f463])).
fof(f3438,plain,(
  sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f3263,f458])).
fof(f3263,plain,(
  sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2230,f476])).
fof(f6750,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6000,f5692])).
fof(f6000,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f571,f546,f5692,f463])).
fof(f6749,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6001,f5692])).
fof(f6001,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f588,f546,f5692,f463])).
fof(f6748,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6006,f6046])).
fof(f6006,plain,(
  ( ! [X0] : (~sP4(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5692,f463])).
fof(f6747,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6017,f6046])).
fof(f6017,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692,f5692,f463])).
fof(f6746,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6019,f5692])).
fof(f6019,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f570,f546,f5692,f463])).
fof(f6745,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6020,f5692])).
fof(f6020,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f587,f546,f5692,f463])).
fof(f6744,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6021,f307])).
fof(f6021,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1914,f546,f5692,f463])).
fof(f6743,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6022,f307])).
fof(f6022,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2455,f546,f5692,f463])).
fof(f6742,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6023,f307])).
fof(f6023,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3438,f546,f5692,f463])).
fof(f6741,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6025,f5692])).
fof(f6025,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f588,f546,f5692,f463])).
fof(f6740,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6033,f6046])).
fof(f6033,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5692,f463])).
fof(f6739,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6044,f6046])).
fof(f6044,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692,f5692,f463])).
fof(f6738,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6049,f5692])).
fof(f6049,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f589,f546,f5692,f467])).
fof(f6737,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6050,f5692])).
fof(f6050,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f574,f546,f5692,f467])).
fof(f6736,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6051,f5692])).
fof(f6051,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f590,f546,f5692,f467])).
fof(f590,plain,(
  ( ! [X0,X1] : (sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f566,f470])).
fof(f6735,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6055,f6080])).
fof(f6055,plain,(
  ( ! [X0] : (~sP3(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5692,f467])).
fof(f6734,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6062,f6080])).
fof(f6062,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692,f5692,f467])).
fof(f6733,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6064,f5692])).
fof(f6064,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f573,f546,f5692,f467])).
fof(f6732,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6065,f5692])).
fof(f6065,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f589,f546,f5692,f467])).
fof(f6731,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6067,f5692])).
fof(f6067,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f590,f546,f5692,f467])).
fof(f6730,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6071,f6080])).
fof(f6071,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5692,f467])).
fof(f6729,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6078,f6080])).
fof(f6078,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692,f5692,f467])).
fof(f6727,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6095,f607])).
fof(f6095,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f827,f5692,f476])).
fof(f827,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f607,f478])).
fof(f6726,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6096,f1673])).
fof(f6096,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1725,f5692,f476])).
fof(f1725,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f1673,f478])).
fof(f6725,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6097,f2187])).
fof(f6097,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2234,f5692,f476])).
fof(f2234,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2187,f478])).
fof(f6724,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6098,f2230])).
fof(f6098,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3262,f5692,f476])).
fof(f3262,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2230,f478])).
fof(f6723,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6099,f607])).
fof(f6099,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f826,f5692,f476])).
fof(f826,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f607,f478])).
fof(f6722,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6100,f1673])).
fof(f6100,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1723,f5692,f476])).
fof(f1723,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f1673,f478])).
fof(f6721,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6101,f2187])).
fof(f6101,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2232,f5692,f476])).
fof(f2232,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2187,f478])).
fof(f6720,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6102,f2230])).
fof(f6102,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3260,f5692,f476])).
fof(f3260,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2230,f478])).
fof(f6719,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6103,f398])).
fof(f6103,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1381,f5692,f476])).
fof(f6718,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6104,f398])).
fof(f6104,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1383,f5692,f476])).
fof(f6717,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6105,f399])).
fof(f6105,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1436,f5692,f476])).
fof(f6716,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6106,f399])).
fof(f6106,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1438,f5692,f476])).
fof(f6715,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6109,f600])).
fof(f6109,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f797,f5692,f477])).
fof(f6714,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6110,f600])).
fof(f6110,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f799,f5692,f477])).
fof(f6713,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6111,f601])).
fof(f6111,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f800,f5692,f477])).
fof(f800,plain,(
  sP5('F_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f601,f475])).
fof(f6712,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6112,f601])).
fof(f6112,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f798,f5692,f477])).
fof(f798,plain,(
  sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f601,f475])).
fof(f6711,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6128,f600])).
fof(f6128,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1275,f5692,f477])).
fof(f1275,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f574,f588,f577,f307,f460])).
fof(f6710,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6144,f6281])).
fof(f6281,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f5692,f5692,f492])).
fof(f6144,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692,f480])).
fof(f6709,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6151,f6281])).
fof(f6151,plain,(
  ( ! [X0,X1] : (~sP7(X0,'F_0',X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692,f481])).
fof(f6708,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6154,f5692])).
fof(f6154,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f580,f630,f5692,f482])).
fof(f6707,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6155,f5692])).
fof(f6155,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f578,f630,f5692,f482])).
fof(f6706,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6157,f5692])).
fof(f6157,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f591,f630,f5692,f482])).
fof(f591,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f566,f488])).
fof(f6705,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6161,f6227])).
fof(f6227,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f487])).
fof(f6161,plain,(
  ~sP6('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f630,f307,f5692,f482])).
fof(f6704,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6169,f6227])).
fof(f6169,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f627,f5692,f5692,f482])).
fof(f6703,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6171,f6227])).
fof(f6171,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f630,f5692,f5692,f482])).
fof(f6702,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6174,f5692])).
fof(f6174,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f579,f627,f5692,f482])).
fof(f6701,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6175,f5692])).
fof(f6175,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f578,f627,f5692,f482])).
fof(f6700,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6177,f5692])).
fof(f6177,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f592,f627,f5692,f482])).
fof(f592,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f566,f489])).
fof(f6699,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6182,f6227])).
fof(f6182,plain,(
  ~sP6('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f627,f307,f5692,f482])).
fof(f6698,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6190,f6227])).
fof(f6190,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f627,f5692,f5692,f482])).
fof(f6697,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6192,f6227])).
fof(f6192,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f630,f5692,f5692,f482])).
fof(f6696,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6196,f5692])).
fof(f6196,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f591,f546,f5692,f486])).
fof(f6695,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6197,f5692])).
fof(f6197,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f580,f546,f5692,f486])).
fof(f6694,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6198,f5692])).
fof(f6198,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f592,f546,f5692,f486])).
fof(f6693,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6202,f6227])).
fof(f6202,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5692,f486])).
fof(f6692,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6209,f6227])).
fof(f6209,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692,f5692,f486])).
fof(f6691,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6211,f5692])).
fof(f6211,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f579,f546,f5692,f486])).
fof(f6690,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6212,f5692])).
fof(f6212,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f591,f546,f5692,f486])).
fof(f6689,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6214,f5692])).
fof(f6214,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f592,f546,f5692,f486])).
fof(f6688,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6218,f6227])).
fof(f6218,plain,(
  ( ! [X0] : (~sP6('T_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5692,f486])).
fof(f6687,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6225,f6227])).
fof(f6225,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692,f5692,f486])).
fof(f6686,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6230,f5692])).
fof(f6230,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f630,f307,f5692,f492])).
fof(f6685,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6232,f5692])).
fof(f6232,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f630,f5692,f5692,f492])).
fof(f6684,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6233,f307])).
fof(f6233,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f916,f5692,f5692,f492])).
fof(f6683,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6234,f307])).
fof(f6234,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f630,f307,f5692,f492])).
fof(f6682,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6240,f307])).
fof(f6240,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f627,f5692,f5692,f492])).
fof(f6681,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6241,f307])).
fof(f6241,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f630,f5692,f5692,f492])).
fof(f6680,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6247,f5692])).
fof(f6247,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f627,f307,f5692,f492])).
fof(f6679,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6249,f5692])).
fof(f6249,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f627,f5692,f5692,f492])).
fof(f6678,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6251,f307])).
fof(f6251,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f916,f5692,f5692,f492])).
fof(f6677,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6252,f307])).
fof(f6252,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f627,f307,f5692,f492])).
fof(f6676,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6258,f307])).
fof(f6258,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f627,f5692,f5692,f492])).
fof(f6675,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6259,f307])).
fof(f6259,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f630,f5692,f5692,f492])).
fof(f6674,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6266,f307])).
fof(f6266,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f916,f307,f5692,f492])).
fof(f6673,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6267,f5692])).
fof(f6267,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f627,f5692,f5692,f492])).
fof(f6672,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6269,f307])).
fof(f6269,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f916,f5692,f5692,f492])).
fof(f6671,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6271,f307])).
fof(f6271,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f916,f307,f5692,f492])).
fof(f6670,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6277,f5692])).
fof(f6277,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f630,f5692,f5692,f492])).
fof(f6669,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6278,f307])).
fof(f6278,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f916,f5692,f5692,f492])).
fof(f6668,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6284,f5692])).
fof(f6284,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f709,f546,f5692,f495])).
fof(f709,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f591,f546,f532])).
fof(f6667,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6285,f5692])).
fof(f6285,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f650,f546,f5692,f495])).
fof(f6666,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6286,f5692])).
fof(f6286,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f711,f546,f5692,f495])).
fof(f711,plain,(
  ( ! [X0] : (sP7('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f592,f546,f532])).
fof(f6665,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6299,f5692])).
fof(f6299,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f644,f546,f5692,f495])).
fof(f6664,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6300,f5692])).
fof(f6300,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f709,f546,f5692,f495])).
fof(f6663,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6302,f5692])).
fof(f6302,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f711,f546,f5692,f495])).
fof(f6662,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6319,f6595])).
fof(f6595,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f583,f5692,f531])).
fof(f6319,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692,f496])).
fof(f6661,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6325,f6595])).
fof(f6325,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692,f497])).
fof(f6660,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6334,f6359])).
fof(f6359,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f505])).
fof(f6334,plain,(
  ~sP8('F_0','T_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f655,f546,f5692,f500])).
fof(f6659,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6342,f6359])).
fof(f6342,plain,(
  ~sP8('F_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f651,f546,f5692,f501])).
fof(f6658,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6355,f6359])).
fof(f6355,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f546,f5692,f503])).
fof(f6657,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6363,f307])).
fof(f6363,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f958,f5692,f508])).
fof(f6656,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6368,f307])).
fof(f6368,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f651,f5692,f508])).
fof(f6655,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6369,f5692])).
fof(f6369,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f655,f5692,f508])).
fof(f6654,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6377,f307])).
fof(f6377,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f958,f5692,f509])).
fof(f6653,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6382,f5692])).
fof(f6382,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f651,f5692,f509])).
fof(f6652,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6383,f307])).
fof(f6383,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f655,f5692,f509])).
fof(f6651,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6404,f6585])).
fof(f6585,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f586,f5692,f530])).
fof(f6404,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692,f514])).
fof(f6650,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6407,f5692])).
fof(f6407,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f586,f671,f5692,f517])).
fof(f6649,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6408,f5692])).
fof(f6408,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f585,f671,f5692,f517])).
fof(f6648,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6410,f5692])).
fof(f6410,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f594,f671,f5692,f517])).
fof(f594,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f566,f521])).
fof(f6647,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6414,f6509])).
fof(f6509,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f522])).
fof(f6414,plain,(
  ~sP10('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f671,f307,f5692,f517])).
fof(f6646,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6422,f6509])).
fof(f6422,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f670,f5692,f5692,f517])).
fof(f6645,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6424,f6509])).
fof(f6424,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f671,f5692,f5692,f517])).
fof(f6644,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6437,f6509])).
fof(f6437,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f670,f5692,f5692,f517])).
fof(f6643,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6438,f6509])).
fof(f6438,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f671,f5692,f5692,f517])).
fof(f6642,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6455,f6509])).
fof(f6455,plain,(
  ~sP10('F_0','F_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f670,f546,f5692,f518])).
fof(f6641,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6469,f6509])).
fof(f6469,plain,(
  ~sP10('F_0','F_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f670,f546,f5692,f519])).
fof(f6640,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6486,f5692])).
fof(f6486,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f594,f546,f5692,f520])).
fof(f6639,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6487,f5692])).
fof(f6487,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f586,f546,f5692,f520])).
fof(f6638,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6488,f5692])).
fof(f6488,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f595,f546,f5692,f520])).
fof(f595,plain,(
  ( ! [X0,X1] : (sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f566,f523])).
fof(f6637,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6489,f6509])).
fof(f6489,plain,(
  ( ! [X0] : (~sP10(X0,'F_0','T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f307,f546,f5692,f520])).
fof(f6636,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6491,f6509])).
fof(f6491,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f546,f5692,f520])).
fof(f6635,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6501,f5692])).
fof(f6501,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f584,f546,f5692,f520])).
fof(f6634,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6502,f5692])).
fof(f6502,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f594,f546,f5692,f520])).
fof(f6633,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6504,f5692])).
fof(f6504,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f595,f546,f5692,f520])).
fof(f6632,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6505,f6509])).
fof(f6505,plain,(
  ( ! [X0] : (~sP10('T_0','F_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f307,f546,f5692,f520])).
fof(f6631,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6507,f6509])).
fof(f6507,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f546,f5692,f520])).
fof(f6630,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6524,f307])).
fof(f6524,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f1009,f5692,f526])).
fof(f6629,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6529,f5692])).
fof(f6529,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f670,f5692,f526])).
fof(f6628,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6530,f307])).
fof(f6530,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f671,f5692,f526])).
fof(f6627,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6536,f307])).
fof(f6536,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f673,f546,f5692,f527])).
fof(f6626,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6537,f307])).
fof(f6537,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f693,f546,f5692,f527])).
fof(f693,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f594,f307,f530])).
fof(f6625,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6539,f5692])).
fof(f6539,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f990,f546,f5692,f527])).
fof(f990,plain,(
  sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f595,f307,f546,f519])).
fof(f6624,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6540,f5692])).
fof(f6540,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3587,f546,f5692,f527])).
fof(f3587,plain,(
  sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f595,f2300,f519])).
fof(f6623,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6541,f5692])).
fof(f6541,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3613,f546,f5692,f527])).
fof(f3613,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f594,f2300,f530])).
fof(f6622,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6557,f5692])).
fof(f6557,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f693,f546,f5692,f527])).
fof(f6621,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6558,f307])).
fof(f6558,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f989,f546,f5692,f527])).
fof(f6620,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6559,f307])).
fof(f6559,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f990,f546,f5692,f527])).
fof(f6619,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6560,f5692])).
fof(f6560,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3587,f546,f5692,f527])).
fof(f6618,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6561,f5692])).
fof(f6561,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3613,f546,f5692,f527])).
fof(f6617,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6576,f584])).
fof(f6576,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f670,f5692,f530])).
fof(f6616,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6577,f585])).
fof(f6577,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f671,f5692,f530])).
fof(f6615,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6588,f581])).
fof(f6588,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f651,f5692,f531])).
fof(f6614,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6589,f582])).
fof(f6589,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f655,f5692,f531])).
fof(f6596,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f651,f583,f5692,f531])).
fof(f6594,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f655,f582,f5692,f531])).
fof(f6592,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f651,f581,f5692,f531])).
fof(f6586,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f586,f5692,f530])).
fof(f6584,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f671,f585,f5692,f530])).
fof(f6582,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f671,f594,f5692,f530])).
fof(f6580,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f584,f5692,f530])).
fof(f6575,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f673,f546,f5692,f5692,f527])).
fof(f6566,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f989,f546,f307,f5692,f527])).
fof(f6555,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f673,f546,f5692,f5692,f527])).
fof(f6546,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f673,f546,f307,f5692,f527])).
fof(f6535,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f5692,f5692,f526])).
fof(f6533,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f671,f307,f5692,f526])).
fof(f6528,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f5692,f5692,f526])).
fof(f6526,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f1009,f307,f5692,f526])).
fof(f6521,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f546,f5692,f525])).
fof(f6516,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f546,f546,f5692,f524])).
fof(f6510,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f1000,f5692,f522])).
fof(f6508,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f584,f5692,f546,f5692,f520])).
fof(f6506,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f584,f307,f546,f5692,f520])).
fof(f6492,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f584,f5692,f546,f5692,f520])).
fof(f6490,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f586,f307,f546,f5692,f520])).
fof(f6475,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f586,f546,f5692,f519])).
fof(f6472,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f584,f546,f5692,f519])).
fof(f6461,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f586,f546,f5692,f518])).
fof(f6458,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f584,f546,f5692,f518])).
fof(f6447,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f586,f5692,f5692,f517])).
fof(f6445,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f671,f585,f5692,f5692,f517])).
fof(f6443,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f671,f594,f5692,f5692,f517])).
fof(f6441,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f584,f5692,f5692,f517])).
fof(f6425,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f586,f671,f5692,f5692,f517])).
fof(f6423,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f586,f670,f5692,f5692,f517])).
fof(f6415,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f586,f671,f307,f5692,f517])).
fof(f6405,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f989,f546,f5692,f514])).
fof(f6396,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f944,f546,f546,f5692,f510])).
fof(f6388,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f651,f5692,f5692,f509])).
fof(f6386,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f655,f307,f5692,f509])).
fof(f6381,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f651,f5692,f5692,f509])).
fof(f6379,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f958,f307,f5692,f509])).
fof(f6374,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f651,f5692,f5692,f508])).
fof(f6372,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f651,f307,f5692,f508])).
fof(f6367,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f651,f5692,f5692,f508])).
fof(f6365,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f958,f307,f5692,f508])).
fof(f6360,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f950,f5692,f505])).
fof(f6358,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f950,f5692,f504])).
fof(f6356,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f583,f546,f546,f5692,f503])).
fof(f6347,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f651,f583,f546,f5692,f501])).
fof(f6345,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f651,f581,f546,f5692,f501])).
fof(f6335,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f583,f655,f546,f5692,f500])).
fof(f6326,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f944,f546,f5692,f497])).
fof(f6320,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f944,f546,f5692,f496])).
fof(f6314,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f644,f546,f5692,f5692,f495])).
fof(f6307,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f644,f546,f307,f5692,f495])).
fof(f6298,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f644,f546,f5692,f5692,f495])).
fof(f6291,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f650,f546,f307,f5692,f495])).
fof(f6282,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f627,f5692,f5692,f5692,f492])).
fof(f6280,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f630,f307,f5692,f5692,f492])).
fof(f6275,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f627,f5692,f307,f5692,f492])).
fof(f6273,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f916,f307,f307,f5692,f492])).
fof(f6264,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f627,f5692,f5692,f5692,f492])).
fof(f6262,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f627,f307,f5692,f5692,f492])).
fof(f6257,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f627,f5692,f307,f5692,f492])).
fof(f6255,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f627,f307,f307,f5692,f492])).
fof(f6246,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f627,f5692,f5692,f5692,f492])).
fof(f6244,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f627,f307,f5692,f5692,f492])).
fof(f6239,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f630,f5692,f307,f5692,f492])).
fof(f6237,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f630,f307,f307,f5692,f492])).
fof(f6228,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f907,f5692,f487])).
fof(f6226,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f579,f546,f5692,f5692,f486])).
fof(f6219,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f579,f546,f307,f5692,f486])).
fof(f6210,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f579,f546,f5692,f5692,f486])).
fof(f6203,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f580,f546,f307,f5692,f486])).
fof(f6193,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f580,f630,f5692,f5692,f482])).
fof(f6191,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f579,f627,f5692,f5692,f482])).
fof(f6183,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f579,f627,f307,f5692,f482])).
fof(f6172,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f580,f630,f5692,f5692,f482])).
fof(f6170,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f579,f627,f5692,f5692,f482])).
fof(f6162,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f580,f630,f307,f5692,f482])).
fof(f6152,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f650,f546,f5692,f481])).
fof(f6145,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f644,f546,f5692,f480])).
fof(f6138,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f798,f601,f5692,f477])).
fof(f6136,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f797,f600,f5692,f477])).
fof(f6134,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f800,f601,f5692,f477])).
fof(f6132,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f799,f600,f5692,f477])).
fof(f6613,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f6094])).
fof(f6094,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f892,f5692,f476])).
fof(f6612,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f6093])).
fof(f6093,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f883,f5692,f476])).
fof(f6087,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f887,f5692,f474])).
fof(f6085,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f887,f5692,f473])).
fof(f6083,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f887,f5692,f472])).
fof(f6081,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f874,f5692,f468])).
fof(f6079,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f573,f546,f5692,f5692,f467])).
fof(f6072,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f573,f546,f307,f5692,f467])).
fof(f6063,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f573,f546,f5692,f5692,f467])).
fof(f6056,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f574,f546,f307,f5692,f467])).
fof(f6047,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f857,f5692,f464])).
fof(f6045,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f570,f546,f5692,f5692,f463])).
fof(f6034,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1914,f546,f307,f5692,f463])).
fof(f6018,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f570,f546,f5692,f5692,f463])).
fof(f6007,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f571,f546,f307,f5692,f463])).
fof(f5976,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f572,f587,f576,f5692,f460])).
fof(f5974,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f589,f587,f576,f5692,f460])).
fof(f5970,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f817,f572,f570,f576,f5692,f460])).
fof(f5968,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f817,f573,f570,f576,f5692,f460])).
fof(f5965,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f589,f569,f576,f5692,f460])).
fof(f5963,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f817,f573,f569,f576,f5692,f460])).
fof(f5961,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f817,f572,f569,f576,f5692,f460])).
fof(f5948,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f572,f587,f575,f5692,f460])).
fof(f5946,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f589,f587,f575,f5692,f460])).
fof(f5941,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f817,f572,f570,f575,f5692,f460])).
fof(f5939,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f817,f573,f570,f575,f5692,f460])).
fof(f5934,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f3465,f573,f3479,f575,f5692,f460])).
fof(f5932,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f3465,f572,f3479,f575,f5692,f460])).
fof(f5929,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f818,f589,f569,f575,f5692,f460])).
fof(f5927,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f817,f573,f569,f575,f5692,f460])).
fof(f5925,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f817,f572,f569,f575,f5692,f460])).
fof(f5867,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f799,f546,f5692,f5692,f459])).
fof(f5843,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f804,f546,f307,f5692,f459])).
fof(f5819,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f799,f546,f5692,f5692,f459])).
fof(f5795,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f799,f546,f307,f5692,f459])).
fof(f5782,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f610,f546,f5692,f450])).
fof(f5776,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f610,f546,f5692,f447])).
fof(f5770,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f614,f5692,f5692,f446])).
fof(f5768,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f614,f307,f5692,f446])).
fof(f5762,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f614,f5692,f5692,f446])).
fof(f5760,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f619,f307,f5692,f446])).
fof(f5754,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f622,f5692,f432])).
fof(f5752,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f622,f5692,f431])).
fof(f5750,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f5692,f601,f5692,f409])).
fof(f5748,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f600,f5692,f409])).
fof(f5746,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f307,f601,f5692,f409])).
fof(f5744,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f307,f600,f5692,f409])).
fof(f5740,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f5692,f601,f5692,f409])).
fof(f5738,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f5692,f600,f5692,f409])).
fof(f5736,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f307,f601,f5692,f409])).
fof(f5734,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f307,f600,f5692,f409])).
fof(f5729,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f600,f5692,f333])).
fof(f5727,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5692])).
fof(f5726,plain,(
  sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | $spl44),
  inference(cnf_transformation,[],[f5726_D])).
fof(f5726_D,plain,(
  sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5724,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0')) ) | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f5719,f332])).
fof(f5719,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl8)),
  inference(superposition,[],[f4924,f1039])).
fof(f5691,plain,(
  sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl40),
  inference(cnf_transformation,[],[f5691_D])).
fof(f5691_D,plain,(
  sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f5689,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5684,f332])).
fof(f5684,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl4)),
  inference(superposition,[],[f4920,f1039])).
fof(f5366,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347,f5349,f5351,f5353,f5355,f5357,f5363,f5365])).
fof(f5365,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5245,f307])).
fof(f5245,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'T_0'))) ) | $spl6),
  inference(condensation,[],[f5232])).
fof(f5232,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f2220])).
fof(f5363,plain,(
  ( ! [X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5230,f307])).
fof(f5230,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f2227])).
fof(f5357,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5218,f307])).
fof(f5218,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f2709])).
fof(f5355,plain,(
  ( ! [X10,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5213,f307])).
fof(f5213,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f2792])).
fof(f5353,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5212,f307])).
fof(f5212,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4)))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f2789])).
fof(f5351,plain,(
  ( ! [X10,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5211,f307])).
fof(f5211,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f2718])).
fof(f5349,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5210,f307])).
fof(f5210,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f2715])).
fof(f5347,plain,(
  ( ! [X92,X90,X91,X89] : (~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92)) ) | $spl6),
  inference(subsumption_resolution,[],[f5209,f307])).
fof(f5209,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1470])).
fof(f5345,plain,(
  ( ! [X88,X87,X85,X86] : (~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88)) ) | $spl6),
  inference(subsumption_resolution,[],[f5208,f307])).
fof(f5208,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1469])).
fof(f5343,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X83)) | p(s(bool,X84)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X81))),s('type.integer.int',X82)),X84,X83)) ) | $spl6),
  inference(subsumption_resolution,[],[f5207,f307])).
fof(f5207,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X81))),s('type.integer.int',X82)),X84,X83)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1468])).
fof(f5341,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X79)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X77))),s('type.integer.int',X78)),X79,X80)) ) | $spl6),
  inference(subsumption_resolution,[],[f5206,f307])).
fof(f5206,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,'T_0')) | ~p(s(bool,X79)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X77))),s('type.integer.int',X78)),X79,X80)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1467])).
fof(f5339,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X75)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X73))),s('type.integer.int',X74)),X76,X75)) ) | $spl6),
  inference(subsumption_resolution,[],[f5205,f307])).
fof(f5205,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X75)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X73))),s('type.integer.int',X74)),X76,X75)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1466])).
fof(f5337,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP10(X72,X71,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X69))),s('type.integer.int',X70)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5204,f307])).
fof(f5204,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP10(X72,X71,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X69))),s('type.integer.int',X70)))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1465])).
fof(f5335,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,X67)) | p(s(bool,X68)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X65))),s('type.integer.int',X66)),X68,X67)) ) | $spl6),
  inference(subsumption_resolution,[],[f5203,f307])).
fof(f5203,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X67)) | p(s(bool,X68)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X65))),s('type.integer.int',X66)),X68,X67)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1464])).
fof(f5333,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X63)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X61))),s('type.integer.int',X62)),X63,X64)) ) | $spl6),
  inference(subsumption_resolution,[],[f5202,f307])).
fof(f5202,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X63)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X61))),s('type.integer.int',X62)),X63,X64)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1463])).
fof(f5331,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X59)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X58)),X60,X59)) ) | $spl6),
  inference(subsumption_resolution,[],[f5201,f307])).
fof(f5201,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X59)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X58)),X60,X59)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1462])).
fof(f5329,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,X55)) | p(s(bool,X56)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X53))),s('type.integer.int',X54)),X55,X56)) ) | $spl6),
  inference(subsumption_resolution,[],[f5200,f307])).
fof(f5200,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | p(s(bool,X56)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X53))),s('type.integer.int',X54)),X55,X56)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1461])).
fof(f5327,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP7(X52,X51,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X49))),s('type.integer.int',X50)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5199,f307])).
fof(f5199,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP7(X52,X51,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X49))),s('type.integer.int',X50)))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1460])).
fof(f5325,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP6(X47,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X45))),s('type.integer.int',X46)),X48)) ) | $spl6),
  inference(subsumption_resolution,[],[f5198,f307])).
fof(f5198,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP6(X47,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X45))),s('type.integer.int',X46)),X48)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1459])).
fof(f5323,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X43)) | p(s(bool,X44)) | ~sP3(X43,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X41))),s('type.integer.int',X42)),X44)) ) | $spl6),
  inference(subsumption_resolution,[],[f5197,f307])).
fof(f5197,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP3(X43,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X41))),s('type.integer.int',X42)),X44)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1458])).
fof(f5321,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | ~sP4(X39,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)),X40)) ) | $spl6),
  inference(subsumption_resolution,[],[f5196,f307])).
fof(f5196,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP4(X39,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)),X40)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1457])).
fof(f5319,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | p(s(bool,X36)) | ~sP5(X35,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33))),s('type.integer.int',X34)),X36)) ) | $spl6),
  inference(subsumption_resolution,[],[f5195,f307])).
fof(f5195,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP5(X35,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33))),s('type.integer.int',X34)),X36)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1456])).
fof(f5317,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X32)) | sP1(X32,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30))),s('type.integer.int',X31)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5194,f307])).
fof(f5194,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | sP1(X32,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30))),s('type.integer.int',X31)))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1455])).
fof(f5315,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X29)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27))),s('type.integer.int',X28)),X29)) ) | $spl6),
  inference(subsumption_resolution,[],[f5193,f307])).
fof(f5193,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27))),s('type.integer.int',X28)),X29)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1454])).
fof(f5313,plain,(
  ( ! [X26,X24,X25] : (sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))),s('type.integer.int',X25)),X26)) ) | $spl6),
  inference(subsumption_resolution,[],[f5192,f307])).
fof(f5192,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))),s('type.integer.int',X25)),X26)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1453])).
fof(f5311,plain,(
  ( ! [X23,X21,X22] : (sP0(X23,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X21))),s('type.integer.int',X22)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5191,f307])).
fof(f5191,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | sP0(X23,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X21))),s('type.integer.int',X22)))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1452])).
fof(f5309,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | sP9('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90)) ) | $spl6),
  inference(subsumption_resolution,[],[f5190,f307])).
fof(f5190,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | sP9('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1413])).
fof(f5307,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | sP11('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86)) ) | $spl6),
  inference(subsumption_resolution,[],[f5189,f307])).
fof(f5189,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | sP11('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1412])).
fof(f5305,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X81)) | p(s(bool,X82)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X80)))),X82,X81)) ) | $spl6),
  inference(subsumption_resolution,[],[f5188,f307])).
fof(f5188,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X80)))),X82,X81)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1411])).
fof(f5303,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X77)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X76)))),X77,X78)) ) | $spl6),
  inference(subsumption_resolution,[],[f5187,f307])).
fof(f5187,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X76)))),X77,X78)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1410])).
fof(f5301,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X73)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X72)))),X74,X73)) ) | $spl6),
  inference(subsumption_resolution,[],[f5186,f307])).
fof(f5186,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X73)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X72)))),X74,X73)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1409])).
fof(f5299,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP10(X70,X69,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X68)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5185,f307])).
fof(f5185,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP10(X70,X69,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X68)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1408])).
fof(f5297,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | p(s(bool,X66)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X64)))),X66,X65)) ) | $spl6),
  inference(subsumption_resolution,[],[f5184,f307])).
fof(f5184,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | p(s(bool,X66)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X64)))),X66,X65)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1407])).
fof(f5295,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X60)))),X61,X62)) ) | $spl6),
  inference(subsumption_resolution,[],[f5183,f307])).
fof(f5183,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X60)))),X61,X62)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1406])).
fof(f5293,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X56)))),X58,X57)) ) | $spl6),
  inference(subsumption_resolution,[],[f5182,f307])).
fof(f5182,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X56)))),X58,X57)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1405])).
fof(f5291,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X53)) | p(s(bool,X54)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X52)))),X53,X54)) ) | $spl6),
  inference(subsumption_resolution,[],[f5181,f307])).
fof(f5181,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | p(s(bool,X54)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X52)))),X53,X54)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1404])).
fof(f5289,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X50,X49,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X48)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5180,f307])).
fof(f5180,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X50,X49,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X48)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1403])).
fof(f5287,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP6(X45,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X44)))),X46)) ) | $spl6),
  inference(subsumption_resolution,[],[f5179,f307])).
fof(f5179,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP6(X45,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X44)))),X46)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1402])).
fof(f5285,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP3(X41,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40)))),X42)) ) | $spl6),
  inference(subsumption_resolution,[],[f5178,f307])).
fof(f5178,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP3(X41,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40)))),X42)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1401])).
fof(f5283,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP4(X37,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X36)))),X38)) ) | $spl6),
  inference(subsumption_resolution,[],[f5177,f307])).
fof(f5177,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP4(X37,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X36)))),X38)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1400])).
fof(f5281,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP5(X33,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))),X34)) ) | $spl6),
  inference(subsumption_resolution,[],[f5176,f307])).
fof(f5176,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP5(X33,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))),X34)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1399])).
fof(f5279,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | sP1(X30,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5175,f307])).
fof(f5175,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1398])).
fof(f5277,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | sP1('const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X26)))),X27)) ) | $spl6),
  inference(subsumption_resolution,[],[f5174,f307])).
fof(f5174,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X26)))),X27)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1397])).
fof(f5275,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X23)))),X24)) ) | $spl6),
  inference(subsumption_resolution,[],[f5173,f307])).
fof(f5173,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0('const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X23)))),X24)) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1396])).
fof(f5273,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5172,f307])).
fof(f5172,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f1395])).
fof(f5271,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5169,f307])).
fof(f5169,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f413])).
fof(f5269,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5168,f307])).
fof(f5168,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f412])).
fof(f5267,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5167,f307])).
fof(f5167,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f410])).
fof(f5258,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f5040,f4921])).
fof(f5040,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f4921,f336])).
fof(f5364,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347,f5349,f5351,f5353,f5355,f5357,f5363])).
fof(f5362,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5226,f307])).
fof(f5226,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4921,f2224])).
fof(f5361,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5225,f307])).
fof(f5225,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4921,f4110])).
fof(f5360,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5224,f307])).
fof(f5224,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4921,f2223])).
fof(f5359,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5223,f307])).
fof(f5223,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4921,f4174])).
fof(f5358,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347,f5349,f5351,f5353,f5355,f5357])).
fof(f5356,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347,f5349,f5351,f5353,f5355])).
fof(f5354,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347,f5349,f5351,f5353])).
fof(f5352,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347,f5349,f5351])).
fof(f5350,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347,f5349])).
fof(f5348,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347])).
fof(f5346,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345])).
fof(f5344,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343])).
fof(f5342,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341])).
fof(f5340,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339])).
fof(f5338,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337])).
fof(f5336,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335])).
fof(f5334,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333])).
fof(f5332,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331])).
fof(f5330,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329])).
fof(f5328,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327])).
fof(f5326,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325])).
fof(f5324,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323])).
fof(f5322,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321])).
fof(f5320,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319])).
fof(f5318,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317])).
fof(f5316,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315])).
fof(f5314,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313])).
fof(f5312,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311])).
fof(f5310,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309])).
fof(f5308,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307])).
fof(f5306,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305])).
fof(f5304,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303])).
fof(f5302,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301])).
fof(f5300,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299])).
fof(f5298,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297])).
fof(f5296,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295])).
fof(f5294,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293])).
fof(f5292,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291])).
fof(f5290,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289])).
fof(f5288,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287])).
fof(f5286,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285])).
fof(f5284,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283])).
fof(f5282,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281])).
fof(f5280,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277,f5279])).
fof(f5278,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275,f5277])).
fof(f5276,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273,f5275])).
fof(f5274,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271,f5273])).
fof(f5272,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269,f5271])).
fof(f5270,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267,f5269])).
fof(f5268,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258,f5267])).
fof(f5260,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5042,f307])).
fof(f5042,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4921,f546])).
fof(f5259,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f586,f600,f610,f611,f614,f619,f622,f627,f630,f644,f650,f651,f655,f658,f662,f670,f671,f673,f681,f731,f782,f799,f797,f817,f813,f857,f874,f883,f887,f892,f896,f907,f916,f944,f950,f958,f989,f1000,f1009,f1046,f1053,f1062,f1061,f1063,f1066,f1067,f1068,f1219,f1225,f1247,f1246,f1283,f1317,f1384,f1383,f1382,f1381,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1413,f1439,f1438,f1437,f1436,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1470,f1570,f1555,f1573,f1562,f1577,f1643,f1642,f1653,f1654,f1655,f1656,f1659,f1658,f1660,f1662,f1663,f1664,f1665,f1810,f1788,f1811,f1812,f1815,f1881,f1885,f1932,f1999,f2129,f2138,f2140,f2223,f2224,f2227,f2220,f2300,f2307,f2296,f2349,f2352,f2353,f2354,f2393,f2396,f2401,f2399,f2398,f2414,f2412,f2411,f2478,f2489,f2595,f2606,f2709,f2715,f2718,f2789,f2792,f2848,f2869,f2872,f2879,f2880,f2940,f2953,f3128,f3141,f3362,f3363,f3365,f3328,f3338,f3395,f3396,f3433,f3399,f3409,f3616,f3615,f3614,f3612,f3611,f3610,f3607,f3606,f3602,f3599,f3598,f3594,f3593,f3588,f3573,f3572,f3571,f3567,f3566,f3562,f3561,f3560,f3557,f3553,f3551,f3548,f3547,f3543,f3542,f3541,f3540,f3539,f3536,f3534,f3529,f3528,f3524,f3523,f3518,f3511,f3510,f3508,f3506,f3505,f3504,f3499,f3498,f3493,f3492,f3491,f3490,f3486,f3485,f3480,f3479,f3475,f3474,f3469,f3468,f3466,f3465,f3462,f3461,f3460,f3459,f3457,f3456,f3455,f3453,f3658,f3772,f3857,f3900,f3903,f3950,f3978,f4007,f4010,f4063,f4110,f4092,f4123,f4122,f4174,f4156,f4190,f4246,f4247,f4248,f4298,f4261,f4304,f4313,f4334,f4335,f4336,f4381,f4389,f4394,f4416,f4419,f4442,f4455,f4533,f4534,f4535,f4536,f4594,f4919,f4923,f4926,f4929,f4932,f4935,f4938,f4941,f4944,f4947,f4950,f4953,f4956,f4959,f4962,f4965,f4968,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4995,f4996,f4997,f4998,f4675,f4676,f4677,f4678,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f4921,f5258])).
fof(f4969,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f4969_D])).
fof(f4969_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4921,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f4921_D])).
fof(f4921_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f5036,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f5035])).
fof(f5035,plain,(
  ( ! [X94,X92,X95,X93] : (~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f4706,f782])).
fof(f4706,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f531,f731])).
fof(f4966,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl36),
  inference(cnf_transformation,[],[f4966_D])).
fof(f4966_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5034,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f5033])).
fof(f5033,plain,(
  ( ! [X90,X88,X91,X89] : (~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f4705,f782])).
fof(f4705,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f530,f731])).
fof(f4963,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f4963_D])).
fof(f4963_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f5032,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f5031])).
fof(f5031,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f4704,f782])).
fof(f4704,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f527,f731])).
fof(f4960,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f4960_D])).
fof(f4960_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5030,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f5029])).
fof(f5029,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f4703,f782])).
fof(f4703,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f526,f731])).
fof(f4957,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f4957_D])).
fof(f4957_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f5028,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f5027])).
fof(f5027,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f4702,f782])).
fof(f4702,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f525,f731])).
fof(f4954,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f4954_D])).
fof(f4954_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5026,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f5025])).
fof(f5025,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f4701,f782])).
fof(f4701,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f520,f731])).
fof(f4951,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f4951_D])).
fof(f4951_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f5024,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f5023])).
fof(f5023,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f4700,f782])).
fof(f4700,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f510,f731])).
fof(f4948,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f4948_D])).
fof(f4948_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f5022,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f5021])).
fof(f5021,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f4699,f782])).
fof(f4699,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f509,f731])).
fof(f4945,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f4945_D])).
fof(f4945_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f5020,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f5019])).
fof(f5019,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f4698,f782])).
fof(f4698,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f508,f731])).
fof(f4942,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f4942_D])).
fof(f4942_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f5018,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f5017])).
fof(f5017,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f4697,f782])).
fof(f4697,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f503,f731])).
fof(f4939,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f4939_D])).
fof(f4939_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5016,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f5015])).
fof(f5015,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f4696,f782])).
fof(f4696,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f495,f731])).
fof(f4936,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f4936_D])).
fof(f4936_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5014,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f5013])).
fof(f5013,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f4695,f782])).
fof(f4695,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f486,f731])).
fof(f4933,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl14),
  inference(cnf_transformation,[],[f4933_D])).
fof(f4933_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5012,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f5011])).
fof(f5011,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f4694,f782])).
fof(f4694,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f467,f731])).
fof(f4930,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f4930_D])).
fof(f4930_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f5010,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f5009])).
fof(f5009,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f4693,f782])).
fof(f4693,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f463,f731])).
fof(f4927,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl10),
  inference(cnf_transformation,[],[f4927_D])).
fof(f4927_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5008,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f5007])).
fof(f5007,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f4692,f782])).
fof(f4692,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f459,f731])).
fof(f4924,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl8),
  inference(cnf_transformation,[],[f4924_D])).
fof(f4924_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f5006,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f5005])).
fof(f5005,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f4691,f782])).
fof(f4691,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f450,f731])).
fof(f4920,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl4),
  inference(cnf_transformation,[],[f4920_D])).
fof(f4920_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f5004,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f5003])).
fof(f5003,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f4690,f782])).
fof(f4690,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f447,f731])).
fof(f4968,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4967])).
fof(f4967,plain,(
  ( ! [X94,X92,X95,X93] : (~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f4621,f782])).
fof(f4621,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f531,f731])).
fof(f4965,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4964])).
fof(f4964,plain,(
  ( ! [X90,X88,X91,X89] : (~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f4620,f782])).
fof(f4620,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f530,f731])).
fof(f4962,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4961])).
fof(f4961,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f4619,f782])).
fof(f4619,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f527,f731])).
fof(f4959,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4958])).
fof(f4958,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f4618,f782])).
fof(f4618,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f526,f731])).
fof(f4956,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4955])).
fof(f4955,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f4617,f782])).
fof(f4617,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f525,f731])).
fof(f4953,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f4952])).
fof(f4952,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f4616,f782])).
fof(f4616,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f520,f731])).
fof(f4950,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4949])).
fof(f4949,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f4615,f782])).
fof(f4615,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f510,f731])).
fof(f4947,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4946])).
fof(f4946,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f4614,f782])).
fof(f4614,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f509,f731])).
fof(f4944,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4943])).
fof(f4943,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f4613,f782])).
fof(f4613,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f508,f731])).
fof(f4941,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f4940])).
fof(f4940,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f4612,f782])).
fof(f4612,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f503,f731])).
fof(f4938,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4937])).
fof(f4937,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f4611,f782])).
fof(f4611,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f495,f731])).
fof(f4935,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f4934])).
fof(f4934,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f4610,f782])).
fof(f4610,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f486,f731])).
fof(f4932,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f4931])).
fof(f4931,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f4609,f782])).
fof(f4609,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f467,f731])).
fof(f4929,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4928])).
fof(f4928,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f4608,f782])).
fof(f4608,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f463,f731])).
fof(f4926,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f4925])).
fof(f4925,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f4607,f782])).
fof(f4607,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f459,f731])).
fof(f4923,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4922])).
fof(f4922,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f4606,f782])).
fof(f4606,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f450,f731])).
fof(f4919,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4918])).
fof(f4918,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f4605,f782])).
fof(f4605,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f447,f731])).
fof(f542,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) | $spl0),
  inference(cnf_transformation,[],[f542_D])).
fof(f542_D,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f544,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | $spl2),
  inference(cnf_transformation,[],[f544_D])).
fof(f544_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f305,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))))))))))),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12)))))))) & (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f202])).
fof(f202,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) & (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ? [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) <~> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|11','thm.integer.INT_DIVIDES_REDUCE|split|11')).
fof(f543,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) | $spl1),
  inference(cnf_transformation,[],[f543_D])).
fof(f543_D,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f545,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | $spl3),
  inference(cnf_transformation,[],[f545_D])).
fof(f545_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f306,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))))))))))),
  inference(cnf_transformation,[],[f203])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 790
% Passive clauses: 2405
% Generated clauses: 21491
% Final active clauses: 543
% Final passive clauses: 1177
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 192
% Trivial inequalities: 136
% Fw subsumption resolutions: 820
% Fw demodulations: 5581
% Bw demodulations: 244
% Condensations: 215
% Global subsumptions: 51
% 
% Simple tautologies: 113
% Equational tautologies: 71
% Forward subsumptions: 11270
% Fw demodulations to eq. taut.: 147
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 5406
% Factoring: 1
% Forward superposition: 4798
% Backward superposition: 3419
% Self superposition: 49
% Equality factoring: 28
% Equality resolution: 26
% 
% Split clauses: 57
% Split components: 119
% SAT solver clauses: 3005
% SAT solver unit clauses: 713
% SAT solver binary clauses: 860
% 
% Sat splits: 57
% Sat splitting refutations: 433
% 
% Memory used [KB]: 13560
% Time elapsed: 0.739 s
% ------------------------------
% ------------------------------
% Success in time 2.058 s
