Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.res_quan
% (78683)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1091
% Passive clauses: 9197
% Generated clauses: 26183
% Final active clauses: 1020
% Final passive clauses: 7634
% Input formulas: 60
% Initial clauses: 225
% 
% Duplicate literals: 482
% Trivial inequalities: 78
% Fw subsumption resolutions: 2932
% Fw demodulations: 1716
% Bw demodulations: 63
% 
% Simple tautologies: 303
% Equational tautologies: 41
% Forward subsumptions: 10728
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 10661
% Factoring: 1376
% Forward superposition: 3760
% Backward superposition: 4488
% Self superposition: 32
% Equality factoring: 26
% Equality resolution: 13
% 
% Split clauses: 657
% Split components: 1320
% SAT solver clauses: 1074
% SAT solver unit clauses: 28
% SAT solver binary clauses: 213
% 
% Sat splits: 657
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1072
% TWLsolver calls for satisfiability: 545
% 
% Memory used [KB]: 12409
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.res_quan
% (78694)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 431
% Passive clauses: 2709
% Generated clauses: 13404
% Final active clauses: 421
% Final passive clauses: 2254
% Input formulas: 60
% Initial clauses: 225
% 
% Duplicate literals: 283
% Trivial inequalities: 60
% Fw subsumption resolutions: 995
% Fw demodulations: 1094
% Bw demodulations: 29
% Condensations: 321
% Global subsumptions: 22
% 
% Simple tautologies: 346
% Equational tautologies: 21
% Forward subsumptions: 4123
% 
% Binary resolution: 3112
% Factoring: 126
% Forward superposition: 2201
% Backward superposition: 1987
% Self superposition: 43
% Equality factoring: 30
% Equality resolution: 12
% 
% Split clauses: 86
% Split components: 172
% SAT solver clauses: 3538
% SAT solver unit clauses: 220
% SAT solver binary clauses: 99
% SAT solver learnt clauses: 45
% SAT solver learnt literals: 12
% 
% Sat splits: 86
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 3449
% TWLsolver calls for satisfiability: 1826
% 
% Memory used [KB]: 14328
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.res_quan
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.res_quan
% SZS output start Proof for thm.res_quan
fof(f4011,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f312,f550,f311,f552,f548,f310,f553,f547,f1285,f1287,f1469,f1288,f1471,f1486,f1488,f1489,f1491,f1493,f1498,f1500,f1504,f1506,f1513,f1515,f1522,f1524,f1530,f1536,f1549,f1575,f1579,f1584,f1591,f1632,f1636,f1641,f1648,f1676,f1678,f1680,f1686,f1688,f1690,f1699,f1701,f1706,f1708,f1717,f1719,f1724,f1726,f1735,f1737,f1742,f1744,f1755,f1757,f2117,f2118,f2119,f1768,f1770,f1772,f1774,f1779,f1784,f1790,f1792,f1794,f1799,f1801,f1806,f1808,f1813,f1815,f1820,f1822,f1828,f1835,f1842,f1863,f1869,f1871,f1881,f1883,f1888,f1890,f1892,f1898,f1903,f1909,f1914,f1916,f1925,f1927,f1932,f1934,f1943,f1945,f1950,f1952,f1961,f1963,f1968,f1970,f1977,f1984,f1996,f1998,f2005,f2007,f2014,f2016,f2023,f2036,f2043,f2045,f2047,f2056,f2058,f2064,f2069,f2074,f2081,f2086,f2088,f2095,f2097,f2103,f2105,f2120,f2121,f2122,f2123,f2124,f2125,f2126,f2127,f2128,f2129,f2130,f2131,f2132,f2133,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2141,f2142,f2143,f2144,f2145,f2146,f2147,f2148,f2149,f2150,f2151,f2152,f2153,f2154,f2155,f2156,f2157,f2158,f2159,f2160,f2161,f2162,f2163,f2164,f2165,f2166,f2167,f2168,f2169,f2170,f2171,f2172,f2173,f2174,f2175,f2176,f2177,f2178,f2179,f2180,f2181,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2189,f2190,f2191,f2192,f2193,f2194,f2195,f2196,f2197,f2198,f2199,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2218,f2219,f2220,f2221,f2222,f2223,f2224,f2225,f2226,f2227,f2228,f2229,f2230,f2243,f2244,f1487,f2917,f3120,f4007,f4009])).
fof(f4009,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4008,f550])).
fof(f4008,plain,(
  s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,sK17))) = s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,sK17))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f3998,f2919])).
fof(f2919,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK17))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(sK13,bool),sK14),s(fun(sK13,X0),X1),s(sK13,sK17)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f552,f452])).
fof(f452,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))))),
  inference(ennf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2)))) => s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2)))) => s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X6,X7] : ! [X8,X20,X11] : (p(s(bool,'const.bool.IN_2'(s(X7,X11),s(fun(X7,bool),X8)))) => s(X6,happ(s(fun(X7,X6),X20),s(X7,X11))) = s(X6,'const.bool.RES_ABSTRACT_3'(s(fun(X7,bool),X8),s(fun(X7,X6),X20),s(X7,X11))))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.bool.RES_ABSTRACT_DEF|split|0')).
fof(f3998,plain,(
  s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,sK17))) = s(sK12,'const.bool.RES_ABSTRACT_3'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15),s(sK13,sK17))) | ($spl0 | $spl4)),
  inference(superposition,[],[f2919,f3139])).
fof(f3139,plain,(
  ( ! [X0] : (s(sK12,'const.bool.RES_ABSTRACT_3'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15),s(sK13,X0))) = s(sK12,'const.bool.RES_ABSTRACT_3'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK16),s(sK13,X0)))) ) | $spl0),
  inference(forward_demodulation,[],[f3127,f451])).
fof(f451,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))),s(X1,X4)))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))),s(X1,X4)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6,X7] : ! [X23,X24,X25] : s(X6,'const.bool.RES_ABSTRACT_3'(s(fun(X7,bool),X23),s(fun(X7,X6),X24),s(X7,X25))) = s(X6,happ(s(fun(X7,X6),'const.bool.RES_ABSTRACT_2'(s(fun(X7,bool),X23),s(fun(X7,X6),X24))),s(X7,X25)))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','ARITY_2_3_const.bool.RES_ABSTRACT')).
fof(f3127,plain,(
  ( ! [X0] : (s(sK12,'const.bool.RES_ABSTRACT_3'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK16),s(sK13,X0))) = s(sK12,happ(s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15))),s(sK13,X0)))) ) | $spl0),
  inference(superposition,[],[f451,f547])).
fof(f4007,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4006,f550])).
fof(f4006,plain,(
  s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,sK17))) = s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,sK17))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f3997,f2919])).
fof(f3997,plain,(
  s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,sK17))) = s(sK12,'const.bool.RES_ABSTRACT_3'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15),s(sK13,sK17))) | ($spl0 | $spl4)),
  inference(superposition,[],[f3139,f2919])).
fof(f3120,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2918,f550])).
fof(f2918,plain,(
  s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,sK17))) = s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,sK17))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f552,f553])).
fof(f2917,plain,(
  $false | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f2715,f2683])).
fof(f2683,plain,(
  s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,sK29(sK16,sK15,sK14,sK12,sK13)))) != s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,sK29(sK16,sK15,sK14,sK12,sK13)))) | $spl1),
  inference(unit_resulting_resolution,[],[f548,f454])).
fof(f454,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,happ(s(fun(X0,X1),X3),s(X0,sK29(X4,X3,X2,X1,X0)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK29(X4,X3,X2,X1,X0)))) | s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4)))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.bool.IN_2'(s(X0,sK29(X4,X3,X2,X1,X0)),s(fun(X0,bool),X2)))) & s(X1,happ(s(fun(X0,X1),X3),s(X0,sK29(X4,X3,X2,X1,X0)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK29(X4,X3,X2,X1,X0))))) | s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29])],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4] : (? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X2)))) & s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,X5)))) | s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4))))),
  inference(ennf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3,X4] : (! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X2)))) => s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))) = s(X1,happ(s(fun(X0,X1),X4),s(X0,X5)))) => s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X2)))) => s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))) = s(X1,happ(s(fun(X0,X1),X4),s(X0,X5)))) => s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X7,X6] : ! [X8,X9,X10] : (! [X11] : (p(s(bool,'const.bool.IN_2'(s(X7,X11),s(fun(X7,bool),X8)))) => s(X6,happ(s(fun(X7,X6),X9),s(X7,X11))) = s(X6,happ(s(fun(X7,X6),X10),s(X7,X11)))) => s(fun(X7,X6),'const.bool.RES_ABSTRACT_2'(s(fun(X7,bool),X8),s(fun(X7,X6),X9))) = s(fun(X7,X6),'const.bool.RES_ABSTRACT_2'(s(fun(X7,bool),X8),s(fun(X7,X6),X10))))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.bool.RES_ABSTRACT_DEF|split|1')).
fof(f2715,plain,(
  s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,sK29(sK16,sK15,sK14,sK12,sK13)))) = s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,sK29(sK16,sK15,sK14,sK12,sK13)))) | ($spl1 | $spl6)),
  inference(unit_resulting_resolution,[],[f2584,f553])).
fof(f2584,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK29(sK16,sK15,sK14,sK12,sK13)),s(fun(sK13,bool),sK14)))) | $spl1),
  inference(unit_resulting_resolution,[],[f548,f453])).
fof(f453,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK29(X4,X3,X2,X1,X0)),s(fun(X0,bool),X2)))) | s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X3))) = s(fun(X0,X1),'const.bool.RES_ABSTRACT_2'(s(fun(X0,bool),X2),s(fun(X0,X1),X4)))) )),
  inference(cnf_transformation,[],[f272])).
fof(f1487,plain,(
  ( ! [X4] : (~p(s(bool,X4))) ) | $spl14),
  inference(cnf_transformation,[],[f1487_D])).
fof(f1487_D,plain,(
  ( ! [X4] : (~p(s(bool,X4))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2244,plain,(
  ( ! [X5] : (p(s(bool,sK21(X5,'const.combin.I_0',bool))) | p(s(bool,X5))) ) | $spl18),
  inference(cnf_transformation,[],[f2244_D])).
fof(f2244_D,plain,(
  ( ! [X5] : (p(s(bool,sK21(X5,'const.combin.I_0',bool))) | p(s(bool,X5))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2243,plain,(
  ( ! [X4,X5] : (p(s(bool,sK21(X5,'const.combin.I_0',bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f2240,f555])).
fof(f555,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f390,f389])).
fof(f389,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X7] : ! [X11] : s(X7,X11) = s(X7,'const.combin.I_1'(s(X7,X11)))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.combin.I_THM')).
fof(f390,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X7] : ! [X23] : s(X7,'const.combin.I_1'(s(X7,X23))) = s(X7,happ(s(fun(X7,X7),'const.combin.I_0'),s(X7,X23)))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','ARITY_0_1_const.combin.I')).
fof(f2240,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK21(X5,'const.combin.I_0',bool))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f401,f555])).
fof(f401,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0))))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21])],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X7] : ! [X16,X17] : ((? [X11] : p(s(bool,happ(s(fun(X7,bool),X16),s(X7,X11)))) | p(s(bool,X17))) <=> ? [X11] : (p(s(bool,happ(s(fun(X7,bool),X16),s(X7,X11)))) | p(s(bool,X17))))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.bool.LEFT_OR_EXISTS_THM')).
fof(f2230,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1490,f1502])).
fof(f1502,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1487,f1487,f407])).
fof(f407,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.bool.IMP_ANTISYM_AX')).
fof(f1490,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1487,f332])).
fof(f332,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.bool.EQ_CLAUSES|split|1')).
fof(f2229,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1499,f554])).
fof(f554,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f316,f313])).
fof(f313,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.bool.TRUTH')).
fof(f316,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f32])).
fof(f32,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.bool.NOT_CLAUSES|split|2')).
fof(f1499,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f580,f1487,f407])).
fof(f580,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f554,f332])).
fof(f2228,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1505,f554])).
fof(f1505,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f580,f1487,f407])).
fof(f2227,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1507,f554])).
fof(f1507,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f606,f1487,f430])).
fof(f430,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f261])).
fof(f261,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f606,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f554,f313,f432])).
fof(f432,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f2226,plain,(
  $false | ($spl9 | $spl14)),
  inference(subsumption_resolution,[],[f1508,f1487])).
fof(f1508,plain,(
  p(s(bool,sK18('const.combin.I_0',bool))) | ($spl9 | $spl14)),
  inference(unit_resulting_resolution,[],[f1299,f1487,f430])).
fof(f1299,plain,(
  sP1('T_0',sK18('const.combin.I_0',bool)) | $spl9),
  inference(unit_resulting_resolution,[],[f313,f1287,f432])).
fof(f2225,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1510,f1487])).
fof(f1510,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f760,f1487,f430])).
fof(f760,plain,(
  sP1('F_0','const.combin.I_1'(s(bool,'T_0')))),
  inference(unit_resulting_resolution,[],[f313,f554,f614])).
fof(f614,plain,(
  ( ! [X0,X1] : (sP1(X1,'const.combin.I_1'(s(bool,X0))) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f435,f389])).
fof(f435,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f2224,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1511,f554])).
fof(f1511,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f740,f1487,f430])).
fof(f740,plain,(
  sP1('const.combin.I_1'(s(bool,'T_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f313,f554,f607])).
fof(f607,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1('const.combin.I_1'(s(bool,X0)),X1)) )),
  inference(superposition,[],[f432,f389])).
fof(f2223,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1518,f554])).
fof(f1518,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f613,f1487,f430])).
fof(f613,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f554,f313,f435])).
fof(f2222,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1519,f554])).
fof(f1519,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f760,f1487,f430])).
fof(f2221,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1520,f1487])).
fof(f1520,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f740,f1487,f430])).
fof(f2220,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1548,f1629])).
fof(f1629,plain,(
  ( ! [X0] : (sP5('T_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f565,f559,f562,f1487,f467])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f276])).
fof(f276,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f275])).
fof(f275,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f562,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f313,f473])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f282])).
fof(f282,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f559,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f313,f469])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f279])).
fof(f279,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f565,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f313,f477])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f1548,plain,(
  ( ! [X0] : (~sP5('T_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f313,f1487,f464])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f277])).
fof(f2219,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1550,f1689])).
fof(f1689,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1487,f475])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f283])).
fof(f1550,plain,(
  ~sP3('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f559,f1487,f466])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f277])).
fof(f788,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f313,f580,f482])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f144,f185,f184,f183,f182])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X21,X22,X8] : ((p(s(bool,X8)) <=> s(bool,X21) = s(bool,X22)) <=> ((p(s(bool,X8)) | p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X8)) | ~p(s(bool,X21)) | ~p(s(bool,X22))) & (p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X8))) & (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X8)))))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.sat.sat_2|49|_prove_ltb')).
fof(f566,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f313,f478])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f285])).
fof(f2218,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1554,f1689])).
fof(f1554,plain,(
  ~sP3('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f560,f1487,f466])).
fof(f560,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f554,f470])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f280])).
fof(f2217,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1561,f1679])).
fof(f1679,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1487,f471])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f1561,plain,(
  ~sP4('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f562,f1487,f466])).
fof(f2216,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1580,f1679])).
fof(f1580,plain,(
  ~sP4('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f563,f1487,f466])).
fof(f563,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f554,f474])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f283])).
fof(f2215,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1607,f1689])).
fof(f1607,plain,(
  ~sP3('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f559,f1487,f467])).
fof(f2214,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1611,f1689])).
fof(f1611,plain,(
  ~sP3('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f560,f1487,f467])).
fof(f2213,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1618,f1679])).
fof(f1618,plain,(
  ~sP4('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f562,f1487,f467])).
fof(f2212,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1637,f1679])).
fof(f1637,plain,(
  ~sP4('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f563,f1487,f467])).
fof(f2211,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1675,f1679])).
fof(f1675,plain,(
  ( ! [X0] : (~sP4('T_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f313,f1487,f468])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f280])).
fof(f2210,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1685,f1689])).
fof(f1685,plain,(
  ( ! [X0] : (~sP3('T_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f313,f1487,f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f283])).
fof(f2209,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1692,f1487])).
fof(f1692,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f567,f554,f1487,f476])).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f285])).
fof(f567,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f313,f479])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f285])).
fof(f2208,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1693,f1487])).
fof(f1693,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f565,f1487,f1487,f476])).
fof(f2207,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1695,f1487])).
fof(f1695,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f567,f1487,f1487,f476])).
fof(f2206,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1697,f1487])).
fof(f1697,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f567,f554,f1487,f476])).
fof(f2205,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1703,f1487])).
fof(f1703,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f566,f1487,f1487,f476])).
fof(f2204,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1704,f1487])).
fof(f1704,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f567,f1487,f1487,f476])).
fof(f2203,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1710,f1487])).
fof(f1710,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f566,f554,f1487,f476])).
fof(f2202,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1712,f1487])).
fof(f1712,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f566,f1487,f1487,f476])).
fof(f2201,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1713,f1487])).
fof(f1713,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f567,f1487,f1487,f476])).
fof(f2200,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1714,f1487])).
fof(f1714,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f566,f554,f1487,f476])).
fof(f2199,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1720,f1487])).
fof(f1720,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f565,f1487,f1487,f476])).
fof(f2198,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1721,f1487])).
fof(f1721,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f566,f1487,f1487,f476])).
fof(f2197,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1727,f1487])).
fof(f1727,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f565,f554,f1487,f476])).
fof(f2196,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1729,f1487])).
fof(f1729,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f565,f1487,f1487,f476])).
fof(f2195,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1731,f1487])).
fof(f1731,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f567,f1487,f1487,f476])).
fof(f2194,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1732,f1487])).
fof(f1732,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f565,f554,f1487,f476])).
fof(f2193,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1738,f1487])).
fof(f1738,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f565,f1487,f1487,f476])).
fof(f2192,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1739,f1487])).
fof(f1739,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f566,f1487,f1487,f476])).
fof(f2191,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1745,f580])).
fof(f1745,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f786,f1487,f480])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f786,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f313,f580,f482])).
fof(f2190,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1746,f580])).
fof(f1746,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f788,f1487,f480])).
fof(f2189,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1760,f389])).
fof(f1760,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f778,f1487,f483])).
fof(f483,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f778,plain,(
  ( ! [X0] : (sP5('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f313,f389,f481])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f2188,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1761,f555])).
fof(f1761,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f818,f1487,f483])).
fof(f818,plain,(
  ( ! [X0] : (sP5('T_0',X0,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f313,f555,f481])).
fof(f2187,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1764,f555])).
fof(f1764,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f816,f1487,f483])).
fof(f816,plain,(
  ( ! [X0] : (sP5('T_0',happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f313,f555,f481])).
fof(f2186,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1765,f389])).
fof(f1765,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f777,f1487,f483])).
fof(f777,plain,(
  ( ! [X0] : (sP5('T_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f313,f389,f481])).
fof(f2185,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1778,f2082])).
fof(f2082,plain,(
  ( ! [X0,X1] : (sP7('T_0',X0,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f568,f1487,f537])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f490])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f288])).
fof(f288,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f568,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f313,f492])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f1778,plain,(
  ( ! [X0,X1] : (~sP7('T_0',X0,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f1487,f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f289])).
fof(f2184,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1783,f2082])).
fof(f1783,plain,(
  ( ! [X0,X1] : (~sP7('T_0',X0,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f1487,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f289])).
fof(f2183,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1789,f1793])).
fof(f1793,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1487,f494])).
fof(f494,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f1789,plain,(
  ( ! [X0] : (~sP6('T_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f313,f1487,f491])).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f2182,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1797,f554])).
fof(f1797,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f958,f1487,f495])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f146,f188,f187])).
fof(f146,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X21,X22,X8] : ((p(s(bool,X8)) <=> (p(s(bool,X22)) & p(s(bool,X21)))) <=> ((p(s(bool,X8)) | ~p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X8))) & (p(s(bool,X21)) | ~p(s(bool,X8)))))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.sat.sat_23_prove_ltb')).
fof(f958,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f953,f484])).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f289])).
fof(f953,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f313,f554,f313,f491])).
fof(f2181,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1802,f554])).
fof(f1802,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f624,f1487,f495])).
fof(f624,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f313,f554,f485])).
fof(f2180,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1803,f1487])).
fof(f1803,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f627,f1487,f495])).
fof(f627,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f313,f554,f486])).
fof(f2179,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1811,f554])).
fof(f1811,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f958,f1487,f496])).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f293])).
fof(f2178,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1816,f1487])).
fof(f1816,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f624,f1487,f496])).
fof(f2177,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1817,f554])).
fof(f1817,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f627,f1487,f496])).
fof(f2176,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1834,f1969])).
fof(f1969,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1487,f1487,f1487,f512])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f148,f191,f190])).
fof(f190,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f191,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f148,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X21,X22,X8] : ((p(s(bool,X8)) <=> (p(s(bool,X22)) | p(s(bool,X21)))) <=> ((p(s(bool,X8)) | ~p(s(bool,X22))) & (p(s(bool,X8)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | p(s(bool,X21)) | ~p(s(bool,X8)))))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.sat.sat_25_prove_ltb')).
fof(f1834,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f1487,f501])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f295])).
fof(f295,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f191])).
fof(f2175,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1841,f1969])).
fof(f1841,plain,(
  ( ! [X0,X1] : (~sP9(X0,'T_0',X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f1487,f502])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f296])).
fof(f2174,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1857,f1487])).
fof(f1857,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f573,f641,f1487,f507])).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f296])).
fof(f641,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f313,f554,f502])).
fof(f573,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f554,f511])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f298])).
fof(f298,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f190])).
fof(f2173,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1858,f1487])).
fof(f1858,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f572,f641,f1487,f507])).
fof(f572,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f313,f510])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f299])).
fof(f2172,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1862,f1915])).
fof(f1915,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1487,f511])).
fof(f1862,plain,(
  ~sP8('F_0','T_0','F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f641,f554,f1487,f507])).
fof(f2171,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1868,f1915])).
fof(f1868,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f640,f1487,f1487,f507])).
fof(f640,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f313,f554,f501])).
fof(f2170,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1870,f1915])).
fof(f1870,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f641,f1487,f1487,f507])).
fof(f2169,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1873,f1487])).
fof(f1873,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f640,f571,f1487,f507])).
fof(f571,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f313,f509])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f2168,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1876,f1487])).
fof(f1876,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f640,f573,f1487,f507])).
fof(f2167,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1878,f1915])).
fof(f1878,plain,(
  ~sP8('F_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f640,f554,f1487,f507])).
fof(f2166,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1884,f1915])).
fof(f1884,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f640,f1487,f1487,f507])).
fof(f2165,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1885,f1915])).
fof(f1885,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f641,f1487,f1487,f507])).
fof(f2164,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1894,f1487])).
fof(f1894,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f572,f313,f1487,f508])).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f299])).
fof(f2163,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1897,f1915])).
fof(f1897,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f554,f1487,f508])).
fof(f2162,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1902,f1915])).
fof(f1902,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f1487,f1487,f508])).
fof(f2161,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1904,f1487])).
fof(f1904,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f571,f313,f1487,f508])).
fof(f2160,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1908,f1915])).
fof(f1908,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f554,f1487,f508])).
fof(f2159,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1913,f1915])).
fof(f1913,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f1487,f1487,f508])).
fof(f2158,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1918,f1487])).
fof(f1918,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f641,f554,f1487,f512])).
fof(f2157,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1920,f1487])).
fof(f1920,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f641,f1487,f1487,f512])).
fof(f2156,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1921,f554])).
fof(f1921,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1054,f1487,f1487,f512])).
fof(f1054,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1049,f503])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f296])).
fof(f1049,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f313,f554,f554,f508])).
fof(f2155,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1922,f554])).
fof(f1922,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f641,f554,f1487,f512])).
fof(f2154,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1928,f554])).
fof(f1928,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f640,f1487,f1487,f512])).
fof(f2153,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1929,f554])).
fof(f1929,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f641,f1487,f1487,f512])).
fof(f2152,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1935,f1487])).
fof(f1935,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f640,f554,f1487,f512])).
fof(f2151,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1937,f1487])).
fof(f1937,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f640,f1487,f1487,f512])).
fof(f2150,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1939,f554])).
fof(f1939,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1054,f1487,f1487,f512])).
fof(f2149,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1940,f554])).
fof(f1940,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f640,f554,f1487,f512])).
fof(f2148,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1946,f554])).
fof(f1946,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f640,f1487,f1487,f512])).
fof(f2147,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1947,f554])).
fof(f1947,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f641,f1487,f1487,f512])).
fof(f2146,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1954,f554])).
fof(f1954,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1054,f554,f1487,f512])).
fof(f2145,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1955,f1487])).
fof(f1955,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f640,f1487,f1487,f512])).
fof(f2144,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1957,f554])).
fof(f1957,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1054,f1487,f1487,f512])).
fof(f2143,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1959,f554])).
fof(f1959,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1054,f554,f1487,f512])).
fof(f2142,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1965,f1487])).
fof(f1965,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f641,f1487,f1487,f512])).
fof(f2141,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1966,f554])).
fof(f1966,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1054,f1487,f1487,f512])).
fof(f2140,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1972,f1487])).
fof(f1972,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f651,f313,f1487,f515])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f301])).
fof(f651,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f313,f313,f514])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f301])).
fof(f2139,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1993,f1487])).
fof(f1993,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f643,f313,f1487,f515])).
fof(f643,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f313,f313,f513])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f301])).
fof(f2138,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1999,f554])).
fof(f1999,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f660,f1487,f518])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f303])).
fof(f303,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f660,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f554,f313,f530])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f177,f194,f193])).
fof(f193,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f177,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X21,X22,X8] : ((p(s(bool,X8)) <=> (p(s(bool,X22)) => p(s(bool,X21)))) <=> ((p(s(bool,X8)) | p(s(bool,X22))) & (p(s(bool,X8)) | ~p(s(bool,X21))) & (~p(s(bool,X22)) | p(s(bool,X21)) | ~p(s(bool,X8)))))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.sat.sat_27_prove_ltb')).
fof(f2137,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2000,f1487])).
fof(f2000,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f667,f1487,f518])).
fof(f667,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f313,f313,f531])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f309])).
fof(f2136,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2002,f554])).
fof(f2002,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1157,f1487,f518])).
fof(f1157,plain,(
  ( ! [X0] : (sP11('const.combin.I_1'(s(bool,'T_0')),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f313,f554,f661])).
fof(f661,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP11('const.combin.I_1'(s(bool,X0)),X2,X1)) )),
  inference(superposition,[],[f530,f389])).
fof(f2135,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2003,f1487])).
fof(f2003,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1174,f1487,f518])).
fof(f1174,plain,(
  ( ! [X0] : (sP11('const.combin.I_1'(s(bool,'T_0')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f313,f313,f668])).
fof(f668,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | sP11('const.combin.I_1'(s(bool,X0)),X1,X2)) )),
  inference(superposition,[],[f531,f389])).
fof(f2134,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2010,f554])).
fof(f2010,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1116,f1487,f518])).
fof(f1116,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f576,f313,f554,f524])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f304])).
fof(f576,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f554,f528])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f306])).
fof(f306,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f193])).
fof(f2133,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2011,f1487])).
fof(f2011,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1157,f1487,f518])).
fof(f2132,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2012,f1487])).
fof(f2012,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1174,f1487,f518])).
fof(f2131,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2035,f2046])).
fof(f2046,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1487,f528])).
fof(f2035,plain,(
  ~sP10('F_0','T_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f657,f313,f1487,f524])).
fof(f657,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f313,f554,f519])).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f304])).
fof(f2130,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2042,f2046])).
fof(f2042,plain,(
  ( ! [X0] : (~sP10('T_0',X0,'T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f313,f313,f1487,f525])).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f307])).
fof(f2129,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2053,f554])).
fof(f2053,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f657,f313,f1487,f529])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f309])).
fof(f2128,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2060,f554])).
fof(f2060,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1132,f313,f1487,f529])).
fof(f1132,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1127,f520])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f304])).
fof(f1127,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f313,f313,f554,f525])).
fof(f2127,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2085,f568])).
fof(f2085,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f624,f1487,f537])).
fof(f2126,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2087,f568])).
fof(f2087,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f627,f1487,f537])).
fof(f2125,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2091,f1487])).
fof(f2091,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f592,f1487,f556])).
fof(f556,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f415,f554])).
fof(f415,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f592,plain,(
  ( ! [X0] : (sP0(X0,'const.combin.I_1'(s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f313,f579])).
fof(f579,plain,(
  ( ! [X2,X3] : (sP0(X3,'const.combin.I_1'(s(bool,X2))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f416,f389])).
fof(f416,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f255])).
fof(f2124,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2092,f1487])).
fof(f2092,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f558,f1487,f556])).
fof(f558,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f313,f417])).
fof(f417,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f255])).
fof(f2123,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2093,f1487])).
fof(f2093,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f586,f1487,f556])).
fof(f586,plain,(
  ( ! [X0] : (sP0('const.combin.I_1'(s(bool,'T_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f313,f578])).
fof(f578,plain,(
  ( ! [X0,X1] : (sP0('const.combin.I_1'(s(bool,X0)),X1) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f417,f389])).
fof(f2122,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2098,f1487])).
fof(f2098,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f557,f1487,f556])).
fof(f557,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f313,f416])).
fof(f2121,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2099,f1487])).
fof(f2099,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f592,f1487,f556])).
fof(f2120,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2101,f1487])).
fof(f2101,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f586,f1487,f556])).
fof(f2105,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f557,f1487,f1487,f556])).
fof(f2103,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f557,f554,f1487,f556])).
fof(f2097,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f557,f1487,f1487,f556])).
fof(f2095,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f558,f554,f1487,f556])).
fof(f2088,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f570,f627,f1487,f537])).
fof(f570,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f554,f494])).
fof(f2086,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f569,f624,f1487,f537])).
fof(f569,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f554,f493])).
fof(f493,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f291])).
fof(f2081,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f667,f313,f313,f1487,f532])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f309])).
fof(f2074,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1132,f313,f1487,f530])).
fof(f2069,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f657,f313,f1487,f1487,f529])).
fof(f2064,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1132,f313,f554,f1487,f529])).
fof(f2058,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f657,f1487,f313,f1487,f529])).
fof(f2056,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f657,f554,f313,f1487,f529])).
fof(f2047,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1127,f1487,f528])).
fof(f2045,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1127,f1487,f526])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f307])).
fof(f2043,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f575,f313,f313,f1487,f525])).
fof(f575,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f313,f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f307])).
fof(f2036,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f576,f657,f313,f1487,f524])).
fof(f2023,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f667,f313,f1487,f519])).
fof(f2016,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f660,f1487,f1487,f518])).
fof(f2014,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1116,f554,f1487,f518])).
fof(f2007,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f660,f1487,f1487,f518])).
fof(f2005,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f660,f554,f1487,f518])).
fof(f1998,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f643,f1487,f313,f1487,f515])).
fof(f1996,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f643,f554,f313,f1487,f515])).
fof(f1984,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f643,f313,f1487,f1487,f515])).
fof(f1977,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f651,f313,f554,f1487,f515])).
fof(f1970,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f1487,f1487,f1487,f512])).
fof(f1968,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f641,f554,f1487,f1487,f512])).
fof(f1963,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f1487,f554,f1487,f512])).
fof(f1961,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1054,f554,f554,f1487,f512])).
fof(f1952,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f1487,f1487,f1487,f512])).
fof(f1950,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f554,f1487,f1487,f512])).
fof(f1945,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f1487,f554,f1487,f512])).
fof(f1943,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f554,f554,f1487,f512])).
fof(f1934,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f1487,f1487,f1487,f512])).
fof(f1932,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f554,f1487,f1487,f512])).
fof(f1927,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f641,f1487,f554,f1487,f512])).
fof(f1925,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f641,f554,f554,f1487,f512])).
fof(f1916,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1049,f1487,f511])).
fof(f1914,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f571,f313,f1487,f1487,f508])).
fof(f1909,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f571,f313,f554,f1487,f508])).
fof(f1903,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f571,f313,f1487,f1487,f508])).
fof(f1898,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f572,f313,f554,f1487,f508])).
fof(f1892,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f573,f1487,f1487,f507])).
fof(f1890,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f641,f572,f1487,f1487,f507])).
fof(f1888,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f571,f1487,f1487,f507])).
fof(f1883,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f573,f554,f1487,f507])).
fof(f1881,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f640,f571,f554,f1487,f507])).
fof(f1871,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f573,f641,f1487,f1487,f507])).
fof(f1869,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f573,f640,f1487,f1487,f507])).
fof(f1863,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f573,f641,f554,f1487,f507])).
fof(f1842,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f651,f313,f1487,f502])).
fof(f1835,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f643,f313,f1487,f501])).
fof(f1828,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f941,f313,f313,f1487,f500])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f293])).
fof(f941,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f568,f313,f313,f487])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f289])).
fof(f1822,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f624,f1487,f1487,f496])).
fof(f1820,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f627,f554,f1487,f496])).
fof(f1815,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f624,f1487,f1487,f496])).
fof(f1813,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f958,f554,f1487,f496])).
fof(f1808,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f624,f1487,f1487,f495])).
fof(f1806,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f624,f554,f1487,f495])).
fof(f1801,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f624,f1487,f1487,f495])).
fof(f1799,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f958,f554,f1487,f495])).
fof(f1794,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f953,f1487,f494])).
fof(f1792,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f953,f1487,f493])).
fof(f1790,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f568,f313,f313,f1487,f491])).
fof(f1784,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f941,f313,f1487,f486])).
fof(f1779,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f941,f313,f1487,f485])).
fof(f1774,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f818,f555,f1487,f483])).
fof(f1772,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f778,f389,f1487,f483])).
fof(f1770,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f816,f555,f1487,f483])).
fof(f1768,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f777,f389,f1487,f483])).
fof(f2119,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f1766])).
fof(f1766,plain,(
  ( ! [X0] : (s(bool,'const.combin.I_1'(s(bool,X0))) != s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f804,f1487,f483])).
fof(f804,plain,(
  ( ! [X2] : (sP5('T_0','const.combin.I_1'(s(bool,X2)),'const.combin.I_1'(s(bool,X2)))) )),
  inference(superposition,[],[f777,f389])).
fof(f2118,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f1763])).
fof(f1763,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1225,f1487,f483])).
fof(f1225,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f565,f560,f563,f554,f466])).
fof(f2117,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f1762])).
fof(f1762,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1195,f1487,f483])).
fof(f1195,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f567,f559,f562,f313,f465])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f277])).
fof(f1757,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f788,f580,f1487,f480])).
fof(f1755,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f786,f580,f1487,f480])).
fof(f1744,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f565,f1487,f1487,f1487,f476])).
fof(f1742,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f565,f554,f1487,f1487,f476])).
fof(f1737,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f565,f1487,f554,f1487,f476])).
fof(f1735,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f565,f554,f554,f1487,f476])).
fof(f1726,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f565,f1487,f1487,f1487,f476])).
fof(f1724,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f565,f554,f1487,f1487,f476])).
fof(f1719,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f566,f1487,f554,f1487,f476])).
fof(f1717,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f566,f554,f554,f1487,f476])).
fof(f1708,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f565,f1487,f1487,f1487,f476])).
fof(f1706,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f566,f554,f1487,f1487,f476])).
fof(f1701,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f565,f1487,f554,f1487,f476])).
fof(f1699,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f567,f554,f554,f1487,f476])).
fof(f1690,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f923,f1487,f475])).
fof(f923,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f313,f554,f313,f472])).
fof(f1688,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f923,f1487,f474])).
fof(f1686,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f562,f313,f313,f1487,f472])).
fof(f1680,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f906,f1487,f471])).
fof(f906,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f554,f313,f313,f468])).
fof(f1678,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f906,f1487,f470])).
fof(f1676,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f799,f313,f313,f1487,f468])).
fof(f799,plain,(
  ( ! [X0] : (sP4('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f777,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f277])).
fof(f1648,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f559,f563,f1487,f467])).
fof(f1641,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f560,f563,f1487,f467])).
fof(f1636,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f560,f562,f1487,f467])).
fof(f1632,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f559,f562,f1487,f467])).
fof(f1591,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f559,f563,f1487,f466])).
fof(f1584,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f560,f563,f1487,f466])).
fof(f1579,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f560,f562,f1487,f466])).
fof(f1575,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f566,f788,f559,f562,f1487,f466])).
fof(f1549,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1195,f313,f313,f1487,f464])).
fof(f1536,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f604,f313,f1487,f435])).
fof(f604,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f313,f313,f431])).
fof(f431,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f262])).
fof(f1530,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f604,f313,f1487,f432])).
fof(f1524,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f606,f1487,f1487,f430])).
fof(f1522,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f613,f554,f1487,f430])).
fof(f1515,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f606,f1487,f1487,f430])).
fof(f1513,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f606,f554,f1487,f430])).
fof(f1506,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f554,f580,f1487,f407])).
fof(f1504,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1487,f580,f1487,f407])).
fof(f1500,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f554,f580,f1487,f407])).
fof(f1498,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1487,f580,f1487,f407])).
fof(f1493,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f580,f1487,f337])).
fof(f337,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.bool.EQ_CLAUSES|split|3')).
fof(f1491,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f389,f1487,f332])).
fof(f1489,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f313,f1487])).
fof(f1488,plain,(
  ( ! [X5] : (p(s(bool,sK20(X5,'const.combin.I_0',bool))) | p(s(bool,X5))) ) | $spl16),
  inference(cnf_transformation,[],[f1488_D])).
fof(f1488_D,plain,(
  ( ! [X5] : (p(s(bool,sK20(X5,'const.combin.I_0',bool))) | p(s(bool,X5))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1486,plain,(
  ( ! [X4,X5] : (p(s(bool,sK20(X5,'const.combin.I_0',bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f1483,f555])).
fof(f1483,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK20(X5,'const.combin.I_0',bool))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f399,f555])).
fof(f399,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f247])).
fof(f1471,plain,(
  ~p(s(bool,sK19('const.combin.I_0',bool))) | $spl13),
  inference(cnf_transformation,[],[f1471_D])).
fof(f1471_D,plain,(
  ~p(s(bool,sK19('const.combin.I_0',bool))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f1288,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl10),
  inference(cnf_transformation,[],[f1288_D])).
fof(f1288_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1469,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,sK19('const.combin.I_0',bool)))) )),
  inference(forward_demodulation,[],[f1466,f555])).
fof(f1466,plain,(
  ( ! [X3] : (~p(s(bool,sK19('const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X3))))) )),
  inference(superposition,[],[f394,f555])).
fof(f394,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19])],[f239])).
fof(f239,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X7] : ! [X16] : (~! [X11] : p(s(bool,happ(s(fun(X7,bool),X16),s(X7,X11)))) <=> ? [X11] : ~p(s(bool,happ(s(fun(X7,bool),X16),s(X7,X11)))))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.bool.NOT_FORALL_THM')).
fof(f1287,plain,(
  ~p(s(bool,sK18('const.combin.I_0',bool))) | $spl9),
  inference(cnf_transformation,[],[f1287_D])).
fof(f1287_D,plain,(
  ~p(s(bool,sK18('const.combin.I_0',bool))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1285,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,sK18('const.combin.I_0',bool)))) )),
  inference(forward_demodulation,[],[f1282,f555])).
fof(f1282,plain,(
  ( ! [X3] : (~p(s(bool,sK18('const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X3))))) )),
  inference(superposition,[],[f393,f555])).
fof(f393,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f240])).
fof(f547,plain,(
  s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15))) = s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK16))) | $spl0),
  inference(cnf_transformation,[],[f547_D])).
fof(f547_D,plain,(
  s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15))) = s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK16))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f553,plain,(
  ( ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X6),s(fun(sK13,bool),sK14)))) | s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,X6))) = s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,X6)))) ) | $spl6),
  inference(cnf_transformation,[],[f553_D])).
fof(f553_D,plain,(
  ( ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X6),s(fun(sK13,bool),sK14)))) | s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,X6))) = s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,X6)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f310,plain,(
  ( ! [X6] : (s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,X6))) = s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,X6))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,X6),s(fun(sK13,bool),sK14)))) | s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15))) = s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK16)))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  (s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15))) = s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK16))) | ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X6),s(fun(sK13,bool),sK14)))) | s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,X6))) = s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,X6))))) & (s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15))) != s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK16))) | (p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK14)))) & s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,sK17))) != s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,sK17)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16,sK17])],[f197])).
fof(f197,plain,(
  ? [X0,X1,X2,X3,X4] : ((s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))) = s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X4))) | ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X2)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))) & (s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))) != s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X4))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X2)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(rectify,[],[f196])).
fof(f196,plain,(
  ? [X0,X1,X2,X3,X4] : ((s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))) = s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X4))) | ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X2)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))) & (s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))) != s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X4))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X2)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ? [X0,X1,X2,X3,X4] : (s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))) = s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X4))) <~> ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X2)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ~! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))) = s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X4))) <=> ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X2)))) => s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))) = s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X4))) <=> ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X2)))) => s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (s(fun(X7,X6),'const.bool.RES_ABSTRACT_2'(s(fun(X7,bool),X8),s(fun(X7,X6),X9))) = s(fun(X7,X6),'const.bool.RES_ABSTRACT_2'(s(fun(X7,bool),X8),s(fun(X7,X6),X10))) <=> ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X7,X11),s(fun(X7,bool),X8)))) => s(X6,happ(s(fun(X7,X6),X9),s(X7,X11))) = s(X6,happ(s(fun(X7,X6),X10),s(X7,X11)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (s(fun(X7,X6),'const.bool.RES_ABSTRACT_2'(s(fun(X7,bool),X8),s(fun(X7,X6),X9))) = s(fun(X7,X6),'const.bool.RES_ABSTRACT_2'(s(fun(X7,bool),X8),s(fun(X7,X6),X10))) <=> ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X7,X11),s(fun(X7,bool),X8)))) => s(X6,happ(s(fun(X7,X6),X9),s(X7,X11))) = s(X6,happ(s(fun(X7,X6),X10),s(X7,X11)))))),
  file('thm.res_quan.RES_ABSTRACT_EQUAL_EQ','thm.res_quan.RES_ABSTRACT_EQUAL_EQ')).
fof(f548,plain,(
  s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15))) != s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK16))) | $spl1),
  inference(cnf_transformation,[],[f548_D])).
fof(f548_D,plain,(
  s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15))) != s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK16))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f552,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK14)))) | $spl4),
  inference(cnf_transformation,[],[f552_D])).
fof(f552_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK14)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f311,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK14)))) | s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15))) != s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK16)))),
  inference(cnf_transformation,[],[f198])).
fof(f550,plain,(
  s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,sK17))) != s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,sK17))) | $spl3),
  inference(cnf_transformation,[],[f550_D])).
fof(f550_D,plain,(
  s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,sK17))) != s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,sK17))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f312,plain,(
  s(sK12,happ(s(fun(sK13,sK12),sK15),s(sK13,sK17))) != s(sK12,happ(s(fun(sK13,sK12),sK16),s(sK13,sK17))) | s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK15))) != s(fun(sK13,sK12),'const.bool.RES_ABSTRACT_2'(s(fun(sK13,bool),sK14),s(fun(sK13,sK12),sK16)))),
  inference(cnf_transformation,[],[f198])).
% SZS output end Proof for thm.res_quan
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 353
% Passive clauses: 867
% Generated clauses: 3695
% Final active clauses: 346
% Final passive clauses: 378
% Input formulas: 60
% Initial clauses: 225
% 
% Duplicate literals: 12
% Trivial inequalities: 6
% Fw subsumption resolutions: 141
% Fw demodulations: 410
% Bw demodulations: 26
% Condensations: 6
% 
% Simple tautologies: 81
% Equational tautologies: 27
% Forward subsumptions: 1908
% 
% Unit resulting resolution: 2255
% Factoring: 1
% Forward superposition: 311
% Backward superposition: 280
% Self superposition: 8
% Equality factoring: 2
% Equality resolution: 2
% 
% Split clauses: 7
% Split components: 14
% SAT solver clauses: 1113
% SAT solver unit clauses: 346
% SAT solver binary clauses: 567
% 
% Sat splits: 7
% Sat splitting refutations: 230
% 
% Memory used [KB]: 3582
% Time elapsed: 0.108 s
% ------------------------------
% ------------------------------
% Success in time 1.428 s
