Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.numpair
% (103727)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1227
% Passive clauses: 7520
% Generated clauses: 42258
% Final active clauses: 1113
% Final passive clauses: 5738
% Input formulas: 516
% Initial clauses: 898
% 
% Duplicate literals: 177
% Trivial inequalities: 145
% Fw subsumption resolutions: 2018
% Fw demodulations: 5913
% Bw demodulations: 67
% 
% Simple tautologies: 330
% Equational tautologies: 22
% Forward subsumptions: 25963
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 8019
% Factoring: 506
% Forward superposition: 12424
% Backward superposition: 11878
% Self superposition: 45
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 297
% Split components: 596
% SAT solver clauses: 507
% SAT solver unit clauses: 26
% SAT solver binary clauses: 104
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 2
% 
% Sat splits: 297
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 503
% TWLsolver calls for satisfiability: 394
% 
% Memory used [KB]: 19189
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numpair
% (103744)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 677
% Passive clauses: 4411
% Generated clauses: 34561
% Final active clauses: 640
% Final passive clauses: 3507
% Input formulas: 516
% Initial clauses: 898
% 
% Duplicate literals: 106
% Trivial inequalities: 108
% Fw subsumption resolutions: 1064
% Fw demodulations: 5673
% Bw demodulations: 261
% Condensations: 85
% Global subsumptions: 39
% 
% Simple tautologies: 362
% Equational tautologies: 16
% Forward subsumptions: 17388
% Fw demodulations to eq. taut.: 10
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 2898
% Factoring: 101
% Forward superposition: 11287
% Backward superposition: 7611
% Self superposition: 83
% Equality factoring: 2
% Equality resolution: 18
% 
% Split clauses: 25
% Split components: 50
% SAT solver clauses: 6266
% SAT solver unit clauses: 1200
% SAT solver binary clauses: 280
% SAT solver learnt clauses: 78
% SAT solver learnt literals: 19
% 
% Sat splits: 25
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 6111
% TWLsolver calls for satisfiability: 5250
% 
% Memory used [KB]: 29551
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numpair
% (103762)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1548
% Passive clauses: 12777
% Generated clauses: 148359
% Final active clauses: 1355
% Final passive clauses: 3501
% Input formulas: 516
% Initial clauses: 898
% Discarded non-redundant clauses: 47535
% 
% Duplicate literals: 26
% Trivial inequalities: 216
% Fw subsumption resolutions: 2741
% Fw demodulations: 37747
% Bw demodulations: 5577
% Condensations: 52
% Global subsumptions: 1554
% 
% Simple tautologies: 270
% Equational tautologies: 798
% Forward subsumptions: 69910
% Fw demodulations to eq. taut.: 1654
% Bw demodulations to eq. taut.: 84
% 
% Unit resulting resolution: 47824
% Factoring: 2
% Forward superposition: 24803
% Backward superposition: 16977
% Self superposition: 127
% Equality factoring: 2
% Equality resolution: 47
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 17150
% SAT solver unit clauses: 8195
% SAT solver binary clauses: 7481
% 
% Sat splits: 23
% Sat splitting refutations: 2771
% 
% Memory used [KB]: 265837
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.numpair
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numpair
% SZS output start Proof for thm.numpair
fof(f30340,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f821,f1105,f1103,f820,f1104,f1102,f1370,f1540,f1374,f1372,f1547,f1552,f1553,f1554,f1555,f1558,f1559,f1560,f1561,f1570,f1699,f1700,f1702,f1703,f1704,f1707,f1806,f2472,f2473,f2475,f2476,f2477,f2480,f2579,f2812,f2799,f2890,f2891,f2892,f2893,f2996,f2997,f2998,f2999,f3002,f3003,f3004,f3005,f3014,f3542,f3543,f3544,f3545,f3618,f3619,f3620,f3621,f3657,f3658,f3808,f3853,f3854,f3855,f3930,f3931,f3932,f4007,f4008,f4009,f4085,f4086,f4087,f4161,f4162,f4163,f4280,f4350,f4396,f4397,f4472,f4473,f4548,f4549,f4622,f4623,f4696,f4697,f5844,f6449,f6452,f6451,f6875,f6876,f6877,f6878,f6893,f6962,f6963,f6964,f6965,f6980,f7043,f7044,f7045,f7046,f7061,f7317,f7321,f7319,f7096,f7098,f7099,f7100,f7103,f7202,f7356,f7380,f7360,f5530,f7390,f7391,f7393,f7397,f7398,f7399,f7400,f5845,f7401,f7422,f7403,f4873,f4304,f4348,f4296,f7436,f7437,f7438,f7439,f4278,f7440,f1975,f1974,f1973,f1972,f7395,f7457,f1942,f1941,f2204,f2202,f2200,f1580,f2198,f2196,f1578,f2194,f2188,f2190,f2192,f7536,f7549,f2668,f7550,f7552,f7553,f7554,f7557,f7656,f2669,f8848,f8849,f8850,f8851,f8866,f8923,f8924,f8925,f8926,f8941,f8998,f8999,f9000,f9001,f9016,f9074,f9075,f9076,f9077,f9092,f9150,f9151,f9152,f9153,f9168,f1582,f9188,f9189,f9190,f9191,f9194,f9195,f3285,f9196,f9197,f3287,f9206,f2666,f9728,f9729,f9730,f9731,f9803,f9804,f9805,f9806,f9878,f9879,f9880,f9881,f9953,f9954,f9955,f9956,f10028,f10029,f10030,f10031,f10101,f10102,f10103,f10104,f10174,f1573,f10175,f10176,f10177,f10287,f1574,f10288,f10437,f10525,f1575,f10594,f10607,f10608,f10610,f10611,f10612,f10615,f10714,f10976,f10977,f10979,f10980,f10981,f10984,f11083,f10962,f10549,f10593,f10393,f10436,f10265,f11228,f10255,f10254,f11230,f10248,f11233,f10247,f11234,f10246,f7357,f11235,f11236,f11244,f11252,f11238,f11255,f11245,f7358,f11492,f11493,f11682,f1371,f12046,f12075,f12241,f7442,f7394,f12329,f11199,f11197,f12330,f12332,f12333,f12334,f12337,f12436,f7322,f12625,f12700,f12775,f12851,f12927,f13004,f13079,f13154,f13229,f13414,f13487,f13560,f13600,f1895,f1893,f13601,f13603,f13604,f13605,f13608,f13707,f1896,f14397,f14542,f14541,f14540,f14539,f14535,f14636,f14637,f14676,f14732,f14790,f14848,f14906,f14963,f15020,f15077,f15280,f15279,f15278,f15277,f15273,f15374,f15375,f15557,f15544,f15635,f15636,f15637,f15638,f15820,f15807,f15898,f15899,f15900,f15901,f16464,f16465,f16466,f16467,f16482,f16545,f16546,f16547,f16548,f16563,f16626,f16627,f16628,f16629,f16644,f16708,f16709,f16710,f16711,f16726,f16791,f16792,f16793,f16794,f16809,f16874,f16875,f16876,f16877,f16892,f16955,f16956,f16957,f16958,f16973,f17036,f17037,f17038,f17039,f17054,f17117,f17118,f17119,f17120,f17135,f17199,f17200,f17201,f17202,f17217,f17278,f17279,f17280,f17281,f17296,f17357,f6914,f17358,f17359,f17360,f17375,f18372,f18374,f18375,f18376,f18379,f18478,f18554,f18565,f6915,f18566,f18718,f18582,f3807,f18955,f18959,f18957,f18761,f18763,f18764,f18765,f18768,f18867,f18960,f19189,f6916,f19258,f3809,f19205,f4349,f6919,f20113,f20100,f20191,f20192,f20193,f20194,f20398,f20394,f20393,f20392,f20391,f20387,f20489,f20490,f20498,f20500,f20501,f20503,f20505,f20605,f20680,f20723,f20747,f20784,f20808,f20845,f20870,f20907,f20932,f20969,f20994,f21032,f21055,f21093,f21116,f21154,f21177,f21215,f21239,f21276,f21298,f21335,f21357,f21394,f21416,f21459,f21460,f21461,f21462,f21477,f3575,f21544,f21545,f21546,f21547,f21562,f3651,f21629,f21630,f21631,f21632,f21647,f3886,f21715,f21716,f21717,f21718,f21733,f3963,f21801,f21802,f21803,f21804,f21819,f4040,f21888,f21889,f21890,f21891,f21906,f4116,f21973,f21974,f21975,f21976,f21991,f4192,f22029,f22031,f22032,f22033,f22036,f22135,f22277,f22330,f22331,f22332,f22333,f22348,f4426,f22415,f22416,f22417,f22418,f22433,f4503,f22500,f22501,f22502,f22503,f22518,f4577,f22583,f22584,f22585,f22586,f22601,f4651,f22666,f22667,f22668,f22669,f22684,f4725,f22720,f22722,f22723,f22724,f22727,f22826,f23005,f23040,f23227,f23225,f23041,f23043,f23044,f23045,f23048,f23147,f23228,f23278,f23464,f23462,f23279,f23281,f23282,f23283,f23286,f23385,f23465,f23474,f23476,f23477,f23479,f23481,f23581,f23654,f23693,f23750,f23778,f24100,f24337,f24395,f24453,f24512,f24569,f24626,f24683,f24740,f24795,f24850,f1576,f24883,f24885,f24886,f24887,f24890,f24989,f25170,f25220,f25276,f25335,f25392,f25451,f25508,f25565,f25624,f25681,f6917,f25713,f25715,f25716,f25717,f25720,f25819,f25895,f25964,f25965,f25966,f25967,f25982,f25239,f26023,f26024,f26025,f26026,f26041,f25297,f26084,f26085,f26086,f26087,f26102,f25355,f26144,f26145,f26146,f26147,f26162,f25413,f26206,f26207,f26208,f26209,f26224,f25470,f26266,f26267,f26268,f26269,f26284,f25527,f26306,f26308,f26309,f26310,f26313,f26412,f26593,f26641,f26642,f26643,f26644,f26659,f25585,f26702,f26703,f26704,f26705,f26720,f25643,f26762,f26763,f26764,f26765,f26780,f25700,f26795,f26797,f26798,f26800,f26802,f26902,f26976,f27020,f27021,f27022,f27023,f27038,f23713,f27080,f27081,f27082,f27083,f27098,f23770,f27140,f27141,f27142,f27143,f27158,f24358,f27201,f27202,f27203,f27204,f27219,f24416,f27262,f27263,f27264,f27265,f27280,f24474,f27324,f27325,f27326,f27327,f27342,f24531,f27384,f27385,f27386,f27387,f27402,f24588,f27444,f27445,f27446,f27447,f27462,f24645,f27504,f27505,f27506,f27507,f27522,f24703,f27564,f27565,f27566,f27567,f27582,f24758,f27623,f27624,f27625,f27626,f27641,f24813,f27681,f3280,f27682,f3281,f27683,f3282,f27684,f3283,f27699,f24868,f3293,f28243,f28245,f28246,f1897,f28248,f1899,f28250,f1898,f28350,f1900,f28426,f28999,f29014,f30322,f30325])).
fof(f30325,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f30324,f1045])).
fof(f1045,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f604])).
fof(f604,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))),
  inference(flattening,[],[f603])).
fof(f603,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X0))))),
  inference(rectify,[],[f133])).
fof(f133,axiom,(
  ! [X7] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X7))))),
  file('thm.numpair.tri_|49||49|','thm.prim_rec.LESS_REFL')).
fof(f30324,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f30323,f2248])).
fof(f2248,plain,(
  ( ! [X2] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',X2))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',X2)))) ) | $spl0),
  inference(forward_demodulation,[],[f1970,f901])).
fof(f901,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f564])).
fof(f564,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X7,X6] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.numpair.tri_1'(s('type.num.num',X7))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',X6)))))),
  file('thm.numpair.tri_|49||49|','thm.numpair.tri_LT')).
fof(f1970,plain,(
  ( ! [X2] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',X2))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',X2)))))) ) | $spl0),
  inference(superposition,[],[f901,f1102])).
fof(f30323,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f30136,f1105])).
fof(f30136,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | s('type.num.num',sK14) = s('type.num.num',sK15) | $spl0),
  inference(resolution,[],[f28692,f879])).
fof(f879,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f539])).
fof(f539,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f302])).
fof(f302,axiom,(
  ! [X6,X7] : (s('type.num.num',X6) = s('type.num.num',X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.LESS_LESS_CASES')).
fof(f28692,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) | $spl0),
  inference(superposition,[],[f1045,f2248])).
fof(f30322,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f30321,f1045])).
fof(f30321,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f30320,f2248])).
fof(f30320,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f30135,f1105])).
fof(f30135,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | s('type.num.num',sK14) = s('type.num.num',sK15) | $spl0),
  inference(resolution,[],[f28692,f879])).
fof(f29014,plain,(
  $false | ($spl0 | $spl20 | $spl136)),
  inference(subsumption_resolution,[],[f28692,f23223])).
fof(f23223,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) | ($spl20 | $spl136)),
  inference(forward_demodulation,[],[f23039,f1582])).
fof(f23039,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15))))))) | $spl136),
  inference(resolution,[],[f18957,f1062])).
fof(f1062,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f813])).
fof(f813,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) & ((s('type.num.num',X0) != s('type.num.num',X1) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(flattening,[],[f812])).
fof(f812,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))))) & ((s('type.num.num',X0) != s('type.num.num',X1) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f618])).
fof(f618,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) <=> (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f276])).
fof(f276,axiom,(
  ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))) <=> (s('type.num.num',X6) = s('type.num.num',X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))))))),
  file('thm.numpair.tri_|49||49|','thm.prim_rec.LESS_THM')).
fof(f28999,plain,(
  $false | ($spl0 | $spl5 | $spl136)),
  inference(subsumption_resolution,[],[f28564,f18957])).
fof(f28564,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) | ($spl0 | $spl5)),
  inference(backward_demodulation,[],[f2248,f18366])).
fof(f18366,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))))) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f1012])).
fof(f1012,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f671])).
fof(f671,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f585])).
fof(f585,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f365])).
fof(f365,axiom,(
  ! [X6,X7] : ~(p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.LESS_ANTISYM')).
fof(f3067,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) | ($spl0 | $spl5)),
  inference(superposition,[],[f1091,f2205])).
fof(f2205,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f2204,f1609])).
fof(f1609,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | $spl5),
  inference(trivial_inequality_removal,[],[f1596])).
fof(f1596,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | $spl5),
  inference(superposition,[],[f1371,f904])).
fof(f904,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f747])).
fof(f747,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f566])).
fof(f566,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X7))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7))),
  file('thm.numpair.tri_|49||49|','thm.numpair.tri_eq_|48||split|1')).
fof(f1091,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f647])).
fof(f647,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f391])).
fof(f391,axiom,(
  ! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))),
  file('thm.numpair.tri_|49||49|','thm.prim_rec.LESS_SUC_REFL')).
fof(f28426,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK14))))) ) | $spl238),
  inference(cnf_transformation,[],[f28426_D])).
fof(f28426_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK14))))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f1900,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) | $spl34),
  inference(cnf_transformation,[],[f1900_D])).
fof(f1900_D,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f28350,plain,(
  ( ! [X210,X209] : (p(s(bool,X209)) | ~p(s(bool,X210)) | s(bool,X210) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK14))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f25707,f1093])).
fof(f1093,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f826])).
fof(f826,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f717])).
fof(f717,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f716])).
fof(f716,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f650])).
fof(f650,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f649])).
fof(f649,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X13,X22,X14,X23] : ((s(bool,X13) = s(bool,X22) & (p(s(bool,X22)) => s(bool,X14) = s(bool,X23))) => ((p(s(bool,X13)) => p(s(bool,X14))) <=> (p(s(bool,X22)) => p(s(bool,X23)))))),
  file('thm.numpair.tri_|49||49|','thm.bool.IMP_CONG')).
fof(f25707,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK14))))) | ($spl0 | $spl5)),
  inference(resolution,[],[f19334,f1012])).
fof(f19334,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK14)),s('type.num.num',sK14)))) | ($spl0 | $spl5)),
  inference(superposition,[],[f1091,f7348])).
fof(f7348,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14)))) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f7347,f1045])).
fof(f7347,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14)))) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f7264,f911])).
fof(f911,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.numpair.tri_|49||49|','thm.numpair.tri_def|split|0')).
fof(f7264,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14)))) | ($spl0 | $spl5)),
  inference(superposition,[],[f2225,f1086])).
fof(f1086,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f819])).
fof(f819,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26])],[f642])).
fof(f642,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | ? [X7] : s('type.num.num',X6) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.num_CASES')).
fof(f2225,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f2224,f1102])).
fof(f2224,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))))) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f1953,f1609])).
fof(f1953,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))))) | $spl0),
  inference(superposition,[],[f1102,f892])).
fof(f892,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f745])).
fof(f745,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)))),
  inference(nnf_transformation,[],[f557])).
fof(f557,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(flattening,[],[f556])).
fof(f556,plain,(
  ! [X0] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f207])).
fof(f207,axiom,(
  ! [X7] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.NOT_ZERO_LT_ZERO')).
fof(f1898,plain,(
  ( ! [X12,X11] : (s(bool,X11) != s(bool,X12) | ~p(s(bool,X11))) ) | $spl30),
  inference(cnf_transformation,[],[f1898_D])).
fof(f1898_D,plain,(
  ( ! [X12,X11] : (s(bool,X11) != s(bool,X12) | ~p(s(bool,X11))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f28250,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK14)))) | s(bool,X22) != s(bool,X24)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f25707,f829])).
fof(f829,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f717])).
fof(f1899,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | $spl32),
  inference(cnf_transformation,[],[f1899_D])).
fof(f1899_D,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f28248,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK14))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f25707,f828])).
fof(f828,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f717])).
fof(f1897,plain,(
  ( ! [X6,X4] : (s(bool,X4) != s(bool,X6) | p(s(bool,X4))) ) | $spl28),
  inference(cnf_transformation,[],[f1897_D])).
fof(f1897_D,plain,(
  ( ! [X6,X4] : (s(bool,X4) != s(bool,X6) | p(s(bool,X4))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f28246,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK14))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f25707,f827])).
fof(f827,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f717])).
fof(f28245,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK14)))) | s(bool,X7) != s(bool,X9)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f25707,f825])).
fof(f825,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f717])).
fof(f28243,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK14))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f25707,f823])).
fof(f823,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f717])).
fof(f3293,plain,(
  s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15))))) | $spl54),
  inference(cnf_transformation,[],[f3293_D])).
fof(f3293_D,plain,(
  s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f24868,plain,(
  ( ! [X26,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26)) ) | $spl210),
  inference(cnf_transformation,[],[f24868_D])).
fof(f24868_D,plain,(
  ( ! [X26,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f27699,plain,(
  ( ! [X68,X67] : (sP12(X67,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23136,f1086])).
fof(f23136,plain,(
  ( ! [X196,X197] : (sP12(X196,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))),X197)) ) | $spl136),
  inference(resolution,[],[f18957,f995])).
fof(f995,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f787])).
fof(f787,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))))),
  inference(flattening,[],[f786])).
fof(f786,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | (p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f711])).
fof(f711,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f581,f710])).
fof(f710,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f581,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(rectify,[],[f307])).
fof(f307,axiom,(
  ! [X8,X19,X27] : ((p(s(bool,X8)) | (p(s(bool,X19)) & p(s(bool,X27)))) <=> ((p(s(bool,X8)) | p(s(bool,X19))) & (p(s(bool,X8)) | p(s(bool,X27)))))),
  file('thm.numpair.tri_|49||49|','thm.bool.LEFT_OR_OVER_AND')).
fof(f3283,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK26(sK15)))))))) ) | $spl48),
  inference(cnf_transformation,[],[f3283_D])).
fof(f3283_D,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK26(sK15)))))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f27684,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23136,f872])).
fof(f872,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f653])).
fof(f653,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))))),
  inference(ennf_transformation,[],[f534])).
fof(f534,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))))),
  inference(flattening,[],[f533])).
fof(f533,plain,(
  ! [X0,X1] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))))),
  inference(rectify,[],[f390])).
fof(f390,axiom,(
  ! [X6,X7] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7)))))))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.LESS_ADD_NONZERO')).
fof(f3282,plain,(
  ( ! [X10] : (s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',sK26(sK15))))) ) | $spl46),
  inference(cnf_transformation,[],[f3282_D])).
fof(f3282_D,plain,(
  ( ! [X10] : (s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',sK26(sK15))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f27683,plain,(
  ( ! [X24,X23,X22] : (sP12(X22,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23136,f864])).
fof(f864,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f734])).
fof(f734,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f531])).
fof(f531,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))),
  inference(rectify,[],[f308])).
fof(f308,axiom,(
  ! [X6,X7] : (s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.ADD_INV_|48|_EQ')).
fof(f3281,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK26(sK15))))) ) | $spl44),
  inference(cnf_transformation,[],[f3281_D])).
fof(f3281_D,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK26(sK15))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f27682,plain,(
  ( ! [X21,X19,X20] : (sP12(X19,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23136,f862])).
fof(f862,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f733])).
fof(f733,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(flattening,[],[f732])).
fof(f732,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f530])).
fof(f530,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6,X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.ADD_EQ_|48|')).
fof(f3280,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X8)))) ) | $spl42),
  inference(cnf_transformation,[],[f3280_D])).
fof(f3280_D,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X8)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f27681,plain,(
  ( ! [X17,X18,X16] : (sP12(X16,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23136,f861])).
fof(f861,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f733])).
fof(f24813,plain,(
  ( ! [X26,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26)) ) | $spl208),
  inference(cnf_transformation,[],[f24813_D])).
fof(f24813_D,plain,(
  ( ! [X26,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f27641,plain,(
  ( ! [X68,X67] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23129,f1086])).
fof(f23129,plain,(
  ( ! [X182,X183] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))),X182,X183)) ) | $spl136),
  inference(resolution,[],[f18957,f986])).
fof(f986,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X2,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f782])).
fof(f782,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP11(X0,X2,X1)) & (~sP11(X0,X2,X1) | (p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))))),
  inference(flattening,[],[f781])).
fof(f781,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP11(X0,X2,X1)) & (~sP11(X0,X2,X1) | ((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f709])).
fof(f709,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> sP11(X0,X2,X1))),
  inference(definition_folding,[],[f580,f708])).
fof(f708,plain,(
  ! [X0,X2,X1] : (sP11(X0,X2,X1) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f580,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  inference(rectify,[],[f264])).
fof(f264,axiom,(
  ! [X8,X19,X27] : (((p(s(bool,X19)) & p(s(bool,X27))) | p(s(bool,X8))) <=> ((p(s(bool,X19)) | p(s(bool,X8))) & (p(s(bool,X27)) | p(s(bool,X8)))))),
  file('thm.numpair.tri_|49||49|','thm.bool.RIGHT_OR_OVER_AND')).
fof(f27626,plain,(
  ( ! [X26,X27,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23129,f872])).
fof(f27625,plain,(
  ( ! [X24,X23,X22] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23129,f864])).
fof(f27624,plain,(
  ( ! [X21,X19,X20] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23129,f862])).
fof(f27623,plain,(
  ( ! [X17,X18,X16] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23129,f861])).
fof(f24758,plain,(
  ( ! [X26,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26)) ) | $spl206),
  inference(cnf_transformation,[],[f24758_D])).
fof(f24758_D,plain,(
  ( ! [X26,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f27582,plain,(
  ( ! [X68,X67] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23118,f1086])).
fof(f23118,plain,(
  ( ! [X161,X160] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))),X160,X161)) ) | $spl136),
  inference(resolution,[],[f18957,f970])).
fof(f970,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f775])).
fof(f775,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X0)) | p(s(bool,X1))) & ~p(s(bool,X2))) | sP10(X2,X1,X0)) & (~sP10(X2,X1,X0) | (~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))))),
  inference(flattening,[],[f774])).
fof(f774,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X0)) | p(s(bool,X1))) & ~p(s(bool,X2))) | sP10(X2,X1,X0)) & (~sP10(X2,X1,X0) | ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f707])).
fof(f707,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))) <=> sP10(X2,X1,X0))),
  inference(definition_folding,[],[f666,f706])).
fof(f706,plain,(
  ! [X2,X1,X0] : (sP10(X2,X1,X0) <=> ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f666,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))) <=> ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f578])).
fof(f578,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X0)) | p(s(bool,X1))) => p(s(bool,X2))) <=> ((p(s(bool,X0)) => p(s(bool,X2))) & (p(s(bool,X1)) => p(s(bool,X2)))))),
  inference(rectify,[],[f376])).
fof(f376,axiom,(
  ! [X20,X28,X52] : (((p(s(bool,X20)) | p(s(bool,X28))) => p(s(bool,X52))) <=> ((p(s(bool,X20)) => p(s(bool,X52))) & (p(s(bool,X28)) => p(s(bool,X52)))))),
  file('thm.numpair.tri_|49||49|','thm.bool.DISJ_IMP_THM')).
fof(f27567,plain,(
  ( ! [X26,X27,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23118,f872])).
fof(f27566,plain,(
  ( ! [X24,X23,X22] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23118,f864])).
fof(f27565,plain,(
  ( ! [X21,X19,X20] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23118,f862])).
fof(f27564,plain,(
  ( ! [X17,X18,X16] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23118,f861])).
fof(f24703,plain,(
  ( ! [X26,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl204),
  inference(cnf_transformation,[],[f24703_D])).
fof(f24703_D,plain,(
  ( ! [X26,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f27522,plain,(
  ( ! [X68,X67] : (sP8(X67,X68,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23106,f1086])).
fof(f23106,plain,(
  ( ! [X136,X137] : (sP8(X136,X137,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(resolution,[],[f18957,f954])).
fof(f954,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f768])).
fof(f768,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(flattening,[],[f767])).
fof(f767,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(nnf_transformation,[],[f703])).
fof(f703,plain,(
  ! [X0,X1,X2] : (sP8(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f27507,plain,(
  ( ! [X26,X27,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23106,f872])).
fof(f27506,plain,(
  ( ! [X24,X23,X22] : (sP8(X22,X23,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23106,f864])).
fof(f27505,plain,(
  ( ! [X21,X19,X20] : (sP8(X19,X20,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23106,f862])).
fof(f27504,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,X17,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23106,f861])).
fof(f24645,plain,(
  ( ! [X26,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26)) ) | $spl202),
  inference(cnf_transformation,[],[f24645_D])).
fof(f24645_D,plain,(
  ( ! [X26,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f27462,plain,(
  ( ! [X68,X67] : (sP6(X67,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23090,f1086])).
fof(f23090,plain,(
  ( ! [X105,X104] : (sP6(X104,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))),X105)) ) | $spl136),
  inference(resolution,[],[f18957,f938])).
fof(f938,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f761])).
fof(f761,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f760])).
fof(f760,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f759])).
fof(f759,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f700])).
fof(f700,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f27447,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23090,f872])).
fof(f27446,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23090,f864])).
fof(f27445,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23090,f862])).
fof(f27444,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23090,f861])).
fof(f24588,plain,(
  ( ! [X26,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl200),
  inference(cnf_transformation,[],[f24588_D])).
fof(f24588_D,plain,(
  ( ! [X26,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f27402,plain,(
  ( ! [X68,X67] : (sP6(X67,X68,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23089,f1086])).
fof(f23089,plain,(
  ( ! [X103,X102] : (sP6(X102,X103,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(resolution,[],[f18957,f937])).
fof(f937,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f761])).
fof(f27387,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23089,f872])).
fof(f27386,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,X23,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23089,f864])).
fof(f27385,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,X20,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23089,f862])).
fof(f27384,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,X17,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23089,f861])).
fof(f24531,plain,(
  ( ! [X26,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26)) ) | $spl198),
  inference(cnf_transformation,[],[f24531_D])).
fof(f24531_D,plain,(
  ( ! [X26,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f27342,plain,(
  ( ! [X68,X67] : (sP4(X67,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23076,f1086])).
fof(f23076,plain,(
  ( ! [X76,X77] : (sP4(X76,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))),X77)) ) | $spl136),
  inference(resolution,[],[f18957,f921])).
fof(f921,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f753])).
fof(f753,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f752])).
fof(f752,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f751])).
fof(f751,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f697])).
fof(f697,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f27327,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23076,f872])).
fof(f27326,plain,(
  ( ! [X24,X23,X22] : (sP4(X22,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23076,f864])).
fof(f27325,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23076,f862])).
fof(f27324,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23076,f861])).
fof(f24474,plain,(
  ( ! [X26,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26)) ) | $spl196),
  inference(cnf_transformation,[],[f24474_D])).
fof(f24474_D,plain,(
  ( ! [X26,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f27280,plain,(
  ( ! [X68,X67] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23060,f1086])).
fof(f23060,plain,(
  ( ! [X50,X51] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))),X50,X51)) ) | $spl136),
  inference(resolution,[],[f18957,f848])).
fof(f848,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f728])).
fof(f728,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f727])).
fof(f727,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f692])).
fof(f692,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f27265,plain,(
  ( ! [X26,X27,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23060,f872])).
fof(f27264,plain,(
  ( ! [X24,X23,X22] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23060,f864])).
fof(f27263,plain,(
  ( ! [X21,X19,X20] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23060,f862])).
fof(f27262,plain,(
  ( ! [X17,X18,X16] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23060,f861])).
fof(f24416,plain,(
  ( ! [X26,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26)) ) | $spl194),
  inference(cnf_transformation,[],[f24416_D])).
fof(f24416_D,plain,(
  ( ! [X26,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f27219,plain,(
  ( ! [X68,X67] : (sP0(X67,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23059,f1086])).
fof(f23059,plain,(
  ( ! [X48,X49] : (sP0(X48,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))),X49)) ) | $spl136),
  inference(resolution,[],[f18957,f847])).
fof(f847,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f728])).
fof(f27204,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23059,f872])).
fof(f27203,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23059,f864])).
fof(f27202,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23059,f862])).
fof(f27201,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23059,f861])).
fof(f24358,plain,(
  ( ! [X26,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl192),
  inference(cnf_transformation,[],[f24358_D])).
fof(f24358_D,plain,(
  ( ! [X26,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f27158,plain,(
  ( ! [X68,X67] : (sP0(X67,X68,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23058,f1086])).
fof(f23058,plain,(
  ( ! [X47,X46] : (sP0(X46,X47,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(resolution,[],[f18957,f846])).
fof(f846,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f728])).
fof(f27143,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23058,f872])).
fof(f27142,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23058,f864])).
fof(f27141,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,X20,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23058,f862])).
fof(f27140,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,X17,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23058,f861])).
fof(f23770,plain,(
  ( ! [X26,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl188),
  inference(cnf_transformation,[],[f23770_D])).
fof(f23770_D,plain,(
  ( ! [X26,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f27098,plain,(
  ( ! [X68,X67] : (sP1(X67,X68,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23057,f1086])).
fof(f23057,plain,(
  ( ! [X45,X44] : (sP1(X44,X45,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(resolution,[],[f18957,f842])).
fof(f842,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f726])).
fof(f726,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f725])).
fof(f725,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f724])).
fof(f724,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f693])).
fof(f693,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f27083,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23057,f872])).
fof(f27082,plain,(
  ( ! [X24,X23,X22] : (sP1(X22,X23,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23057,f864])).
fof(f27081,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,X20,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23057,f862])).
fof(f27080,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,X17,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23057,f861])).
fof(f23713,plain,(
  ( ! [X26,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl186),
  inference(cnf_transformation,[],[f23713_D])).
fof(f23713_D,plain,(
  ( ! [X26,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f27038,plain,(
  ( ! [X68,X67] : (sP2(X67,X68,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl136),
  inference(superposition,[],[f23054,f1086])).
fof(f23054,plain,(
  ( ! [X39,X38] : (sP2(X38,X39,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(resolution,[],[f18957,f838])).
fof(f838,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f723])).
fof(f723,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f722])).
fof(f722,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f721])).
fof(f721,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f694])).
fof(f694,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f27023,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl136),
  inference(superposition,[],[f23054,f872])).
fof(f27022,plain,(
  ( ! [X24,X23,X22] : (sP2(X22,X23,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23054,f864])).
fof(f27021,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,X20,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(superposition,[],[f23054,f862])).
fof(f27020,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl136),
  inference(superposition,[],[f23054,f861])).
fof(f26976,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK14)))) ) | $spl236),
  inference(cnf_transformation,[],[f26976_D])).
fof(f26976_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK14)))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f26902,plain,(
  ( ! [X210,X209] : (p(s(bool,X209)) | ~p(s(bool,X210)) | s(bool,X210) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK14)))) ) | $spl136),
  inference(resolution,[],[f23035,f1093])).
fof(f23035,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK14)))) | $spl136),
  inference(resolution,[],[f18957,f1012])).
fof(f26802,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl136),
  inference(resolution,[],[f23035,f829])).
fof(f26800,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK14)))) ) | $spl136),
  inference(resolution,[],[f23035,f828])).
fof(f26798,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK14)))) ) | $spl136),
  inference(resolution,[],[f23035,f827])).
fof(f26797,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl136),
  inference(resolution,[],[f23035,f825])).
fof(f26795,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK14)))) ) | $spl136),
  inference(resolution,[],[f23035,f823])).
fof(f25700,plain,(
  ( ! [X26,X25] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26)) ) | $spl230),
  inference(cnf_transformation,[],[f25700_D])).
fof(f25700_D,plain,(
  ( ! [X26,X25] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f26780,plain,(
  ( ! [X68,X67] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20563,f1086])).
fof(f20563,plain,(
  ( ! [X134,X133] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))),X133,X134)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f956])).
fof(f956,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f768])).
fof(f26765,plain,(
  ( ! [X26,X27,X25] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20563,f872])).
fof(f26764,plain,(
  ( ! [X24,X23,X22] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20563,f864])).
fof(f26763,plain,(
  ( ! [X21,X19,X20] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20563,f862])).
fof(f26762,plain,(
  ( ! [X17,X18,X16] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20563,f861])).
fof(f25643,plain,(
  ( ! [X26,X25] : (sP8(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26)) ) | $spl228),
  inference(cnf_transformation,[],[f25643_D])).
fof(f25643_D,plain,(
  ( ! [X26,X25] : (sP8(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f26720,plain,(
  ( ! [X68,X67] : (sP8(X67,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20562,f1086])).
fof(f20562,plain,(
  ( ! [X132,X131] : (sP8(X131,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))),X132)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f955])).
fof(f955,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f768])).
fof(f26705,plain,(
  ( ! [X26,X27,X25] : (sP8(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20562,f872])).
fof(f26704,plain,(
  ( ! [X24,X23,X22] : (sP8(X22,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20562,f864])).
fof(f26703,plain,(
  ( ! [X21,X19,X20] : (sP8(X19,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20562,f862])).
fof(f26702,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20562,f861])).
fof(f25585,plain,(
  ( ! [X26,X25] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26)) ) | $spl226),
  inference(cnf_transformation,[],[f25585_D])).
fof(f25585_D,plain,(
  ( ! [X26,X25] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f26659,plain,(
  ( ! [X68,X67] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20549,f1086])).
fof(f20549,plain,(
  ( ! [X105,X106] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))),X105,X106)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f939])).
fof(f939,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f761])).
fof(f26644,plain,(
  ( ! [X26,X27,X25] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20549,f872])).
fof(f26643,plain,(
  ( ! [X24,X23,X22] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20549,f864])).
fof(f26642,plain,(
  ( ! [X21,X19,X20] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20549,f862])).
fof(f26641,plain,(
  ( ! [X17,X18,X16] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20549,f861])).
fof(f26593,plain,(
  ( ! [X20,X18] : (p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num',sK14)))))) ) | $spl234),
  inference(cnf_transformation,[],[f26593_D])).
fof(f26593_D,plain,(
  ( ! [X20,X18] : (p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num',sK14)))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f26412,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,X328)) | p(s(bool,X329)) | s(bool,X329) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X330),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X330),s('type.num.num',sK14)))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(resolution,[],[f18169,f1093])).
fof(f18169,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK14))))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f18168,f17397])).
fof(f17397,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | ($spl0 | $spl22)),
  inference(backward_demodulation,[],[f14166,f1102])).
fof(f14166,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | $spl22),
  inference(forward_demodulation,[],[f14165,f896])).
fof(f896,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f560])).
fof(f560,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.numpair.tri_|49||49|','thm.numeral.numeral_distrib|split|1')).
fof(f14165,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | $spl22),
  inference(forward_demodulation,[],[f13872,f911])).
fof(f13872,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))))) | $spl22),
  inference(superposition,[],[f910,f1893])).
fof(f910,plain,(
  ( ! [X0] : (s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f571])).
fof(f571,plain,(
  ! [X0] : s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',X7)))))),
  file('thm.numpair.tri_|49||49|','thm.numpair.tri_def|split|1')).
fof(f18168,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))))))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17912,f1371])).
fof(f17912,plain,(
  ( ! [X11] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))))))))) ) | ($spl0 | $spl22)),
  inference(superposition,[],[f17397,f872])).
fof(f26313,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X46),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',sK14))))) | s(bool,X45) != s(bool,X47)) ) | ($spl0 | $spl5 | $spl22)),
  inference(resolution,[],[f18169,f829])).
fof(f26310,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num',sK14)))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(resolution,[],[f18169,f828])).
fof(f26309,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK14)))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(resolution,[],[f18169,f827])).
fof(f26308,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))) | s(bool,X26) != s(bool,X28)) ) | ($spl0 | $spl5 | $spl22)),
  inference(resolution,[],[f18169,f825])).
fof(f26306,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num',sK14)))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(resolution,[],[f18169,f823])).
fof(f25527,plain,(
  ( ! [X26,X25] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26)) ) | $spl224),
  inference(cnf_transformation,[],[f25527_D])).
fof(f25527_D,plain,(
  ( ! [X26,X25] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f26284,plain,(
  ( ! [X68,X67] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20533,f1086])).
fof(f20533,plain,(
  ( ! [X74,X73] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))),X73,X74)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f922])).
fof(f922,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f753])).
fof(f26269,plain,(
  ( ! [X26,X27,X25] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20533,f872])).
fof(f26268,plain,(
  ( ! [X24,X23,X22] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20533,f864])).
fof(f26267,plain,(
  ( ! [X21,X19,X20] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20533,f862])).
fof(f26266,plain,(
  ( ! [X17,X18,X16] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20533,f861])).
fof(f25470,plain,(
  ( ! [X26,X25] : (sP4(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl222),
  inference(cnf_transformation,[],[f25470_D])).
fof(f25470_D,plain,(
  ( ! [X26,X25] : (sP4(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f26224,plain,(
  ( ! [X68,X67] : (sP4(X67,X68,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20532,f1086])).
fof(f20532,plain,(
  ( ! [X72,X71] : (sP4(X71,X72,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f920])).
fof(f920,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f753])).
fof(f26209,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20532,f872])).
fof(f26208,plain,(
  ( ! [X24,X23,X22] : (sP4(X22,X23,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20532,f864])).
fof(f26207,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20532,f862])).
fof(f26206,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20532,f861])).
fof(f25413,plain,(
  ( ! [X26,X25] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26)) ) | $spl220),
  inference(cnf_transformation,[],[f25413_D])).
fof(f25413_D,plain,(
  ( ! [X26,X25] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f26162,plain,(
  ( ! [X68,X67] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20514,f1086])).
fof(f20514,plain,(
  ( ! [X41,X42] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))),X41,X42)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f844])).
fof(f844,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f726])).
fof(f26147,plain,(
  ( ! [X26,X27,X25] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20514,f872])).
fof(f26146,plain,(
  ( ! [X24,X23,X22] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20514,f864])).
fof(f26145,plain,(
  ( ! [X21,X19,X20] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20514,f862])).
fof(f26144,plain,(
  ( ! [X17,X18,X16] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20514,f861])).
fof(f25355,plain,(
  ( ! [X26,X25] : (sP1(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26)) ) | $spl218),
  inference(cnf_transformation,[],[f25355_D])).
fof(f25355_D,plain,(
  ( ! [X26,X25] : (sP1(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f26102,plain,(
  ( ! [X68,X67] : (sP1(X67,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20513,f1086])).
fof(f20513,plain,(
  ( ! [X39,X40] : (sP1(X39,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))),X40)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f843])).
fof(f843,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f726])).
fof(f26087,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20513,f872])).
fof(f26086,plain,(
  ( ! [X24,X23,X22] : (sP1(X22,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20513,f864])).
fof(f26085,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20513,f862])).
fof(f26084,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20513,f861])).
fof(f25297,plain,(
  ( ! [X26,X25] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26)) ) | $spl216),
  inference(cnf_transformation,[],[f25297_D])).
fof(f25297_D,plain,(
  ( ! [X26,X25] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f26041,plain,(
  ( ! [X68,X67] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20511,f1086])).
fof(f20511,plain,(
  ( ! [X35,X36] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))),X35,X36)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f840])).
fof(f840,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f723])).
fof(f26026,plain,(
  ( ! [X26,X27,X25] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20511,f872])).
fof(f26025,plain,(
  ( ! [X24,X23,X22] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20511,f864])).
fof(f26024,plain,(
  ( ! [X21,X19,X20] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20511,f862])).
fof(f26023,plain,(
  ( ! [X17,X18,X16] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20511,f861])).
fof(f25239,plain,(
  ( ! [X26,X25] : (sP2(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26)) ) | $spl214),
  inference(cnf_transformation,[],[f25239_D])).
fof(f25239_D,plain,(
  ( ! [X26,X25] : (sP2(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f25982,plain,(
  ( ! [X68,X67] : (sP2(X67,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20510,f1086])).
fof(f20510,plain,(
  ( ! [X33,X34] : (sP2(X33,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))),X34)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f839])).
fof(f839,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f723])).
fof(f25967,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20510,f872])).
fof(f25966,plain,(
  ( ! [X24,X23,X22] : (sP2(X22,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20510,f864])).
fof(f25965,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20510,f862])).
fof(f25964,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f20510,f861])).
fof(f25895,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK14)),s('type.num.num',sK14)))) ) | $spl232),
  inference(cnf_transformation,[],[f25895_D])).
fof(f25895_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK14)),s('type.num.num',sK14)))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f25819,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK14)),s('type.num.num',sK14)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f19334,f1093])).
fof(f25720,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK14)),s('type.num.num',sK14))) | s(bool,X26) != s(bool,X27)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f19334,f829])).
fof(f25717,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK14)),s('type.num.num',sK14)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f19334,f828])).
fof(f25716,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK14)),s('type.num.num',sK14)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f19334,f827])).
fof(f25715,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK14)),s('type.num.num',sK14))) | s(bool,X11) != s(bool,X12)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f19334,f825])).
fof(f25713,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK14)),s('type.num.num',sK14)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f19334,f823])).
fof(f6917,plain,(
  ( ! [X27] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl96),
  inference(cnf_transformation,[],[f6917_D])).
fof(f6917_D,plain,(
  ( ! [X27] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f25681,plain,(
  ( ! [X26,X27,X25] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23539,f872])).
fof(f23539,plain,(
  ( ! [X134,X133] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X133,X134)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f956])).
fof(f23273,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f1012])).
fof(f25624,plain,(
  ( ! [X26,X27,X25] : (sP8(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23538,f872])).
fof(f23538,plain,(
  ( ! [X132,X131] : (sP8(X131,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X132)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f955])).
fof(f25565,plain,(
  ( ! [X26,X27,X25] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23525,f872])).
fof(f23525,plain,(
  ( ! [X105,X106] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X105,X106)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f939])).
fof(f25508,plain,(
  ( ! [X26,X27,X25] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23509,f872])).
fof(f23509,plain,(
  ( ! [X74,X73] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X73,X74)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f922])).
fof(f25451,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23508,f872])).
fof(f23508,plain,(
  ( ! [X72,X71] : (sP4(X71,X72,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f920])).
fof(f25392,plain,(
  ( ! [X26,X27,X25] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23490,f872])).
fof(f23490,plain,(
  ( ! [X41,X42] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X41,X42)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f844])).
fof(f25335,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23489,f872])).
fof(f23489,plain,(
  ( ! [X39,X40] : (sP1(X39,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X40)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f843])).
fof(f25276,plain,(
  ( ! [X26,X27,X25] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23487,f872])).
fof(f23487,plain,(
  ( ! [X35,X36] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X35,X36)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f840])).
fof(f25220,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23486,f872])).
fof(f23486,plain,(
  ( ! [X33,X34] : (sP2(X33,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X34)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f839])).
fof(f25170,plain,(
  ( ! [X20,X18] : (p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num',sK15)))))) ) | $spl212),
  inference(cnf_transformation,[],[f25170_D])).
fof(f25170_D,plain,(
  ( ! [X20,X18] : (p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num',sK15)))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f24989,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,X328)) | p(s(bool,X329)) | s(bool,X329) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X330),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X330),s('type.num.num',sK15)))))) ) | $spl14),
  inference(resolution,[],[f1576,f1093])).
fof(f24890,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X46),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',sK15))))) | s(bool,X45) != s(bool,X47)) ) | $spl14),
  inference(resolution,[],[f1576,f829])).
fof(f24887,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num',sK15)))))) ) | $spl14),
  inference(resolution,[],[f1576,f828])).
fof(f24886,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK15)))))) ) | $spl14),
  inference(resolution,[],[f1576,f827])).
fof(f24885,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))) | s(bool,X26) != s(bool,X28)) ) | $spl14),
  inference(resolution,[],[f1576,f825])).
fof(f24883,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num',sK15)))))) ) | $spl14),
  inference(resolution,[],[f1576,f823])).
fof(f1576,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK15))))))) ) | $spl14),
  inference(cnf_transformation,[],[f1576_D])).
fof(f1576_D,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK15))))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f24850,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23374,f872])).
fof(f23374,plain,(
  ( ! [X196,X197] : (sP12(X196,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)),X197)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f995])).
fof(f24795,plain,(
  ( ! [X26,X27,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23367,f872])).
fof(f23367,plain,(
  ( ! [X182,X183] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)),X182,X183)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f986])).
fof(f24740,plain,(
  ( ! [X26,X27,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23356,f872])).
fof(f23356,plain,(
  ( ! [X161,X160] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)),X160,X161)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f970])).
fof(f24683,plain,(
  ( ! [X26,X27,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23344,f872])).
fof(f23344,plain,(
  ( ! [X136,X137] : (sP8(X136,X137,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f954])).
fof(f24626,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23328,f872])).
fof(f23328,plain,(
  ( ! [X105,X104] : (sP6(X104,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)),X105)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f938])).
fof(f24569,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23327,f872])).
fof(f23327,plain,(
  ( ! [X103,X102] : (sP6(X102,X103,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f937])).
fof(f24512,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23314,f872])).
fof(f23314,plain,(
  ( ! [X76,X77] : (sP4(X76,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)),X77)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f921])).
fof(f24453,plain,(
  ( ! [X26,X27,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23298,f872])).
fof(f23298,plain,(
  ( ! [X50,X51] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)),X50,X51)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f848])).
fof(f24395,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23297,f872])).
fof(f23297,plain,(
  ( ! [X48,X49] : (sP0(X48,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)),X49)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f847])).
fof(f24337,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23296,f872])).
fof(f23296,plain,(
  ( ! [X47,X46] : (sP0(X46,X47,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f846])).
fof(f24100,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num','const.num.|48|_0')))))) | $spl190),
  inference(cnf_transformation,[],[f24100_D])).
fof(f24100_D,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f23778,plain,(
  ( ! [X11] : (s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(superposition,[],[f18195,f872])).
fof(f18195,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | ($spl0 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f18194,f17397])).
fof(f18194,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | ($spl0 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17927,f1371])).
fof(f17927,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | ($spl0 | $spl22)),
  inference(superposition,[],[f17397,f1086])).
fof(f23750,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23295,f872])).
fof(f23295,plain,(
  ( ! [X45,X44] : (sP1(X44,X45,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f842])).
fof(f23693,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl20 | $spl136)),
  inference(superposition,[],[f23292,f872])).
fof(f23292,plain,(
  ( ! [X39,X38] : (sP2(X38,X39,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f838])).
fof(f23654,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl184),
  inference(cnf_transformation,[],[f23654_D])).
fof(f23654_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f23581,plain,(
  ( ! [X210,X209] : (p(s(bool,X209)) | ~p(s(bool,X210)) | s(bool,X210) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f1093])).
fof(f23481,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | s(bool,X22) != s(bool,X24)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f829])).
fof(f23479,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f828])).
fof(f23477,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f827])).
fof(f23476,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | s(bool,X7) != s(bool,X9)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f825])).
fof(f23474,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23273,f823])).
fof(f23465,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) ) | $spl182),
  inference(cnf_transformation,[],[f23465_D])).
fof(f23465_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f23385,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f1093])).
fof(f23286,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15))) | s(bool,X26) != s(bool,X27)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f829])).
fof(f23283,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f828])).
fof(f23282,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f827])).
fof(f23281,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15))) | s(bool,X11) != s(bool,X12)) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f825])).
fof(f23279,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK15)))) ) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f823])).
fof(f23462,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))) | $spl178),
  inference(cnf_transformation,[],[f23462_D])).
fof(f23462_D,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f23464,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK15)))) | $spl180),
  inference(cnf_transformation,[],[f23464_D])).
fof(f23464_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK15)))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f23278,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK15)))) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))) | ($spl20 | $spl136)),
  inference(resolution,[],[f23223,f1067])).
fof(f1067,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f686])).
fof(f686,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))))),
  inference(flattening,[],[f685])).
fof(f685,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f622])).
fof(f622,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X1) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))))),
  inference(flattening,[],[f621])).
fof(f621,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & ~s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))))),
  inference(rectify,[],[f398])).
fof(f398,axiom,(
  ! [X6,X7] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) & ~s('type.num.num',X7) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7)))))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.LESS_SUC_EQ_COR')).
fof(f23228,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) ) | $spl176),
  inference(cnf_transformation,[],[f23228_D])).
fof(f23228_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f23147,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(resolution,[],[f18957,f1093])).
fof(f23048,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15)))) | s(bool,X26) != s(bool,X27)) ) | $spl136),
  inference(resolution,[],[f18957,f829])).
fof(f23045,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(resolution,[],[f18957,f828])).
fof(f23044,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(resolution,[],[f18957,f827])).
fof(f23043,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15)))) | s(bool,X11) != s(bool,X12)) ) | $spl136),
  inference(resolution,[],[f18957,f825])).
fof(f23041,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) ) | $spl136),
  inference(resolution,[],[f18957,f823])).
fof(f23225,plain,(
  s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))) | $spl172),
  inference(cnf_transformation,[],[f23225_D])).
fof(f23225_D,plain,(
  s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f23227,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK26(sK15))))) | $spl174),
  inference(cnf_transformation,[],[f23227_D])).
fof(f23227_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK26(sK15))))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f23040,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))),s('type.num.num',sK26(sK15))))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK14))) | $spl136),
  inference(resolution,[],[f18957,f1067])).
fof(f23005,plain,(
  ( ! [X20,X18] : (p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))) ) | $spl170),
  inference(cnf_transformation,[],[f23005_D])).
fof(f23005_D,plain,(
  ( ! [X20,X18] : (p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f22826,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,X328)) | p(s(bool,X329)) | s(bool,X329) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X330)))))) )),
  inference(resolution,[],[f8367,f1093])).
fof(f8367,plain,(
  ( ! [X207] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X207))))))) )),
  inference(subsumption_resolution,[],[f8366,f1082])).
fof(f1082,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f638])).
fof(f638,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f637])).
fof(f637,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X7] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))),
  file('thm.numpair.tri_|49||49|','thm.num.NOT_SUC')).
fof(f8366,plain,(
  ( ! [X207] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X207),s('type.num.num',X207))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X207))))))) )),
  inference(forward_demodulation,[],[f8104,f1075])).
fof(f1075,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f630])).
fof(f630,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f281])).
fof(f281,axiom,(
  ! [X6,X7] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.ADD_SUC')).
fof(f8104,plain,(
  ( ! [X207] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X207)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X207),s('type.num.num','const.num.SUC_1'(s('type.num.num',X207)))))) )),
  inference(superposition,[],[f892,f7467])).
fof(f7467,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f7466,f896])).
fof(f7466,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f1087,f1075])).
fof(f1087,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(cnf_transformation,[],[f643])).
fof(f643,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(rectify,[],[f227])).
fof(f227,axiom,(
  ! [X7] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.BIT|49|')).
fof(f22727,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X46))))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f8367,f829])).
fof(f22724,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X36)))))) )),
  inference(resolution,[],[f8367,f828])).
fof(f22723,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X32)))))) )),
  inference(resolution,[],[f8367,f827])).
fof(f22722,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X27))))) | s(bool,X26) != s(bool,X28)) )),
  inference(resolution,[],[f8367,f825])).
fof(f22720,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))) )),
  inference(resolution,[],[f8367,f823])).
fof(f4725,plain,(
  ( ! [X23,X22] : (sP12(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23)) ) | $spl84),
  inference(cnf_transformation,[],[f4725_D])).
fof(f4725_D,plain,(
  ( ! [X23,X22] : (sP12(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f22684,plain,(
  ( ! [X68,X67] : (sP12(X67,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18856,f1086])).
fof(f18856,plain,(
  ( ! [X196,X197] : (sP12(X196,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))),X197)) ) | $spl50),
  inference(resolution,[],[f3285,f995])).
fof(f22669,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18856,f872])).
fof(f22668,plain,(
  ( ! [X24,X23,X22] : (sP12(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18856,f864])).
fof(f22667,plain,(
  ( ! [X21,X19,X20] : (sP12(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18856,f862])).
fof(f22666,plain,(
  ( ! [X17,X18,X16] : (sP12(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18856,f861])).
fof(f4651,plain,(
  ( ! [X23,X22] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23)) ) | $spl82),
  inference(cnf_transformation,[],[f4651_D])).
fof(f4651_D,plain,(
  ( ! [X23,X22] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f22601,plain,(
  ( ! [X68,X67] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18849,f1086])).
fof(f18849,plain,(
  ( ! [X182,X183] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))),X182,X183)) ) | $spl50),
  inference(resolution,[],[f3285,f986])).
fof(f22586,plain,(
  ( ! [X26,X27,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18849,f872])).
fof(f22585,plain,(
  ( ! [X24,X23,X22] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18849,f864])).
fof(f22584,plain,(
  ( ! [X21,X19,X20] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18849,f862])).
fof(f22583,plain,(
  ( ! [X17,X18,X16] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18849,f861])).
fof(f4577,plain,(
  ( ! [X23,X22] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23)) ) | $spl80),
  inference(cnf_transformation,[],[f4577_D])).
fof(f4577_D,plain,(
  ( ! [X23,X22] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f22518,plain,(
  ( ! [X68,X67] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18838,f1086])).
fof(f18838,plain,(
  ( ! [X161,X160] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))),X160,X161)) ) | $spl50),
  inference(resolution,[],[f3285,f970])).
fof(f22503,plain,(
  ( ! [X26,X27,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18838,f872])).
fof(f22502,plain,(
  ( ! [X24,X23,X22] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18838,f864])).
fof(f22501,plain,(
  ( ! [X21,X19,X20] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18838,f862])).
fof(f22500,plain,(
  ( ! [X17,X18,X16] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18838,f861])).
fof(f4503,plain,(
  ( ! [X23,X22] : (sP8(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl78),
  inference(cnf_transformation,[],[f4503_D])).
fof(f4503_D,plain,(
  ( ! [X23,X22] : (sP8(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f22433,plain,(
  ( ! [X68,X67] : (sP8(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18826,f1086])).
fof(f18826,plain,(
  ( ! [X136,X137] : (sP8(X136,X137,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(resolution,[],[f3285,f954])).
fof(f22418,plain,(
  ( ! [X26,X27,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18826,f872])).
fof(f22417,plain,(
  ( ! [X24,X23,X22] : (sP8(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18826,f864])).
fof(f22416,plain,(
  ( ! [X21,X19,X20] : (sP8(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18826,f862])).
fof(f22415,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18826,f861])).
fof(f4426,plain,(
  ( ! [X23,X22] : (sP6(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23)) ) | $spl76),
  inference(cnf_transformation,[],[f4426_D])).
fof(f4426_D,plain,(
  ( ! [X23,X22] : (sP6(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f22348,plain,(
  ( ! [X68,X67] : (sP6(X67,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18810,f1086])).
fof(f18810,plain,(
  ( ! [X105,X104] : (sP6(X104,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))),X105)) ) | $spl50),
  inference(resolution,[],[f3285,f938])).
fof(f22333,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18810,f872])).
fof(f22332,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18810,f864])).
fof(f22331,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18810,f862])).
fof(f22330,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18810,f861])).
fof(f22277,plain,(
  ( ! [X20,X18] : (p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))) ) | $spl168),
  inference(cnf_transformation,[],[f22277_D])).
fof(f22277_D,plain,(
  ( ! [X20,X18] : (p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f22135,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,X328)) | p(s(bool,X329)) | s(bool,X329) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X330),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X330)))))) )),
  inference(resolution,[],[f8310,f1093])).
fof(f8310,plain,(
  ( ! [X31] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X31),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X31))))))) )),
  inference(subsumption_resolution,[],[f8028,f1082])).
fof(f8028,plain,(
  ( ! [X31] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X31),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X31)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))) )),
  inference(superposition,[],[f872,f7467])).
fof(f22036,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X46),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X46))))) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f8310,f829])).
fof(f22033,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X36)))))) )),
  inference(resolution,[],[f8310,f828])).
fof(f22032,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X32)))))) )),
  inference(resolution,[],[f8310,f827])).
fof(f22031,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X27))))) | s(bool,X26) != s(bool,X28)) )),
  inference(resolution,[],[f8310,f825])).
fof(f22029,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))) )),
  inference(resolution,[],[f8310,f823])).
fof(f4192,plain,(
  ( ! [X19,X20] : (sP6(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl72),
  inference(cnf_transformation,[],[f4192_D])).
fof(f4192_D,plain,(
  ( ! [X19,X20] : (sP6(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f21991,plain,(
  ( ! [X68,X67] : (sP6(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18809,f1086])).
fof(f18809,plain,(
  ( ! [X103,X102] : (sP6(X102,X103,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(resolution,[],[f3285,f937])).
fof(f21976,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18809,f872])).
fof(f21975,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18809,f864])).
fof(f21974,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18809,f862])).
fof(f21973,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18809,f861])).
fof(f4116,plain,(
  ( ! [X19,X20] : (sP4(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20)) ) | $spl70),
  inference(cnf_transformation,[],[f4116_D])).
fof(f4116_D,plain,(
  ( ! [X19,X20] : (sP4(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f21906,plain,(
  ( ! [X68,X67] : (sP4(X67,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18796,f1086])).
fof(f18796,plain,(
  ( ! [X76,X77] : (sP4(X76,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))),X77)) ) | $spl50),
  inference(resolution,[],[f3285,f921])).
fof(f21891,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18796,f872])).
fof(f21890,plain,(
  ( ! [X24,X23,X22] : (sP4(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18796,f864])).
fof(f21889,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18796,f862])).
fof(f21888,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18796,f861])).
fof(f4040,plain,(
  ( ! [X19,X20] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20)) ) | $spl68),
  inference(cnf_transformation,[],[f4040_D])).
fof(f4040_D,plain,(
  ( ! [X19,X20] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f21819,plain,(
  ( ! [X68,X67] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18780,f1086])).
fof(f18780,plain,(
  ( ! [X50,X51] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))),X50,X51)) ) | $spl50),
  inference(resolution,[],[f3285,f848])).
fof(f21804,plain,(
  ( ! [X26,X27,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18780,f872])).
fof(f21803,plain,(
  ( ! [X24,X23,X22] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18780,f864])).
fof(f21802,plain,(
  ( ! [X21,X19,X20] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18780,f862])).
fof(f21801,plain,(
  ( ! [X17,X18,X16] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18780,f861])).
fof(f3963,plain,(
  ( ! [X19,X20] : (sP0(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20)) ) | $spl66),
  inference(cnf_transformation,[],[f3963_D])).
fof(f3963_D,plain,(
  ( ! [X19,X20] : (sP0(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f21733,plain,(
  ( ! [X68,X67] : (sP0(X67,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18779,f1086])).
fof(f18779,plain,(
  ( ! [X48,X49] : (sP0(X48,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))),X49)) ) | $spl50),
  inference(resolution,[],[f3285,f847])).
fof(f21718,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18779,f872])).
fof(f21717,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18779,f864])).
fof(f21716,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18779,f862])).
fof(f21715,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18779,f861])).
fof(f3886,plain,(
  ( ! [X19,X20] : (sP0(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f3886_D])).
fof(f3886_D,plain,(
  ( ! [X19,X20] : (sP0(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f21647,plain,(
  ( ! [X68,X67] : (sP0(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18778,f1086])).
fof(f18778,plain,(
  ( ! [X47,X46] : (sP0(X46,X47,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(resolution,[],[f3285,f846])).
fof(f21632,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18778,f872])).
fof(f21631,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18778,f864])).
fof(f21630,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18778,f862])).
fof(f21629,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18778,f861])).
fof(f3651,plain,(
  ( ! [X17,X16] : (sP1(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl58),
  inference(cnf_transformation,[],[f3651_D])).
fof(f3651_D,plain,(
  ( ! [X17,X16] : (sP1(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f21562,plain,(
  ( ! [X68,X67] : (sP1(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18777,f1086])).
fof(f18777,plain,(
  ( ! [X45,X44] : (sP1(X44,X45,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(resolution,[],[f3285,f842])).
fof(f21547,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18777,f872])).
fof(f21546,plain,(
  ( ! [X24,X23,X22] : (sP1(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18777,f864])).
fof(f21545,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18777,f862])).
fof(f21544,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18777,f861])).
fof(f3575,plain,(
  ( ! [X17,X16] : (sP2(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl56),
  inference(cnf_transformation,[],[f3575_D])).
fof(f3575_D,plain,(
  ( ! [X17,X16] : (sP2(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f21477,plain,(
  ( ! [X68,X67] : (sP2(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15)))))) ) | $spl50),
  inference(superposition,[],[f18774,f1086])).
fof(f18774,plain,(
  ( ! [X39,X38] : (sP2(X38,X39,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(resolution,[],[f3285,f838])).
fof(f21462,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK26(sK15)))))))) ) | $spl50),
  inference(superposition,[],[f18774,f872])).
fof(f21461,plain,(
  ( ! [X24,X23,X22] : (sP2(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18774,f864])).
fof(f21460,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(superposition,[],[f18774,f862])).
fof(f21459,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X18)))) ) | $spl50),
  inference(superposition,[],[f18774,f861])).
fof(f21416,plain,(
  ( ! [X26,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26)) ) | $spl166),
  inference(cnf_transformation,[],[f21416_D])).
fof(f21416_D,plain,(
  ( ! [X26,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f21394,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18467,f872])).
fof(f18467,plain,(
  ( ! [X196,X197] : (sP12(X196,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)),X197)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f995])).
fof(f21357,plain,(
  ( ! [X26,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X25,X26)) ) | $spl164),
  inference(cnf_transformation,[],[f21357_D])).
fof(f21357_D,plain,(
  ( ! [X26,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X25,X26)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f21335,plain,(
  ( ! [X26,X27,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18460,f872])).
fof(f18460,plain,(
  ( ! [X182,X183] : (sP11('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)),X182,X183)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f986])).
fof(f21298,plain,(
  ( ! [X26,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X25,X26)) ) | $spl162),
  inference(cnf_transformation,[],[f21298_D])).
fof(f21298_D,plain,(
  ( ! [X26,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X25,X26)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f21276,plain,(
  ( ! [X26,X27,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18449,f872])).
fof(f18449,plain,(
  ( ! [X161,X160] : (sP10('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)),X160,X161)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f970])).
fof(f21239,plain,(
  ( ! [X26,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')))) ) | $spl160),
  inference(cnf_transformation,[],[f21239_D])).
fof(f21239_D,plain,(
  ( ! [X26,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f21215,plain,(
  ( ! [X26,X27,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18437,f872])).
fof(f18437,plain,(
  ( ! [X136,X137] : (sP8(X136,X137,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f954])).
fof(f21177,plain,(
  ( ! [X26,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26)) ) | $spl158),
  inference(cnf_transformation,[],[f21177_D])).
fof(f21177_D,plain,(
  ( ! [X26,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f21154,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18421,f872])).
fof(f18421,plain,(
  ( ! [X105,X104] : (sP6(X104,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)),X105)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f938])).
fof(f21116,plain,(
  ( ! [X26,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')))) ) | $spl156),
  inference(cnf_transformation,[],[f21116_D])).
fof(f21116_D,plain,(
  ( ! [X26,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f21093,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18420,f872])).
fof(f18420,plain,(
  ( ! [X103,X102] : (sP6(X102,X103,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f937])).
fof(f21055,plain,(
  ( ! [X26,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26)) ) | $spl154),
  inference(cnf_transformation,[],[f21055_D])).
fof(f21055_D,plain,(
  ( ! [X26,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f21032,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18407,f872])).
fof(f18407,plain,(
  ( ! [X76,X77] : (sP4(X76,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)),X77)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f921])).
fof(f20994,plain,(
  ( ! [X26,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X25,X26)) ) | $spl152),
  inference(cnf_transformation,[],[f20994_D])).
fof(f20994_D,plain,(
  ( ! [X26,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X25,X26)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f20969,plain,(
  ( ! [X26,X27,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18391,f872])).
fof(f18391,plain,(
  ( ! [X50,X51] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)),X50,X51)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f848])).
fof(f20932,plain,(
  ( ! [X26,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26)) ) | $spl150),
  inference(cnf_transformation,[],[f20932_D])).
fof(f20932_D,plain,(
  ( ! [X26,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f20907,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18390,f872])).
fof(f18390,plain,(
  ( ! [X48,X49] : (sP0(X48,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)),X49)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f847])).
fof(f20870,plain,(
  ( ! [X26,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')))) ) | $spl148),
  inference(cnf_transformation,[],[f20870_D])).
fof(f20870_D,plain,(
  ( ! [X26,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f20845,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18389,f872])).
fof(f18389,plain,(
  ( ! [X47,X46] : (sP0(X46,X47,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f846])).
fof(f20808,plain,(
  ( ! [X26,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')))) ) | $spl146),
  inference(cnf_transformation,[],[f20808_D])).
fof(f20808_D,plain,(
  ( ! [X26,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f20784,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18388,f872])).
fof(f18388,plain,(
  ( ! [X45,X44] : (sP1(X44,X45,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f842])).
fof(f20747,plain,(
  ( ! [X26,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')))) ) | $spl144),
  inference(cnf_transformation,[],[f20747_D])).
fof(f20747_D,plain,(
  ( ! [X26,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f20723,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f18385,f872])).
fof(f18385,plain,(
  ( ! [X39,X38] : (sP2(X38,X39,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f838])).
fof(f20680,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))))) ) | $spl142),
  inference(cnf_transformation,[],[f20680_D])).
fof(f20680_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f20605,plain,(
  ( ! [X210,X209] : (p(s(bool,X209)) | ~p(s(bool,X210)) | s(bool,X210) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f1093])).
fof(f20505,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15)))) | s(bool,X22) != s(bool,X24)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f829])).
fof(f20503,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f828])).
fof(f20501,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f827])).
fof(f20500,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15)))) | s(bool,X7) != s(bool,X9)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f825])).
fof(f20498,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18366,f823])).
fof(f20490,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f20397,f20251])).
fof(f20251,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(equality_resolution,[],[f1898])).
fof(f20397,plain,(
  ( ! [X350,X351] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X350),s('type.num.num',X351))))) ) | $spl30),
  inference(resolution,[],[f20251,f1059])).
fof(f1059,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f811])).
fof(f811,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))))))))))),
  inference(nnf_transformation,[],[f617])).
fof(f617,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f616])).
fof(f616,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f506])).
fof(f506,axiom,(
  ! [X13,X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num',X13))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X14)))))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num',X13)))))),
  file('thm.numpair.tri_|49||49|','thm.numeral.numeral_|49||48|_prove_ltb|split|2')).
fof(f20489,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f20381,f20251])).
fof(f20381,plain,(
  ( ! [X331,X330] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X330),s('type.num.num',X330))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X331),s('type.num.num',X331))))))))) ) | $spl30),
  inference(resolution,[],[f20251,f1059])).
fof(f20387,plain,(
  $false | $spl30),
  inference(resolution,[],[f20251,f1091])).
fof(f20391,plain,(
  $false | $spl30),
  inference(resolution,[],[f20251,f1090])).
fof(f1090,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f646])).
fof(f646,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f177])).
fof(f177,axiom,(
  ! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))),
  file('thm.numpair.tri_|49||49|','thm.prim_rec.LESS_|48|')).
fof(f20392,plain,(
  $false | ($spl30 | $spl111)),
  inference(resolution,[],[f20251,f8575])).
fof(f8575,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl111),
  inference(subsumption_resolution,[],[f8518,f911])).
fof(f8518,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl111),
  inference(superposition,[],[f7395,f892])).
fof(f20393,plain,(
  $false | ($spl16 | $spl30)),
  inference(resolution,[],[f20251,f1578])).
fof(f20394,plain,(
  $false | ($spl30 | $spl50)),
  inference(resolution,[],[f20251,f3285])).
fof(f20398,plain,(
  $false | ($spl0 | $spl5 | $spl30)),
  inference(resolution,[],[f20251,f3067])).
fof(f20194,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f20118,f19972])).
fof(f19972,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl28),
  inference(equality_resolution,[],[f1897])).
fof(f20118,plain,(
  ( ! [X364,X365] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X364),s('type.num.num',X365))))) ) | $spl28),
  inference(resolution,[],[f19972,f1058])).
fof(f1058,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))))))))) )),
  inference(cnf_transformation,[],[f811])).
fof(f20193,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f20110,f19972])).
fof(f20110,plain,(
  ( ! [X352,X351] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X351),s('type.num.num',X351))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X352),s('type.num.num',X352))))))))) ) | $spl28),
  inference(resolution,[],[f19972,f1058])).
fof(f20192,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f20108,f19972])).
fof(f20108,plain,(
  ( ! [X346,X347] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X346),s('type.num.num',X347))))) ) | $spl28),
  inference(resolution,[],[f19972,f1012])).
fof(f20191,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f20107,f19972])).
fof(f20107,plain,(
  ( ! [X344,X345] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X344),s('type.num.num',X345))))) ) | $spl28),
  inference(resolution,[],[f19972,f1012])).
fof(f20100,plain,(
  $false | $spl28),
  inference(resolution,[],[f19972,f1045])).
fof(f20113,plain,(
  $false | $spl28),
  inference(resolution,[],[f19972,f1046])).
fof(f1046,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f606])).
fof(f606,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f605])).
fof(f605,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f178])).
fof(f178,axiom,(
  ! [X7] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0'))))),
  file('thm.numpair.tri_|49||49|','thm.prim_rec.NOT_LESS_|48|')).
fof(f6919,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14)))) | $spl98),
  inference(cnf_transformation,[],[f6919_D])).
fof(f6919_D,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14)))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f4349,plain,(
  ( ! [X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.num.|48|_0')))) ) | $spl74),
  inference(cnf_transformation,[],[f4349_D])).
fof(f4349_D,plain,(
  ( ! [X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f19205,plain,(
  ( ! [X45] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(superposition,[],[f18165,f1086])).
fof(f18165,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK14)))) ) | ($spl0 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f18164,f17397])).
fof(f18164,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17910,f1371])).
fof(f17910,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))))) ) | ($spl0 | $spl22)),
  inference(superposition,[],[f17397,f862])).
fof(f3809,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X18)) ) | $spl62),
  inference(cnf_transformation,[],[f3809_D])).
fof(f3809_D,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X18)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f19258,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK14))))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f19190,f896])).
fof(f19190,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK14))))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(superposition,[],[f18165,f872])).
fof(f6916,plain,(
  ( ! [X24] : (s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl94),
  inference(cnf_transformation,[],[f6916_D])).
fof(f6916_D,plain,(
  ( ! [X24] : (s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f19189,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',sK14)))) ) | ($spl0 | $spl5 | $spl22)),
  inference(superposition,[],[f18165,f864])).
fof(f18960,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) ) | $spl140),
  inference(cnf_transformation,[],[f18960_D])).
fof(f18960_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f18867,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(resolution,[],[f3285,f1093])).
fof(f18768,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15)))) | s(bool,X26) != s(bool,X27)) ) | $spl50),
  inference(resolution,[],[f3285,f829])).
fof(f18765,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(resolution,[],[f3285,f828])).
fof(f18764,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(resolution,[],[f3285,f827])).
fof(f18763,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15)))) | s(bool,X11) != s(bool,X12)) ) | $spl50),
  inference(resolution,[],[f3285,f825])).
fof(f18761,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) ) | $spl50),
  inference(resolution,[],[f3285,f823])).
fof(f18957,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) | $spl136),
  inference(cnf_transformation,[],[f18957_D])).
fof(f18957_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f18959,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK26(sK15)) | $spl138),
  inference(cnf_transformation,[],[f18959_D])).
fof(f18959_D,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK26(sK15)) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f18955,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK26(sK15)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) | ($spl22 | $spl50)),
  inference(forward_demodulation,[],[f18954,f1893])).
fof(f18954,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK26(sK15))))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl22 | $spl50)),
  inference(forward_demodulation,[],[f18760,f1893])).
fof(f18760,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK26(sK15))))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl50),
  inference(resolution,[],[f3285,f1067])).
fof(f3807,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14)))) ) | $spl60),
  inference(cnf_transformation,[],[f3807_D])).
fof(f3807_D,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f18582,plain,(
  ( ! [X45] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X45))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(superposition,[],[f18163,f1086])).
fof(f18163,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X8)))) ) | ($spl0 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f18162,f17397])).
fof(f18162,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))),s('type.num.num',X8)))) ) | ($spl0 | $spl5 | $spl22)),
  inference(subsumption_resolution,[],[f17909,f1371])).
fof(f17909,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))),s('type.num.num',X8)))) ) | ($spl0 | $spl22)),
  inference(superposition,[],[f17397,f861])).
fof(f18718,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK14))))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(forward_demodulation,[],[f18567,f895])).
fof(f895,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f559])).
fof(f559,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f123])).
fof(f123,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.ADD_CLAUSES|split|0')).
fof(f18567,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK14))))))) ) | ($spl0 | $spl5 | $spl22)),
  inference(superposition,[],[f18163,f872])).
fof(f18566,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X16))) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',sK14)))) ) | ($spl0 | $spl5 | $spl22)),
  inference(superposition,[],[f18163,f864])).
fof(f6915,plain,(
  ( ! [X21] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl92),
  inference(cnf_transformation,[],[f6915_D])).
fof(f6915_D,plain,(
  ( ! [X21] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f18565,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',sK14)))) ) | ($spl0 | $spl5 | $spl22)),
  inference(superposition,[],[f18163,f862])).
fof(f18554,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) ) | $spl134),
  inference(cnf_transformation,[],[f18554_D])).
fof(f18554_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f18478,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f1093])).
fof(f18379,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15))) | s(bool,X26) != s(bool,X27)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f829])).
fof(f18376,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f828])).
fof(f18375,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f827])).
fof(f18374,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15))) | s(bool,X11) != s(bool,X12)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f825])).
fof(f18372,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK26(sK15)),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f3067,f823])).
fof(f17375,plain,(
  ( ! [X68,X67] : (sP12(X67,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13696,f1086])).
fof(f13696,plain,(
  ( ! [X196,X197] : (sP12(X196,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X197)) ) | $spl111),
  inference(resolution,[],[f8575,f995])).
fof(f17360,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13696,f872])).
fof(f17359,plain,(
  ( ! [X24,X23,X22] : (sP12(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13696,f864])).
fof(f17358,plain,(
  ( ! [X21,X19,X20] : (sP12(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13696,f862])).
fof(f6914,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl90),
  inference(cnf_transformation,[],[f6914_D])).
fof(f6914_D,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f17357,plain,(
  ( ! [X17,X18,X16] : (sP12(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13696,f861])).
fof(f17296,plain,(
  ( ! [X68,X67] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13689,f1086])).
fof(f13689,plain,(
  ( ! [X182,X183] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X182,X183)) ) | $spl111),
  inference(resolution,[],[f8575,f986])).
fof(f17281,plain,(
  ( ! [X26,X27,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13689,f872])).
fof(f17280,plain,(
  ( ! [X24,X23,X22] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13689,f864])).
fof(f17279,plain,(
  ( ! [X21,X19,X20] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13689,f862])).
fof(f17278,plain,(
  ( ! [X17,X18,X16] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13689,f861])).
fof(f17217,plain,(
  ( ! [X68,X67] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13678,f1086])).
fof(f13678,plain,(
  ( ! [X161,X160] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X160,X161)) ) | $spl111),
  inference(resolution,[],[f8575,f970])).
fof(f17202,plain,(
  ( ! [X26,X27,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13678,f872])).
fof(f17201,plain,(
  ( ! [X24,X23,X22] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13678,f864])).
fof(f17200,plain,(
  ( ! [X21,X19,X20] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13678,f862])).
fof(f17199,plain,(
  ( ! [X17,X18,X16] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13678,f861])).
fof(f17135,plain,(
  ( ! [X68,X67] : (sP8(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13666,f1086])).
fof(f13666,plain,(
  ( ! [X136,X137] : (sP8(X136,X137,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl111),
  inference(resolution,[],[f8575,f954])).
fof(f17120,plain,(
  ( ! [X26,X27,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13666,f872])).
fof(f17119,plain,(
  ( ! [X24,X23,X22] : (sP8(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13666,f864])).
fof(f17118,plain,(
  ( ! [X21,X19,X20] : (sP8(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13666,f862])).
fof(f17117,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13666,f861])).
fof(f17054,plain,(
  ( ! [X68,X67] : (sP6(X67,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13650,f1086])).
fof(f13650,plain,(
  ( ! [X105,X104] : (sP6(X104,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X105)) ) | $spl111),
  inference(resolution,[],[f8575,f938])).
fof(f17039,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13650,f872])).
fof(f17038,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13650,f864])).
fof(f17037,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13650,f862])).
fof(f17036,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13650,f861])).
fof(f16973,plain,(
  ( ! [X68,X67] : (sP6(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13649,f1086])).
fof(f13649,plain,(
  ( ! [X103,X102] : (sP6(X102,X103,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl111),
  inference(resolution,[],[f8575,f937])).
fof(f16958,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13649,f872])).
fof(f16957,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13649,f864])).
fof(f16956,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13649,f862])).
fof(f16955,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13649,f861])).
fof(f16892,plain,(
  ( ! [X68,X67] : (sP4(X67,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13636,f1086])).
fof(f13636,plain,(
  ( ! [X76,X77] : (sP4(X76,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X77)) ) | $spl111),
  inference(resolution,[],[f8575,f921])).
fof(f16877,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13636,f872])).
fof(f16876,plain,(
  ( ! [X24,X23,X22] : (sP4(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13636,f864])).
fof(f16875,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13636,f862])).
fof(f16874,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13636,f861])).
fof(f16809,plain,(
  ( ! [X68,X67] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13620,f1086])).
fof(f13620,plain,(
  ( ! [X50,X51] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X50,X51)) ) | $spl111),
  inference(resolution,[],[f8575,f848])).
fof(f16794,plain,(
  ( ! [X26,X27,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13620,f872])).
fof(f16793,plain,(
  ( ! [X24,X23,X22] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13620,f864])).
fof(f16792,plain,(
  ( ! [X21,X19,X20] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13620,f862])).
fof(f16791,plain,(
  ( ! [X17,X18,X16] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13620,f861])).
fof(f16726,plain,(
  ( ! [X68,X67] : (sP0(X67,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13619,f1086])).
fof(f13619,plain,(
  ( ! [X48,X49] : (sP0(X48,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X49)) ) | $spl111),
  inference(resolution,[],[f8575,f847])).
fof(f16711,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13619,f872])).
fof(f16710,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13619,f864])).
fof(f16709,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13619,f862])).
fof(f16708,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13619,f861])).
fof(f16644,plain,(
  ( ! [X68,X67] : (sP0(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13618,f1086])).
fof(f13618,plain,(
  ( ! [X47,X46] : (sP0(X46,X47,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl111),
  inference(resolution,[],[f8575,f846])).
fof(f16629,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13618,f872])).
fof(f16628,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13618,f864])).
fof(f16627,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13618,f862])).
fof(f16626,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13618,f861])).
fof(f16563,plain,(
  ( ! [X68,X67] : (sP1(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13617,f1086])).
fof(f13617,plain,(
  ( ! [X45,X44] : (sP1(X44,X45,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl111),
  inference(resolution,[],[f8575,f842])).
fof(f16548,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13617,f872])).
fof(f16547,plain,(
  ( ! [X24,X23,X22] : (sP1(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13617,f864])).
fof(f16546,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13617,f862])).
fof(f16545,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13617,f861])).
fof(f16482,plain,(
  ( ! [X68,X67] : (sP2(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl111),
  inference(superposition,[],[f13614,f1086])).
fof(f13614,plain,(
  ( ! [X39,X38] : (sP2(X38,X39,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl111),
  inference(resolution,[],[f8575,f838])).
fof(f16467,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl111),
  inference(superposition,[],[f13614,f872])).
fof(f16466,plain,(
  ( ! [X24,X23,X22] : (sP2(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13614,f864])).
fof(f16465,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl111),
  inference(superposition,[],[f13614,f862])).
fof(f16464,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl111),
  inference(superposition,[],[f13614,f861])).
fof(f15901,plain,(
  $false | $spl84),
  inference(subsumption_resolution,[],[f15825,f15677])).
fof(f15677,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl84),
  inference(subsumption_resolution,[],[f15639,f1045])).
fof(f15639,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl84),
  inference(resolution,[],[f4725,f997])).
fof(f997,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f787])).
fof(f15825,plain,(
  ( ! [X364,X365] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X364),s('type.num.num',X365))))) ) | $spl84),
  inference(resolution,[],[f15677,f1058])).
fof(f15900,plain,(
  $false | $spl84),
  inference(subsumption_resolution,[],[f15817,f15677])).
fof(f15817,plain,(
  ( ! [X352,X351] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X351),s('type.num.num',X351))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X352),s('type.num.num',X352))))))))) ) | $spl84),
  inference(resolution,[],[f15677,f1058])).
fof(f15899,plain,(
  $false | $spl84),
  inference(subsumption_resolution,[],[f15815,f15677])).
fof(f15815,plain,(
  ( ! [X346,X347] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X346),s('type.num.num',X347))))) ) | $spl84),
  inference(resolution,[],[f15677,f1012])).
fof(f15898,plain,(
  $false | $spl84),
  inference(subsumption_resolution,[],[f15814,f15677])).
fof(f15814,plain,(
  ( ! [X344,X345] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X344),s('type.num.num',X345))))) ) | $spl84),
  inference(resolution,[],[f15677,f1012])).
fof(f15807,plain,(
  $false | $spl84),
  inference(resolution,[],[f15677,f1045])).
fof(f15820,plain,(
  $false | $spl84),
  inference(resolution,[],[f15677,f1046])).
fof(f15638,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15562,f15414])).
fof(f15414,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl82),
  inference(subsumption_resolution,[],[f15376,f1045])).
fof(f15376,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0))) ) | $spl82),
  inference(resolution,[],[f4651,f987])).
fof(f987,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f782])).
fof(f15562,plain,(
  ( ! [X364,X365] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X364),s('type.num.num',X365))))) ) | $spl82),
  inference(resolution,[],[f15414,f1058])).
fof(f15637,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15554,f15414])).
fof(f15554,plain,(
  ( ! [X352,X351] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X351),s('type.num.num',X351))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X352),s('type.num.num',X352))))))))) ) | $spl82),
  inference(resolution,[],[f15414,f1058])).
fof(f15636,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15552,f15414])).
fof(f15552,plain,(
  ( ! [X346,X347] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X346),s('type.num.num',X347))))) ) | $spl82),
  inference(resolution,[],[f15414,f1012])).
fof(f15635,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15551,f15414])).
fof(f15551,plain,(
  ( ! [X344,X345] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X344),s('type.num.num',X345))))) ) | $spl82),
  inference(resolution,[],[f15414,f1012])).
fof(f15544,plain,(
  $false | $spl82),
  inference(resolution,[],[f15414,f1045])).
fof(f15557,plain,(
  $false | $spl82),
  inference(resolution,[],[f15414,f1046])).
fof(f15375,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f15283,f15135])).
fof(f15135,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl80),
  inference(subsumption_resolution,[],[f15097,f1045])).
fof(f15097,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X0))) ) | $spl80),
  inference(resolution,[],[f4577,f971])).
fof(f971,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~sP10(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f775])).
fof(f15283,plain,(
  ( ! [X350,X351] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X350),s('type.num.num',X351))))) ) | $spl80),
  inference(resolution,[],[f15135,f1059])).
fof(f15374,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f15267,f15135])).
fof(f15267,plain,(
  ( ! [X331,X330] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X330),s('type.num.num',X330))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X331),s('type.num.num',X331))))))))) ) | $spl80),
  inference(resolution,[],[f15135,f1059])).
fof(f15273,plain,(
  $false | $spl80),
  inference(resolution,[],[f15135,f1091])).
fof(f15277,plain,(
  $false | ($spl0 | $spl5 | $spl80)),
  inference(resolution,[],[f15135,f2225])).
fof(f15278,plain,(
  $false | $spl80),
  inference(resolution,[],[f15135,f1090])).
fof(f15279,plain,(
  $false | ($spl80 | $spl111)),
  inference(resolution,[],[f15135,f8575])).
fof(f15280,plain,(
  $false | ($spl16 | $spl80)),
  inference(resolution,[],[f15135,f1578])).
fof(f15077,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl78),
  inference(subsumption_resolution,[],[f15038,f1045])).
fof(f15038,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl78),
  inference(resolution,[],[f4503,f953])).
fof(f953,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f768])).
fof(f15020,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl76),
  inference(subsumption_resolution,[],[f14981,f1045])).
fof(f14981,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | $spl76),
  inference(resolution,[],[f4426,f936])).
fof(f936,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f761])).
fof(f14963,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl72),
  inference(subsumption_resolution,[],[f14924,f1045])).
fof(f14924,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl72),
  inference(resolution,[],[f4192,f936])).
fof(f14906,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl70),
  inference(subsumption_resolution,[],[f14867,f1045])).
fof(f14867,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X1))) ) | $spl70),
  inference(resolution,[],[f4116,f919])).
fof(f919,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f753])).
fof(f14848,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1))) ) | $spl68),
  inference(subsumption_resolution,[],[f14810,f1045])).
fof(f14810,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl68),
  inference(resolution,[],[f4040,f845])).
fof(f845,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f728])).
fof(f14790,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1))) ) | $spl66),
  inference(subsumption_resolution,[],[f14752,f1045])).
fof(f14752,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | $spl66),
  inference(resolution,[],[f3963,f845])).
fof(f14732,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1))) ) | $spl64),
  inference(subsumption_resolution,[],[f14694,f1045])).
fof(f14694,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl64),
  inference(resolution,[],[f3886,f845])).
fof(f14676,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl58),
  inference(subsumption_resolution,[],[f14638,f1045])).
fof(f14638,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl58),
  inference(resolution,[],[f3651,f841])).
fof(f841,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f726])).
fof(f14637,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f14545,f1899])).
fof(f14545,plain,(
  ( ! [X350,X351] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X350),s('type.num.num',X351))))) ) | $spl32),
  inference(resolution,[],[f1899,f1059])).
fof(f14636,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f14529,f1899])).
fof(f14529,plain,(
  ( ! [X331,X330] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X330),s('type.num.num',X330))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X331),s('type.num.num',X331))))))))) ) | $spl32),
  inference(resolution,[],[f1899,f1059])).
fof(f14535,plain,(
  $false | $spl32),
  inference(resolution,[],[f1899,f1091])).
fof(f14539,plain,(
  $false | ($spl0 | $spl5 | $spl32)),
  inference(resolution,[],[f1899,f2225])).
fof(f14540,plain,(
  $false | $spl32),
  inference(resolution,[],[f1899,f1090])).
fof(f14541,plain,(
  $false | ($spl32 | $spl111)),
  inference(resolution,[],[f1899,f8575])).
fof(f14542,plain,(
  $false | ($spl16 | $spl32)),
  inference(resolution,[],[f1899,f1578])).
fof(f14397,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl56),
  inference(subsumption_resolution,[],[f14359,f1045])).
fof(f14359,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl56),
  inference(resolution,[],[f3575,f837])).
fof(f837,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f723])).
fof(f1896,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl26),
  inference(cnf_transformation,[],[f1896_D])).
fof(f1896_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f13707,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl111),
  inference(resolution,[],[f8575,f1093])).
fof(f13608,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | s(bool,X26) != s(bool,X27)) ) | $spl111),
  inference(resolution,[],[f8575,f829])).
fof(f13605,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl111),
  inference(resolution,[],[f8575,f828])).
fof(f13604,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl111),
  inference(resolution,[],[f8575,f827])).
fof(f13603,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | s(bool,X11) != s(bool,X12)) ) | $spl111),
  inference(resolution,[],[f8575,f825])).
fof(f13601,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl111),
  inference(resolution,[],[f8575,f823])).
fof(f1893,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl22),
  inference(cnf_transformation,[],[f1893_D])).
fof(f1893_D,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1895,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK14)))) | $spl24),
  inference(cnf_transformation,[],[f1895_D])).
fof(f1895_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK14)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f13600,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK14)))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl111),
  inference(resolution,[],[f8575,f1067])).
fof(f13560,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2568,f872])).
fof(f2568,plain,(
  ( ! [X196,X197] : (sP12(X196,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)),X197)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f995])).
fof(f2197,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f2196,f1609])).
fof(f13487,plain,(
  ( ! [X26,X27,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2561,f872])).
fof(f2561,plain,(
  ( ! [X182,X183] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)),X182,X183)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f986])).
fof(f13414,plain,(
  ( ! [X26,X27,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2550,f872])).
fof(f2550,plain,(
  ( ! [X161,X160] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)),X160,X161)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f970])).
fof(f13229,plain,(
  ( ! [X26,X27,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2538,f872])).
fof(f2538,plain,(
  ( ! [X136,X137] : (sP8(X136,X137,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f954])).
fof(f13154,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2522,f872])).
fof(f2522,plain,(
  ( ! [X105,X104] : (sP6(X104,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)),X105)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f938])).
fof(f13079,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2521,f872])).
fof(f2521,plain,(
  ( ! [X103,X102] : (sP6(X102,X103,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f937])).
fof(f13004,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2508,f872])).
fof(f2508,plain,(
  ( ! [X76,X77] : (sP4(X76,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)),X77)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f921])).
fof(f12927,plain,(
  ( ! [X26,X27,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2492,f872])).
fof(f2492,plain,(
  ( ! [X50,X51] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)),X50,X51)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f848])).
fof(f12851,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2491,f872])).
fof(f2491,plain,(
  ( ! [X48,X49] : (sP0(X48,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)),X49)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f847])).
fof(f12775,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2490,f872])).
fof(f2490,plain,(
  ( ! [X47,X46] : (sP0(X46,X47,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f846])).
fof(f12700,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2489,f872])).
fof(f2489,plain,(
  ( ! [X45,X44] : (sP1(X44,X45,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f842])).
fof(f12625,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2486,f872])).
fof(f2486,plain,(
  ( ! [X39,X38] : (sP2(X38,X39,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f838])).
fof(f7322,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | $spl104),
  inference(cnf_transformation,[],[f7322_D])).
fof(f7322_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f12436,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f1093])).
fof(f12337,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))))) | s(bool,X26) != s(bool,X27)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f829])).
fof(f12334,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f828])).
fof(f12333,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f827])).
fof(f12332,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))))) | s(bool,X11) != s(bool,X12)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f825])).
fof(f12330,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f823])).
fof(f11197,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl116),
  inference(cnf_transformation,[],[f11197_D])).
fof(f11197_D,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f11199,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) | $spl118),
  inference(cnf_transformation,[],[f11199_D])).
fof(f11199_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f12329,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) | s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f1067])).
fof(f7394,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | $spl110),
  inference(cnf_transformation,[],[f7394_D])).
fof(f7394_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f7442,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | $spl114),
  inference(cnf_transformation,[],[f7442_D])).
fof(f7442_D,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f12241,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | $spl0),
  inference(forward_demodulation,[],[f12020,f1102])).
fof(f12020,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | $spl0),
  inference(superposition,[],[f1086,f1102])).
fof(f12075,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | $spl0),
  inference(forward_demodulation,[],[f11817,f1102])).
fof(f11817,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | $spl0),
  inference(superposition,[],[f1102,f1086])).
fof(f12046,plain,(
  ( ! [X11] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK15))))))) ) | $spl0),
  inference(forward_demodulation,[],[f11775,f911])).
fof(f11775,plain,(
  ( ! [X11] : (s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK15))))))) ) | $spl0),
  inference(superposition,[],[f1102,f872])).
fof(f1371,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | $spl5),
  inference(cnf_transformation,[],[f1371_D])).
fof(f1371_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f11682,plain,(
  ( ! [X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK15))))))) ) | $spl3),
  inference(superposition,[],[f1105,f872])).
fof(f11493,plain,(
  $false | ($spl4 | $spl111)),
  inference(subsumption_resolution,[],[f11261,f911])).
fof(f11261,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) | ($spl4 | $spl111)),
  inference(backward_demodulation,[],[f1372,f7395])).
fof(f11492,plain,(
  $false | ($spl0 | $spl4 | $spl19)),
  inference(subsumption_resolution,[],[f11491,f1580])).
fof(f11491,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f11259,f911])).
fof(f11259,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f1372,f1102])).
fof(f7358,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl106),
  inference(cnf_transformation,[],[f7358_D])).
fof(f7358_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f11245,plain,(
  ( ! [X97] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X97),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X97),s('type.num.num',sK26(sK15)))))))))) ) | $spl132),
  inference(cnf_transformation,[],[f11245_D])).
fof(f11245_D,plain,(
  ( ! [X97] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X97),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X97),s('type.num.num',sK26(sK15)))))))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f11255,plain,(
  ( ! [X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',sK26(sK15))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15)) ) | $spl20),
  inference(forward_demodulation,[],[f9219,f1075])).
fof(f9219,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))))))))) ) | $spl20),
  inference(superposition,[],[f1582,f872])).
fof(f11238,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num',sK26(sK15))))))) | $spl130),
  inference(cnf_transformation,[],[f11238_D])).
fof(f11238_D,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num',sK26(sK15))))))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f11252,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num',sK26(sK15))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl20),
  inference(forward_demodulation,[],[f9235,f1582])).
fof(f9235,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num',sK26(sK15))))))) | $spl20),
  inference(superposition,[],[f1582,f1086])).
fof(f11244,plain,(
  ( ! [X97] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X97),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X97),s('type.num.num',sK26(sK15))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15)) ) | $spl20),
  inference(forward_demodulation,[],[f9309,f1075])).
fof(f9309,plain,(
  ( ! [X97] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X97),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X97),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))))))))) ) | $spl20),
  inference(superposition,[],[f872,f1582])).
fof(f11236,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num',sK26(sK15))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl20),
  inference(forward_demodulation,[],[f9452,f1582])).
fof(f9452,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num',sK26(sK15))))))) | $spl20),
  inference(superposition,[],[f1086,f1582])).
fof(f11235,plain,(
  ( ! [X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X11)))) ) | $spl128),
  inference(cnf_transformation,[],[f11235_D])).
fof(f11235_D,plain,(
  ( ! [X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X11)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7357,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | $spl107),
  inference(cnf_transformation,[],[f7357_D])).
fof(f7357_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f10246,plain,(
  ( ! [X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X11)))) ) | $spl37),
  inference(superposition,[],[f2665,f861])).
fof(f2665,plain,(
  s('type.num.num',sK15) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl37),
  inference(cnf_transformation,[],[f2665_D])).
fof(f2665_D,plain,(
  s('type.num.num',sK15) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f11234,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl126),
  inference(cnf_transformation,[],[f11234_D])).
fof(f11234_D,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f10247,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl37),
  inference(superposition,[],[f2665,f862])).
fof(f11233,plain,(
  ( ! [X13] : (s('type.num.num',X13) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl124),
  inference(cnf_transformation,[],[f11233_D])).
fof(f11233_D,plain,(
  ( ! [X13] : (s('type.num.num',X13) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f10248,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl37),
  inference(superposition,[],[f2665,f864])).
fof(f11230,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl123),
  inference(cnf_transformation,[],[f11230_D])).
fof(f11230_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f10254,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl37),
  inference(superposition,[],[f2665,f902])).
fof(f902,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f746])).
fof(f746,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f565])).
fof(f565,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X7))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7))),
  file('thm.numpair.tri_|49||49|','thm.numpair.tri_eq_|48||split|0')).
fof(f10255,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl37),
  inference(superposition,[],[f2665,f904])).
fof(f11228,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl120),
  inference(cnf_transformation,[],[f11228_D])).
fof(f11228_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f10265,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl37),
  inference(superposition,[],[f2665,f1086])).
fof(f10436,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X8) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15)) ) | $spl8),
  inference(trivial_inequality_removal,[],[f10392])).
fof(f10392,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num',X8) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15)) ) | $spl8),
  inference(superposition,[],[f1573,f863])).
fof(f863,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f733])).
fof(f10393,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9)) ) | $spl8),
  inference(superposition,[],[f1573,f865])).
fof(f865,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f734])).
fof(f10593,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X8)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10548])).
fof(f10548,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X8)) ) | $spl10),
  inference(superposition,[],[f1574,f863])).
fof(f10549,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15)) ) | $spl10),
  inference(superposition,[],[f1574,f865])).
fof(f10962,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X8) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15)) ) | $spl12),
  inference(duplicate_literal_removal,[],[f10925])).
fof(f10925,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X8) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X8)) ) | $spl12),
  inference(superposition,[],[f1575,f863])).
fof(f11083,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f1093])).
fof(f10984,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))))) | s(bool,X26) != s(bool,X27)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f829])).
fof(f10981,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f828])).
fof(f10980,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f827])).
fof(f10979,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))))) | s(bool,X11) != s(bool,X12)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f825])).
fof(f10977,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f823])).
fof(f10976,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) | s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f1067])).
fof(f10714,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl5),
  inference(resolution,[],[f7496,f1093])).
fof(f7496,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl5),
  inference(trivial_inequality_removal,[],[f7479])).
fof(f7479,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl5),
  inference(superposition,[],[f1371,f894])).
fof(f894,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f558])).
fof(f558,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f378])).
fof(f378,axiom,(
  ! [X6] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.LESS_|48|_CASES')).
fof(f10615,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | s(bool,X26) != s(bool,X27)) ) | $spl5),
  inference(resolution,[],[f7496,f829])).
fof(f10612,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl5),
  inference(resolution,[],[f7496,f828])).
fof(f10611,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl5),
  inference(resolution,[],[f7496,f827])).
fof(f10610,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | s(bool,X11) != s(bool,X12)) ) | $spl5),
  inference(resolution,[],[f7496,f825])).
fof(f10608,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl5),
  inference(resolution,[],[f7496,f823])).
fof(f10607,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK14)))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl5),
  inference(resolution,[],[f7496,f1067])).
fof(f10594,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK15))))))) ) | $spl10),
  inference(forward_demodulation,[],[f10526,f896])).
fof(f10526,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK15))))))) ) | $spl10),
  inference(superposition,[],[f1574,f872])).
fof(f1575,plain,(
  ( ! [X10] : (s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',sK15)))) ) | $spl12),
  inference(cnf_transformation,[],[f1575_D])).
fof(f1575_D,plain,(
  ( ! [X10] : (s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',sK15)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f10525,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',sK15)))) ) | $spl10),
  inference(superposition,[],[f1574,f864])).
fof(f10437,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK15))))))) ) | $spl8),
  inference(forward_demodulation,[],[f10289,f895])).
fof(f10289,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK15))))))) ) | $spl8),
  inference(superposition,[],[f1573,f872])).
fof(f10288,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X16))) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',sK15)))) ) | $spl8),
  inference(superposition,[],[f1573,f864])).
fof(f1574,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK15)))) ) | $spl10),
  inference(cnf_transformation,[],[f1574_D])).
fof(f1574_D,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK15)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f10287,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',sK15)))) ) | $spl8),
  inference(superposition,[],[f1573,f862])).
fof(f10177,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2568,f872])).
fof(f10176,plain,(
  ( ! [X24,X23,X22] : (sP12(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2568,f864])).
fof(f10175,plain,(
  ( ! [X21,X19,X20] : (sP12(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2568,f862])).
fof(f1573,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X8)))) ) | $spl8),
  inference(cnf_transformation,[],[f1573_D])).
fof(f1573_D,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X8)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f10174,plain,(
  ( ! [X17,X18,X16] : (sP12(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2568,f861])).
fof(f10104,plain,(
  ( ! [X26,X27,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2561,f872])).
fof(f10103,plain,(
  ( ! [X24,X23,X22] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2561,f864])).
fof(f10102,plain,(
  ( ! [X21,X19,X20] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2561,f862])).
fof(f10101,plain,(
  ( ! [X17,X18,X16] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2561,f861])).
fof(f10031,plain,(
  ( ! [X26,X27,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2550,f872])).
fof(f10030,plain,(
  ( ! [X24,X23,X22] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2550,f864])).
fof(f10029,plain,(
  ( ! [X21,X19,X20] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2550,f862])).
fof(f10028,plain,(
  ( ! [X17,X18,X16] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2550,f861])).
fof(f9956,plain,(
  ( ! [X26,X27,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2538,f872])).
fof(f9955,plain,(
  ( ! [X24,X23,X22] : (sP8(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2538,f864])).
fof(f9954,plain,(
  ( ! [X21,X19,X20] : (sP8(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2538,f862])).
fof(f9953,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2538,f861])).
fof(f9881,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2522,f872])).
fof(f9880,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2522,f864])).
fof(f9879,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2522,f862])).
fof(f9878,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2522,f861])).
fof(f9806,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2521,f872])).
fof(f9805,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2521,f864])).
fof(f9804,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2521,f862])).
fof(f9803,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2521,f861])).
fof(f9731,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2508,f872])).
fof(f9730,plain,(
  ( ! [X24,X23,X22] : (sP4(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2508,f864])).
fof(f9729,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2508,f862])).
fof(f9728,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2508,f861])).
fof(f2666,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(cnf_transformation,[],[f2666_D])).
fof(f2666_D,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f9206,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15))))) | $spl20),
  inference(superposition,[],[f1582,f1086])).
fof(f3287,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK26(sK15)))) | $spl53),
  inference(cnf_transformation,[],[f3287_D])).
fof(f3287_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK26(sK15)))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f9197,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK26(sK15)))) | $spl20),
  inference(superposition,[],[f1582,f904])).
fof(f9196,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK26(sK15)))) | $spl20),
  inference(superposition,[],[f1582,f902])).
fof(f3285,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) | $spl50),
  inference(cnf_transformation,[],[f3285_D])).
fof(f3285_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f9195,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) | $spl20),
  inference(superposition,[],[f1582,f894])).
fof(f9194,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) | $spl20),
  inference(superposition,[],[f1582,f892])).
fof(f9191,plain,(
  ( ! [X11] : (s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK26(sK15)))))))) ) | $spl20),
  inference(superposition,[],[f1582,f872])).
fof(f9190,plain,(
  ( ! [X10] : (s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',sK26(sK15))))) ) | $spl20),
  inference(superposition,[],[f1582,f864])).
fof(f9189,plain,(
  ( ! [X9] : (s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK26(sK15))))) ) | $spl20),
  inference(superposition,[],[f1582,f862])).
fof(f9188,plain,(
  ( ! [X8] : (s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X8)))) ) | $spl20),
  inference(superposition,[],[f1582,f861])).
fof(f1582,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) | $spl20),
  inference(cnf_transformation,[],[f1582_D])).
fof(f1582_D,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f9168,plain,(
  ( ! [X68,X67] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X67,X68) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2492,f1086])).
fof(f9153,plain,(
  ( ! [X26,X27,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2492,f872])).
fof(f9152,plain,(
  ( ! [X24,X23,X22] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2492,f864])).
fof(f9151,plain,(
  ( ! [X21,X19,X20] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2492,f862])).
fof(f9150,plain,(
  ( ! [X17,X18,X16] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2492,f861])).
fof(f9092,plain,(
  ( ! [X68,X67] : (sP0(X67,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X68) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2491,f1086])).
fof(f9077,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2491,f872])).
fof(f9076,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2491,f864])).
fof(f9075,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2491,f862])).
fof(f9074,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2491,f861])).
fof(f9016,plain,(
  ( ! [X68,X67] : (sP0(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2490,f1086])).
fof(f9001,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2490,f872])).
fof(f9000,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2490,f864])).
fof(f8999,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2490,f862])).
fof(f8998,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2490,f861])).
fof(f8941,plain,(
  ( ! [X68,X67] : (sP1(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2489,f1086])).
fof(f8926,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2489,f872])).
fof(f8925,plain,(
  ( ! [X24,X23,X22] : (sP1(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2489,f864])).
fof(f8924,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2489,f862])).
fof(f8923,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2489,f861])).
fof(f8866,plain,(
  ( ! [X68,X67] : (sP2(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2486,f1086])).
fof(f8851,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2486,f872])).
fof(f8850,plain,(
  ( ! [X24,X23,X22] : (sP2(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2486,f864])).
fof(f8849,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2486,f862])).
fof(f8848,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2486,f861])).
fof(f2669,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | $spl40),
  inference(cnf_transformation,[],[f2669_D])).
fof(f2669_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f7656,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | $spl16),
  inference(resolution,[],[f1578,f1093])).
fof(f7557,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))) | s(bool,X26) != s(bool,X27)) ) | $spl16),
  inference(resolution,[],[f1578,f829])).
fof(f7554,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | $spl16),
  inference(resolution,[],[f1578,f828])).
fof(f7553,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | $spl16),
  inference(resolution,[],[f1578,f827])).
fof(f7552,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))) | s(bool,X11) != s(bool,X12)) ) | $spl16),
  inference(resolution,[],[f1578,f825])).
fof(f7550,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | $spl16),
  inference(resolution,[],[f1578,f823])).
fof(f2668,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK15)))) | $spl38),
  inference(cnf_transformation,[],[f2668_D])).
fof(f2668_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK15)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f7549,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK15)))) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl16),
  inference(resolution,[],[f1578,f1067])).
fof(f7536,plain,(
  $false | ($spl0 | $spl5 | $spl106)),
  inference(subsumption_resolution,[],[f7358,f2206])).
fof(f2206,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f1941,f1609])).
fof(f2192,plain,(
  ( ! [X10] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',sK15)))) ) | $spl0),
  inference(forward_demodulation,[],[f1922,f911])).
fof(f1922,plain,(
  ( ! [X10] : (s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',sK15)))) ) | $spl0),
  inference(superposition,[],[f1102,f864])).
fof(f2190,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK15)))) ) | $spl0),
  inference(forward_demodulation,[],[f1921,f911])).
fof(f1921,plain,(
  ( ! [X9] : (s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK15)))) ) | $spl0),
  inference(superposition,[],[f1102,f862])).
fof(f2188,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X8)))) ) | $spl0),
  inference(forward_demodulation,[],[f1920,f911])).
fof(f1920,plain,(
  ( ! [X8] : (s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X8)))) ) | $spl0),
  inference(superposition,[],[f1102,f861])).
fof(f2194,plain,(
  ( ! [X11] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK15))))))) ) | $spl0),
  inference(forward_demodulation,[],[f1923,f911])).
fof(f1923,plain,(
  ( ! [X11] : (s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK15))))))) ) | $spl0),
  inference(superposition,[],[f1102,f872])).
fof(f1578,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) | $spl16),
  inference(cnf_transformation,[],[f1578_D])).
fof(f1578_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2196,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) | $spl0),
  inference(forward_demodulation,[],[f1926,f911])).
fof(f1926,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) | $spl0),
  inference(superposition,[],[f1102,f892])).
fof(f2198,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) | $spl0),
  inference(forward_demodulation,[],[f1927,f911])).
fof(f1927,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) | $spl0),
  inference(superposition,[],[f1102,f894])).
fof(f1580,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | $spl19),
  inference(cnf_transformation,[],[f1580_D])).
fof(f1580_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f2200,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | $spl0),
  inference(forward_demodulation,[],[f1928,f911])).
fof(f1928,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | $spl0),
  inference(superposition,[],[f1102,f902])).
fof(f2202,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | $spl0),
  inference(forward_demodulation,[],[f1929,f911])).
fof(f1929,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | $spl0),
  inference(superposition,[],[f1102,f904])).
fof(f2204,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) | $spl0),
  inference(forward_demodulation,[],[f1938,f911])).
fof(f1938,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) | $spl0),
  inference(superposition,[],[f1102,f1086])).
fof(f1941,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | $spl0),
  inference(superposition,[],[f1102,f905])).
fof(f905,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f747])).
fof(f1942,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | $spl0),
  inference(superposition,[],[f1102,f903])).
fof(f903,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f746])).
fof(f7457,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | $spl0),
  inference(forward_demodulation,[],[f1965,f1102])).
fof(f1965,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | $spl0),
  inference(superposition,[],[f1102,f1086])).
fof(f7395,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | $spl111),
  inference(cnf_transformation,[],[f7395_D])).
fof(f7395_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f1972,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl0),
  inference(superposition,[],[f902,f1102])).
fof(f1973,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | $spl0),
  inference(superposition,[],[f903,f1102])).
fof(f1974,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl0),
  inference(superposition,[],[f904,f1102])).
fof(f1975,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | $spl0),
  inference(superposition,[],[f905,f1102])).
fof(f7440,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | $spl0),
  inference(forward_demodulation,[],[f2167,f1102])).
fof(f2167,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.numpair.tri_1'(s('type.num.num',sK15)))))) | $spl0),
  inference(superposition,[],[f1086,f1102])).
fof(f4278,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X13)))) ) | $spl10),
  inference(superposition,[],[f1574,f861])).
fof(f7439,plain,(
  ( ! [X24] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X24) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))))) ) | $spl10),
  inference(forward_demodulation,[],[f4284,f896])).
fof(f4284,plain,(
  ( ! [X24] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))))) ) | $spl10),
  inference(superposition,[],[f1574,f892])).
fof(f7438,plain,(
  ( ! [X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X25) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))))) ) | $spl10),
  inference(forward_demodulation,[],[f4285,f896])).
fof(f4285,plain,(
  ( ! [X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))))) ) | $spl10),
  inference(superposition,[],[f1574,f894])).
fof(f7437,plain,(
  ( ! [X26] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f4286,f1102])).
fof(f4286,plain,(
  ( ! [X26] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))) ) | $spl10),
  inference(superposition,[],[f1574,f902])).
fof(f7436,plain,(
  ( ! [X27] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f4287,f1102])).
fof(f4287,plain,(
  ( ! [X27] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))) ) | $spl10),
  inference(superposition,[],[f1574,f904])).
fof(f4296,plain,(
  ( ! [X45] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15))))) ) | $spl10),
  inference(superposition,[],[f1574,f1086])).
fof(f4348,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X8)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f4303])).
fof(f4303,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X8)) ) | $spl10),
  inference(superposition,[],[f1574,f863])).
fof(f4304,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15)) ) | $spl10),
  inference(superposition,[],[f1574,f865])).
fof(f4873,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X8) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15)) ) | $spl12),
  inference(duplicate_literal_removal,[],[f4836])).
fof(f4836,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X8) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X8)) ) | $spl12),
  inference(superposition,[],[f1575,f863])).
fof(f7403,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl112),
  inference(cnf_transformation,[],[f7403_D])).
fof(f7403_D,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f7422,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl36),
  inference(forward_demodulation,[],[f4936,f2666])).
fof(f4936,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl36),
  inference(superposition,[],[f2666,f1086])).
fof(f7401,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl36),
  inference(forward_demodulation,[],[f5152,f2666])).
fof(f5152,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl36),
  inference(superposition,[],[f1086,f2666])).
fof(f5845,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X18))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num',X18)))) ) | $spl86),
  inference(cnf_transformation,[],[f5845_D])).
fof(f5845_D,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X18))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num',X18)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f7400,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X12))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X12))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X13)))) ) | $spl36),
  inference(forward_demodulation,[],[f5504,f896])).
fof(f5504,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X12))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X13)))) ) | $spl36),
  inference(superposition,[],[f4875,f861])).
fof(f4875,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK15)))))) ) | $spl36),
  inference(backward_demodulation,[],[f2666,f1087])).
fof(f7399,plain,(
  ( ! [X24] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X24))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))))) ) | $spl36),
  inference(forward_demodulation,[],[f5510,f896])).
fof(f5510,plain,(
  ( ! [X24] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X24))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))))) ) | $spl36),
  inference(superposition,[],[f4875,f892])).
fof(f7398,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X25))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))))) ) | $spl36),
  inference(forward_demodulation,[],[f5511,f896])).
fof(f5511,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X25))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))))) ) | $spl36),
  inference(superposition,[],[f4875,f894])).
fof(f7397,plain,(
  ( ! [X26] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X26))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X26)))) ) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f7396,f1102])).
fof(f7396,plain,(
  ( ! [X26] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X26))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X26))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))) ) | $spl36),
  inference(forward_demodulation,[],[f5512,f896])).
fof(f5512,plain,(
  ( ! [X26] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X26))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))) ) | $spl36),
  inference(superposition,[],[f4875,f902])).
fof(f7393,plain,(
  ( ! [X27] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X27))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27)))) ) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f7392,f1102])).
fof(f7392,plain,(
  ( ! [X27] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X27))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))) ) | $spl36),
  inference(forward_demodulation,[],[f5513,f896])).
fof(f5513,plain,(
  ( ! [X27] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X27))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))) ) | $spl36),
  inference(superposition,[],[f4875,f904])).
fof(f7391,plain,(
  ( ! [X45] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X45))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X45))) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15))))) ) | $spl36),
  inference(forward_demodulation,[],[f5522,f896])).
fof(f5522,plain,(
  ( ! [X45] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X45))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15))))) ) | $spl36),
  inference(superposition,[],[f4875,f1086])).
fof(f7390,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X8)) ) | ($spl10 | $spl36)),
  inference(subsumption_resolution,[],[f7389,f5903])).
fof(f5903,plain,(
  ( ! [X27] : (s('type.num.num',X27) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X27)))) ) | ($spl10 | $spl36)),
  inference(subsumption_resolution,[],[f5617,f1574])).
fof(f5617,plain,(
  ( ! [X27] : (s('type.num.num',X27) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X27))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15)))) ) | $spl36),
  inference(superposition,[],[f864,f4875])).
fof(f7389,plain,(
  ( ! [X8] : (s('type.num.num',X8) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X8)) ) | $spl36),
  inference(forward_demodulation,[],[f5529,f896])).
fof(f5529,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X8)) ) | $spl36),
  inference(superposition,[],[f4875,f863])).
fof(f5530,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15)) ) | $spl36),
  inference(superposition,[],[f4875,f865])).
fof(f7360,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | $spl108),
  inference(cnf_transformation,[],[f7360_D])).
fof(f7360_D,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f7380,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl36),
  inference(forward_demodulation,[],[f6136,f5846])).
fof(f5846,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | $spl36),
  inference(forward_demodulation,[],[f5533,f895])).
fof(f5533,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))) | $spl36),
  inference(superposition,[],[f4875,f897])).
fof(f897,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f561])).
fof(f561,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.numpair.tri_|49||49|','thm.arithmetic.ADD|split|0')).
fof(f6136,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | $spl36),
  inference(superposition,[],[f5846,f1086])).
fof(f7356,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl36),
  inference(forward_demodulation,[],[f6340,f5846])).
fof(f6340,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | $spl36),
  inference(superposition,[],[f1086,f5846])).
fof(f7202,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f1093])).
fof(f7103,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))))) | s(bool,X26) != s(bool,X27)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f829])).
fof(f7100,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f828])).
fof(f7099,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f827])).
fof(f7098,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))))) | s(bool,X11) != s(bool,X12)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f825])).
fof(f7096,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f823])).
fof(f7319,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) | $spl100),
  inference(cnf_transformation,[],[f7319_D])).
fof(f7319_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f7321,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | $spl102),
  inference(cnf_transformation,[],[f7321_D])).
fof(f7321_D,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f7317,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) | ($spl0 | $spl5 | $spl36)),
  inference(forward_demodulation,[],[f7316,f2666])).
fof(f7316,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) | s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl5 | $spl36)),
  inference(forward_demodulation,[],[f7095,f2666])).
fof(f7095,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14)))))) | s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl5)),
  inference(resolution,[],[f2225,f1067])).
fof(f7061,plain,(
  ( ! [X68,X67] : (sP0(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl5),
  inference(superposition,[],[f1717,f1086])).
fof(f1717,plain,(
  ( ! [X47,X46] : (sP0(X46,X47,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl5),
  inference(resolution,[],[f1611,f846])).
fof(f1611,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl5),
  inference(trivial_inequality_removal,[],[f1594])).
fof(f1594,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | $spl5),
  inference(superposition,[],[f1371,f894])).
fof(f7046,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl5),
  inference(superposition,[],[f1717,f872])).
fof(f7045,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl5),
  inference(superposition,[],[f1717,f864])).
fof(f7044,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl5),
  inference(superposition,[],[f1717,f862])).
fof(f7043,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl5),
  inference(superposition,[],[f1717,f861])).
fof(f6980,plain,(
  ( ! [X68,X67] : (sP1(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl5),
  inference(superposition,[],[f1716,f1086])).
fof(f1716,plain,(
  ( ! [X45,X44] : (sP1(X44,X45,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl5),
  inference(resolution,[],[f1611,f842])).
fof(f6965,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl5),
  inference(superposition,[],[f1716,f872])).
fof(f6964,plain,(
  ( ! [X24,X23,X22] : (sP1(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl5),
  inference(superposition,[],[f1716,f864])).
fof(f6963,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl5),
  inference(superposition,[],[f1716,f862])).
fof(f6962,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl5),
  inference(superposition,[],[f1716,f861])).
fof(f6893,plain,(
  ( ! [X68,X67] : (sP2(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK14))))) ) | $spl5),
  inference(superposition,[],[f1713,f1086])).
fof(f1713,plain,(
  ( ! [X39,X38] : (sP2(X38,X39,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl5),
  inference(resolution,[],[f1611,f838])).
fof(f6878,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK14))))))) ) | $spl5),
  inference(superposition,[],[f1713,f872])).
fof(f6877,plain,(
  ( ! [X24,X23,X22] : (sP2(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK14)))) ) | $spl5),
  inference(superposition,[],[f1713,f864])).
fof(f6876,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK14)))) ) | $spl5),
  inference(superposition,[],[f1713,f862])).
fof(f6875,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X18)))) ) | $spl5),
  inference(superposition,[],[f1713,f861])).
fof(f6451,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | $spl89),
  inference(cnf_transformation,[],[f6451_D])).
fof(f6451_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f6452,plain,(
  ( ! [X77] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X77)) ) | ($spl10 | $spl36)),
  inference(subsumption_resolution,[],[f6186,f1574])).
fof(f6186,plain,(
  ( ! [X77] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X77),s('type.num.num',sK15))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X77)) ) | $spl36),
  inference(superposition,[],[f863,f5846])).
fof(f6449,plain,(
  ( ! [X76] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X76) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl5 | $spl36)),
  inference(subsumption_resolution,[],[f6185,f2189])).
fof(f2189,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X8)))) ) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f2188,f1609])).
fof(f6185,plain,(
  ( ! [X76] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X76))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X76) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl36),
  inference(superposition,[],[f863,f5846])).
fof(f5844,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X18))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num',X18))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK15))))))) ) | $spl36),
  inference(forward_demodulation,[],[f5507,f896])).
fof(f5507,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X18))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK15))))))) ) | $spl36),
  inference(superposition,[],[f4875,f872])).
fof(f4697,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2568,f872])).
fof(f4696,plain,(
  ( ! [X24,X23,X22] : (sP12(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2568,f864])).
fof(f4623,plain,(
  ( ! [X26,X27,X25] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2561,f872])).
fof(f4622,plain,(
  ( ! [X24,X23,X22] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2561,f864])).
fof(f4549,plain,(
  ( ! [X26,X27,X25] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2550,f872])).
fof(f4548,plain,(
  ( ! [X24,X23,X22] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2550,f864])).
fof(f4473,plain,(
  ( ! [X26,X27,X25] : (sP8(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2538,f872])).
fof(f4472,plain,(
  ( ! [X24,X23,X22] : (sP8(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2538,f864])).
fof(f4397,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2522,f872])).
fof(f4396,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2522,f864])).
fof(f4350,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK15))))))) ) | $spl10),
  inference(forward_demodulation,[],[f4281,f896])).
fof(f4281,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK15))))))) ) | $spl10),
  inference(superposition,[],[f1574,f872])).
fof(f4280,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',sK15)))) ) | $spl10),
  inference(superposition,[],[f1574,f864])).
fof(f4163,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2521,f872])).
fof(f4162,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2521,f864])).
fof(f4161,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2521,f862])).
fof(f4087,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2508,f872])).
fof(f4086,plain,(
  ( ! [X24,X23,X22] : (sP4(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2508,f864])).
fof(f4085,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2508,f862])).
fof(f4009,plain,(
  ( ! [X26,X27,X25] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2492,f872])).
fof(f4008,plain,(
  ( ! [X24,X23,X22] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2492,f864])).
fof(f4007,plain,(
  ( ! [X21,X19,X20] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2492,f862])).
fof(f3932,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2491,f872])).
fof(f3931,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2491,f864])).
fof(f3930,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2491,f862])).
fof(f3855,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2490,f872])).
fof(f3854,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2490,f864])).
fof(f3853,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2490,f862])).
fof(f3808,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X18) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f3659,f895])).
fof(f3659,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2189,f872])).
fof(f3658,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X16))) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2189,f864])).
fof(f3657,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2189,f862])).
fof(f3621,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2489,f872])).
fof(f3620,plain,(
  ( ! [X24,X23,X22] : (sP1(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2489,f864])).
fof(f3619,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2489,f862])).
fof(f3618,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2489,f861])).
fof(f3545,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,X26,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK15))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2486,f872])).
fof(f3544,plain,(
  ( ! [X24,X23,X22] : (sP2(X22,X23,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2486,f864])).
fof(f3543,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,X20,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2486,f862])).
fof(f3542,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X18)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2486,f861])).
fof(f3014,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK26(sK15)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK26(sK15))))) | ($spl0 | $spl5)),
  inference(superposition,[],[f2205,f1086])).
fof(f3005,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK26(sK15)))) | ($spl0 | $spl5)),
  inference(superposition,[],[f2205,f904])).
fof(f3004,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK26(sK15)))) | ($spl0 | $spl5)),
  inference(superposition,[],[f2205,f902])).
fof(f3003,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) | ($spl0 | $spl5)),
  inference(superposition,[],[f2205,f894])).
fof(f3002,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK26(sK15))))) | ($spl0 | $spl5)),
  inference(superposition,[],[f2205,f892])).
fof(f2999,plain,(
  ( ! [X11] : (s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK26(sK15)))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2205,f872])).
fof(f2998,plain,(
  ( ! [X10] : (s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2205,f864])).
fof(f2997,plain,(
  ( ! [X9] : (s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK26(sK15))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2205,f862])).
fof(f2996,plain,(
  ( ! [X8] : (s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK26(sK15)),s('type.num.num',X8)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2205,f861])).
fof(f2893,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f2817,f1900])).
fof(f2817,plain,(
  ( ! [X364,X365] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X364),s('type.num.num',X365))))) ) | $spl34),
  inference(resolution,[],[f1900,f1058])).
fof(f2892,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f2809,f1900])).
fof(f2809,plain,(
  ( ! [X352,X351] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X351),s('type.num.num',X351))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X352),s('type.num.num',X352))))))))) ) | $spl34),
  inference(resolution,[],[f1900,f1058])).
fof(f2891,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f2807,f1900])).
fof(f2807,plain,(
  ( ! [X346,X347] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X346),s('type.num.num',X347))))) ) | $spl34),
  inference(resolution,[],[f1900,f1012])).
fof(f2890,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f2806,f1900])).
fof(f2806,plain,(
  ( ! [X344,X345] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X344),s('type.num.num',X345))))) ) | $spl34),
  inference(resolution,[],[f1900,f1012])).
fof(f2799,plain,(
  $false | $spl34),
  inference(resolution,[],[f1900,f1045])).
fof(f2812,plain,(
  $false | $spl34),
  inference(resolution,[],[f1900,f1046])).
fof(f2579,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f1093])).
fof(f2480,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))) | s(bool,X26) != s(bool,X27)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f829])).
fof(f2477,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f828])).
fof(f2476,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f827])).
fof(f2475,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15))) | s(bool,X11) != s(bool,X12)) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f825])).
fof(f2473,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f823])).
fof(f2472,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK15)))) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl5)),
  inference(resolution,[],[f2197,f1067])).
fof(f1806,plain,(
  ( ! [X210,X211] : (p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl5),
  inference(resolution,[],[f1611,f1093])).
fof(f1707,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | s(bool,X26) != s(bool,X27)) ) | $spl5),
  inference(resolution,[],[f1611,f829])).
fof(f1704,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl5),
  inference(resolution,[],[f1611,f828])).
fof(f1703,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl5),
  inference(resolution,[],[f1611,f827])).
fof(f1702,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | s(bool,X11) != s(bool,X12)) ) | $spl5),
  inference(resolution,[],[f1611,f825])).
fof(f1700,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl5),
  inference(resolution,[],[f1611,f823])).
fof(f1699,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK14)))) | s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl5),
  inference(resolution,[],[f1611,f1067])).
fof(f1570,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) | $spl3),
  inference(superposition,[],[f1105,f1086])).
fof(f1561,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | $spl3),
  inference(superposition,[],[f1105,f904])).
fof(f1560,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | $spl3),
  inference(superposition,[],[f1105,f902])).
fof(f1559,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) | $spl3),
  inference(superposition,[],[f1105,f894])).
fof(f1558,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK15)))) | $spl3),
  inference(superposition,[],[f1105,f892])).
fof(f1555,plain,(
  ( ! [X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK15))))))) ) | $spl3),
  inference(superposition,[],[f1105,f872])).
fof(f1554,plain,(
  ( ! [X10] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | s('type.num.num',X10) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',sK15)))) ) | $spl3),
  inference(superposition,[],[f1105,f864])).
fof(f1553,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK15)))) ) | $spl3),
  inference(superposition,[],[f1105,f862])).
fof(f1552,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK14) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num',X8)))) ) | $spl3),
  inference(superposition,[],[f1105,f861])).
fof(f1547,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f1546])).
fof(f1546,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1103,f1104])).
fof(f1372,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl4),
  inference(cnf_transformation,[],[f1372_D])).
fof(f1372_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1374,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) | $spl6),
  inference(cnf_transformation,[],[f1374_D])).
fof(f1374_D,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1540,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl2),
  inference(forward_demodulation,[],[f1321,f1104])).
fof(f1321,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) | $spl2),
  inference(superposition,[],[f1086,f1104])).
fof(f1370,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | $spl2),
  inference(forward_demodulation,[],[f1128,f1104])).
fof(f1128,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK14) | s('type.num.num',sK15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK26(sK15)))) | $spl2),
  inference(superposition,[],[f1104,f1086])).
fof(f1102,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | $spl0),
  inference(cnf_transformation,[],[f1102_D])).
fof(f1102_D,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1104,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | $spl2),
  inference(cnf_transformation,[],[f1104_D])).
fof(f1104_D,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f820,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))),
  inference(cnf_transformation,[],[f715])).
fof(f715,plain,(
  (s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | s('type.num.num',sK14) = s('type.num.num',sK15)) & (s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | s('type.num.num',sK14) != s('type.num.num',sK15))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f714])).
fof(f714,plain,(
  ? [X0,X1] : ((s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))) != s('type.num.num','const.numpair.tri_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f648])).
fof(f648,plain,(
  ? [X0,X1] : (s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f518])).
fof(f518,plain,(
  ~! [X0,X1] : (s('type.num.num','const.numpair.tri_1'(s('type.num.num',X0))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.num.num','const.numpair.tri_1'(s('type.num.num',X6))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.num.num','const.numpair.tri_1'(s('type.num.num',X6))) = s('type.num.num','const.numpair.tri_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.numpair.tri_|49||49|','thm.numpair.tri_|49||49|')).
fof(f1103,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) | $spl1),
  inference(cnf_transformation,[],[f1103_D])).
fof(f1103_D,plain,(
  s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1105,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) | $spl3),
  inference(cnf_transformation,[],[f1105_D])).
fof(f1105_D,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f821,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) | s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK14))) != s('type.num.num','const.numpair.tri_1'(s('type.num.num',sK15)))),
  inference(cnf_transformation,[],[f715])).
% SZS output end Proof for thm.numpair
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 423
% Passive clauses: 11331
% Generated clauses: 29857
% Final active clauses: 267
% Final passive clauses: 6096
% Input formulas: 516
% Initial clauses: 273
% 
% Selected by SInE selection: 100
% SInE iterations: 1
% 
% Duplicate literals: 138
% Trivial inequalities: 228
% Fw subsumption resolutions: 1651
% Fw demodulations: 5041
% Bw demodulations: 49
% 
% Simple tautologies: 376
% Equational tautologies: 2
% Forward subsumptions: 9847
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 4235
% Forward superposition: 13030
% Backward superposition: 4716
% Self superposition: 2
% Equality resolution: 2
% 
% Split clauses: 626
% Split components: 1273
% SAT solver clauses: 688
% SAT solver unit clauses: 36
% SAT solver binary clauses: 30
% 
% Sat splits: 626
% Sat splitting refutations: 54
% 
% Memory used [KB]: 12537
% Time elapsed: 0.548 s
% ------------------------------
% ------------------------------
% Success in time 12.305 s
