Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (60456)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 181
% Passive clauses: 1940
% Generated clauses: 27312
% Final active clauses: 176
% Final passive clauses: 1543
% Input formulas: 27
% Initial clauses: 57
% 
% Duplicate literals: 1538
% Trivial inequalities: 643
% Fw subsumption resolutions: 3651
% Fw demodulations: 4066
% Bw demodulations: 4
% 
% Simple tautologies: 580
% Equational tautologies: 47
% Forward subsumptions: 10613
% 
% Binary resolution: 201
% Forward superposition: 8211
% Backward superposition: 8407
% Self superposition: 204
% Equality factoring: 145
% Equality resolution: 54
% 
% Split clauses: 591
% Split components: 1199
% SAT solver clauses: 1067
% SAT solver unit clauses: 25
% SAT solver binary clauses: 120
% 
% Sat splits: 591
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1065
% TWLsolver calls for satisfiability: 170
% 
% Memory used [KB]: 11769
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (60475)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 112
% Passive clauses: 269
% Generated clauses: 15360
% Final active clauses: 105
% Final passive clauses: 149
% Input formulas: 27
% Initial clauses: 57
% 
% Duplicate literals: 832
% Trivial inequalities: 204
% Fw subsumption resolutions: 1982
% Fw demodulations: 1827
% Bw demodulations: 11
% Condensations: 1905
% Global subsumptions: 4
% 
% Simple tautologies: 296
% Equational tautologies: 28
% Forward subsumptions: 7497
% 
% Binary resolution: 34
% Forward superposition: 4549
% Backward superposition: 3538
% Self superposition: 148
% Equality factoring: 73
% Equality resolution: 57
% 
% Split clauses: 43
% Split components: 86
% SAT solver clauses: 580
% SAT solver unit clauses: 67
% SAT solver binary clauses: 31
% SAT solver learnt clauses: 8
% 
% Sat splits: 43
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 564
% TWLsolver calls for satisfiability: 256
% 
% Memory used [KB]: 8187
% Time elapsed: 0.719 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f7914,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f7913])).
fof(f7913,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',sK0),s(fun('type.num.num',bool),sK1))))),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',sK0),s(fun('type.num.num',bool),sK1))))),s('type.num.num',X0)))) )),
  inference(backward_demodulation,[],[f7911,f658])).
fof(f658,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',sK0),s(fun('type.num.num',bool),sK1))))),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',sK0))))))),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f99,f139])).
fof(f139,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : ((s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8,X9,X10] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X10))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))) <=> s('type.num.num',X8) = s('type.num.num',X9))),
  file('thm.pred_set.SUM_SET_THM|split|1','thm.arithmetic.EQ_MONO_ADD_EQ')).
fof(f99,plain,(
  s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',sK0),s(fun('type.num.num',bool),sK1))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',sK0)))))))),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),sK1)))) & s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',sK0),s(fun('type.num.num',bool),sK1))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',sK0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f63])).
fof(f63,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X1)))) & s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),X1))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X1),s('type.num.num',X0))))))))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X1)))) => s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X1),s('type.num.num',X0))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X7)))) => s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X6),s(fun('type.num.num',bool),X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X7),s('type.num.num',X6))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X7)))) => s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X6),s(fun('type.num.num',bool),X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X7),s('type.num.num',X6))))))))),
  file('thm.pred_set.SUM_SET_THM|split|1','thm.pred_set.SUM_SET_THM|split|1')).
fof(f7911,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),sK1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f7910,f989])).
fof(f989,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),X0))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun('type.num.num','type.num.num'),'const.combin.I_0'),s(fun('type.num.num',bool),X0)))) )),
  inference(forward_demodulation,[],[f979,f104])).
fof(f104,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),X0))) = s('type.num.num',happ(s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.SUM_SET_0'),s(fun('type.num.num',bool),X0)))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),X0))) = s('type.num.num',happ(s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.SUM_SET_0'),s(fun('type.num.num',bool),X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X21] : s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),X21))) = s('type.num.num',happ(s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.SUM_SET_0'),s(fun('type.num.num',bool),X21)))),
  file('thm.pred_set.SUM_SET_THM|split|1','ARITY_0_1_const.pred_set.SUM_SET')).
fof(f979,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun('type.num.num','type.num.num'),'const.combin.I_0'),s(fun('type.num.num',bool),X0))) = s('type.num.num',happ(s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.SUM_SET_0'),s(fun('type.num.num',bool),X0)))) )),
  inference(superposition,[],[f134,f102])).
fof(f102,plain,(
  s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.SUM_SET_0') = s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.SUM_IMAGE_1'(s(fun('type.num.num','type.num.num'),'const.combin.I_0')))),
  inference(cnf_transformation,[],[f25])).
fof(f25,axiom,(
  s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.SUM_SET_0') = s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.SUM_IMAGE_1'(s(fun('type.num.num','type.num.num'),'const.combin.I_0')))),
  file('thm.pred_set.SUM_SET_THM|split|1','thm.pred_set.SUM_SET_DEF')).
fof(f134,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))) = s('type.num.num',happ(s(fun(fun(X0,bool),'type.num.num'),'const.pred_set.SUM_IMAGE_1'(s(fun(X0,'type.num.num'),X1))),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))) = s('type.num.num',happ(s(fun(fun(X0,bool),'type.num.num'),'const.pred_set.SUM_IMAGE_1'(s(fun(X0,'type.num.num'),X1))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))) = s('type.num.num',happ(s(fun(fun(X0,bool),'type.num.num'),'const.pred_set.SUM_IMAGE_1'(s(fun(X0,'type.num.num'),X1))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X12] : ! [X21,X22] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X12,'type.num.num'),X21),s(fun(X12,bool),X22))) = s('type.num.num',happ(s(fun(fun(X12,bool),'type.num.num'),'const.pred_set.SUM_IMAGE_1'(s(fun(X12,'type.num.num'),X21))),s(fun(X12,bool),X22)))),
  file('thm.pred_set.SUM_SET_THM|split|1','ARITY_1_2_const.pred_set.SUM_IMAGE')).
fof(f7910,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun('type.num.num','type.num.num'),'const.combin.I_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),sK1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.SUM_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f7873,f989])).
fof(f7873,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun('type.num.num','type.num.num'),'const.combin.I_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),sK1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun('type.num.num','type.num.num'),'const.combin.I_0'),s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f1032,f172])).
fof(f172,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f131,f129])).
fof(f129,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : ! [X6] : s(X12,X6) = s(X12,'const.combin.I_1'(s(X12,X6)))),
  file('thm.pred_set.SUM_SET_THM|split|1','thm.combin.I_THM')).
fof(f131,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X12] : ! [X21] : s(X12,'const.combin.I_1'(s(X12,X21))) = s(X12,happ(s(fun(X12,X12),'const.combin.I_0'),s(X12,X21)))),
  file('thm.pred_set.SUM_SET_THM|split|1','ARITY_0_1_const.combin.I')).
fof(f1032,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun('type.num.num','type.num.num'),X0),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X1),s(fun('type.num.num',bool),sK1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun('type.num.num','type.num.num'),X0),s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',X1)))))))) )),
  inference(subsumption_resolution,[],[f1024,f100])).
fof(f100,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SUM_SET_THM|split|1',aHL_TRUTH)).
fof(f1024,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun('type.num.num','type.num.num'),X0),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X1),s(fun('type.num.num',bool),sK1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun('type.num.num','type.num.num'),X0),s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f137,f180])).
fof(f180,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),sK1)))),
  inference(unit_resulting_resolution,[],[f98,f107])).
fof(f107,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.pred_set.SUM_SET_THM|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f98,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),sK1))))),
  inference(cnf_transformation,[],[f75])).
fof(f137,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3)))) | s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2)))))))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3)))) | s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))))))))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3)))) => s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3)))) => s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X12] : ! [X19,X20,X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X12,bool),X7)))) => s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X12,'type.num.num'),X19),s(fun(X12,bool),'const.pred_set.INSERT_2'(s(X12,X20),s(fun(X12,bool),X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',happ(s(fun(X12,'type.num.num'),X19),s(X12,X20))),s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X12,'type.num.num'),X19),s(fun(X12,bool),'const.pred_set.DELETE_2'(s(fun(X12,bool),X7),s(X12,X20))))))))),
  file('thm.pred_set.SUM_SET_THM|split|1','thm.pred_set.SUM_IMAGE_THM|split|1')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 125
% Passive clauses: 419
% Generated clauses: 7950
% Final active clauses: 67
% Final passive clauses: 119
% Input formulas: 27
% Initial clauses: 57
% 
% Duplicate literals: 437
% Trivial inequalities: 138
% Fw subsumption resolutions: 548
% Fw demodulations: 970
% Bw demodulations: 206
% Condensations: 544
% Global subsumptions: 21
% 
% Simple tautologies: 193
% Equational tautologies: 164
% Forward subsumptions: 4379
% Fw demodulations to eq. taut.: 19
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 713
% Forward superposition: 2395
% Backward superposition: 1464
% Self superposition: 83
% Equality factoring: 52
% Equality resolution: 18
% 
% Split clauses: 34
% Split components: 76
% SAT solver clauses: 559
% SAT solver unit clauses: 28
% SAT solver binary clauses: 142
% 
% Sat splits: 34
% Sat splitting refutations: 49
% 
% Memory used [KB]: 3326
% Time elapsed: 0.179 s
% ------------------------------
% ------------------------------
% Success in time 1.527 s
