Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% (74571)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 155
% Passive clauses: 3829
% Generated clauses: 19475
% Final active clauses: 147
% Final passive clauses: 3542
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 543
% Trivial inequalities: 312
% Fw subsumption resolutions: 895
% Fw demodulations: 1124
% Bw demodulations: 4
% 
% Simple tautologies: 326
% Equational tautologies: 22
% Forward subsumptions: 11890
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 31
% Forward superposition: 9009
% Backward superposition: 7081
% Self superposition: 208
% Equality factoring: 54
% Equality resolution: 27
% 
% Split clauses: 219
% Split components: 438
% SAT solver clauses: 432
% SAT solver unit clauses: 5
% SAT solver binary clauses: 100
% 
% Sat splits: 219
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 432
% TWLsolver calls for satisfiability: 54
% 
% Memory used [KB]: 20212
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numRing
% (74597)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 126
% Passive clauses: 1303
% Generated clauses: 15637
% Final active clauses: 114
% Final passive clauses: 1033
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 579
% Trivial inequalities: 198
% Fw subsumption resolutions: 555
% Fw demodulations: 1018
% Bw demodulations: 121
% Condensations: 889
% 
% Simple tautologies: 175
% Equational tautologies: 25
% Forward subsumptions: 8825
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 26
% 
% Binary resolution: 43
% Forward superposition: 5928
% Backward superposition: 4599
% Self superposition: 211
% Equality factoring: 74
% Equality resolution: 18
% 
% Split clauses: 101
% Split components: 202
% SAT solver clauses: 2025
% SAT solver unit clauses: 235
% SAT solver binary clauses: 49
% 
% Sat splits: 101
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 2025
% TWLsolver calls for satisfiability: 260
% 
% Memory used [KB]: 22387
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f121072,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f120996])).
fof(f120996,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f307,f120941])).
fof(f120941,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f120862,f372])).
fof(f372,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,X2) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X7] : ! [X33,X42,X43,X44] : s(X7,X42) = s(X7,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X7),'const.semi_ring.recordtype|46|semi_ring_4'(s(X7,X33),s(X7,X42),s(fun(X7,fun(X7,X7)),X43),s(fun(X7,fun(X7,X7)),X44)))))),
  file('thm.numRing.num_ring_thms|split|40','thm.semi_ring.semi_ring_SR|49|')).
fof(f120862,plain,(
  ( ! [X2] : (s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))) = s('type.num.num','const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))) )),
  inference(superposition,[],[f25203,f338])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X7] : ! [X12,X6] : s(X7,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X7),X12))) = s(X7,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X7),X12),s('type.quote.varmap'(X7),X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  file('thm.numRing.num_ring_thms|split|40','thm.canonical.interp_vl_def|split|0')).
fof(f25203,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.list.list'('type.quote.index'),X1))) = s('type.num.num','const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.quote.varmap'('type.num.num'),X0),s('type.list.list'('type.quote.index'),X1)))) )),
  inference(forward_demodulation,[],[f25169,f417])).
fof(f417,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.list.list'('type.quote.index'),X1))) = s('type.num.num',happ(s(fun('type.list.list'('type.quote.index'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.list.list'('type.quote.index'),'type.num.num')),'const.numRing.num_interp_vl_0'),s('type.quote.varmap'('type.num.num'),X0))),s('type.list.list'('type.quote.index'),X1)))) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.list.list'('type.quote.index'),X1))) = s('type.num.num',happ(s(fun('type.list.list'('type.quote.index'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.list.list'('type.quote.index'),'type.num.num')),'const.numRing.num_interp_vl_0'),s('type.quote.varmap'('type.num.num'),X0))),s('type.list.list'('type.quote.index'),X1)))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X45,X46] : s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),X45),s('type.list.list'('type.quote.index'),X46))) = s('type.num.num',happ(s(fun('type.list.list'('type.quote.index'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.list.list'('type.quote.index'),'type.num.num')),'const.numRing.num_interp_vl_0'),s('type.quote.varmap'('type.num.num'),X45))),s('type.list.list'('type.quote.index'),X46)))),
  file('thm.numRing.num_ring_thms|split|40','ARITY_0_2_const.numRing.num_interp_vl')).
fof(f25169,plain,(
  ( ! [X0,X1] : (s('type.num.num',happ(s(fun('type.list.list'('type.quote.index'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.list.list'('type.quote.index'),'type.num.num')),'const.numRing.num_interp_vl_0'),s('type.quote.varmap'('type.num.num'),X0))),s('type.list.list'('type.quote.index'),X1))) = s('type.num.num','const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.quote.varmap'('type.num.num'),X0),s('type.list.list'('type.quote.index'),X1)))) )),
  inference(superposition,[],[f361,f316])).
fof(f316,plain,(
  s(fun('type.quote.varmap'('type.num.num'),fun('type.list.list'('type.quote.index'),'type.num.num')),'const.numRing.num_interp_vl_0') = s(fun('type.quote.varmap'('type.num.num'),fun('type.list.list'('type.quote.index'),'type.num.num')),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  inference(cnf_transformation,[],[f65])).
fof(f65,axiom,(
  s(fun('type.quote.varmap'('type.num.num'),fun('type.list.list'('type.quote.index'),'type.num.num')),'const.numRing.num_interp_vl_0') = s(fun('type.quote.varmap'('type.num.num'),fun('type.list.list'('type.quote.index'),'type.num.num')),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  file('thm.numRing.num_ring_thms|split|40','thm.numRing.num_interp_vl_def')).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.list.list'('type.quote.index'),X3))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.list.list'('type.quote.index'),X0)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.list.list'('type.quote.index'),X3)))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X7] : ! [X45,X46,X47] : s(X7,'const.canonical.interp_vl_3'(s('type.semi_ring.semi_ring'(X7),X45),s('type.quote.varmap'(X7),X46),s('type.list.list'('type.quote.index'),X47))) = s(X7,happ(s(fun('type.list.list'('type.quote.index'),X7),happ(s(fun('type.quote.varmap'(X7),fun('type.list.list'('type.quote.index'),X7)),'const.canonical.interp_vl_1'(s('type.semi_ring.semi_ring'(X7),X45))),s('type.quote.varmap'(X7),X46))),s('type.list.list'('type.quote.index'),X47)))),
  file('thm.numRing.num_ring_thms|split|40','ARITY_1_3_const.canonical.interp_vl')).
fof(f307,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),sK0),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),sK0),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f298])).
fof(f298,plain,(
  ? [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ~! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.numRing.num_interp_vl_2'(s('type.quote.varmap'('type.num.num'),X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  file('thm.numRing.num_ring_thms|split|40','thm.numRing.num_ring_thms|split|40')).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 540
% Passive clauses: 9575
% Generated clauses: 124094
% Final active clauses: 223
% Final passive clauses: 1898
% Input formulas: 118
% Initial clauses: 123
% Discarded non-redundant clauses: 3367
% 
% Duplicate literals: 3472
% Trivial inequalities: 1176
% Fw subsumption resolutions: 3148
% Fw demodulations: 13979
% Bw demodulations: 1815
% Condensations: 4773
% Global subsumptions: 207
% 
% Simple tautologies: 599
% Equational tautologies: 265
% Forward subsumptions: 85819
% Fw demodulations to eq. taut.: 23
% Bw demodulations to eq. taut.: 116
% 
% Unit resulting resolution: 2734
% Forward superposition: 46065
% Backward superposition: 41257
% Self superposition: 721
% Equality factoring: 216
% Equality resolution: 48
% 
% Split clauses: 338
% Split components: 681
% SAT solver clauses: 10694
% SAT solver unit clauses: 1209
% SAT solver binary clauses: 2740
% 
% Sat splits: 338
% Sat splitting refutations: 234
% 
% Memory used [KB]: 157225
% Time elapsed: 7.011 s
% ------------------------------
% ------------------------------
% Success in time 8.356 s
