Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (90704)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 475
% Passive clauses: 5487
% Generated clauses: 15104
% Final active clauses: 471
% Final passive clauses: 4732
% Input formulas: 68
% Initial clauses: 173
% 
% Duplicate literals: 205
% Trivial inequalities: 22
% Fw subsumption resolutions: 788
% Fw demodulations: 882
% Bw demodulations: 3
% 
% Simple tautologies: 173
% Equational tautologies: 31
% Forward subsumptions: 7740
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 4712
% Factoring: 977
% Forward superposition: 2790
% Backward superposition: 4503
% Self superposition: 95
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 28
% SAT solver unit clauses: 4
% SAT solver binary clauses: 18
% 
% Sat splits: 23
% 
% TWLsolver clauses: 28
% TWLsolver calls for satisfiability: 17
% 
% Memory used [KB]: 20212
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f8496,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f8460,f6460,f6462,f6461,f290,f467,f327,f482,f5987,f6170,f265,f445,f331,f486,f8229,f266,f447,f452,f454,f6483,f1935,f6484,f347,f498,f433,f570,f1948])).
fof(f1948,plain,(
  $false | ($spl2 | $spl566)),
  inference(resolution,[],[f1935,f447])).
fof(f570,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) ) | $spl224),
  inference(cnf_transformation,[],[f570_D])).
fof(f570_D,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X8,X9,X33,X10] : ! [X11,X12,X34,X13] : s(X8,happ(s(fun(X10,X8),'_dst_x0x0x1_3'(s(fun(X9,fun(X33,X8)),X11),s(X9,X12),s(fun(X10,X33),X34))),s(X10,X13))) = s(X8,happ(s(fun(X33,X8),happ(s(fun(X9,fun(X33,X8)),X11),s(X9,X12))),s(X33,happ(s(fun(X10,X33),X34),s(X10,X13)))))),
  file('thm.extreal.extreal_distinct|split|1','_dst_x0x0x1_ax')).
fof(f498,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X1))) ) | $spl80),
  inference(cnf_transformation,[],[f498_D])).
fof(f498_D,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X1))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f347,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X38,X39] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X38))),s(bool,X39)))) <=> (p(s(bool,X38)) => p(s(bool,X39))))),
  file('thm.extreal.extreal_distinct|split|1','==>_HL_0')).
fof(f6484,plain,(
  ( ! [X327,X326] : (p(s(bool,happ(s(fun('type.realax.real',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X327),s(fun('type.realax.real',bool),X326))),s('type.realax.real',sK9))))) ) | $spl2804),
  inference(cnf_transformation,[],[f6484_D])).
fof(f6484_D,plain,(
  ( ! [X327,X326] : (p(s(bool,happ(s(fun('type.realax.real',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X327),s(fun('type.realax.real',bool),X326))),s('type.realax.real',sK9))))) ) <=> ~$spl2804),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2804])])).
fof(f1935,plain,(
  ( ! [X5] : (~p(s(bool,X5))) ) | $spl566),
  inference(cnf_transformation,[],[f1935_D])).
fof(f1935_D,plain,(
  ( ! [X5] : (~p(s(bool,X5))) ) <=> ~$spl566),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl566])])).
fof(f6483,plain,(
  ( ! [X327,X324,X326] : (~p(s(bool,X324)) | p(s(bool,happ(s(fun('type.realax.real',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X327),s(fun('type.realax.real',bool),X326))),s('type.realax.real',sK9))))) ) | ($spl48 | $spl80 | $spl224 | $spl2650)),
  inference(forward_demodulation,[],[f6482,f482])).
fof(f6482,plain,(
  ( ! [X325,X327,X324,X326] : (p(s(bool,happ(s(fun('type.realax.real',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X327),s(fun('type.realax.real',bool),X326))),s('type.realax.real',sK9)))) | ~p(s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X324),s(bool,X325),s(fun('type.realax.real',bool),X326))))) ) | ($spl80 | $spl224 | $spl2650)),
  inference(forward_demodulation,[],[f6371,f570])).
fof(f6371,plain,(
  ( ! [X325,X327,X324,X326] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X327))),s(bool,happ(s(fun('type.realax.real',bool),X326),s('type.realax.real',sK9)))))) | ~p(s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X324),s(bool,X325),s(fun('type.realax.real',bool),X326))))) ) | ($spl80 | $spl2650)),
  inference(superposition,[],[f498,f6170])).
fof(f454,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f454_D])).
fof(f454_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f452,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f271,f266])).
fof(f271,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f32])).
fof(f32,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.extreal.extreal_distinct|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f447,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f447_D])).
fof(f447_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f266,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.extreal_distinct|split|1',aHL_TRUTH)).
fof(f8229,plain,(
  $false | ($spl2 | $spl9 | $spl2784)),
  inference(subsumption_resolution,[],[f7722,f447])).
fof(f7722,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl2784)),
  inference(backward_demodulation,[],[f6462,f454])).
fof(f486,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) ) | $spl56),
  inference(cnf_transformation,[],[f486_D])).
fof(f486_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f331,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X16] : ! [X6,X29,X30,X20] : s(X16,happ(s(fun('type.realax.real',X16),X20),s('type.realax.real',X6))) = s(X16,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))),s(X16,X29),s(X16,X30),s(fun('type.realax.real',X16),X20)))),
  file('thm.extreal.extreal_distinct|split|1','thm.extreal.extreal_case_def|split|2')).
fof(f445,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK9))) | $spl0),
  inference(cnf_transformation,[],[f445_D])).
fof(f445_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK9))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f265,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK9)))),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK9)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f169])).
fof(f169,plain,(
  ? [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),
  file('thm.extreal.extreal_distinct|split|1','thm.extreal.extreal_distinct|split|1')).
fof(f6170,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X1),s('type.realax.real',sK9))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X1)))) ) | $spl2650),
  inference(cnf_transformation,[],[f6170_D])).
fof(f6170_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X1),s('type.realax.real',sK9))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X1)))) ) <=> ~$spl2650),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2650])])).
fof(f5987,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X1),s('type.realax.real',sK9))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X1)))) ) | ($spl0 | $spl56)),
  inference(superposition,[],[f486,f445])).
fof(f482,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) ) | $spl48),
  inference(cnf_transformation,[],[f482_D])).
fof(f482_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f327,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X16] : ! [X29,X30,X20] : s(X16,X29) = s(X16,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X16,X29),s(X16,X30),s(fun('type.realax.real',X16),X20)))),
  file('thm.extreal.extreal_distinct|split|1','thm.extreal.extreal_case_def|split|0')).
fof(f467,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(cnf_transformation,[],[f467_D])).
fof(f467_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f290,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X18] : (s(bool,'T_0') = s(bool,X18) <=> p(s(bool,X18)))),
  file('thm.extreal.extreal_distinct|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f6461,plain,(
  ( ! [X278] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X278),s('type.realax.real',sK9))))) ) | $spl2782),
  inference(cnf_transformation,[],[f6461_D])).
fof(f6461_D,plain,(
  ( ! [X278] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X278),s('type.realax.real',sK9))))) ) <=> ~$spl2782),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2782])])).
fof(f6462,plain,(
  ( ! [X276] : (s(bool,'T_0') = s(bool,X276)) ) | $spl2784),
  inference(cnf_transformation,[],[f6462_D])).
fof(f6462_D,plain,(
  ( ! [X276] : (s(bool,'T_0') = s(bool,X276)) ) <=> ~$spl2784),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2784])])).
fof(f6460,plain,(
  ( ! [X276,X278] : (s(bool,'T_0') = s(bool,X276) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X278),s('type.realax.real',sK9))))) ) | ($spl26 | $spl48 | $spl2650)),
  inference(forward_demodulation,[],[f6358,f482])).
fof(f6358,plain,(
  ( ! [X277,X276,X278] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X278),s('type.realax.real',sK9)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X276),s(bool,X277),s(fun('type.realax.real',bool),X278)))) ) | ($spl26 | $spl2650)),
  inference(superposition,[],[f467,f6170])).
fof(f8460,plain,(
  $false | ($spl2782 | $spl2804)),
  inference(subsumption_resolution,[],[f6484,f6461])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 226
% Passive clauses: 1886
% Generated clauses: 9119
% Final active clauses: 195
% Final passive clauses: 733
% Input formulas: 68
% Initial clauses: 173
% 
% Duplicate literals: 96
% Trivial inequalities: 33
% Fw subsumption resolutions: 694
% Fw demodulations: 373
% Bw demodulations: 946
% Condensations: 220
% Global subsumptions: 8
% 
% Simple tautologies: 143
% Equational tautologies: 17
% Forward subsumptions: 3177
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 8
% 
% Binary resolution: 1525
% Factoring: 178
% Forward superposition: 1269
% Backward superposition: 1119
% Self superposition: 36
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 54
% Split components: 108
% SAT solver clauses: 1906
% SAT solver unit clauses: 202
% SAT solver binary clauses: 69
% SAT solver learnt clauses: 17
% 
% Sat splits: 54
% Sat splitting refutations: 22
% 
% TWLsolver clauses: 1870
% TWLsolver calls for satisfiability: 1340
% 
% Memory used [KB]: 8955
% Time elapsed: 0.440 s
% ------------------------------
% ------------------------------
% Success in time 1.06 s
