Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.relation
% (67983)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 173
% Passive clauses: 2907
% Generated clauses: 25132
% Final active clauses: 109
% Final passive clauses: 2501
% Input formulas: 24
% Initial clauses: 57
% 
% Duplicate literals: 811
% Trivial inequalities: 520
% Fw subsumption resolutions: 3370
% Fw demodulations: 3818
% Bw demodulations: 50
% 
% Simple tautologies: 486
% Equational tautologies: 39
% Forward subsumptions: 12071
% Bw demodulations to eq. taut.: 9
% 
% Binary resolution: 437
% Forward superposition: 9318
% Backward superposition: 6265
% Self superposition: 192
% Equality factoring: 89
% Equality resolution: 38
% 
% Split clauses: 722
% Split components: 1474
% SAT solver clauses: 1237
% SAT solver unit clauses: 18
% SAT solver binary clauses: 152
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 722
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 1233
% TWLsolver calls for satisfiability: 240
% 
% Memory used [KB]: 13560
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (68018)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 125
% Passive clauses: 781
% Generated clauses: 16400
% Final active clauses: 109
% Final passive clauses: 625
% Input formulas: 24
% Initial clauses: 57
% 
% Duplicate literals: 1116
% Trivial inequalities: 179
% Fw subsumption resolutions: 1316
% Fw demodulations: 1856
% Bw demodulations: 38
% Condensations: 1597
% Global subsumptions: 4
% 
% Simple tautologies: 294
% Equational tautologies: 49
% Forward subsumptions: 7220
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 177
% Forward superposition: 5112
% Backward superposition: 5915
% Self superposition: 211
% Equality factoring: 142
% Equality resolution: 23
% 
% Split clauses: 299
% Split components: 606
% SAT solver clauses: 1400
% SAT solver unit clauses: 58
% SAT solver binary clauses: 36
% SAT solver learnt clauses: 9
% SAT solver learnt literals: 1
% 
% Sat splits: 299
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1383
% TWLsolver calls for satisfiability: 299
% 
% Memory used [KB]: 14456
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 287 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f2575,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f102,f178,f176,f101,f177,f175,f749,f750,f757,f754,f751,f752,f760,f753,f762,f755,f759,f783,f797,f795,f811,f812,f820,f813,f814,f823,f815,f1004,f1005,f1006,f1007,f1008,f1009,f824,f1015,f1016,f1017,f1027,f822,f1018,f819,f1020,f825,f818,f1428,f1430,f1432,f1445,f1447,f1450,f1265,f1452,f1264,f1453,f1455,f1262,f1456,f1312,f1459,f1335,f1451,f1333,f1318,f1454,f1319,f1329,f761,f1457,f1796,f1799,f1840,f1842,f1849,f1854,f2572,f2574])).
fof(f2574,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2573,f179])).
fof(f179,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(X0,X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f137,f136])).
fof(f136,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : s(bool,'T_0') = s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ! [X2,X3] : s(bool,'T_0') = s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6,X7] : ! [X10,X11] : s(bool,'T_0') = s(bool,'const.relation.RUNIV_2'(s(X6,X10),s(X7,X11)))),
  file('thm.relation.RUNIV_SUBSET|split|0','thm.relation.RUNIV')).
fof(f137,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ! [X2,X3] : s(bool,'const.relation.RUNIV_2'(s(X0,X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : ! [X21,X22] : s(bool,'const.relation.RUNIV_2'(s(X6,X21),s(X7,X22))) = s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),'const.relation.RUNIV_0'),s(X6,X21))),s(X7,X22)))),
  file('thm.relation.RUNIV_SUBSET|split|0','ARITY_0_2_const.relation.RUNIV')).
fof(f2573,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool))))),s(sK1,sK4(happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),sK1,bool)))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f2565,f2360])).
fof(f2360,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,X0))),s(sK1,X1)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2299,f108])).
fof(f108,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.relation.RUNIV_SUBSET|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f2299,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,X0))),s(sK1,X1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2298,f103])).
fof(f103,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.RUNIV_SUBSET|split|0',aHL_TRUTH)).
fof(f2298,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,X0))),s(sK1,X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f2297,f179])).
fof(f2297,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,X0))),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,X0))),s(sK1,X1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2263,f103])).
fof(f2263,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,X0))),s(sK1,X1)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,X0))),s(sK1,X1))))) ) | $spl0),
  inference(superposition,[],[f141,f1653])).
fof(f1653,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),sK2))) | $spl0),
  inference(unit_resulting_resolution,[],[f175,f108])).
fof(f141,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X6))),s(X1,X7))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X6))),s(X1,X7)))))) & ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK5(X3,X2,X1,X0)))),s(X1,sK6(X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,sK5(X3,X2,X1,X0)))),s(X1,sK6(X3,X2,X1,X0)))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X6))),s(X1,X7)))))) & (? [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))))),
  inference(rectify,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ! [X4,X5] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))))) & (? [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) <=> ! [X4,X5] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) <=> ! [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) <=> ! [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6,X7] : ! [X19,X20] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X7,bool)),X19),s(fun(X6,fun(X7,bool)),X20)))) <=> ! [X10,X11] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X19),s(X6,X10))),s(X7,X11)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X20),s(X6,X10))),s(X7,X11))))))),
  file('thm.relation.RUNIV_SUBSET|split|0','thm.relation.RSUBSET')).
fof(f2565,plain,(
  s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool))))),s(sK1,sK4(happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),sK1,bool)))) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool))))),s(sK1,sK4(happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),sK1,bool)))) | $spl3),
  inference(unit_resulting_resolution,[],[f1552,f140])).
fof(f140,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK4(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK4(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK4(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK4(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7,X6] : ! [X12,X13] : (s(fun(X6,X7),X12) = s(fun(X6,X7),X13) <=> ! [X10] : s(X7,happ(s(fun(X6,X7),X12),s(X6,X10))) = s(X7,happ(s(fun(X6,X7),X13),s(X6,X10))))),
  file('thm.relation.RUNIV_SUBSET|split|0','thm.bool.FUN_EQ_THM')).
fof(f1552,plain,(
  s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool))))) != s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool))))) | $spl3),
  inference(unit_resulting_resolution,[],[f178,f140])).
fof(f2572,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2571,f179])).
fof(f2571,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool))))),s(sK1,sK4(happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),sK1,bool)))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f2568,f2360])).
fof(f2568,plain,(
  s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool))))),s(sK1,sK4(happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),sK1,bool)))) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool))))),s(sK1,sK4(happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK4(sK2,'const.relation.RUNIV_0',sK0,fun(sK1,bool)))),sK1,bool)))) | $spl3),
  inference(unit_resulting_resolution,[],[f1552,f140])).
fof(f1854,plain,(
  $false | ($spl20 | $spl24)),
  inference(global_subsumption,[],[f102,f101,f103,f105,f106,f108,f107,f112,f111,f136,f179,f138,f140,f139,f143,f142,f141,f169,f181,f218,f285,f339,f340,f346,f347,f506,f590,f604,f695,f730,f732,f736,f738,f864,f866,f868,f873,f875,f878,f880,f882,f971,f973,f977,f979,f822,f818,f1082,f1084,f1088,f1090,f1105,f1107,f1110,f1208,f1472,f734,f740,f870,f877,f884,f975,f981,f1086,f1092,f1679,f1785,f1841,f1848,f1853])).
fof(f1853,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1757,f1785])).
fof(f1757,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | $spl20),
  inference(equality_factoring,[],[f818])).
fof(f1848,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1847,f1785])).
fof(f1847,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | p(s(bool,X1))) ) | $spl20),
  inference(subsumption_resolution,[],[f1695,f107])).
fof(f1695,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl20),
  inference(superposition,[],[f107,f818])).
fof(f1841,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1801,f1785])).
fof(f1801,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl20 | $spl24)),
  inference(backward_demodulation,[],[f1785,f111])).
fof(f1785,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl20 | $spl24)),
  inference(duplicate_literal_removal,[],[f1686])).
fof(f1686,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl20 | $spl24)),
  inference(superposition,[],[f818,f822])).
fof(f1679,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl20),
  inference(superposition,[],[f818,f218])).
fof(f1092,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X21,fun(bool,bool)),X19),s(X21,sK5(X19,X20,bool,X21)))),s(bool,'F_0')))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X21,fun(bool,bool)),X20),s(fun(X21,fun(bool,bool)),X19)))) | s(bool,'T_0') = s(bool,sK6(X19,X20,bool,X21))) )),
  inference(superposition,[],[f143,f106])).
fof(f1086,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,happ(s(fun(X21,bool),happ(s(fun(bool,fun(X21,bool)),X19),s(bool,'F_0'))),s(X21,sK6(X19,X20,X21,bool))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(bool,fun(X21,bool)),X20),s(fun(bool,fun(X21,bool)),X19)))) | s(bool,'T_0') = s(bool,sK5(X19,X20,X21,bool))) )),
  inference(superposition,[],[f143,f106])).
fof(f981,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X21,fun(bool,bool)),X20),s(X21,sK5(X19,X20,bool,X21)))),s(bool,'F_0')))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X21,fun(bool,bool)),X20),s(fun(X21,fun(bool,bool)),X19)))) | s(bool,'T_0') = s(bool,sK6(X19,X20,bool,X21))) )),
  inference(superposition,[],[f142,f106])).
fof(f975,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,happ(s(fun(X21,bool),happ(s(fun(bool,fun(X21,bool)),X20),s(bool,'F_0'))),s(X21,sK6(X19,X20,X21,bool))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(bool,fun(X21,bool)),X20),s(fun(bool,fun(X21,bool)),X19)))) | s(bool,'T_0') = s(bool,sK5(X19,X20,X21,bool))) )),
  inference(superposition,[],[f142,f106])).
fof(f884,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') != s(bool,happ(s(fun(X23,bool),X25),s(X23,sK4(X24,X25,X23,bool)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'T_0') = s(bool,happ(s(fun(X23,bool),X24),s(X23,sK4(X24,X25,X23,bool))))) )),
  inference(superposition,[],[f140,f106])).
fof(f877,plain,(
  ( ! [X21,X19,X20] : (s(X21,happ(s(fun(bool,X21),X19),s(bool,'F_0'))) != s(X21,happ(s(fun(bool,X21),X20),s(bool,'F_0'))) | s(fun(bool,X21),X19) = s(fun(bool,X21),X20) | s(bool,'T_0') = s(bool,sK4(X19,X20,bool,X21))) )),
  inference(superposition,[],[f140,f106])).
fof(f870,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') != s(bool,happ(s(fun(X23,bool),X25),s(X23,sK4(X25,X24,X23,bool)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'T_0') = s(bool,happ(s(fun(X23,bool),X24),s(X23,sK4(X25,X24,X23,bool))))) )),
  inference(superposition,[],[f140,f106])).
fof(f740,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,f(s(X19,sK3(X20,X21,bool,X19)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,g(s(X19,sK3(X20,X21,bool,X19))))) )),
  inference(superposition,[],[f138,f106])).
fof(f734,plain,(
  ( ! [X21,X19,X20] : (s(X21,f(s(bool,'F_0'))) != s(X21,g(s(bool,'F_0'))) | s(fun(bool,X21),X19) = s(fun(bool,X21),X20) | s(bool,'T_0') = s(bool,sK3(X19,X20,X21,bool))) )),
  inference(superposition,[],[f138,f106])).
fof(f1472,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X28,fun(X27,bool)),X29),s(fun(X28,fun(X27,bool)),X30)))) | s(bool,'T_0') = s(bool,happ(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X29),s(X28,sK5(X30,X29,X27,X28)))),s(X27,sK6(X30,X29,X27,X28))))) )),
  inference(subsumption_resolution,[],[f988,f105])).
fof(f988,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X28,fun(X27,bool)),X29),s(fun(X28,fun(X27,bool)),X30)))) | s(bool,'T_0') = s(bool,happ(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X29),s(X28,sK5(X30,X29,X27,X28)))),s(X27,sK6(X30,X29,X27,X28))))) )),
  inference(superposition,[],[f142,f106])).
fof(f1208,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X13),s(X11,X15))),s(X12,X16)))) | p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X14),s(X11,X15))),s(X12,X16)))) | s(bool,X17) = s(bool,'const.relation.RSUBSET_2'(s(fun(X11,fun(X12,bool)),X13),s(fun(X11,fun(X12,bool)),X14))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f1198,f103])).
fof(f1198,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X13),s(X11,X15))),s(X12,X16)))) | p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X14),s(X11,X15))),s(X12,X16)))) | s(bool,X17) = s(bool,'const.relation.RSUBSET_2'(s(fun(X11,fun(X12,bool)),X13),s(fun(X11,fun(X12,bool)),X14))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f141,f218])).
fof(f1110,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(superposition,[],[f822,f218])).
fof(f1107,plain,(
  ( ! [X10,X8,X7,X11,X9] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X8,fun(X7,bool)),X10),s(fun(X8,fun(X7,bool)),X9)))) | s(bool,X11) = s(bool,happ(s(fun(X7,bool),happ(s(fun(X8,fun(X7,bool)),X9),s(X8,sK5(X9,X10,X7,X8)))),s(X7,sK6(X9,X10,X7,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f1095,f103])).
fof(f1095,plain,(
  ( ! [X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X8,fun(X7,bool)),X10),s(fun(X8,fun(X7,bool)),X9)))) | s(bool,X11) = s(bool,happ(s(fun(X7,bool),happ(s(fun(X8,fun(X7,bool)),X9),s(X8,sK5(X9,X10,X7,X8)))),s(X7,sK6(X9,X10,X7,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f143,f218])).
fof(f1105,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,fun(X0,bool)),'const.relation.RUNIV_0'))))) )),
  inference(subsumption_resolution,[],[f1093,f103])).
fof(f1093,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X1,fun(X0,bool)),X2),s(fun(X1,fun(X0,bool)),'const.relation.RUNIV_0'))))) )),
  inference(superposition,[],[f143,f179])).
fof(f1090,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X13,fun(bool,bool)),X11),s(X13,sK5(X11,X12,bool,X13)))),s(bool,X14)))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X13,fun(bool,bool)),X12),s(fun(X13,fun(bool,bool)),X11)))) | s(bool,'T_0') = s(bool,sK6(X11,X12,bool,X13)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f143,f218])).
fof(f1088,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X5,fun(bool,bool)),X3),s(X5,sK5(X3,X4,bool,X5)))),s(bool,'T_0')))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X5,fun(bool,bool)),X4),s(fun(X5,fun(bool,bool)),X3)))) | s(bool,X6) = s(bool,sK6(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f143,f218])).
fof(f1084,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,happ(s(fun(X13,bool),happ(s(fun(bool,fun(X13,bool)),X11),s(bool,X14))),s(X13,sK6(X11,X12,X13,bool))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(bool,fun(X13,bool)),X12),s(fun(bool,fun(X13,bool)),X11)))) | s(bool,'T_0') = s(bool,sK5(X11,X12,X13,bool)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f143,f218])).
fof(f1082,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X5,bool),happ(s(fun(bool,fun(X5,bool)),X3),s(bool,'T_0'))),s(X5,sK6(X3,X4,X5,bool))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(bool,fun(X5,bool)),X4),s(fun(bool,fun(X5,bool)),X3)))) | s(bool,X6) = s(bool,sK5(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f143,f218])).
fof(f979,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X13,fun(bool,bool)),X12),s(X13,sK5(X11,X12,bool,X13)))),s(bool,X14)))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X13,fun(bool,bool)),X12),s(fun(X13,fun(bool,bool)),X11)))) | s(bool,'T_0') = s(bool,sK6(X11,X12,bool,X13)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f142,f218])).
fof(f977,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X5,fun(bool,bool)),X4),s(X5,sK5(X3,X4,bool,X5)))),s(bool,'T_0')))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X5,fun(bool,bool)),X4),s(fun(X5,fun(bool,bool)),X3)))) | s(bool,X6) = s(bool,sK6(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f142,f218])).
fof(f973,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,happ(s(fun(X13,bool),happ(s(fun(bool,fun(X13,bool)),X12),s(bool,X14))),s(X13,sK6(X11,X12,X13,bool))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(bool,fun(X13,bool)),X12),s(fun(bool,fun(X13,bool)),X11)))) | s(bool,'T_0') = s(bool,sK5(X11,X12,X13,bool)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f142,f218])).
fof(f971,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(X5,bool),happ(s(fun(bool,fun(X5,bool)),X4),s(bool,'T_0'))),s(X5,sK6(X3,X4,X5,bool))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(bool,fun(X5,bool)),X4),s(fun(bool,fun(X5,bool)),X3)))) | s(bool,X6) = s(bool,sK5(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f142,f218])).
fof(f882,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK4(X16,X17,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,'T_0') = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK4(X16,X17,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f140,f218])).
fof(f880,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') != s(bool,happ(s(fun(X7,bool),X9),s(X7,sK4(X8,X9,X7,bool)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,X10) = s(bool,happ(s(fun(X7,bool),X8),s(X7,sK4(X8,X9,X7,bool)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f140,f218])).
fof(f878,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X3),s(X0,sK4(happ(s(fun(X1,fun(X0,bool)),'const.relation.RUNIV_0'),s(X1,X2)),X3,X0,bool)))) | s(fun(X0,bool),X3) = s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'const.relation.RUNIV_0'),s(X1,X2)))) )),
  inference(superposition,[],[f140,f179])).
fof(f875,plain,(
  ( ! [X14,X12,X13,X11] : (s(X13,happ(s(fun(bool,X13),X11),s(bool,X14))) != s(X13,happ(s(fun(bool,X13),X12),s(bool,X14))) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool,X13)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f140,f218])).
fof(f873,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,happ(s(fun(bool,X5),X3),s(bool,'T_0'))) != s(X5,happ(s(fun(bool,X5),X4),s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK4(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f140,f218])).
fof(f868,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK4(X17,X16,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,'T_0') = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK4(X17,X16,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f140,f218])).
fof(f866,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') != s(bool,happ(s(fun(X7,bool),X9),s(X7,sK4(X9,X8,X7,bool)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,X10) = s(bool,happ(s(fun(X7,bool),X8),s(X7,sK4(X9,X8,X7,bool)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f140,f218])).
fof(f864,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X3),s(X0,sK4(X3,happ(s(fun(X1,fun(X0,bool)),'const.relation.RUNIV_0'),s(X1,X2)),X0,bool)))) | s(fun(X0,bool),X3) = s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'const.relation.RUNIV_0'),s(X1,X2)))) )),
  inference(superposition,[],[f140,f179])).
fof(f738,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) != s(bool,f(s(X11,sK3(X12,X13,bool,X11)))) | s(fun(X11,bool),X12) = s(fun(X11,bool),X13) | s(bool,'T_0') = s(bool,g(s(X11,sK3(X12,X13,bool,X11)))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f138,f218])).
fof(f736,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK3(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK3(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f138,f218])).
fof(f732,plain,(
  ( ! [X14,X12,X13,X11] : (s(X13,f(s(bool,X14))) != s(X13,g(s(bool,X14))) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(bool,'T_0') = s(bool,sK3(X11,X12,X13,bool)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f138,f218])).
fof(f730,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK3(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f138,f218])).
fof(f695,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f506,f106])).
fof(f604,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) = s(bool,happ(s(fun(bool,bool),happ(s(fun(X25,fun(bool,bool)),'const.relation.RUNIV_0'),s(X25,X26))),s(bool,X24))) | s(bool,X24) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f179,f218])).
fof(f590,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) = s(bool,happ(s(fun(X25,bool),happ(s(fun(bool,fun(X25,bool)),'const.relation.RUNIV_0'),s(bool,X24))),s(X25,X26))) | s(bool,X24) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f179,f218])).
fof(f506,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f351,f103])).
fof(f351,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f111,f218])).
fof(f347,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X22) = s(bool,'const.relation.RUNIV_2'(s(bool,X22),s(X23,X24))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f136,f218])).
fof(f346,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,X18) = s(bool,'const.relation.RUNIV_2'(s(X19,X20),s(bool,X18))) | s(bool,X18) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f136,f218])).
fof(f340,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f103,f218])).
fof(f339,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f107,f218])).
fof(f285,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f218,f218])).
fof(f218,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f106,f106])).
fof(f181,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f105,f107])).
fof(f169,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f168])).
fof(f168,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f157])).
fof(f157,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X10,X17,X11,X18] : ((s(bool,X10) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X11) = s(bool,X18))) => ((p(s(bool,X10)) => p(s(bool,X11))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.relation.RUNIV_SUBSET|split|0','thm.bool.IMP_CONG')).
fof(f142,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK5(X3,X2,X1,X0)))),s(X1,sK6(X3,X2,X1,X0))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f143,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,sK5(X3,X2,X1,X0)))),s(X1,sK6(X3,X2,X1,X0))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f139,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f93])).
fof(f138,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.relation.RUNIV_SUBSET|split|0',aHL_EXT)).
fof(f111,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.relation.RUNIV_SUBSET|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f112,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f74])).
fof(f107,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f72])).
fof(f106,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.RUNIV_SUBSET|split|0',aHL_BOOL_CASES)).
fof(f105,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.RUNIV_SUBSET|split|0',aHL_FALSITY)).
fof(f1849,plain,(
  $false | ($spl20 | $spl24)),
  inference(global_subsumption,[],[f102,f101,f103,f105,f106,f108,f107,f112,f111,f136,f179,f138,f140,f139,f143,f142,f141,f169,f181,f218,f285,f339,f340,f346,f347,f506,f590,f604,f695,f730,f732,f736,f738,f864,f866,f868,f873,f875,f878,f880,f882,f971,f973,f977,f979,f822,f818,f1082,f1084,f1088,f1090,f1105,f1107,f1110,f1208,f1472,f734,f740,f870,f877,f884,f975,f981,f1086,f1092,f1679,f1785,f1841,f1848])).
fof(f1842,plain,(
  $false | ($spl20 | $spl24)),
  inference(global_subsumption,[],[f102,f101,f103,f105,f106,f108,f107,f112,f111,f136,f179,f138,f140,f139,f143,f142,f141,f169,f181,f218,f285,f339,f340,f346,f347,f506,f590,f604,f695,f730,f732,f736,f738,f864,f866,f868,f873,f875,f878,f880,f882,f971,f973,f977,f979,f822,f818,f1082,f1084,f1088,f1090,f1105,f1107,f1110,f1208,f1472,f734,f740,f870,f877,f884,f975,f981,f1086,f1092,f1679,f1785,f1841])).
fof(f1840,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1800,f103])).
fof(f1800,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl24)),
  inference(backward_demodulation,[],[f1785,f105])).
fof(f1799,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1798,f181])).
fof(f1798,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl24)),
  inference(condensation,[],[f1797])).
fof(f1797,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl20 | $spl24)),
  inference(forward_demodulation,[],[f1794,f822])).
fof(f1794,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f1667])).
fof(f1667,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(superposition,[],[f818,f106])).
fof(f1796,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1795,f181])).
fof(f1795,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl24)),
  inference(forward_demodulation,[],[f1666,f822])).
fof(f1666,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | $spl20),
  inference(unit_resulting_resolution,[],[f181,f818])).
fof(f1457,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,X20)) ) | $spl42),
  inference(cnf_transformation,[],[f1457_D])).
fof(f1457_D,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,X20)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f761,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl14),
  inference(cnf_transformation,[],[f761_D])).
fof(f761_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1329,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1213])).
fof(f1213,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f857,f106])).
fof(f857,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(trivial_inequality_removal,[],[f851])).
fof(f851,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f757,f218])).
fof(f1319,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1240])).
fof(f1240,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f106,f857])).
fof(f1454,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl37),
  inference(cnf_transformation,[],[f1454_D])).
fof(f1454_D,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f1318,plain,(
  ( ! [X2] : (~p(s(bool,f(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X2)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1242])).
fof(f1242,plain,(
  ( ! [X2] : (~p(s(bool,f(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl9),
  inference(superposition,[],[f108,f857])).
fof(f1333,plain,(
  ( ! [X3] : (~p(s(bool,f(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X3)) ) | $spl9),
  inference(subsumption_resolution,[],[f1243,f106])).
fof(f1243,plain,(
  ( ! [X3] : (~p(s(bool,f(s(bool,'F_0')))) | s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl9),
  inference(superposition,[],[f111,f857])).
fof(f1451,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) | $spl32),
  inference(cnf_transformation,[],[f1451_D])).
fof(f1451_D,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1335,plain,(
  ( ! [X6,X5] : (~p(s(bool,f(s(bool,'F_0')))) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X5)) ) | $spl9),
  inference(subsumption_resolution,[],[f1245,f340])).
fof(f1245,plain,(
  ( ! [X6,X5] : (~p(s(bool,f(s(bool,'F_0')))) | s(bool,X5) != s(bool,X6) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X5)) ) | $spl9),
  inference(superposition,[],[f169,f857])).
fof(f1459,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'F_0'))))) | $spl45),
  inference(cnf_transformation,[],[f1459_D])).
fof(f1459_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'F_0'))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1312,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X20)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1253])).
fof(f1253,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | $spl9),
  inference(superposition,[],[f845,f857])).
fof(f845,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f757,f106])).
fof(f1456,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun(X24,bool),happ(s(fun(X23,fun(X24,bool)),X25),s(X23,X27))),s(X24,X28)))) | s(bool,'T_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(X23,fun(X24,bool)),X25),s(fun(X23,fun(X24,bool)),X26))) | p(s(bool,happ(s(fun(X24,bool),happ(s(fun(X23,fun(X24,bool)),X26),s(X23,X27))),s(X24,X28))))) ) | $spl40),
  inference(cnf_transformation,[],[f1456_D])).
fof(f1456_D,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun(X24,bool),happ(s(fun(X23,fun(X24,bool)),X25),s(X23,X27))),s(X24,X28)))) | s(bool,'T_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(X23,fun(X24,bool)),X25),s(fun(X23,fun(X24,bool)),X26))) | p(s(bool,happ(s(fun(X24,bool),happ(s(fun(X23,fun(X24,bool)),X26),s(X23,X27))),s(X24,X28))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1262,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (~p(s(bool,f(s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(X24,bool),happ(s(fun(X23,fun(X24,bool)),X25),s(X23,X27))),s(X24,X28)))) | p(s(bool,happ(s(fun(X24,bool),happ(s(fun(X23,fun(X24,bool)),X26),s(X23,X27))),s(X24,X28)))) | s(bool,'T_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(X23,fun(X24,bool)),X25),s(fun(X23,fun(X24,bool)),X26)))) ) | $spl9),
  inference(superposition,[],[f141,f857])).
fof(f1455,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X34,fun(X33,bool)),X35),s(fun(X34,fun(X33,bool)),X36)))) | s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),happ(s(fun(X34,fun(X33,bool)),X35),s(X34,sK5(X36,X35,X33,X34)))),s(X33,sK6(X36,X35,X33,X34))))) ) | $spl38),
  inference(cnf_transformation,[],[f1455_D])).
fof(f1455_D,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X34,fun(X33,bool)),X35),s(fun(X34,fun(X33,bool)),X36)))) | s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),happ(s(fun(X34,fun(X33,bool)),X35),s(X34,sK5(X36,X35,X33,X34)))),s(X33,sK6(X36,X35,X33,X34))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1453,plain,(
  p(s(bool,f(s(bool,'F_0')))) | $spl36),
  inference(cnf_transformation,[],[f1453_D])).
fof(f1453_D,plain,(
  p(s(bool,f(s(bool,'F_0')))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1264,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,f(s(bool,'F_0')))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X34,fun(X33,bool)),X35),s(fun(X34,fun(X33,bool)),X36)))) | s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),happ(s(fun(X34,fun(X33,bool)),X35),s(X34,sK5(X36,X35,X33,X34)))),s(X33,sK6(X36,X35,X33,X34))))) ) | $spl9),
  inference(superposition,[],[f142,f857])).
fof(f1452,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X38,fun(X37,bool)),X40),s(fun(X38,fun(X37,bool)),X39)))) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),happ(s(fun(X38,fun(X37,bool)),X39),s(X38,sK5(X39,X40,X37,X38)))),s(X37,sK6(X39,X40,X37,X38))))) ) | $spl34),
  inference(cnf_transformation,[],[f1452_D])).
fof(f1452_D,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X38,fun(X37,bool)),X40),s(fun(X38,fun(X37,bool)),X39)))) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),happ(s(fun(X38,fun(X37,bool)),X39),s(X38,sK5(X39,X40,X37,X38)))),s(X37,sK6(X39,X40,X37,X38))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1265,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,f(s(bool,'F_0')))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X38,fun(X37,bool)),X40),s(fun(X38,fun(X37,bool)),X39)))) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),happ(s(fun(X38,fun(X37,bool)),X39),s(X38,sK5(X39,X40,X37,X38)))),s(X37,sK6(X39,X40,X37,X38))))) ) | $spl9),
  inference(superposition,[],[f143,f857])).
fof(f1450,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,f(s(bool,'F_0')))) ) | $spl9),
  inference(subsumption_resolution,[],[f1288,f107])).
fof(f1288,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl9),
  inference(superposition,[],[f112,f857])).
fof(f1447,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1427,f105])).
fof(f1427,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f189,f1102,f169])).
fof(f1102,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),X0),s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'))))) ) | $spl2),
  inference(forward_demodulation,[],[f1101,f177])).
fof(f1101,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),X0),s(fun(sK0,fun(sK1,bool)),sK2))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1100,f103])).
fof(f1100,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),X0),s(fun(sK0,fun(sK1,bool)),sK2))))) ) | $spl2),
  inference(forward_demodulation,[],[f1080,f179])).
fof(f1080,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(sK0,sK5(sK2,X0,sK1,sK0)))),s(sK1,sK6(sK2,X0,sK1,sK0))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),X0),s(fun(sK0,fun(sK1,bool)),sK2))))) ) | $spl2),
  inference(superposition,[],[f143,f177])).
fof(f189,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f180,f112])).
fof(f180,plain,(
  ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f176,f177])).
fof(f1445,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1443,f181])).
fof(f1443,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f1431,f189])).
fof(f1431,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),X0),s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1102,f108])).
fof(f1432,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f182,f1102,f108])).
fof(f182,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f180,f107])).
fof(f1430,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f189,f1102,f111])).
fof(f1428,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f105,f189,f1102,f169])).
fof(f818,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f818_D])).
fof(f818_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f825,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl29),
  inference(cnf_transformation,[],[f825_D])).
fof(f825_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f1020,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl18),
  inference(superposition,[],[f506,f796])).
fof(f796,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) | $spl18),
  inference(cnf_transformation,[],[f796_D])).
fof(f796_D,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f819,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl22),
  inference(cnf_transformation,[],[f819_D])).
fof(f819_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1018,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X4) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl18),
  inference(superposition,[],[f218,f796])).
fof(f822,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl24),
  inference(cnf_transformation,[],[f822_D])).
fof(f822_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1027,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl30),
  inference(cnf_transformation,[],[f1027_D])).
fof(f1027_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1017,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl18),
  inference(superposition,[],[f218,f796])).
fof(f1016,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl18),
  inference(superposition,[],[f218,f796])).
fof(f1015,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl18),
  inference(superposition,[],[f218,f796])).
fof(f824,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl28),
  inference(cnf_transformation,[],[f824_D])).
fof(f824_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1009,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl18),
  inference(superposition,[],[f106,f796])).
fof(f1008,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl18),
  inference(superposition,[],[f796,f106])).
fof(f1007,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl18),
  inference(superposition,[],[f796,f218])).
fof(f1006,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl18),
  inference(superposition,[],[f796,f218])).
fof(f1005,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl18),
  inference(superposition,[],[f796,f218])).
fof(f1004,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(superposition,[],[f796,f218])).
fof(f815,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl19),
  inference(superposition,[],[f797,f106])).
fof(f823,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl26),
  inference(cnf_transformation,[],[f823_D])).
fof(f823_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f814,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl19),
  inference(superposition,[],[f797,f218])).
fof(f813,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl19),
  inference(superposition,[],[f797,f218])).
fof(f820,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl23),
  inference(cnf_transformation,[],[f820_D])).
fof(f820_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f812,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl19),
  inference(superposition,[],[f797,f218])).
fof(f811,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl19),
  inference(superposition,[],[f797,f218])).
fof(f795,plain,(
  ( ! [X31,X32] : (s(fun(bool,bool),X31) = s(fun(bool,bool),X32)) ) | $spl16),
  inference(cnf_transformation,[],[f795_D])).
fof(f795_D,plain,(
  ( ! [X31,X32] : (s(fun(bool,bool),X31) = s(fun(bool,bool),X32)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f797,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | $spl19),
  inference(cnf_transformation,[],[f797_D])).
fof(f797_D,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f783,plain,(
  ( ! [X31,X32] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f778])).
fof(f778,plain,(
  ( ! [X31,X32] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32)) ) | $spl6),
  inference(superposition,[],[f138,f755])).
fof(f759,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl10),
  inference(cnf_transformation,[],[f759_D])).
fof(f759_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f755,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl6),
  inference(cnf_transformation,[],[f755_D])).
fof(f755_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f762,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl15),
  inference(cnf_transformation,[],[f762_D])).
fof(f762_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f753,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK3(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f741,f106])).
fof(f741,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,X2,bool))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f729,f189])).
fof(f729,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'))))) != s(X2,g(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'))))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,X2,bool))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f138,f503])).
fof(f503,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f295,f181])).
fof(f295,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X14) = s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f218,f189])).
fof(f760,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(cnf_transformation,[],[f760_D])).
fof(f760_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f752,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK3(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f741,f218])).
fof(f751,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK3(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f741,f218])).
fof(f754,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl4),
  inference(cnf_transformation,[],[f754_D])).
fof(f754_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f757,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl9),
  inference(cnf_transformation,[],[f757_D])).
fof(f757_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f750,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK3(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f741,f218])).
fof(f749,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f741,f218])).
fof(f175,plain,(
  p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),sK2)))) | $spl0),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),sK2)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f177,plain,(
  s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0') = s(fun(sK0,fun(sK1,bool)),sK2) | $spl2),
  inference(cnf_transformation,[],[f177_D])).
fof(f177_D,plain,(
  s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0') = s(fun(sK0,fun(sK1,bool)),sK2) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f101,plain,(
  s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0') = s(fun(sK0,fun(sK1,bool)),sK2) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),sK2))))),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  (p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),sK2)))) | s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0') = s(fun(sK0,fun(sK1,bool)),sK2)) & (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),sK2)))) | s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0') != s(fun(sK0,fun(sK1,bool)),sK2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f70])).
fof(f70,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(fun(X0,fun(X1,bool)),X2)))) | s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0') = s(fun(X0,fun(X1,bool)),X2)) & (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(fun(X0,fun(X1,bool)),X2)))) | s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0') != s(fun(X0,fun(X1,bool)),X2)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(fun(X0,fun(X1,bool)),X2)))) <~> s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0') = s(fun(X0,fun(X1,bool)),X2))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(fun(X0,fun(X1,bool)),X2)))) <=> s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0') = s(fun(X0,fun(X1,bool)),X2))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ~! [X0,X1] : ! [X2] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0'),s(fun(X0,fun(X1,bool)),X2)))) <=> s(fun(X0,fun(X1,bool)),'const.relation.RUNIV_0') = s(fun(X0,fun(X1,bool)),X2))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X7,bool)),'const.relation.RUNIV_0'),s(fun(X6,fun(X7,bool)),X8)))) <=> s(fun(X6,fun(X7,bool)),'const.relation.RUNIV_0') = s(fun(X6,fun(X7,bool)),X8))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X7,bool)),'const.relation.RUNIV_0'),s(fun(X6,fun(X7,bool)),X8)))) <=> s(fun(X6,fun(X7,bool)),'const.relation.RUNIV_0') = s(fun(X6,fun(X7,bool)),X8))),
  file('thm.relation.RUNIV_SUBSET|split|0','thm.relation.RUNIV_SUBSET|split|0')).
fof(f176,plain,(
  ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),sK2)))) | $spl1),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),sK2)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f178,plain,(
  s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0') != s(fun(sK0,fun(sK1,bool)),sK2) | $spl3),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0') != s(fun(sK0,fun(sK1,bool)),sK2) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f102,plain,(
  s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0') != s(fun(sK0,fun(sK1,bool)),sK2) | ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK1,bool)),'const.relation.RUNIV_0'),s(fun(sK0,fun(sK1,bool)),sK2))))),
  inference(cnf_transformation,[],[f71])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 71
% Passive clauses: 216
% Generated clauses: 2484
% Final active clauses: 42
% Final passive clauses: 65
% Input formulas: 24
% Initial clauses: 57
% 
% Duplicate literals: 210
% Trivial inequalities: 37
% Fw subsumption resolutions: 129
% Fw demodulations: 169
% Bw demodulations: 63
% Condensations: 225
% Global subsumptions: 3
% 
% Simple tautologies: 70
% Equational tautologies: 51
% Forward subsumptions: 1325
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 162
% Forward superposition: 701
% Backward superposition: 548
% Self superposition: 38
% Equality factoring: 32
% Equality resolution: 4
% 
% Split clauses: 34
% Split components: 74
% SAT solver clauses: 302
% SAT solver unit clauses: 9
% SAT solver binary clauses: 68
% 
% Sat splits: 34
% Sat splitting refutations: 13
% 
% Memory used [KB]: 2046
% Time elapsed: 0.063 s
% ------------------------------
% ------------------------------
% Success in time 1.361 s
