Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (56093)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 115
% Passive clauses: 3701
% Generated clauses: 12860
% Final active clauses: 108
% Final passive clauses: 3353
% Input formulas: 28
% Initial clauses: 70
% 
% Duplicate literals: 589
% Trivial inequalities: 149
% Fw subsumption resolutions: 2136
% Fw demodulations: 386
% Bw demodulations: 4
% 
% Simple tautologies: 268
% Equational tautologies: 75
% Forward subsumptions: 4317
% 
% Binary resolution: 764
% Factoring: 2
% Forward superposition: 4856
% Backward superposition: 3310
% Self superposition: 207
% Equality factoring: 170
% Equality resolution: 23
% 
% Split clauses: 658
% Split components: 1365
% SAT solver clauses: 1408
% SAT solver unit clauses: 10
% SAT solver binary clauses: 205
% 
% Sat splits: 658
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1396
% TWLsolver calls for satisfiability: 233
% 
% Memory used [KB]: 7164
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (56104)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 115
% Passive clauses: 931
% Generated clauses: 12632
% Final active clauses: 111
% Final passive clauses: 792
% Input formulas: 28
% Initial clauses: 70
% 
% Duplicate literals: 663
% Trivial inequalities: 101
% Fw subsumption resolutions: 1152
% Fw demodulations: 1615
% Bw demodulations: 26
% Condensations: 867
% Global subsumptions: 9
% 
% Simple tautologies: 201
% Equational tautologies: 29
% Forward subsumptions: 5590
% 
% Binary resolution: 558
% Factoring: 1
% Forward superposition: 3136
% Backward superposition: 3337
% Self superposition: 98
% Equality factoring: 57
% Equality resolution: 15
% 
% Split clauses: 97
% Split components: 200
% SAT solver clauses: 1450
% SAT solver unit clauses: 84
% SAT solver binary clauses: 39
% SAT solver learnt clauses: 18
% 
% Sat splits: 97
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1414
% TWLsolver calls for satisfiability: 399
% 
% Memory used [KB]: 10490
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f2864,plain,(
  $false),
  inference(subsumption_resolution,[],[f2863,f212])).
fof(f212,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f131,f128])).
fof(f128,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_EQUATION_FORALL_POINT',aHL_TRUTH)).
fof(f131,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_EQUATION_FORALL_POINT','thm.bool.NOT_CLAUSES|split|2')).
fof(f2863,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f2822,f306])).
fof(f306,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1x1_3'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(fun(sK1,sK2),sK4),s(fun(sK1,sK2),sK5)))))),
  inference(unit_resulting_resolution,[],[f212,f127,f158])).
fof(f158,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_EQUATION_FORALL_POINT','thm.bool.IMP_ANTISYM_AX')).
fof(f127,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1x1_3'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(fun(sK1,sK2),sK4),s(fun(sK1,sK2),sK5))))))),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X6] : s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X6))))) != s(sK2,happ(s(fun(sK1,sK2),sK5),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X6))))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1x1_3'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(fun(sK1,sK2),sK4),s(fun(sK1,sK2),sK5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f67])).
fof(f67,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (! [X6] : s(X2,happ(s(fun(X1,X2),X4),s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))))) != s(X2,happ(s(fun(X1,X2),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1x1_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X1,X2),X4),s(fun(X1,X2),X5)))))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (! [X6] : s(X2,happ(s(fun(X1,X2),X4),s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))))) != s(X2,happ(s(fun(X1,X2),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1x1_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X1,X2),X4),s(fun(X1,X2),X5)))))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5] : (! [X6] : ~s(X2,happ(s(fun(X1,X2),X4),s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))))) = s(X2,happ(s(fun(X1,X2),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1x1_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X1,X2),X4),s(fun(X1,X2),X5)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11] : (! [X12] : ~s(X8,happ(s(fun(X7,X8),X10),s(X7,happ(s(fun(X6,X7),X9),s(X6,X12))))) = s(X8,happ(s(fun(X7,X8),X11),s(X7,happ(s(fun(X6,X7),X9),s(X6,X12))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x0x1x1_3'(s(fun(X8,fun(X8,bool)),'=_0'),s(fun(X7,X8),X10),s(fun(X7,X8),X11)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11] : (! [X12] : ~s(X8,happ(s(fun(X7,X8),X10),s(X7,happ(s(fun(X6,X7),X9),s(X6,X12))))) = s(X8,happ(s(fun(X7,X8),X11),s(X7,happ(s(fun(X6,X7),X9),s(X6,X12))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x0x1x1_3'(s(fun(X8,fun(X8,bool)),'=_0'),s(fun(X7,X8),X10),s(fun(X7,X8),X11)))))))),
  file('thm.quantHeuristics.GUESS_RULES_EQUATION_FORALL_POINT','thm.quantHeuristics.GUESS_RULES_EQUATION_FORALL_POINT')).
fof(f2822,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK0,sK1),sK3),s(fun(sK1,bool),'_dst_x0x1x1_3'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(fun(sK1,sK2),sK4),s(fun(sK1,sK2),sK5))))))),
  inference(unit_resulting_resolution,[],[f403,f178])).
fof(f178,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK16(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK16(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7,X6] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X22),s(fun(X7,bool),X23)))) <=> ! [X12] : ~p(s(bool,happ(s(fun(X7,bool),X23),s(X7,happ(s(fun(X6,X7),X22),s(X6,X12)))))))),
  file('thm.quantHeuristics.GUESS_RULES_EQUATION_FORALL_POINT','thm.quantHeuristics.GUESS_REWRITES|split|3')).
fof(f403,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),'_dst_x0x1x1_3'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(fun(sK1,sK2),sK4),s(fun(sK1,sK2),sK5))),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X0))))))) )),
  inference(forward_demodulation,[],[f385,f195])).
fof(f195,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X27,X28,X29,X30] : ! [X31,X32,X33,X34] : s(X27,happ(s(fun(X30,X27),'_dst_x0x1x1_3'(s(fun(X28,fun(X29,X27)),X31),s(fun(X30,X28),X32),s(fun(X30,X29),X33))),s(X30,X34))) = s(X27,happ(s(fun(X29,X27),happ(s(fun(X28,fun(X29,X27)),X31),s(X28,happ(s(fun(X30,X28),X32),s(X30,X34))))),s(X29,happ(s(fun(X30,X29),X33),s(X30,X34)))))),
  file('thm.quantHeuristics.GUESS_RULES_EQUATION_FORALL_POINT','_dst_x0x1x1_ax')).
fof(f385,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X0))))))),s(sK2,happ(s(fun(sK1,sK2),sK5),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X0))))))))) )),
  inference(unit_resulting_resolution,[],[f126,f179])).
fof(f179,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X24,X25,X26] : (p(s(bool,happ(s(fun(X24,bool),happ(s(fun(X24,fun(X24,bool)),'=_0'),s(X24,X25))),s(X24,X26)))) <=> s(X24,X25) = s(X24,X26))),
  file('thm.quantHeuristics.GUESS_RULES_EQUATION_FORALL_POINT',arity_0_HL_EQ)).
fof(f126,plain,(
  ( ! [X6] : (s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X6))))) != s(sK2,happ(s(fun(sK1,sK2),sK5),s(sK1,happ(s(fun(sK0,sK1),sK3),s(sK0,X6)))))) )),
  inference(cnf_transformation,[],[f85])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 80
% Passive clauses: 312
% Generated clauses: 2721
% Final active clauses: 64
% Final passive clauses: 200
% Input formulas: 28
% Initial clauses: 70
% 
% Duplicate literals: 85
% Trivial inequalities: 6
% Fw subsumption resolutions: 138
% Fw demodulations: 611
% Bw demodulations: 45
% Condensations: 64
% 
% Simple tautologies: 60
% Equational tautologies: 46
% Forward subsumptions: 1384
% 
% Unit resulting resolution: 313
% Factoring: 1
% Forward superposition: 784
% Backward superposition: 548
% Self superposition: 22
% Equality factoring: 8
% Equality resolution: 2
% 
% Split clauses: 4
% Split components: 9
% SAT solver clauses: 325
% SAT solver unit clauses: 63
% SAT solver binary clauses: 128
% 
% Sat splits: 4
% Sat splitting refutations: 3
% 
% Memory used [KB]: 3454
% Time elapsed: 0.119 s
% ------------------------------
% ------------------------------
% Success in time 1.46 s
