Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (61776)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 436
% Passive clauses: 4349
% Generated clauses: 15552
% Final active clauses: 374
% Final passive clauses: 3529
% Input formulas: 45
% Initial clauses: 145
% 
% Duplicate literals: 424
% Trivial inequalities: 80
% Fw subsumption resolutions: 1446
% Fw demodulations: 190
% Bw demodulations: 57
% 
% Simple tautologies: 292
% Equational tautologies: 53
% Forward subsumptions: 8200
% 
% Binary resolution: 8188
% Factoring: 37
% Forward superposition: 2667
% Backward superposition: 1998
% Self superposition: 66
% Equality factoring: 58
% Equality resolution: 20
% 
% Split clauses: 504
% Split components: 1022
% SAT solver clauses: 730
% SAT solver unit clauses: 42
% SAT solver binary clauses: 250
% 
% Sat splits: 504
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 726
% TWLsolver calls for satisfiability: 241
% 
% Memory used [KB]: 8059
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (61793)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 186
% Passive clauses: 1395
% Generated clauses: 4756
% Final active clauses: 184
% Final passive clauses: 1187
% Input formulas: 45
% Initial clauses: 145
% 
% Duplicate literals: 66
% Trivial inequalities: 14
% Fw subsumption resolutions: 137
% Fw demodulations: 75
% Bw demodulations: 22
% Condensations: 187
% Global subsumptions: 13
% 
% Simple tautologies: 130
% Equational tautologies: 30
% Forward subsumptions: 577
% 
% Binary resolution: 1370
% Factoring: 32
% Forward superposition: 467
% Backward superposition: 964
% Self superposition: 62
% Equality factoring: 46
% Equality resolution: 7
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 1653
% SAT solver unit clauses: 152
% SAT solver binary clauses: 67
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 6
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1601
% TWLsolver calls for satisfiability: 1153
% 
% Memory used [KB]: 8827
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f5655,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1298,f1302,f1305,f1308,f1311,f1314,f1317,f1336,f1299,f1338,f1303,f1340,f1306,f1342,f1309,f1344,f1312,f1346,f1315,f1348,f1300,f1318,f1376,f1377,f1383,f1388,f1398,f2370,f2373,f2371,f2376,f2378,f2383,f2385,f2387,f2389,f2393,f2395,f2397,f2399,f2401,f2403,f2412,f2424,f2445,f2447,f2486,f2488,f2490,f2492,f2495,f2497,f2503,f2505,f2511,f2513,f2528,f2530,f2532,f2534,f2537,f2539,f2542,f2544,f2547,f2549,f2587,f2596,f2608,f2617,f2619,f2631,f2640,f2652,f2661,f2663,f2665,f2667,f2669,f2927,f2928,f2692,f2694,f2696,f2698,f2705,f2715,f2723,f2725,f2741,f2743,f2745,f2747,f2758,f2761,f2772,f2775,f2790,f2792,f2806,f2808,f2810,f2816,f2821,f2826,f2828,f2833,f2835,f2844,f2851,f2860,f2867,f2875,f2877,f2879,f2881,f2888,f2894,f2899,f2901,f2906,f2909,f2929,f2930,f2931,f2932,f2933,f2934,f2935,f2936,f2937,f2938,f2939,f2940,f2941,f2942,f2943,f2944,f2945,f2946,f2947,f2948,f2949,f2950,f2951,f2952,f2953,f2954,f2955,f2956,f2957,f2958,f2959,f2960,f2961,f2962,f2963,f2964,f2965,f2966,f2967,f2968,f2969,f2970,f2971,f2972,f2973,f2974,f2975,f2976,f2977,f2978,f2979,f2980,f2981,f2982,f2983,f2984,f2985,f2986,f2987,f2988,f2989,f2990,f2991,f2992,f2993,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3038,f3039,f3040,f3041,f3042,f3043,f3044,f3045,f3046,f3075,f3076,f3077,f3079,f2372,f3078,f5647])).
fof(f5647,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f5646,f396])).
fof(f396,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f238,f235])).
fof(f235,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES',aHL_TRUTH)).
fof(f238,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.bool.NOT_CLAUSES|split|2')).
fof(f5646,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(forward_demodulation,[],[f5645,f426])).
fof(f426,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))),
  inference(unit_resulting_resolution,[],[f234,f254])).
fof(f254,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.bool.EQ_CLAUSES|split|2')).
fof(f234,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))))))),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X4] : (s(sK8,sK9) = s(sK8,X4) | s(sK8,X4) = s(sK8,happ(s(fun(sK7,sK8),sK10),s(sK7,sK11(X4))))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9,sK10,sK11])],[f117])).
fof(f117,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (s(X1,X2) = s(X1,X4) | ? [X5] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X5)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x1x0_2'(s(fun(X1,fun(X1,bool)),'=_0'),s(X1,X2)))))))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0,X1,X2,X3] : (! [X4] : (s(X1,X2) = s(X1,X4) | ? [X5] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X5)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x1x0_2'(s(fun(X1,fun(X1,bool)),'=_0'),s(X1,X2)))))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ~! [X0,X1] : ! [X2,X3] : (! [X4] : (s(X1,X2) = s(X1,X4) | ? [X5] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X5)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x1x0_2'(s(fun(X1,fun(X1,bool)),'=_0'),s(X1,X2)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (! [X10] : (s(X7,X8) = s(X7,X10) | ? [X11] : s(X7,X10) = s(X7,happ(s(fun(X6,X7),X9),s(X6,X11)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x1x0_2'(s(fun(X7,fun(X7,bool)),'=_0'),s(X7,X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (! [X10] : (s(X7,X8) = s(X7,X10) | ? [X11] : s(X7,X10) = s(X7,happ(s(fun(X6,X7),X9),s(X6,X11)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x1x0_2'(s(fun(X7,fun(X7,bool)),'=_0'),s(X7,X8)))))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.quantHeuristics.GUESS_RULES_TWO_CASES')).
fof(f5645,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl20),
  inference(subsumption_resolution,[],[f5644,f235])).
fof(f5644,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl20),
  inference(forward_demodulation,[],[f5643,f4484])).
fof(f4484,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f235,f1902,f284])).
fof(f284,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.bool.IMP_ANTISYM_AX')).
fof(f1902,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(equality_resolution,[],[f322])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X26,X27,X28] : (p(s(bool,happ(s(fun(X26,bool),happ(s(fun(X26,fun(X26,bool)),'=_0'),s(X26,X27))),s(X26,X28)))) <=> s(X26,X27) = s(X26,X28))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES',arity_0_HL_EQ)).
fof(f5643,plain,(
  ~p(s(bool,happ(s(fun(sK8,bool),happ(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))),s(sK8,sK9)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl20),
  inference(forward_demodulation,[],[f5625,f369])).
fof(f369,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X29,X30,X31] : ! [X32,X33,X34] : s(X29,happ(s(fun(X30,X29),'_dst_x1x0_2'(s(fun(X30,fun(X31,X29)),X32),s(X31,X33))),s(X30,X34))) = s(X29,happ(s(fun(X31,X29),happ(s(fun(X30,fun(X31,X29)),X32),s(X30,X34))),s(X31,X33)))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','_dst_x1x0_ax')).
fof(f5625,plain,(
  ~p(s(bool,happ(s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))),s(sK8,sK9)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl20),
  inference(superposition,[],[f314,f5554])).
fof(f5554,plain,(
  s(sK8,sK9) = s(sK8,sK19('_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)),sK10,sK7,sK8)) | $spl20),
  inference(unit_resulting_resolution,[],[f3076,f233])).
fof(f233,plain,(
  ( ! [X4] : (s(sK8,X4) = s(sK8,happ(s(fun(sK7,sK8),sK10),s(sK7,sK11(X4)))) | s(sK8,sK9) = s(sK8,X4)) )),
  inference(cnf_transformation,[],[f163])).
fof(f314,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK19(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK18(X6,X3,X2,X1,X0)))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK19(X3,X2,X1,X0))))) & ! [X5] : s(X0,sK19(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19])],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6,X7] : ! [X21,X22] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X7,X6),X21),s(fun(X6,bool),X22)))) <=> ! [X20] : (~p(s(bool,happ(s(fun(X6,bool),X22),s(X6,X20)))) => ? [X11] : s(X6,X20) = s(X6,happ(s(fun(X7,X6),X21),s(X7,X11)))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.quantHeuristics.GUESS_REWRITES|split|5')).
fof(f3078,plain,(
  ( ! [X14,X19,X17,X15,X16] : (s(X15,sK19(X17,X16,X14,X15)) != s(X15,happ(s(fun(X14,X15),X16),s(X14,X19))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X14,X15),X16),s(fun(X15,bool),X17)))) ) | $spl22),
  inference(cnf_transformation,[],[f3078_D])).
fof(f3078_D,plain,(
  ( ! [X14,X19,X17,X15,X16] : (s(X15,sK19(X17,X16,X14,X15)) != s(X15,happ(s(fun(X14,X15),X16),s(X14,X19))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X14,X15),X16),s(fun(X15,bool),X17)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2372,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) | $spl18),
  inference(cnf_transformation,[],[f2372_D])).
fof(f2372_D,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3079,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,X24)) | s(X21,sK19(X23,X22,X20,X21)) != s(X21,happ(s(fun(X20,X21),X22),s(X20,X25))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X20,X21),X22),s(fun(X21,bool),X23)))) )),
  inference(subsumption_resolution,[],[f3072,f249])).
fof(f249,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.bool.EQ_CLAUSES|split|1')).
fof(f3072,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,X24)) | s(X21,sK19(X23,X22,X20,X21)) != s(X21,happ(s(fun(X20,X21),X22),s(X20,X25))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X20,X21),X22),s(fun(X21,bool),X23)))) )),
  inference(superposition,[],[f315,f469])).
fof(f469,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f242,f242])).
fof(f242,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES',aHL_BOOL_CASES)).
fof(f315,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | s(X0,sK19(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f201])).
fof(f3077,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (p(s(bool,X18)) | s(X15,sK19(X17,X16,X14,X15)) != s(X15,happ(s(fun(X14,X15),X16),s(X14,X19))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X14,X15),X16),s(fun(X15,bool),X17)))) )),
  inference(subsumption_resolution,[],[f3071,f249])).
fof(f3071,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (p(s(bool,X18)) | s(X15,sK19(X17,X16,X14,X15)) != s(X15,happ(s(fun(X14,X15),X16),s(X14,X19))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X14,X15),X16),s(fun(X15,bool),X17))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f315,f469])).
fof(f3076,plain,(
  ( ! [X1] : (s(sK8,happ(s(fun(sK7,sK8),sK10),s(sK7,X1))) != s(sK8,sK19('_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)),sK10,sK7,sK8))) ) | $spl20),
  inference(cnf_transformation,[],[f3076_D])).
fof(f3076_D,plain,(
  ( ! [X1] : (s(sK8,happ(s(fun(sK7,sK8),sK10),s(sK7,X1))) != s(sK8,sK19('_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)),sK10,sK7,sK8))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3075,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(sK8,happ(s(fun(sK7,sK8),sK10),s(sK7,X1))) != s(sK8,sK19('_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)),sK10,sK7,sK8))) )),
  inference(subsumption_resolution,[],[f3068,f249])).
fof(f3068,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(sK8,happ(s(fun(sK7,sK8),sK10),s(sK7,X1))) != s(sK8,sK19('_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)),sK10,sK7,sK8)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f315,f707])).
fof(f707,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f705,f396])).
fof(f705,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f505,f426])).
fof(f505,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f502,f396])).
fof(f502,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f281,f242])).
fof(f281,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f130])).
fof(f3046,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2379,f2908])).
fof(f2908,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f321])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f208])).
fof(f2379,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f253])).
fof(f253,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f3045,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2382,f235])).
fof(f2382,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f413,f2372,f284])).
fof(f413,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f396,f249])).
fof(f3044,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2384,f235])).
fof(f2384,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f425,f2372,f284])).
fof(f425,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))),
  inference(unit_resulting_resolution,[],[f234,f249])).
fof(f3043,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2388,f2372])).
fof(f2388,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f425,f2372,f284])).
fof(f3042,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2392,f235])).
fof(f2392,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f413,f2372,f284])).
fof(f3041,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2394,f235])).
fof(f2394,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f425,f2372,f284])).
fof(f3040,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2398,f2372])).
fof(f2398,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f425,f2372,f284])).
fof(f3039,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2404,f235])).
fof(f2404,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f518,f396,f2372,f332])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (sP1(X2,X1,X0) & sP2(X1,X2,X0) & sP3(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP1(X2,X1,X0) | ~sP2(X1,X2,X0) | ~sP3(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X2,X1,X0] : ((~sP4(X2,X1,X0) | (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP1(X0,X1,X2) | ~sP2(X1,X0,X2) | ~sP3(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP4(X2,X1,X0)))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X2,X1,X0] : ((~sP4(X2,X1,X0) | (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP1(X0,X1,X2) | ~sP2(X1,X0,X2) | ~sP3(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP4(X2,X1,X0)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X2,X1,X0] : (sP4(X2,X1,X0) <=> (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f518,plain,(
  sP4('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f396,f413,f348])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP4(X2,X1,X0)) & (~sP4(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP4(X2,X1,X0))),
  inference(definition_folding,[],[f111,f158,f157,f156,f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : (sP1(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f156,plain,(
  ! [X1,X0,X2] : (sP2(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f157,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X23,X24,X25] : ((p(s(bool,X25)) <=> s(bool,X23) = s(bool,X24)) <=> ((p(s(bool,X25)) | p(s(bool,X24)) | p(s(bool,X23))) & (p(s(bool,X25)) | ~p(s(bool,X23)) | ~p(s(bool,X24))) & (p(s(bool,X24)) | ~p(s(bool,X23)) | ~p(s(bool,X25))) & (p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.sat.sat_2|49|_prove_ltb')).
fof(f3038,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2406,f2372])).
fof(f2406,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))) | $spl18),
  inference(unit_resulting_resolution,[],[f1727,f396,f2372,f332])).
fof(f1727,plain,(
  sP4('T_0',happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f405,f402,f407,f1710,f334])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (~sP1(X2,X1,X0) | p(s(bool,X1)) | ~sP3(X0,X2,X1) | ~sP2(X1,X2,X0) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f213])).
fof(f1710,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))))),
  inference(forward_demodulation,[],[f1709,f369])).
fof(f1709,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(forward_demodulation,[],[f1708,f426])).
fof(f1708,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))))))))),
  inference(forward_demodulation,[],[f1676,f369])).
fof(f1676,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))))))),s(bool,'T_0'))))),
  inference(unit_resulting_resolution,[],[f425,f321])).
fof(f407,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f235,f345])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(nnf_transformation,[],[f155])).
fof(f402,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f396,f338])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f215])).
fof(f215,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f157])).
fof(f405,plain,(
  ( ! [X0,X1] : (sP2(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f396,f342])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f218])).
fof(f218,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0,X2)))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0,X2)))),
  inference(nnf_transformation,[],[f156])).
fof(f3037,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2407,f2372])).
fof(f2407,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f567,f396,f2372,f332])).
fof(f567,plain,(
  sP4('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'F_0')),
  inference(unit_resulting_resolution,[],[f235,f426,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f3036,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2441,f235])).
fof(f2441,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f785,f396,f2372,f332])).
fof(f785,plain,(
  sP4('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f406,f402,f407,f396,f334])).
fof(f406,plain,(
  ( ! [X0,X1] : (sP2('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f396,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f219])).
fof(f3035,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2442,f235])).
fof(f2442,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f1727,f396,f2372,f332])).
fof(f3034,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2443,f235])).
fof(f2443,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f567,f396,f2372,f332])).
fof(f3033,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2448,f2668])).
fof(f2668,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f2448,plain,(
  ~sP1('F_0','T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f404,f401,f2372,f333])).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (~sP1(X2,X1,X0) | ~p(s(bool,X2)) | ~sP3(X0,X2,X1) | ~sP2(X1,X2,X0) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f213])).
fof(f401,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f235,f337])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f216])).
fof(f404,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f235,f341])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f219])).
fof(f543,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f235,f413,f350])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (~sP4(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f222])).
fof(f3032,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2449,f2668])).
fof(f2449,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f404,f401,f2372,f333])).
fof(f544,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))),
  inference(unit_resulting_resolution,[],[f235,f425,f350])).
fof(f3031,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2450,f2668])).
fof(f2450,plain,(
  ~sP1(happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f404,f401,f2372,f333])).
fof(f1726,plain,(
  ~sP4('T_0','T_0',happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))),
  inference(unit_resulting_resolution,[],[f235,f235,f1710,f332])).
fof(f3030,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2451,f2668])).
fof(f2451,plain,(
  ~sP1('F_0','T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f405,f401,f2372,f333])).
fof(f3029,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2452,f2668])).
fof(f2452,plain,(
  ~sP1(happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f1745,f401,f2372,f333])).
fof(f1745,plain,(
  ( ! [X0,X1] : (sP2(X0,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1710,f342])).
fof(f3028,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2453,f2668])).
fof(f2453,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f429,f401,f2372,f333])).
fof(f429,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),X1)) )),
  inference(unit_resulting_resolution,[],[f234,f342])).
fof(f3027,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2454,f2668])).
fof(f2454,plain,(
  ~sP1('F_0','T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f405,f402,f2372,f333])).
fof(f3026,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2457,f2668])).
fof(f2457,plain,(
  ~sP1('F_0','T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f404,f402,f2372,f333])).
fof(f3025,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2460,f2668])).
fof(f2460,plain,(
  ~sP1(happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f1745,f1738,f2372,f333])).
fof(f1738,plain,(
  ( ! [X0,X1] : (sP3(X0,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1710,f338])).
fof(f3024,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2461,f2668])).
fof(f2461,plain,(
  ~sP1(happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f404,f1738,f2372,f333])).
fof(f3023,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2462,f2668])).
fof(f2462,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f429,f427,f2372,f333])).
fof(f427,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),X1)) )),
  inference(unit_resulting_resolution,[],[f234,f338])).
fof(f3022,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2464,f2668])).
fof(f2464,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f404,f427,f2372,f333])).
fof(f3021,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2473,f2618])).
fof(f2618,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f337])).
fof(f2473,plain,(
  ~sP3('T_0','F_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f404,f407,f2372,f333])).
fof(f3020,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2474,f2618])).
fof(f2474,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f404,f407,f2372,f333])).
fof(f3019,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2475,f2618])).
fof(f2475,plain,(
  ~sP3('T_0',happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f404,f407,f2372,f333])).
fof(f3018,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2477,f2618])).
fof(f2477,plain,(
  ~sP3('T_0','F_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f405,f407,f2372,f333])).
fof(f3017,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2478,f2618])).
fof(f2478,plain,(
  ~sP3('T_0',happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f1745,f407,f2372,f333])).
fof(f3016,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2479,f2618])).
fof(f2479,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f429,f407,f2372,f333])).
fof(f3015,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2482,f2662])).
fof(f2662,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f341])).
fof(f2482,plain,(
  ~sP2('T_0','F_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f401,f407,f2372,f333])).
fof(f3014,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2483,f2662])).
fof(f2483,plain,(
  ~sP2('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f401,f407,f2372,f333])).
fof(f3013,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2484,f2662])).
fof(f2484,plain,(
  ~sP2('T_0',happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f401,f407,f2372,f333])).
fof(f3012,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2493,f2662])).
fof(f2493,plain,(
  ~sP2('T_0','F_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f402,f407,f2372,f333])).
fof(f3011,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2501,f2662])).
fof(f2501,plain,(
  ~sP2('T_0',happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f1738,f407,f2372,f333])).
fof(f3010,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2509,f2662])).
fof(f2509,plain,(
  ~sP2('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f427,f407,f2372,f333])).
fof(f3009,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2517,f2618])).
fof(f2517,plain,(
  ~sP3('T_0','F_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f404,f408,f2372,f333])).
fof(f408,plain,(
  ( ! [X0,X1] : (sP1(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f235,f346])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f221])).
fof(f3008,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2518,f2618])).
fof(f2518,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f404,f408,f2372,f333])).
fof(f3007,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2519,f2618])).
fof(f2519,plain,(
  ~sP3('T_0',happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f404,f408,f2372,f333])).
fof(f3006,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2520,f2618])).
fof(f2520,plain,(
  ~sP3('T_0','F_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f405,f408,f2372,f333])).
fof(f3005,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2521,f2618])).
fof(f2521,plain,(
  ~sP3('T_0',happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f1745,f408,f2372,f333])).
fof(f3004,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2522,f2618])).
fof(f2522,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f429,f408,f2372,f333])).
fof(f3003,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2523,f2662])).
fof(f2523,plain,(
  ~sP2('T_0','F_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f401,f408,f2372,f333])).
fof(f3002,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2524,f2662])).
fof(f2524,plain,(
  ~sP2('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f401,f408,f2372,f333])).
fof(f3001,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2525,f2662])).
fof(f2525,plain,(
  ~sP2('T_0',happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f401,f408,f2372,f333])).
fof(f3000,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2529,f2691])).
fof(f2691,plain,(
  ( ! [X0] : (~sP4(X0,'T_0','F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f413,f2372,f350])).
fof(f2529,plain,(
  ( ! [X0] : (sP4(X0,'T_0','F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f405,f401,f408,f2372,f333])).
fof(f2999,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2533,f2693])).
fof(f2693,plain,(
  ( ! [X0] : (~sP4(X0,'T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f425,f2372,f350])).
fof(f2533,plain,(
  ( ! [X0] : (sP4(X0,'T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f429,f401,f408,f2372,f333])).
fof(f2998,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2535,f2662])).
fof(f2535,plain,(
  ~sP2('T_0','F_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f543,f402,f408,f2372,f333])).
fof(f2997,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2536,f2691])).
fof(f2536,plain,(
  ( ! [X0] : (sP4(X0,'T_0','F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f405,f402,f408,f2372,f333])).
fof(f2996,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2538,f543])).
fof(f2538,plain,(
  sP4('T_0','T_0','F_0') | $spl18),
  inference(unit_resulting_resolution,[],[f404,f402,f408,f2372,f333])).
fof(f2995,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2540,f2662])).
fof(f2540,plain,(
  ~sP2('T_0',happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f1738,f408,f2372,f333])).
fof(f2994,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2543,f1726])).
fof(f2543,plain,(
  sP4('T_0','T_0',happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))) | $spl18),
  inference(unit_resulting_resolution,[],[f404,f1738,f408,f2372,f333])).
fof(f2993,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2545,f2662])).
fof(f2545,plain,(
  ~sP2('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f544,f427,f408,f2372,f333])).
fof(f2992,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2546,f2693])).
fof(f2546,plain,(
  ( ! [X0] : (sP4(X0,'T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f429,f427,f408,f2372,f333])).
fof(f2991,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2548,f544])).
fof(f2548,plain,(
  sP4('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))))) | $spl18),
  inference(unit_resulting_resolution,[],[f404,f427,f408,f2372,f333])).
fof(f2990,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2557,f2691])).
fof(f2557,plain,(
  sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0','F_0') | $spl18),
  inference(unit_resulting_resolution,[],[f405,f428,f408,f2372,f333])).
fof(f428,plain,(
  ( ! [X0,X1] : (sP3('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f234,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f216])).
fof(f2989,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2559,f2693])).
fof(f2559,plain,(
  sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))))) | $spl18),
  inference(unit_resulting_resolution,[],[f429,f428,f408,f2372,f333])).
fof(f2988,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2577,f2372])).
fof(f2577,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))) | $spl18),
  inference(unit_resulting_resolution,[],[f1738,f396,f2372,f336])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f216])).
fof(f2987,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2578,f2372])).
fof(f2578,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f427,f396,f2372,f336])).
fof(f2986,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2579,f2372])).
fof(f2579,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f403,f396,f2372,f336])).
fof(f403,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f396,f339])).
fof(f2985,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2580,f2372])).
fof(f2580,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f1739,f396,f2372,f336])).
fof(f1739,plain,(
  ( ! [X0,X1] : (sP3(happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1710,f339])).
fof(f2984,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2581,f2372])).
fof(f2581,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f428,f396,f2372,f336])).
fof(f2983,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2595,f2618])).
fof(f2595,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f396,f2372,f2372,f336])).
fof(f2982,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2597,f2372])).
fof(f2597,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f402,f396,f2372,f336])).
fof(f2981,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2598,f2372])).
fof(f2598,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f1738,f396,f2372,f336])).
fof(f2980,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2599,f2372])).
fof(f2599,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f427,f396,f2372,f336])).
fof(f2979,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2601,f2372])).
fof(f2601,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))) | $spl18),
  inference(unit_resulting_resolution,[],[f1739,f396,f2372,f336])).
fof(f2978,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2602,f2372])).
fof(f2602,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f428,f396,f2372,f336])).
fof(f2977,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2607,f2618])).
fof(f2607,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f396,f235,f2372,f336])).
fof(f2976,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2616,f2618])).
fof(f2616,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f396,f2372,f2372,f336])).
fof(f2975,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2621,f2372])).
fof(f2621,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))) | $spl18),
  inference(unit_resulting_resolution,[],[f1745,f396,f2372,f340])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f219])).
fof(f2974,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2622,f2372])).
fof(f2622,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f429,f396,f2372,f340])).
fof(f2973,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2623,f2372])).
fof(f2623,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f406,f396,f2372,f340])).
fof(f2972,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2624,f2372])).
fof(f2624,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f1746,f396,f2372,f340])).
fof(f1746,plain,(
  ( ! [X0,X1] : (sP2(happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1710,f343])).
fof(f2971,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2625,f2372])).
fof(f2625,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f430,f396,f2372,f340])).
fof(f430,plain,(
  ( ! [X0,X1] : (sP2('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f234,f343])).
fof(f2970,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2639,f2662])).
fof(f2639,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1,'F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f396,f2372,f2372,f340])).
fof(f2969,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2641,f2372])).
fof(f2641,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f405,f396,f2372,f340])).
fof(f2968,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2642,f2372])).
fof(f2642,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f1745,f396,f2372,f340])).
fof(f2967,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2643,f2372])).
fof(f2643,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f429,f396,f2372,f340])).
fof(f2966,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2645,f2372])).
fof(f2645,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))) | $spl18),
  inference(unit_resulting_resolution,[],[f1746,f396,f2372,f340])).
fof(f2965,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2646,f2372])).
fof(f2646,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f430,f396,f2372,f340])).
fof(f2964,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2651,f2662])).
fof(f2651,plain,(
  ( ! [X0] : (~sP2('T_0',X0,'F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f396,f235,f2372,f340])).
fof(f2963,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2660,f2662])).
fof(f2660,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1,'F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f396,f2372,f2372,f340])).
fof(f2962,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2675,f426])).
fof(f2675,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))))) | $spl18),
  inference(unit_resulting_resolution,[],[f566,f2372,f349])).
fof(f566,plain,(
  ~sP4('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'F_0')),
  inference(unit_resulting_resolution,[],[f396,f426,f351])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (~sP4(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f2961,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2676,f426])).
fof(f2676,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9))))) | $spl18),
  inference(unit_resulting_resolution,[],[f569,f2372,f349])).
fof(f569,plain,(
  ~sP4('F_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))),
  inference(unit_resulting_resolution,[],[f396,f426,f351])).
fof(f2960,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2704,f2880])).
fof(f2880,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f412,f2372,f390])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP5(X0,X1,X2) | sP6(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP5(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP5(X0,X1,X2) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP5(X2,X0,X1) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP5(X2,X0,X1)) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f412,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f396,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f2704,plain,(
  ( ! [X0,X1] : (~sP6('F_0',X0,X1)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f396,f2372,f353])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f225])).
fof(f2959,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2707,f2372])).
fof(f2707,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f412,f444,f2372,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f444,plain,(
  ( ! [X0] : (~sP6('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f235,f396,f353])).
fof(f2958,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2708,f2372])).
fof(f2708,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f411,f444,f2372,f356])).
fof(f411,plain,(
  ( ! [X0,X1] : (sP5(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f235,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f228])).
fof(f2957,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2710,f2372])).
fof(f2710,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f431,f444,f2372,f356])).
fof(f431,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) )),
  inference(unit_resulting_resolution,[],[f234,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f2956,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2714,f2809])).
fof(f2809,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f361])).
fof(f2714,plain,(
  ~sP5('F_0','T_0','T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f444,f235,f2372,f356])).
fof(f2955,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2722,f2809])).
fof(f2722,plain,(
  ( ! [X0] : (~sP5('F_0',X0,'F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f440,f2372,f2372,f356])).
fof(f440,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f396,f396,f352])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f2954,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2724,f2809])).
fof(f2724,plain,(
  ( ! [X0] : (~sP5('F_0','T_0',X0)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f444,f2372,f2372,f356])).
fof(f2953,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2737,f2809])).
fof(f2737,plain,(
  ( ! [X0] : (~sP5('F_0',X0,'F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f440,f2372,f2372,f356])).
fof(f2952,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2738,f2809])).
fof(f2738,plain,(
  ( ! [X0] : (~sP5('F_0','T_0',X0)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f444,f2372,f2372,f356])).
fof(f2951,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2755,f2809])).
fof(f2755,plain,(
  ~sP5('F_0','F_0','F_0') | $spl18),
  inference(unit_resulting_resolution,[],[f440,f396,f2372,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f2950,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2769,f2809])).
fof(f2769,plain,(
  ~sP5('F_0','F_0','F_0') | $spl18),
  inference(unit_resulting_resolution,[],[f440,f396,f2372,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f2949,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2786,f2372])).
fof(f2786,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f431,f396,f2372,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f228])).
fof(f2948,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2787,f2372])).
fof(f2787,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f412,f396,f2372,f359])).
fof(f2947,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2788,f2372])).
fof(f2788,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f432,f396,f2372,f359])).
fof(f432,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f234,f362])).
fof(f2946,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2791,f2809])).
fof(f2791,plain,(
  ( ! [X0,X1] : (~sP5(X0,'F_0',X1)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f396,f2372,f359])).
fof(f2945,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2801,f2372])).
fof(f2801,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f410,f396,f2372,f359])).
fof(f410,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f396,f360])).
fof(f2944,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2802,f2372])).
fof(f2802,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f431,f396,f2372,f359])).
fof(f2943,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2804,f2372])).
fof(f2804,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f432,f396,f2372,f359])).
fof(f2942,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2805,f2809])).
fof(f2805,plain,(
  ( ! [X0] : (~sP5('T_0','F_0',X0)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f235,f396,f2372,f359])).
fof(f2941,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2807,f2809])).
fof(f2807,plain,(
  ( ! [X0,X1] : (~sP5(X0,'F_0',X1)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f396,f2372,f359])).
fof(f2940,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2824,f235])).
fof(f2824,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f676,f2372,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP6(X2,X0,X1)) & (~sP6(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP6(X2,X0,X1)) & (~sP6(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP6(X2,X0,X1))),
  inference(definition_folding,[],[f150,f161,f160])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X23,X24,X25] : ((p(s(bool,X25)) <=> (p(s(bool,X24)) => p(s(bool,X23)))) <=> ((p(s(bool,X25)) | p(s(bool,X24))) & (p(s(bool,X25)) | ~p(s(bool,X23))) & (~p(s(bool,X24)) | p(s(bool,X23)) | ~p(s(bool,X25)))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.sat.sat_27_prove_ltb')).
fof(f676,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f667,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f225])).
fof(f667,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f235,f396,f235,f359])).
fof(f2939,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2829,f2372])).
fof(f2829,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f440,f2372,f365])).
fof(f2938,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2830,f235])).
fof(f2830,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f444,f2372,f365])).
fof(f2937,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2836,f235])).
fof(f2836,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f450,f396,f2372,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f450,plain,(
  ( ! [X0] : (sP6('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f396,f235,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f2936,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2837,f235])).
fof(f2837,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f506,f396,f2372,f366])).
fof(f506,plain,(
  ( ! [X0] : (sP6('T_0',X0,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) )),
  inference(unit_resulting_resolution,[],[f235,f431,f390])).
fof(f2935,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2839,f2372])).
fof(f2839,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f657,f396,f2372,f366])).
fof(f657,plain,(
  sP6('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f432,f235,f396,f358])).
fof(f2934,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2853,f2372])).
fof(f2853,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))) | $spl18),
  inference(unit_resulting_resolution,[],[f506,f396,f2372,f366])).
fof(f2933,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2854,f235])).
fof(f2854,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f656,f396,f2372,f366])).
fof(f656,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f412,f235,f396,f358])).
fof(f2932,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2855,f235])).
fof(f2855,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(unit_resulting_resolution,[],[f657,f396,f2372,f366])).
fof(f2931,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2868,f2908])).
fof(f2868,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f396,f2372,f389])).
fof(f389,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f388])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f377])).
fof(f377,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X10,X17,X8,X18] : ((s(bool,X10) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X8) = s(bool,X18))) => ((p(s(bool,X10)) => p(s(bool,X8))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.bool.IMP_CONG')).
fof(f2930,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2871,f410])).
fof(f2871,plain,(
  ( ! [X0] : (~sP5('F_0',X0,'F_0')) ) | $spl18),
  inference(unit_resulting_resolution,[],[f440,f2372,f390])).
fof(f2929,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2872,f411])).
fof(f2872,plain,(
  ( ! [X0] : (~sP5('F_0','T_0',X0)) ) | $spl18),
  inference(unit_resulting_resolution,[],[f444,f2372,f390])).
fof(f2909,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f413,f2372,f321])).
fof(f2906,plain,(
  $false | ($spl12 | $spl18)),
  inference(unit_resulting_resolution,[],[f450,f396,f2372,f1315])).
fof(f2901,plain,(
  $false | ($spl10 | $spl18)),
  inference(unit_resulting_resolution,[],[f676,f2372,f1312])).
fof(f2899,plain,(
  $false | ($spl6 | $spl18)),
  inference(unit_resulting_resolution,[],[f412,f396,f2372,f1306])).
fof(f2894,plain,(
  $false | ($spl4 | $spl18)),
  inference(unit_resulting_resolution,[],[f406,f396,f2372,f1303])).
fof(f2888,plain,(
  $false | ($spl0 | $spl18)),
  inference(unit_resulting_resolution,[],[f403,f396,f2372,f1299])).
fof(f2881,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f412,f2372,f390])).
fof(f2879,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f444,f411,f2372,f390])).
fof(f2877,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f444,f431,f2372,f390])).
fof(f2875,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f410,f2372,f390])).
fof(f2867,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f450,f396,f2372,f2372,f366])).
fof(f2860,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f656,f396,f235,f2372,f366])).
fof(f2851,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f450,f396,f2372,f2372,f366])).
fof(f2844,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f450,f396,f235,f2372,f366])).
fof(f2835,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f2372,f2372,f365])).
fof(f2833,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f444,f235,f2372,f365])).
fof(f2828,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f2372,f2372,f365])).
fof(f2826,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f676,f235,f2372,f365])).
fof(f2821,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f396,f2372,f364])).
fof(f2816,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f396,f396,f2372,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f230])).
fof(f2810,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f667,f2372,f361])).
fof(f2808,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f410,f2372,f396,f2372,f359])).
fof(f2806,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f410,f235,f396,f2372,f359])).
fof(f2792,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f410,f2372,f396,f2372,f359])).
fof(f2790,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f412,f235,f396,f2372,f359])).
fof(f2775,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f412,f396,f2372,f358])).
fof(f2772,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f410,f396,f2372,f358])).
fof(f2761,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f412,f396,f2372,f357])).
fof(f2758,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f410,f396,f2372,f357])).
fof(f2747,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f412,f2372,f2372,f356])).
fof(f2745,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f444,f411,f2372,f2372,f356])).
fof(f2743,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f444,f431,f2372,f2372,f356])).
fof(f2741,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f440,f410,f2372,f2372,f356])).
fof(f2725,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f412,f444,f2372,f2372,f356])).
fof(f2723,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f412,f440,f2372,f2372,f356])).
fof(f2715,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f412,f444,f235,f2372,f356])).
fof(f2705,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f656,f396,f2372,f353])).
fof(f2698,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f527,f425,f2372,f350])).
fof(f527,plain,(
  sP4('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0')),
  inference(unit_resulting_resolution,[],[f396,f425,f348])).
fof(f2696,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f517,f413,f2372,f350])).
fof(f517,plain,(
  sP4('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f396,f413,f348])).
fof(f2694,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f529,f425,f2372,f350])).
fof(f529,plain,(
  sP4('F_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))),
  inference(unit_resulting_resolution,[],[f396,f425,f348])).
fof(f2692,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f518,f413,f2372,f350])).
fof(f2928,plain,(
  $false | $spl18),
  inference(trivial_inequality_removal,[],[f2678])).
fof(f2678,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl18),
  inference(unit_resulting_resolution,[],[f633,f2372,f349])).
fof(f633,plain,(
  ~sP4('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f624,f330])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | sP2(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f624,plain,(
  ~sP2('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f396,f235,f235,f340])).
fof(f2927,plain,(
  $false | $spl18),
  inference(trivial_inequality_removal,[],[f2677])).
fof(f2677,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl18),
  inference(unit_resulting_resolution,[],[f642,f2372,f349])).
fof(f642,plain,(
  ~sP4('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f637,f329])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f637,plain,(
  ~sP1('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f396,f396,f396,f344])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f2669,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f637,f2372,f347])).
fof(f2667,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f637,f2372,f346])).
fof(f2665,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f637,f2372,f345])).
fof(f2663,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f624,f2372,f341])).
fof(f2661,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f405,f396,f2372,f2372,f340])).
fof(f2652,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f405,f396,f235,f2372,f340])).
fof(f2640,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f405,f396,f2372,f2372,f340])).
fof(f2631,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f406,f396,f235,f2372,f340])).
fof(f2619,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f607,f2372,f337])).
fof(f607,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f396,f235,f235,f336])).
fof(f2617,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f402,f396,f2372,f2372,f336])).
fof(f2608,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f402,f396,f235,f2372,f336])).
fof(f2596,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f402,f396,f2372,f2372,f336])).
fof(f2587,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f403,f396,f235,f2372,f336])).
fof(f2549,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f544,f404,f427,f408,f2372,f333])).
fof(f2547,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f544,f429,f427,f408,f2372,f333])).
fof(f2544,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f404,f1738,f408,f2372,f333])).
fof(f2542,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f1745,f1738,f408,f2372,f333])).
fof(f2539,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f543,f404,f402,f408,f2372,f333])).
fof(f2537,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f543,f405,f402,f408,f2372,f333])).
fof(f2534,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f544,f429,f401,f408,f2372,f333])).
fof(f2532,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f1745,f401,f408,f2372,f333])).
fof(f2530,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f543,f405,f401,f408,f2372,f333])).
fof(f2528,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f543,f404,f401,f408,f2372,f333])).
fof(f2513,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f544,f404,f427,f407,f2372,f333])).
fof(f2511,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f544,f429,f427,f407,f2372,f333])).
fof(f2505,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f404,f1738,f407,f2372,f333])).
fof(f2503,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f1745,f1738,f407,f2372,f333])).
fof(f2497,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f543,f404,f402,f407,f2372,f333])).
fof(f2495,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f543,f405,f402,f407,f2372,f333])).
fof(f2492,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f544,f429,f401,f407,f2372,f333])).
fof(f2490,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f1726,f1745,f401,f407,f2372,f333])).
fof(f2488,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f543,f405,f401,f407,f2372,f333])).
fof(f2486,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f543,f404,f401,f407,f2372,f333])).
fof(f2447,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f518,f2372,f396,f2372,f332])).
fof(f2445,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f785,f235,f396,f2372,f332])).
fof(f2424,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f518,f396,f2372,f2372,f332])).
fof(f2412,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f518,f396,f235,f2372,f332])).
fof(f2403,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f422,f2372,f291])).
fof(f291,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f422,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f396,f396,f398])).
fof(f398,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f289,f396])).
fof(f289,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f2401,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f422,f2372,f290])).
fof(f290,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f2399,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f425,f2372,f284])).
fof(f2397,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f413,f2372,f284])).
fof(f2395,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f235,f425,f2372,f284])).
fof(f2393,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f235,f413,f2372,f284])).
fof(f2389,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f425,f2372,f284])).
fof(f2387,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f2372,f413,f2372,f284])).
fof(f2385,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f235,f425,f2372,f284])).
fof(f2383,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f235,f413,f2372,f284])).
fof(f2378,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f413,f2372,f250])).
fof(f250,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f166])).
fof(f2376,plain,(
  $false | $spl18),
  inference(unit_resulting_resolution,[],[f396,f2372])).
fof(f2371,plain,(
  ( ! [X14,X12,X17,X15,X13] : (s(X13,sK15(X15,X14,X12,X13)) != s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15)))) ) | $spl16),
  inference(cnf_transformation,[],[f2371_D])).
fof(f2371_D,plain,(
  ( ! [X14,X12,X17,X15,X13] : (s(X13,sK15(X15,X14,X12,X13)) != s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2373,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X22)) | s(X19,sK15(X21,X20,X18,X19)) != s(X19,happ(s(fun(X18,X19),X20),s(X18,X23))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X18,X19),X20),s(fun(X19,bool),X21)))) )),
  inference(subsumption_resolution,[],[f2367,f249])).
fof(f2367,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X22)) | s(X19,sK15(X21,X20,X18,X19)) != s(X19,happ(s(fun(X18,X19),X20),s(X18,X23))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X18,X19),X20),s(fun(X19,bool),X21)))) )),
  inference(superposition,[],[f309,f469])).
fof(f309,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | s(X0,sK15(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X6,X3,X2,X1,X0)))))) & ((p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK15(X3,X2,X1,X0))))) & ! [X5] : s(X0,sK15(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6,X7] : ! [X21,X22] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X7,X6),X21),s(fun(X6,bool),X22)))) <=> ! [X20] : (p(s(bool,happ(s(fun(X6,bool),X22),s(X6,X20)))) => ? [X11] : s(X6,X20) = s(X6,happ(s(fun(X7,X6),X21),s(X7,X11)))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.quantHeuristics.GUESS_REWRITES|split|4')).
fof(f2370,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,X16)) | s(X13,sK15(X15,X14,X12,X13)) != s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15)))) )),
  inference(subsumption_resolution,[],[f2366,f249])).
fof(f2366,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,X16)) | s(X13,sK15(X15,X14,X12,X13)) != s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f309,f469])).
fof(f1398,plain,(
  $false | $spl2),
  inference(global_subsumption,[],[f234,f233,f235,f396,f242,f250,f249,f254,f253,f277,f397,f284,f281,f291,f290,f398,f304,f306,f305,f309,f308,f307,f312,f311,f310,f315,f314,f313,f318,f317,f316,f320,f319,f322,f321,f335,f334,f333,f332,f331,f330,f329,f339,f338,f337,f336,f343,f342,f341,f340,f347,f346,f345,f344,f351,f350,f349,f348,f358,f357,f356,f390,f354,f353,f352,f362,f361,f360,f359,f366,f365,f364,f363,f369,f389,f399,f400,f401,f402,f403,f404,f405,f406,f407,f408,f409,f410,f411,f412,f413,f422,f437,f436,f433,f432,f431,f430,f429,f428,f427,f426,f425,f440,f444,f450,f456,f469,f505,f506,f518,f517,f529,f527,f544,f543,f540,f539,f570,f569,f567,f566,f607,f624,f633,f637,f642,f657,f656,f667,f676,f707,f741,f738,f734,f733,f785,f782,f824,f967,f1053,f1298,f1302,f1305,f1308,f1311,f1314,f1317,f1319,f1320,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1336,f1338,f1340,f1342,f1344,f1346,f1348,f1300,f1375,f1382,f1397])).
fof(f1397,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl2),
  inference(subsumption_resolution,[],[f1370,f235])).
fof(f1370,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X2)) ) | $spl2),
  inference(backward_demodulation,[],[f1300,f321])).
fof(f1382,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X6,X3,X2,X1,X0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1381,f235])).
fof(f1381,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X6,X3,X2,X1,X0))))) ) | $spl2),
  inference(forward_demodulation,[],[f1380,f1300])).
fof(f1380,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X6,X3,X2,X1,X0)))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1357,f235])).
fof(f1357,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X6,X3,X2,X1,X0)))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) ) | $spl2),
  inference(backward_demodulation,[],[f1300,f307])).
fof(f1375,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f1349,f1300])).
fof(f1349,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl2),
  inference(backward_demodulation,[],[f1300,f253])).
fof(f1330,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,'|dollar|exists_1'(s(fun(X68,bool),X69)))) | s(bool,X71) = s(bool,happ(s(fun(X68,bool),X69),s(X68,X70))) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f1099,f235])).
fof(f1099,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X68,bool),X69)))) | s(bool,X71) = s(bool,happ(s(fun(X68,bool),X69),s(X68,X70))) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f397,f469])).
fof(f1329,plain,(
  ( ! [X54,X52,X55,X53] : (~sP5(X52,X53,X54) | sP6(X52,X53,X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f1091,f235])).
fof(f1091,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,'T_0')) | ~sP5(X52,X53,X54) | sP6(X52,X53,X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f390,f469])).
fof(f1328,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | p(s(bool,X47)) | ~sP6(X45,X47,X46) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f1089,f235])).
fof(f1089,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP6(X45,X47,X46) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f366,f469])).
fof(f1327,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X42)) | sP6(X41,X42,X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f1088,f235])).
fof(f1088,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X42)) | sP6(X41,X42,X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f365,f469])).
fof(f1326,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,X38)) | sP6(X37,X39,X38) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f1087,f235])).
fof(f1087,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | p(s(bool,X38)) | sP6(X37,X39,X38) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f364,f469])).
fof(f1325,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | ~sP5(X35,X34,X33) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f1086,f235])).
fof(f1086,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'T_0')) | p(s(bool,X34)) | ~p(s(bool,X35)) | ~sP5(X35,X34,X33) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f359,f469])).
fof(f1324,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP2(X30,X29,X31) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f1085,f235])).
fof(f1085,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP2(X30,X29,X31) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f340,f469])).
fof(f1323,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,X26)) | p(s(bool,X27)) | ~sP3(X26,X25,X27) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f1084,f235])).
fof(f1084,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~sP3(X26,X25,X27) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f336,f469])).
fof(f1322,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1083,f235])).
fof(f1083,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f291,f469])).
fof(f1321,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1082,f235])).
fof(f1082,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f290,f469])).
fof(f1320,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1081,f235])).
fof(f1081,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f284,f469])).
fof(f1319,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1077,f235])).
fof(f1077,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f253,f469])).
fof(f1053,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f249,f469])).
fof(f967,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f469,f469])).
fof(f824,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'F_0')),
  inference(unit_resulting_resolution,[],[f566,f430,f402,f396,f335])).
fof(f782,plain,(
  ~sP1('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f569,f429,f403,f396,f334])).
fof(f733,plain,(
  sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f406,f428,f409,f235,f333])).
fof(f734,plain,(
  sP4('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0')),
  inference(unit_resulting_resolution,[],[f430,f428,f409,f235,f333])).
fof(f738,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))),
  inference(unit_resulting_resolution,[],[f540,f430,f407,f235,f333])).
fof(f741,plain,(
  sP4('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f404,f401,f408,f235,f333])).
fof(f570,plain,(
  sP4('T_0','F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))),
  inference(unit_resulting_resolution,[],[f235,f426,f349])).
fof(f539,plain,(
  ~sP4('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f235,f413,f350])).
fof(f540,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'T_0')),
  inference(unit_resulting_resolution,[],[f235,f425,f350])).
fof(f456,plain,(
  ( ! [X0] : (sP6('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f235,f235,f365])).
fof(f433,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'F_0')),
  inference(unit_resulting_resolution,[],[f396,f234,f398])).
fof(f436,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))),
  inference(unit_resulting_resolution,[],[f396,f234,f398])).
fof(f437,plain,(
  ~sP0('const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))),'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK7,sK8),sK10),s(fun(sK8,bool),'_dst_x1x0_2'(s(fun(sK8,fun(sK8,bool)),'=_0'),s(sK8,sK9)))))),
  inference(unit_resulting_resolution,[],[f234,f234,f398])).
fof(f409,plain,(
  ( ! [X0,X1] : (sP1('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f235,f347])).
fof(f400,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f235,f291])).
fof(f399,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f235,f290])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | sP3(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (~sP1(X2,X1,X0) | p(s(bool,X0)) | ~sP3(X0,X2,X1) | ~sP2(X1,X2,X0) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f213])).
fof(f319,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK22(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22])],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X7] : ! [X21,X22] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X7,X6),X21),s(fun(X6,bool),X22)))) <=> ! [X11] : ~p(s(bool,happ(s(fun(X6,bool),X22),s(X6,happ(s(fun(X7,X6),X21),s(X7,X11)))))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.quantHeuristics.GUESS_REWRITES|split|3')).
fof(f320,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK22(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f207])).
fof(f316,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK20(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK20(X6,X3,X2,X1,X0))))))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK21(X3,X2,X1,X0))))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21])],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7] : ! [X19,X12] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X7,X6),X19),s(fun(X6,bool),X12)))) <=> ! [X20] : (~p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X20)))) => ? [X11] : ~p(s(bool,happ(s(fun(X6,bool),X12),s(X6,happ(s(fun(X7,X6),X19),s(X7,X11))))))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|1')).
fof(f317,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK21(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f204])).
fof(f318,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f204])).
fof(f313,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK18(X6,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f201])).
fof(f310,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK16(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK16(X6,X3,X2,X1,X0))))))))) & ((p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK17(X3,X2,X1,X0))))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X7] : ! [X19,X12] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X7,X6),X19),s(fun(X6,bool),X12)))) <=> ! [X20] : (p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X20)))) => ? [X11] : p(s(bool,happ(s(fun(X6,bool),X12),s(X6,happ(s(fun(X7,X6),X19),s(X7,X11))))))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|0')).
fof(f311,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK17(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f312,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f307,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X6,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f308,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK15(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f305,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X7] : ! [X21,X22] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X7,X6),X21),s(fun(X6,bool),X22)))) <=> ! [X11] : p(s(bool,happ(s(fun(X6,bool),X22),s(X6,happ(s(fun(X7,X6),X21),s(X7,X11)))))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.quantHeuristics.GUESS_REWRITES|split|2')).
fof(f306,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f192])).
fof(f304,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES',aHL_EXT)).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f278,f277])).
fof(f278,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : ! [X12,X10] : (p(s(bool,happ(s(fun(X7,bool),X12),s(X7,X10)))) => p(s(bool,happ(s(fun(X7,bool),X12),s(X7,'const.min.@_1'(s(fun(X7,bool),X12)))))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.bool.SELECT_AX')).
fof(f277,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X7,bool),X5))) = s(bool,happ(s(fun(X7,bool),X5),s(X7,'const.min.@_1'(s(fun(X7,bool),X5)))))),
  file('thm.quantHeuristics.GUESS_RULES_TWO_CASES','thm.bool.EXISTS_DEF')).
fof(f1388,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1387,f235])).
fof(f1387,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(forward_demodulation,[],[f1386,f1300])).
fof(f1386,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1359,f235])).
fof(f1359,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) ) | $spl2),
  inference(backward_demodulation,[],[f1300,f319])).
fof(f1383,plain,(
  $false | $spl2),
  inference(global_subsumption,[],[f234,f233,f235,f396,f242,f250,f249,f254,f253,f277,f397,f284,f281,f291,f290,f398,f304,f306,f305,f309,f308,f307,f312,f311,f310,f315,f314,f313,f318,f317,f316,f320,f319,f322,f321,f335,f334,f333,f332,f331,f330,f329,f339,f338,f337,f336,f343,f342,f341,f340,f347,f346,f345,f344,f351,f350,f349,f348,f358,f357,f356,f390,f354,f353,f352,f362,f361,f360,f359,f366,f365,f364,f363,f369,f389,f399,f400,f401,f402,f403,f404,f405,f406,f407,f408,f409,f410,f411,f412,f413,f422,f437,f436,f433,f432,f431,f430,f429,f428,f427,f426,f425,f440,f444,f450,f456,f469,f505,f506,f518,f517,f529,f527,f544,f543,f540,f539,f570,f569,f567,f566,f607,f624,f633,f637,f642,f657,f656,f667,f676,f707,f741,f738,f734,f733,f785,f782,f824,f967,f1053,f1298,f1302,f1305,f1308,f1311,f1314,f1317,f1319,f1320,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1336,f1338,f1340,f1342,f1344,f1346,f1348,f1300,f1375,f1382])).
fof(f1377,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1351,f235])).
fof(f1351,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(backward_demodulation,[],[f1300,f396])).
fof(f1376,plain,(
  $false | $spl2),
  inference(global_subsumption,[],[f234,f233,f235,f396,f242,f250,f249,f254,f253,f277,f397,f284,f281,f291,f290,f398,f304,f306,f305,f309,f308,f307,f312,f311,f310,f315,f314,f313,f318,f317,f316,f320,f319,f322,f321,f335,f334,f333,f332,f331,f330,f329,f339,f338,f337,f336,f343,f342,f341,f340,f347,f346,f345,f344,f351,f350,f349,f348,f358,f357,f356,f390,f354,f353,f352,f362,f361,f360,f359,f366,f365,f364,f363,f369,f389,f399,f400,f401,f402,f403,f404,f405,f406,f407,f408,f409,f410,f411,f412,f413,f422,f437,f436,f433,f432,f431,f430,f429,f428,f427,f426,f425,f440,f444,f450,f456,f469,f505,f506,f518,f517,f529,f527,f544,f543,f540,f539,f570,f569,f567,f566,f607,f624,f633,f637,f642,f657,f656,f667,f676,f707,f741,f738,f734,f733,f785,f782,f824,f967,f1053,f1298,f1302,f1305,f1308,f1311,f1314,f1317,f1319,f1320,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1336,f1338,f1340,f1342,f1344,f1346,f1348,f1300,f1375])).
fof(f1318,plain,(
  ( ! [X2,X1] : (sP6('T_0',X1,X2) | ~sP5('T_0',X1,X2)) ) | $spl14),
  inference(cnf_transformation,[],[f1318_D])).
fof(f1318_D,plain,(
  ( ! [X2,X1] : (sP6('T_0',X1,X2) | ~sP5('T_0',X1,X2)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1300,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(cnf_transformation,[],[f1300_D])).
fof(f1300_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1348,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP5('T_0',X1,X2) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f1347])).
fof(f1347,plain,(
  ( ! [X59,X57,X60,X58] : (~sP5('T_0',X58,X59) | sP6('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f1122,f505])).
fof(f1122,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | ~sP5('T_0',X58,X59) | sP6('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f390,f469])).
fof(f1315,plain,(
  ( ! [X2,X1] : (~sP6('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl12),
  inference(cnf_transformation,[],[f1315_D])).
fof(f1315_D,plain,(
  ( ! [X2,X1] : (~sP6('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1346,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f1345])).
fof(f1345,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP6('T_0',X52,X51) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f1120,f505])).
fof(f1120,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP6('T_0',X52,X51) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f366,f469])).
fof(f1312,plain,(
  ( ! [X2,X1] : (sP6('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl10),
  inference(cnf_transformation,[],[f1312_D])).
fof(f1312_D,plain,(
  ( ! [X2,X1] : (sP6('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1344,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f1343])).
fof(f1343,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | sP6('T_0',X47,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f1119,f505])).
fof(f1119,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | sP6('T_0',X47,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f365,f469])).
fof(f1309,plain,(
  ( ! [X2,X1] : (sP6('T_0',X2,X1) | p(s(bool,X1))) ) | $spl8),
  inference(cnf_transformation,[],[f1309_D])).
fof(f1309_D,plain,(
  ( ! [X2,X1] : (sP6('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1342,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f1341])).
fof(f1341,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X43)) | sP6('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f1118,f505])).
fof(f1118,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | sP6('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f364,f469])).
fof(f1306,plain,(
  ( ! [X2,X1] : (~sP5(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl6),
  inference(cnf_transformation,[],[f1306_D])).
fof(f1306_D,plain,(
  ( ! [X2,X1] : (~sP5(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1340,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X1,'T_0')) )),
  inference(condensation,[],[f1339])).
fof(f1339,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | ~sP5(X40,X39,'T_0') | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f1117,f505])).
fof(f1117,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~sP5(X40,X39,'T_0') | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f359,f469])).
fof(f1303,plain,(
  ( ! [X2,X1] : (~sP2(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl4),
  inference(cnf_transformation,[],[f1303_D])).
fof(f1303_D,plain,(
  ( ! [X2,X1] : (~sP2(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1338,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X1,'T_0',X2)) )),
  inference(condensation,[],[f1337])).
fof(f1337,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X35)) | p(s(bool,X36)) | ~sP2(X35,'T_0',X36) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f1116,f505])).
fof(f1116,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP2(X35,'T_0',X36) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f340,f469])).
fof(f1299,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl0),
  inference(cnf_transformation,[],[f1299_D])).
fof(f1299_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1336,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1335])).
fof(f1335,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | ~sP3(X31,'T_0',X32) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f1115,f505])).
fof(f1115,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | p(s(bool,X32)) | ~sP3(X31,'T_0',X32) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f336,f469])).
fof(f1317,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP5('T_0',X1,X2) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f1316])).
fof(f1316,plain,(
  ( ! [X59,X57,X60,X58] : (~sP5('T_0',X58,X59) | sP6('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f1071,f505])).
fof(f1071,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | ~sP5('T_0',X58,X59) | sP6('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f390,f469])).
fof(f1314,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f1313])).
fof(f1313,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP6('T_0',X52,X51) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f1069,f505])).
fof(f1069,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP6('T_0',X52,X51) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f366,f469])).
fof(f1311,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f1310])).
fof(f1310,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | sP6('T_0',X47,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f1068,f505])).
fof(f1068,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | sP6('T_0',X47,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f365,f469])).
fof(f1308,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f1307])).
fof(f1307,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X43)) | sP6('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f1067,f505])).
fof(f1067,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | sP6('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f364,f469])).
fof(f1305,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X1,'T_0')) )),
  inference(condensation,[],[f1304])).
fof(f1304,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | ~sP5(X40,X39,'T_0') | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f1066,f505])).
fof(f1066,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~sP5(X40,X39,'T_0') | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f359,f469])).
fof(f1302,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X1,'T_0',X2)) )),
  inference(condensation,[],[f1301])).
fof(f1301,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X35)) | p(s(bool,X36)) | ~sP2(X35,'T_0',X36) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f1065,f505])).
fof(f1065,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP2(X35,'T_0',X36) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f340,f469])).
fof(f1298,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1297])).
fof(f1297,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | ~sP3(X31,'T_0',X32) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f1064,f505])).
fof(f1064,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | p(s(bool,X32)) | ~sP3(X31,'T_0',X32) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f336,f469])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 247
% Passive clauses: 898
% Generated clauses: 5271
% Final active clauses: 236
% Final passive clauses: 586
% Input formulas: 45
% Initial clauses: 145
% 
% Duplicate literals: 224
% Trivial inequalities: 16
% Fw subsumption resolutions: 240
% Fw demodulations: 470
% Bw demodulations: 39
% Condensations: 147
% Global subsumptions: 3
% 
% Simple tautologies: 72
% Equational tautologies: 63
% Forward subsumptions: 2903
% 
% Unit resulting resolution: 1864
% Factoring: 1
% Forward superposition: 1491
% Backward superposition: 530
% Self superposition: 32
% Equality factoring: 28
% Equality resolution: 4
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 1157
% SAT solver unit clauses: 433
% SAT solver binary clauses: 349
% 
% Sat splits: 19
% Sat splitting refutations: 216
% 
% Memory used [KB]: 4733
% Time elapsed: 0.126 s
% ------------------------------
% ------------------------------
% Success in time 1.449 s
