Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.rich_list
% (103002)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1103
% Passive clauses: 8101
% Generated clauses: 41998
% Final active clauses: 1050
% Final passive clauses: 6583
% Input formulas: 515
% Initial clauses: 910
% 
% Duplicate literals: 163
% Trivial inequalities: 222
% Fw subsumption resolutions: 2031
% Fw demodulations: 5466
% Bw demodulations: 52
% 
% Simple tautologies: 354
% Equational tautologies: 27
% Forward subsumptions: 25199
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 7406
% Factoring: 406
% Forward superposition: 12981
% Backward superposition: 12133
% Self superposition: 51
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 317
% Split components: 636
% SAT solver clauses: 485
% SAT solver unit clauses: 32
% SAT solver binary clauses: 91
% 
% Sat splits: 317
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 485
% TWLsolver calls for satisfiability: 337
% 
% Memory used [KB]: 22259
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rich_list
% (103019)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 701
% Passive clauses: 4063
% Generated clauses: 40939
% Final active clauses: 683
% Final passive clauses: 3194
% Input formulas: 515
% Initial clauses: 910
% 
% Duplicate literals: 107
% Trivial inequalities: 37
% Fw subsumption resolutions: 1272
% Fw demodulations: 5823
% Bw demodulations: 184
% Condensations: 132
% Global subsumptions: 38
% 
% Simple tautologies: 352
% Equational tautologies: 16
% Forward subsumptions: 23971
% Fw demodulations to eq. taut.: 13
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 3074
% Factoring: 253
% Forward superposition: 16143
% Backward superposition: 8671
% Self superposition: 90
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 11
% Split components: 22
% SAT solver clauses: 5961
% SAT solver unit clauses: 1216
% SAT solver binary clauses: 243
% SAT solver learnt clauses: 76
% SAT solver learnt literals: 16
% 
% Sat splits: 11
% 
% TWLsolver clauses: 5809
% TWLsolver calls for satisfiability: 5047
% 
% Memory used [KB]: 37099
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rich_list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rich_list
% SZS output start Proof for thm.rich_list
fof(f99546,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f99545])).
fof(f99545,plain,(
  s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))),s('type.list.list'(sK16),sK18))) != s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))),s('type.list.list'(sK16),sK18)))),
  inference(forward_demodulation,[],[f99397,f1758])).
fof(f1758,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f569])).
fof(f569,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f176])).
fof(f176,axiom,(
  ! [X7] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7)))))),
  file('thm.rich_list.ELL_compute|split|2','thm.numeral.numeral_suc|split|1')).
fof(f99397,plain,(
  s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))),s('type.list.list'(sK16),sK18))) != s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK17))))),s('type.list.list'(sK16),sK18)))),
  inference(superposition,[],[f2651,f2000])).
fof(f2000,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.rich_list.ELL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(X0),X2))) = s(X0,'const.rich_list.ELL_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f806])).
fof(f806,plain,(
  ! [X0,X1,X2] : s(X0,'const.rich_list.ELL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(X0),X2))) = s(X0,'const.rich_list.ELL_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f805])).
fof(f805,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.rich_list.ELL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(X0),X2))) = s(X0,'const.rich_list.ELL_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8] : s(X6,'const.rich_list.ELL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.list.list'(X6),X8))) = s(X6,'const.rich_list.ELL_2'(s('type.num.num',X7),s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),X8)))))),
  file('thm.rich_list.ELL_compute|split|2','thm.rich_list.ELL|split|1')).
fof(f2651,plain,(
  s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))),s('type.list.list'(sK16),sK18))) != s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK17))),s('type.list.list'(sK16),'const.list.FRONT_1'(s('type.list.list'(sK16),sK18)))))),
  inference(forward_demodulation,[],[f2647,f1716])).
fof(f1716,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f527])).
fof(f527,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X15] : s('type.num.num',X15) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15)))),
  file('thm.rich_list.ELL_compute|split|2','thm.arithmetic.NUMERAL_DEF')).
fof(f2647,plain,(
  s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))),s('type.list.list'(sK16),sK18))) != s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK17))),s('type.list.list'(sK16),'const.list.FRONT_1'(s('type.list.list'(sK16),sK18)))))),
  inference(backward_demodulation,[],[f1716,f1689])).
fof(f1689,plain,(
  s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))),s('type.list.list'(sK16),sK18))) != s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK17))))),s('type.list.list'(sK16),'const.list.FRONT_1'(s('type.list.list'(sK16),sK18)))))),
  inference(cnf_transformation,[],[f1375])).
fof(f1375,plain,(
  s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))),s('type.list.list'(sK16),sK18))) != s(sK16,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK17))))),s('type.list.list'(sK16),'const.list.FRONT_1'(s('type.list.list'(sK16),sK18)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17,sK18])],[f1233])).
fof(f1233,plain,(
  ? [X0,X1,X2] : s(X0,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))),s('type.list.list'(X0),X2))) != s(X0,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2)))))),
  inference(ennf_transformation,[],[f518])).
fof(f518,plain,(
  ~! [X0,X1,X2] : s(X0,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))),s('type.list.list'(X0),X2))) = s(X0,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f517])).
fof(f517,plain,(
  ~! [X0] : ! [X1,X2] : s(X0,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))),s('type.list.list'(X0),X2))) = s(X0,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))),s('type.list.list'(X0),'const.list.FRONT_1'(s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(X6,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))))),s('type.list.list'(X6),X8))) = s(X6,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))))),s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(X6,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))))),s('type.list.list'(X6),X8))) = s(X6,'const.rich_list.ELL_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))))),s('type.list.list'(X6),'const.list.FRONT_1'(s('type.list.list'(X6),X8)))))),
  file('thm.rich_list.ELL_compute|split|2','thm.rich_list.ELL_compute|split|2')).
% SZS output end Proof for thm.rich_list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 989
% Passive clauses: 12079
% Generated clauses: 97898
% Final active clauses: 940
% Final passive clauses: 9367
% Input formulas: 515
% Initial clauses: 910
% 
% Duplicate literals: 11
% Trivial inequalities: 63
% Fw subsumption resolutions: 1901
% Fw demodulations: 23092
% Bw demodulations: 565
% Condensations: 88
% 
% Simple tautologies: 224
% Equational tautologies: 65
% Forward subsumptions: 58058
% Fw demodulations to eq. taut.: 1010
% Bw demodulations to eq. taut.: 15
% 
% Unit resulting resolution: 30868
% Factoring: 1
% Forward superposition: 22605
% Backward superposition: 17571
% Self superposition: 124
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 13402
% SAT solver unit clauses: 7051
% SAT solver binary clauses: 5980
% 
% Sat splits: 15
% Sat splitting refutations: 1291
% 
% Memory used [KB]: 52067
% Time elapsed: 2.091 s
% ------------------------------
% ------------------------------
% Success in time 3.451 s
