Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (62458)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 124
% Passive clauses: 2491
% Generated clauses: 22986
% Final active clauses: 104
% Final passive clauses: 2131
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 890
% Trivial inequalities: 295
% Fw subsumption resolutions: 3733
% Fw demodulations: 3351
% Bw demodulations: 15
% 
% Simple tautologies: 345
% Equational tautologies: 53
% Forward subsumptions: 9932
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 382
% Factoring: 5
% Forward superposition: 7475
% Backward superposition: 6222
% Self superposition: 216
% Equality factoring: 134
% Equality resolution: 34
% 
% Split clauses: 788
% Split components: 1603
% SAT solver clauses: 1927
% SAT solver unit clauses: 31
% SAT solver binary clauses: 146
% 
% Sat splits: 788
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1916
% TWLsolver calls for satisfiability: 251
% 
% Memory used [KB]: 9978
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (62479)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 77
% Passive clauses: 419
% Generated clauses: 10763
% Final active clauses: 61
% Final passive clauses: 212
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 888
% Trivial inequalities: 111
% Fw subsumption resolutions: 982
% Fw demodulations: 668
% Bw demodulations: 140
% Condensations: 1890
% 
% Simple tautologies: 169
% Equational tautologies: 69
% Forward subsumptions: 4347
% 
% Binary resolution: 175
% Factoring: 4
% Forward superposition: 2953
% Backward superposition: 3136
% Self superposition: 251
% Equality factoring: 202
% Equality resolution: 12
% 
% Split clauses: 94
% Split components: 188
% SAT solver clauses: 809
% SAT solver unit clauses: 15
% SAT solver binary clauses: 40
% 
% Sat splits: 94
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 809
% TWLsolver calls for satisfiability: 139
% 
% Memory used [KB]: 8443
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (62494)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 413
% Passive clauses: 4436
% Generated clauses: 158549
% Final active clauses: 89
% Final passive clauses: 306
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 6784
% Trivial inequalities: 633
% Fw subsumption resolutions: 8309
% Fw demodulations: 16935
% Bw demodulations: 3439
% Condensations: 17252
% Global subsumptions: 41
% 
% Simple tautologies: 897
% Equational tautologies: 11034
% Forward subsumptions: 90437
% Bw demodulations to eq. taut.: 13
% 
% Unit resulting resolution: 54393
% Factoring: 1
% Forward superposition: 25493
% Backward superposition: 18011
% Self superposition: 723
% Equality factoring: 454
% Equality resolution: 109
% 
% Split clauses: 588
% Split components: 1180
% SAT solver clauses: 6991
% SAT solver unit clauses: 88
% SAT solver binary clauses: 622
% 
% Sat splits: 588
% Sat splitting refutations: 1336
% 
% Memory used [KB]: 60894
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (62775)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 477
% Passive clauses: 23440
% Generated clauses: 91988
% Final active clauses: 222
% Final passive clauses: 18461
% Input formulas: 8
% Initial clauses: 14
% 
% Selected by SInE selection: 8
% SInE iterations: 1
% 
% Duplicate literals: 2563
% Trivial inequalities: 1473
% Fw subsumption resolutions: 16163
% Fw demodulations: 6096
% Bw demodulations: 137
% 
% Simple tautologies: 2086
% Equational tautologies: 103
% Forward subsumptions: 36910
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 5614
% Factoring: 4
% Forward superposition: 33650
% Backward superposition: 23948
% Self superposition: 489
% Equality factoring: 260
% Equality resolution: 148
% 
% Split clauses: 4457
% Split components: 9093
% SAT solver clauses: 4625
% SAT solver unit clauses: 9
% SAT solver binary clauses: 335
% 
% Sat splits: 4457
% Sat splitting refutations: 67
% 
% Memory used [KB]: 35820
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (62897)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 1529
% Generated clauses: 11347
% Final active clauses: 62
% Final passive clauses: 952
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 495
% Trivial inequalities: 256
% Fw subsumption resolutions: 1759
% Fw demodulations: 21
% Bw demodulations: 406
% Global subsumptions: 5
% 
% Simple tautologies: 557
% Equational tautologies: 55
% Forward subsumptions: 6426
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 138
% Forward superposition: 4220
% Backward superposition: 3303
% Self superposition: 110
% Equality factoring: 50
% Equality resolution: 15
% 
% Split clauses: 21
% Split components: 42
% SAT solver clauses: 1592
% SAT solver unit clauses: 22
% SAT solver binary clauses: 55
% 
% Sat splits: 21
% Sat splitting refutations: 35
% 
% Memory used [KB]: 7419
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (62914)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 171
% Passive clauses: 1604
% Generated clauses: 9654
% Final active clauses: 78
% Final passive clauses: 903
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 561
% Trivial inequalities: 146
% Fw subsumption resolutions: 544
% Bw subsumption resolutions: 102
% Fw demodulations: 295
% Bw demodulations: 267
% Condensations: 1089
% Global subsumptions: 4
% 
% Simple tautologies: 469
% Equational tautologies: 53
% Forward subsumptions: 4816
% Backward subsumptions: 162
% 
% Binary resolution: 1000
% Factoring: 5
% Forward superposition: 2685
% Backward superposition: 1829
% Self superposition: 119
% Equality factoring: 82
% Equality resolution: 37
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 1666
% SAT solver unit clauses: 19
% SAT solver binary clauses: 165
% 
% Sat splits: 4
% Sat splitting refutations: 52
% 
% Memory used [KB]: 8443
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 120 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f1538,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f31,f46,f50,f52,f49,f54,f48,f56,f47,f58,f59,f63,f64,f68,f66,f71,f78,f82,f85,f84,f89,f83,f87,f91,f92,f94,f95,f97,f123,f90,f153,f152,f157,f151,f155,f159,f160,f162,f197,f158,f227,f226,f231,f225,f229,f233,f234,f236,f258,f232,f242,f260,f273,f286,f301,f300,f305,f299,f303,f307,f308,f310,f332,f306,f342,f355,f375,f366,f370,f374,f378,f316,f381,f382,f72,f385,f79,f390,f389,f394,f388,f392,f73,f396,f401,f400,f405,f399,f403,f80,f493,f484,f488,f492,f497,f410,f380,f553,f551,f550,f287,f549,f555,f562,f559,f563,f565,f607,f605,f604,f609,f288,f610,f615,f572,f259,f701,f710,f60,f715,f714,f719,f713,f717,f61,f807,f805,f814,f809,f617,f560,f815,f818,f75,f993,f991,f989,f616,f988,f995,f561,f1001,f999,f1004,f1012,f1005,f1014,f1006,f1016,f1007,f1018,f1008,f1020,f1009,f1022,f1010,f1024,f1119,f1130,f1143,f1141,f1140,f1145,f503,f1146,f1151,f1159,f1157,f1156,f1155,f502,f1154,f1161,f1152,f1162,f1167,f1166,f1174,f1172,f1171,f1176,f357,f1153,f1177,f1182,f1216,f1214,f1213,f1218,f1224,f1222,f1282,f1289,f1250,f706,f705,f1384,f290,f1393,f1394,f1399,f1396,f1401,f1397,f1403,f1404,f1406,f1407,f1409,f1410,f1412,f1413,f1415,f1416,f1418,f1419,f1421,f1422,f1424,f1425,f1427,f1428,f1430,f1431,f1433,f1531,f262])).
fof(f262,plain,(
  ~p(s(bool,'T_0')) | $spl49),
  inference(cnf_transformation,[],[f262_D])).
fof(f262_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f1531,plain,(
  ~p(s(bool,'T_0')) | ($spl80 | $spl142)),
  inference(superposition,[],[f1512,f1433])).
fof(f1512,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK1))))))) ) | $spl80),
  inference(forward_demodulation,[],[f1504,f617])).
fof(f1504,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK1))))))) ) | $spl80),
  inference(unit_resulting_resolution,[],[f617,f38])).
fof(f38,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3)) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(nnf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3)))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & ~s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X9,X7,X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X9),s(X6,X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X9)))) & ~s(X6,X7) = s(X6,X8)))),
  file('thm.pred_set.DELETE_COMM','thm.pred_set.IN_DELETE')).
fof(f1433,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl142),
  inference(cnf_transformation,[],[f1433_D])).
fof(f1433_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f1431,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | ($spl18 | $spl80)),
  inference(forward_demodulation,[],[f73,f617])).
fof(f1430,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl140),
  inference(cnf_transformation,[],[f1430_D])).
fof(f1430_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f1428,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | ($spl20 | $spl80)),
  inference(forward_demodulation,[],[f74,f617])).
fof(f74,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl20),
  inference(cnf_transformation,[],[f74_D])).
fof(f74_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1427,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) | $spl139),
  inference(cnf_transformation,[],[f1427_D])).
fof(f1427_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f1425,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) | ($spl41 | $spl80)),
  inference(forward_demodulation,[],[f230,f617])).
fof(f230,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) | $spl41),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f1424,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl137),
  inference(cnf_transformation,[],[f1424_D])).
fof(f1424_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f1422,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | ($spl43 | $spl80)),
  inference(forward_demodulation,[],[f232,f617])).
fof(f1421,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl135),
  inference(cnf_transformation,[],[f1421_D])).
fof(f1421_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f1419,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | ($spl45 | $spl80)),
  inference(forward_demodulation,[],[f236,f617])).
fof(f1418,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) | $spl132),
  inference(cnf_transformation,[],[f1418_D])).
fof(f1418_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f1416,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) | ($spl46 | $spl80)),
  inference(forward_demodulation,[],[f260,f617])).
fof(f1415,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) | $spl130),
  inference(cnf_transformation,[],[f1415_D])).
fof(f1415_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f1413,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) | ($spl68 | $spl80)),
  inference(forward_demodulation,[],[f394,f617])).
fof(f1412,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl128),
  inference(cnf_transformation,[],[f1412_D])).
fof(f1412_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f1410,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | ($spl74 | $spl80)),
  inference(forward_demodulation,[],[f559,f617])).
fof(f1409,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))) | $spl126),
  inference(cnf_transformation,[],[f1409_D])).
fof(f1409_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f1407,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))) | ($spl76 | $spl80)),
  inference(forward_demodulation,[],[f561,f617])).
fof(f1406,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl125),
  inference(cnf_transformation,[],[f1406_D])).
fof(f1406_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f1404,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | ($spl79 | $spl80)),
  inference(forward_demodulation,[],[f614,f617])).
fof(f614,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl79),
  inference(cnf_transformation,[],[f614_D])).
fof(f614_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f1403,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl123),
  inference(cnf_transformation,[],[f1403_D])).
fof(f1403_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f1397,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | ($spl80 | $spl91)),
  inference(backward_demodulation,[],[f617,f998])).
fof(f998,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl91),
  inference(cnf_transformation,[],[f998_D])).
fof(f998_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f1401,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl120),
  inference(cnf_transformation,[],[f1401_D])).
fof(f1401_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f1396,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | ($spl80 | $spl88)),
  inference(backward_demodulation,[],[f617,f814])).
fof(f1399,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl119),
  inference(cnf_transformation,[],[f1399_D])).
fof(f1399_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f1394,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | ($spl5 | $spl80)),
  inference(backward_demodulation,[],[f617,f54])).
fof(f1393,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) | ~p(s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))))),
  introduced(theory_axiom,[])).
fof(f290,plain,(
  p(s(bool,'F_0')) | $spl54),
  inference(cnf_transformation,[],[f290_D])).
fof(f290_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1384,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl42 | $spl88)),
  inference(forward_demodulation,[],[f1364,f814])).
fof(f1364,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) | $spl42),
  inference(superposition,[],[f37,f233])).
fof(f37,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f705,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3)) ) | $spl82),
  inference(cnf_transformation,[],[f705_D])).
fof(f705_D,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f706,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl84),
  inference(cnf_transformation,[],[f706_D])).
fof(f706_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1250,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl30),
  inference(duplicate_literal_removal,[],[f1242])).
fof(f1242,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl30),
  inference(superposition,[],[f874,f34])).
fof(f34,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.DELETE_COMM',aHL_BOOL_CASES)).
fof(f874,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | $spl30),
  inference(resolution,[],[f835,f43])).
fof(f43,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.DELETE_COMM','thm.bool.IMP_ANTISYM_AX')).
fof(f835,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl30),
  inference(forward_demodulation,[],[f827,f824])).
fof(f824,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0')))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f33,f733,f40])).
fof(f40,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f22])).
fof(f733,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))))) ) | $spl30),
  inference(forward_demodulation,[],[f720,f93])).
fof(f93,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl30),
  inference(cnf_transformation,[],[f93_D])).
fof(f93_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f720,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))))))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f93,f38])).
fof(f33,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f12])).
fof(f12,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.DELETE_COMM',aHL_FALSITY)).
fof(f827,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'F_0')))))) ) | $spl30),
  inference(resolution,[],[f733,f40])).
fof(f1289,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | ~p(s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0))))) | p(s(bool,'F_0'))),
  introduced(theory_axiom,[])).
fof(f1282,plain,(
  ~p(s(bool,'T_0')) | ($spl52 | $spl94)),
  inference(superposition,[],[f1127,f1014])).
fof(f1127,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK2))))))) ) | $spl52),
  inference(forward_demodulation,[],[f1118,f288])).
fof(f1118,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK2))))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f288,f38])).
fof(f1222,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl116),
  inference(cnf_transformation,[],[f1222_D])).
fof(f1222_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f1224,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl100),
  inference(global_subsumption,[],[f31,f1020,f1216,f1214,f1219,f1220])).
fof(f1220,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl100),
  inference(superposition,[],[f1020,f34])).
fof(f1219,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3)))) ) | $spl100),
  inference(resolution,[],[f1020,f43])).
fof(f1218,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | ($spl50 | $spl90 | $spl100)),
  inference(resolution,[],[f1020,f1002])).
fof(f1002,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl50 | $spl90)),
  inference(backward_demodulation,[],[f999,f996])).
fof(f996,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) ) | $spl50),
  inference(resolution,[],[f285,f43])).
fof(f285,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) | $spl50),
  inference(cnf_transformation,[],[f285_D])).
fof(f285_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1213,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | ($spl50 | $spl90 | $spl100)),
  inference(unit_resulting_resolution,[],[f1020,f1002])).
fof(f1214,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl100),
  inference(unit_resulting_resolution,[],[f32,f1020,f43])).
fof(f32,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.DELETE_COMM',aHL_TRUTH)).
fof(f1216,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl100),
  inference(unit_resulting_resolution,[],[f32,f1020,f43])).
fof(f1182,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl114),
  inference(cnf_transformation,[],[f1182_D])).
fof(f1182_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f1177,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | ($spl30 | $spl67)),
  inference(resolution,[],[f380,f835])).
fof(f1153,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK3)))) | $spl111),
  inference(cnf_transformation,[],[f1153_D])).
fof(f1153_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK3)))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f357,plain,(
  s(sK0,sK2) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)) | $spl64),
  inference(cnf_transformation,[],[f357_D])).
fof(f357_D,plain,(
  s(sK0,sK2) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1176,plain,(
  s(sK0,sK2) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK3)))) | $spl67),
  inference(resolution,[],[f380,f39])).
fof(f39,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f1171,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | ($spl30 | $spl67)),
  inference(unit_resulting_resolution,[],[f380,f835])).
fof(f1172,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f33,f380,f40])).
fof(f1174,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f33,f380,f40])).
fof(f1166,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl112),
  inference(cnf_transformation,[],[f1166_D])).
fof(f1166_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f1167,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl62),
  inference(global_subsumption,[],[f31,f354,f1159,f1157,f1155,f1154,f1161,f1163,f1164])).
fof(f1164,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl62),
  inference(superposition,[],[f354,f34])).
fof(f1163,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) ) | $spl62),
  inference(resolution,[],[f354,f43])).
fof(f354,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) | $spl62),
  inference(cnf_transformation,[],[f354_D])).
fof(f354_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1162,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | ($spl50 | $spl62 | $spl90)),
  inference(resolution,[],[f354,f1002])).
fof(f1152,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK3)))) | $spl110),
  inference(cnf_transformation,[],[f1152_D])).
fof(f1152_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK3)))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1161,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK3)))) | $spl62),
  inference(resolution,[],[f354,f37])).
fof(f1154,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK3)))) | $spl62),
  inference(unit_resulting_resolution,[],[f354,f37])).
fof(f502,plain,(
  s(sK0,sK1) != s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)) | $spl73),
  inference(cnf_transformation,[],[f502_D])).
fof(f502_D,plain,(
  s(sK0,sK1) != s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f1155,plain,(
  s(sK0,sK1) != s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)) | $spl62),
  inference(unit_resulting_resolution,[],[f354,f38])).
fof(f1156,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | ($spl50 | $spl62 | $spl90)),
  inference(unit_resulting_resolution,[],[f354,f1002])).
fof(f1157,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f32,f354,f43])).
fof(f1159,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f32,f354,f43])).
fof(f1151,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl108),
  inference(cnf_transformation,[],[f1151_D])).
fof(f1151_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1146,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | ($spl30 | $spl63)),
  inference(resolution,[],[f355,f835])).
fof(f503,plain,(
  s(sK0,sK1) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)) | $spl72),
  inference(cnf_transformation,[],[f503_D])).
fof(f503_D,plain,(
  s(sK0,sK1) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1145,plain,(
  s(sK0,sK1) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK3)))) | $spl63),
  inference(resolution,[],[f355,f39])).
fof(f1140,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | ($spl30 | $spl63)),
  inference(unit_resulting_resolution,[],[f355,f835])).
fof(f1141,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl63),
  inference(unit_resulting_resolution,[],[f33,f355,f40])).
fof(f1143,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl63),
  inference(unit_resulting_resolution,[],[f33,f355,f40])).
fof(f1130,plain,(
  s(sK0,sK1) != s(sK0,sK2) | $spl107),
  inference(cnf_transformation,[],[f1130_D])).
fof(f1130_D,plain,(
  s(sK0,sK1) != s(sK0,sK2) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f1119,plain,(
  s(sK0,sK1) != s(sK0,sK2) | ($spl52 | $spl81)),
  inference(superposition,[],[f616,f288])).
fof(f1024,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl104),
  inference(cnf_transformation,[],[f1024_D])).
fof(f1024_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1010,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | ($spl52 | $spl90)),
  inference(backward_demodulation,[],[f288,f999])).
fof(f1022,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl102),
  inference(cnf_transformation,[],[f1022_D])).
fof(f1022_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1009,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | ($spl52 | $spl78)),
  inference(backward_demodulation,[],[f288,f615])).
fof(f1020,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3)))) | $spl100),
  inference(cnf_transformation,[],[f1020_D])).
fof(f1020_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3)))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1008,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3)))) | ($spl52 | $spl76)),
  inference(backward_demodulation,[],[f288,f561])).
fof(f1018,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl99),
  inference(cnf_transformation,[],[f1018_D])).
fof(f1018_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f1007,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | ($spl52 | $spl75)),
  inference(backward_demodulation,[],[f288,f558])).
fof(f558,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl75),
  inference(cnf_transformation,[],[f558_D])).
fof(f558_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f1016,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl96),
  inference(cnf_transformation,[],[f1016_D])).
fof(f1016_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f1006,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | ($spl44 | $spl52)),
  inference(backward_demodulation,[],[f288,f235])).
fof(f235,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl44),
  inference(cnf_transformation,[],[f235_D])).
fof(f235_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1014,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl94),
  inference(cnf_transformation,[],[f1014_D])).
fof(f1014_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f1005,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | ($spl42 | $spl52)),
  inference(backward_demodulation,[],[f288,f233])).
fof(f1012,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl93),
  inference(cnf_transformation,[],[f1012_D])).
fof(f1012_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f1004,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | ($spl19 | $spl52)),
  inference(backward_demodulation,[],[f288,f72])).
fof(f999,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl90),
  inference(cnf_transformation,[],[f999_D])).
fof(f999_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f1001,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl50),
  inference(global_subsumption,[],[f31,f285,f993,f991,f989,f988,f995,f996,f997])).
fof(f997,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl50),
  inference(superposition,[],[f285,f34])).
fof(f561,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | $spl76),
  inference(cnf_transformation,[],[f561_D])).
fof(f561_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f995,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | $spl50),
  inference(resolution,[],[f285,f37])).
fof(f988,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | $spl50),
  inference(unit_resulting_resolution,[],[f285,f37])).
fof(f616,plain,(
  s(sK0,sK1) != s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) | $spl81),
  inference(cnf_transformation,[],[f616_D])).
fof(f616_D,plain,(
  s(sK0,sK1) != s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f989,plain,(
  s(sK0,sK1) != s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) | $spl50),
  inference(unit_resulting_resolution,[],[f285,f38])).
fof(f991,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl50),
  inference(unit_resulting_resolution,[],[f32,f285,f43])).
fof(f993,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl50),
  inference(unit_resulting_resolution,[],[f32,f285,f43])).
fof(f75,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl21),
  inference(cnf_transformation,[],[f75_D])).
fof(f75_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f818,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | ($spl5 | $spl44)),
  inference(backward_demodulation,[],[f235,f54])).
fof(f815,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | ~p(s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))))),
  introduced(theory_axiom,[])).
fof(f560,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | $spl77),
  inference(cnf_transformation,[],[f560_D])).
fof(f560_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f617,plain,(
  s(sK0,sK1) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) | $spl80),
  inference(cnf_transformation,[],[f617_D])).
fof(f617_D,plain,(
  s(sK0,sK1) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f809,plain,(
  s(sK0,sK1) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | $spl51),
  inference(resolution,[],[f286,f39])).
fof(f814,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl88),
  inference(cnf_transformation,[],[f814_D])).
fof(f814_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f805,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl51),
  inference(unit_resulting_resolution,[],[f33,f286,f40])).
fof(f807,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))))) | $spl51),
  inference(unit_resulting_resolution,[],[f33,f286,f40])).
fof(f61,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl10),
  inference(cnf_transformation,[],[f61_D])).
fof(f61_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f717,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl31),
  inference(trivial_inequality_removal,[],[f716])).
fof(f716,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl31),
  inference(superposition,[],[f94,f34])).
fof(f713,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f94,f34])).
fof(f719,plain,(
  p(s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0))))) | $spl86),
  inference(cnf_transformation,[],[f719_D])).
fof(f719_D,plain,(
  p(s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f714,plain,(
  p(s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0))))) | $spl31),
  inference(unit_resulting_resolution,[],[f33,f94,f40])).
fof(f715,plain,(
  p(s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0))))) | $spl31),
  inference(unit_resulting_resolution,[],[f33,f94,f40])).
fof(f60,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f60_D])).
fof(f60_D,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f710,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | ($spl3 | $spl28)),
  inference(backward_demodulation,[],[f91,f52])).
fof(f701,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f693])).
fof(f693,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f523,f34])).
fof(f523,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f176,f43])).
fof(f176,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl10),
  inference(forward_demodulation,[],[f170,f167])).
fof(f167,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'T_0')))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f33,f120,f40])).
fof(f120,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'T_0'))))))) ) | $spl10),
  inference(forward_demodulation,[],[f107,f61])).
fof(f107,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'T_0'))))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f61,f38])).
fof(f170,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'T_0')))))) ) | $spl10),
  inference(resolution,[],[f120,f40])).
fof(f259,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) | $spl47),
  inference(cnf_transformation,[],[f259_D])).
fof(f259_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f572,plain,(
  p(s(bool,'F_0')) | s(sK0,sK1) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) | $spl44),
  inference(superposition,[],[f39,f235])).
fof(f615,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl78),
  inference(cnf_transformation,[],[f615_D])).
fof(f615_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f610,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | ($spl10 | $spl47)),
  inference(resolution,[],[f259,f176])).
fof(f288,plain,(
  s(sK0,sK2) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) | $spl52),
  inference(cnf_transformation,[],[f288_D])).
fof(f288_D,plain,(
  s(sK0,sK2) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f609,plain,(
  s(sK0,sK2) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | $spl47),
  inference(resolution,[],[f259,f39])).
fof(f604,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | ($spl10 | $spl47)),
  inference(unit_resulting_resolution,[],[f259,f176])).
fof(f605,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl47),
  inference(unit_resulting_resolution,[],[f33,f259,f40])).
fof(f607,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl47),
  inference(unit_resulting_resolution,[],[f33,f259,f40])).
fof(f565,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | ($spl5 | $spl44)),
  inference(backward_demodulation,[],[f235,f54])).
fof(f563,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) | p(s(bool,'F_0'))),
  introduced(theory_axiom,[])).
fof(f559,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl74),
  inference(cnf_transformation,[],[f559_D])).
fof(f559_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f562,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl46),
  inference(global_subsumption,[],[f31,f260,f553,f551,f550,f549,f555,f556,f557])).
fof(f557,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl46),
  inference(superposition,[],[f260,f34])).
fof(f556,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) ) | $spl46),
  inference(resolution,[],[f260,f43])).
fof(f555,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | $spl46),
  inference(resolution,[],[f260,f37])).
fof(f549,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | $spl46),
  inference(unit_resulting_resolution,[],[f260,f37])).
fof(f287,plain,(
  s(sK0,sK2) != s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) | $spl53),
  inference(cnf_transformation,[],[f287_D])).
fof(f287_D,plain,(
  s(sK0,sK2) != s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f550,plain,(
  s(sK0,sK2) != s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) | $spl46),
  inference(unit_resulting_resolution,[],[f260,f38])).
fof(f551,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl46),
  inference(unit_resulting_resolution,[],[f32,f260,f43])).
fof(f553,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))))) | $spl46),
  inference(unit_resulting_resolution,[],[f32,f260,f43])).
fof(f380,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) | $spl67),
  inference(cnf_transformation,[],[f380_D])).
fof(f380_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f410,plain,(
  p(s(bool,'F_0')) | s(sK0,sK1) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) | $spl60),
  inference(superposition,[],[f39,f309])).
fof(f309,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl60),
  inference(cnf_transformation,[],[f309_D])).
fof(f309_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f497,plain,(
  $false | ($spl33 | $spl54)),
  inference(global_subsumption,[],[f31,f97,f290,f491,f489,f487])).
fof(f487,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl54)),
  inference(unit_resulting_resolution,[],[f97,f290,f43])).
fof(f489,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl54),
  inference(unit_resulting_resolution,[],[f32,f290,f43])).
fof(f491,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl54)),
  inference(unit_resulting_resolution,[],[f97,f290,f43])).
fof(f492,plain,(
  $false | ($spl33 | $spl54)),
  inference(unit_resulting_resolution,[],[f32,f97,f290,f43])).
fof(f488,plain,(
  $false | ($spl33 | $spl54)),
  inference(unit_resulting_resolution,[],[f32,f97,f290,f43])).
fof(f484,plain,(
  $false | $spl54),
  inference(unit_resulting_resolution,[],[f33,f290])).
fof(f493,plain,(
  $false | $spl54),
  inference(resolution,[],[f290,f33])).
fof(f80,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl22),
  inference(cnf_transformation,[],[f80_D])).
fof(f80_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f403,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl61),
  inference(trivial_inequality_removal,[],[f402])).
fof(f402,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl61),
  inference(superposition,[],[f310,f34])).
fof(f399,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl61),
  inference(unit_resulting_resolution,[],[f310,f34])).
fof(f405,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) | $spl70),
  inference(cnf_transformation,[],[f405_D])).
fof(f405_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f400,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) | $spl61),
  inference(unit_resulting_resolution,[],[f33,f310,f40])).
fof(f401,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) | $spl61),
  inference(unit_resulting_resolution,[],[f33,f310,f40])).
fof(f396,plain,(
  p(s(bool,'F_0')) | ($spl20 | $spl40)),
  inference(backward_demodulation,[],[f74,f231])).
fof(f73,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl18),
  inference(cnf_transformation,[],[f73_D])).
fof(f73_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f392,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl45),
  inference(trivial_inequality_removal,[],[f391])).
fof(f391,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl45),
  inference(superposition,[],[f236,f34])).
fof(f388,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl45),
  inference(unit_resulting_resolution,[],[f236,f34])).
fof(f394,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) | $spl68),
  inference(cnf_transformation,[],[f394_D])).
fof(f394_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f389,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) | $spl45),
  inference(unit_resulting_resolution,[],[f33,f236,f40])).
fof(f390,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))))) | $spl45),
  inference(unit_resulting_resolution,[],[f33,f236,f40])).
fof(f79,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl23),
  inference(cnf_transformation,[],[f79_D])).
fof(f79_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f385,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | ($spl9 | $spl58)),
  inference(backward_demodulation,[],[f307,f58])).
fof(f72,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl19),
  inference(cnf_transformation,[],[f72_D])).
fof(f72_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f382,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | ($spl5 | $spl42)),
  inference(backward_demodulation,[],[f233,f54])).
fof(f381,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) | $spl66),
  inference(cnf_transformation,[],[f381_D])).
fof(f381_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f316,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) | $spl22),
  inference(superposition,[],[f37,f80])).
fof(f378,plain,(
  $false | ($spl33 | $spl49)),
  inference(global_subsumption,[],[f31,f97,f262,f373,f371,f369])).
fof(f369,plain,(
  p(s(bool,'F_0')) | ($spl33 | $spl49)),
  inference(unit_resulting_resolution,[],[f97,f262,f40])).
fof(f371,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl49),
  inference(unit_resulting_resolution,[],[f33,f262,f40])).
fof(f373,plain,(
  p(s(bool,'F_0')) | ($spl33 | $spl49)),
  inference(unit_resulting_resolution,[],[f97,f262,f40])).
fof(f374,plain,(
  $false | ($spl33 | $spl49)),
  inference(unit_resulting_resolution,[],[f33,f97,f262,f40])).
fof(f370,plain,(
  $false | ($spl33 | $spl49)),
  inference(unit_resulting_resolution,[],[f33,f97,f262,f40])).
fof(f366,plain,(
  $false | $spl49),
  inference(unit_resulting_resolution,[],[f32,f262])).
fof(f375,plain,(
  $false | $spl49),
  inference(resolution,[],[f262,f32])).
fof(f355,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) | $spl63),
  inference(cnf_transformation,[],[f355_D])).
fof(f355_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f342,plain,(
  p(s(bool,'F_0')) | s(sK0,sK2) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) | $spl24),
  inference(superposition,[],[f39,f81])).
fof(f81,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl24),
  inference(cnf_transformation,[],[f81_D])).
fof(f81_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f306,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl59),
  inference(cnf_transformation,[],[f306_D])).
fof(f306_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f332,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl22),
  inference(global_subsumption,[],[f31,f80,f315])).
fof(f315,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))) = s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))) | $spl22),
  inference(superposition,[],[f36,f80])).
fof(f36,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f14])).
fof(f14,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X9,X12] : (s(fun(X6,bool),X9) = s(fun(X6,bool),X12) <=> ! [X7] : s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X9))) = s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X12))))),
  file('thm.pred_set.DELETE_COMM','thm.pred_set.EXTENSION')).
fof(f310,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl61),
  inference(cnf_transformation,[],[f310_D])).
fof(f310_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f308,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | ($spl9 | $spl24)),
  inference(backward_demodulation,[],[f81,f58])).
fof(f307,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl58),
  inference(cnf_transformation,[],[f307_D])).
fof(f307_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f303,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl25),
  inference(trivial_inequality_removal,[],[f302])).
fof(f302,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl25),
  inference(superposition,[],[f82,f34])).
fof(f299,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl25),
  inference(unit_resulting_resolution,[],[f82,f34])).
fof(f305,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) | $spl56),
  inference(cnf_transformation,[],[f305_D])).
fof(f305_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f300,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) | $spl25),
  inference(unit_resulting_resolution,[],[f33,f82,f40])).
fof(f301,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) | $spl25),
  inference(unit_resulting_resolution,[],[f33,f82,f40])).
fof(f286,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) | $spl51),
  inference(cnf_transformation,[],[f286_D])).
fof(f286_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f273,plain,(
  p(s(bool,'F_0')) | s(sK0,sK2) = s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)))))) | $spl20),
  inference(superposition,[],[f39,f74])).
fof(f260,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) | $spl46),
  inference(cnf_transformation,[],[f260_D])).
fof(f260_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f242,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2)))))) | $spl18),
  inference(superposition,[],[f37,f73])).
fof(f232,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl43),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f258,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl18),
  inference(global_subsumption,[],[f31,f73,f241])).
fof(f241,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))) = s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))) | $spl18),
  inference(superposition,[],[f36,f73])).
fof(f236,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl45),
  inference(cnf_transformation,[],[f236_D])).
fof(f236_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f234,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | ($spl5 | $spl20)),
  inference(backward_demodulation,[],[f74,f54])).
fof(f233,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl42),
  inference(cnf_transformation,[],[f233_D])).
fof(f233_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f229,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl21),
  inference(trivial_inequality_removal,[],[f228])).
fof(f228,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl21),
  inference(superposition,[],[f75,f34])).
fof(f225,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl21),
  inference(unit_resulting_resolution,[],[f75,f34])).
fof(f231,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) | $spl40),
  inference(cnf_transformation,[],[f231_D])).
fof(f231_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f226,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f33,f75,f40])).
fof(f227,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f33,f75,f40])).
fof(f158,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl37),
  inference(cnf_transformation,[],[f158_D])).
fof(f158_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f197,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl14),
  inference(global_subsumption,[],[f31,f66,f185])).
fof(f185,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))) = s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))) | $spl14),
  inference(superposition,[],[f44,f66])).
fof(f44,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.DELETE_COMM',aHL_EXT)).
fof(f162,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl39),
  inference(cnf_transformation,[],[f162_D])).
fof(f162_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f160,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | ($spl7 | $spl16)),
  inference(backward_demodulation,[],[f67,f56])).
fof(f67,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f67_D])).
fof(f67_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f159,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl36),
  inference(cnf_transformation,[],[f159_D])).
fof(f159_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f155,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl17),
  inference(trivial_inequality_removal,[],[f154])).
fof(f154,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl17),
  inference(superposition,[],[f68,f34])).
fof(f151,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl17),
  inference(unit_resulting_resolution,[],[f68,f34])).
fof(f157,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0))))) | $spl34),
  inference(cnf_transformation,[],[f157_D])).
fof(f157_D,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f152,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f33,f68,f40])).
fof(f153,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f33,f68,f40])).
fof(f90,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl29),
  inference(cnf_transformation,[],[f90_D])).
fof(f90_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f123,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl10),
  inference(global_subsumption,[],[f31,f61,f120,f113])).
fof(f113,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))) = s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))) | $spl10),
  inference(superposition,[],[f44,f61])).
fof(f97,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl33),
  inference(cnf_transformation,[],[f97_D])).
fof(f97_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f95,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl10 | $spl31)),
  inference(forward_demodulation,[],[f94,f61])).
fof(f94,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl31),
  inference(cnf_transformation,[],[f94_D])).
fof(f94_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f92,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | ($spl3 | $spl12)),
  inference(backward_demodulation,[],[f62,f52])).
fof(f62,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f62_D])).
fof(f62_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f91,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl28),
  inference(cnf_transformation,[],[f91_D])).
fof(f91_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f87,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl13),
  inference(trivial_inequality_removal,[],[f86])).
fof(f86,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl13),
  inference(superposition,[],[f63,f34])).
fof(f83,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl13),
  inference(unit_resulting_resolution,[],[f63,f34])).
fof(f89,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0))))) | $spl26),
  inference(cnf_transformation,[],[f89_D])).
fof(f89_D,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f84,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f33,f63,f40])).
fof(f85,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f33,f63,f40])).
fof(f82,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | $spl25),
  inference(cnf_transformation,[],[f82_D])).
fof(f82_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f78,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl9),
  inference(superposition,[],[f58,f34])).
fof(f71,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl5),
  inference(superposition,[],[f54,f34])).
fof(f66,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl14),
  inference(cnf_transformation,[],[f66_D])).
fof(f66_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f68,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl17),
  inference(cnf_transformation,[],[f68_D])).
fof(f68_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f64,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl7),
  inference(superposition,[],[f56,f34])).
fof(f63,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl13),
  inference(cnf_transformation,[],[f63_D])).
fof(f63_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f59,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl3),
  inference(superposition,[],[f52,f34])).
fof(f58,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl9),
  inference(cnf_transformation,[],[f58_D])).
fof(f58_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f47,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f46,f36])).
fof(f56,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f56_D])).
fof(f56_D,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f48,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f46,f44])).
fof(f54,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl5),
  inference(cnf_transformation,[],[f54_D])).
fof(f54_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f49,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f46,f36])).
fof(f52,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl3),
  inference(cnf_transformation,[],[f52_D])).
fof(f52_D,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f50,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2)),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f46,f44])).
fof(f46,plain,(
  s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))) != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))) | $spl1),
  inference(cnf_transformation,[],[f46_D])).
fof(f46_D,plain,(
  s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))) != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f31,plain,(
  s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))) != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1))),s(sK0,sK2))) != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK2))),s(sK0,sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f20])).
fof(f20,plain,(
  ? [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1))),s(X0,X2))) != s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))),s(X0,X1)))),
  inference(ennf_transformation,[],[f11])).
fof(f11,plain,(
  ~! [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1))),s(X0,X2))) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))),s(X0,X1)))),
  inference(flattening,[],[f10])).
fof(f10,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1))),s(X0,X2))) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))),s(X0,X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X9),s(X6,X7))),s(X6,X8))) = s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X9),s(X6,X8))),s(X6,X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X9),s(X6,X7))),s(X6,X8))) = s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X9),s(X6,X8))),s(X6,X7)))),
  file('thm.pred_set.DELETE_COMM','thm.pred_set.DELETE_COMM')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 90
% Passive clauses: 188
% Generated clauses: 1527
% Final active clauses: 33
% Final passive clauses: 26
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 5
% Trivial inequalities: 7
% Fw demodulations: 433
% Bw demodulations: 64
% Global subsumptions: 10
% 
% Simple tautologies: 11
% Equational tautologies: 43
% Forward subsumptions: 685
% Backward subsumptions: 3
% 
% Binary resolution: 90
% Unit resulting resolution: 273
% Forward superposition: 143
% Backward superposition: 317
% 
% Split clauses: 17
% Split components: 38
% SAT solver clauses: 478
% SAT solver unit clauses: 6
% SAT solver binary clauses: 292
% 
% Sat splits: 17
% Sat splitting refutations: 10
% 
% Memory used [KB]: 1918
% Time elapsed: 0.063 s
% ------------------------------
% ------------------------------
% Success in time 18.068 s
