Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.int_arith
% (84624)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 745
% Passive clauses: 8777
% Generated clauses: 20266
% Final active clauses: 687
% Final passive clauses: 7424
% Input formulas: 51
% Initial clauses: 161
% 
% Duplicate literals: 477
% Trivial inequalities: 77
% Fw subsumption resolutions: 1625
% Fw demodulations: 785
% Bw demodulations: 57
% 
% Simple tautologies: 179
% Equational tautologies: 64
% Forward subsumptions: 8050
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 10309
% Factoring: 1316
% Forward superposition: 3210
% Backward superposition: 1914
% Self superposition: 60
% Equality factoring: 26
% Equality resolution: 45
% 
% Split clauses: 358
% Split components: 722
% SAT solver clauses: 748
% SAT solver unit clauses: 29
% SAT solver binary clauses: 165
% 
% Sat splits: 358
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 746
% TWLsolver calls for satisfiability: 474
% 
% Memory used [KB]: 11129
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.int_arith
% (84634)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 401
% Passive clauses: 2475
% Generated clauses: 11177
% Final active clauses: 395
% Final passive clauses: 2052
% Input formulas: 51
% Initial clauses: 161
% 
% Duplicate literals: 167
% Trivial inequalities: 13
% Fw subsumption resolutions: 728
% Fw demodulations: 506
% Bw demodulations: 24
% Condensations: 241
% Global subsumptions: 13
% 
% Simple tautologies: 188
% Equational tautologies: 22
% Forward subsumptions: 4017
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 3809
% Factoring: 333
% Forward superposition: 1685
% Backward superposition: 936
% Self superposition: 45
% Equality factoring: 26
% Equality resolution: 16
% 
% Split clauses: 51
% Split components: 107
% SAT solver clauses: 2950
% SAT solver unit clauses: 184
% SAT solver binary clauses: 76
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 51
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2898
% TWLsolver calls for satisfiability: 1471
% 
% Memory used [KB]: 12665
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.int_arith
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.int_arith
% SZS output start Proof for thm.int_arith
fof(f11021,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1100,f1104,f1107,f1110,f1113,f1116,f1119,f1122,f1125,f1128,f1131,f1134,f1137,f1140,f1143,f1169,f1101,f1171,f1105,f1173,f1108,f1175,f1111,f1177,f1114,f1179,f1117,f1181,f1120,f1183,f1123,f1185,f1126,f1187,f1129,f1189,f1132,f1191,f1135,f1193,f1138,f1195,f1141,f1197,f1144,f1208,f1209,f4363,f4364,f4371,f4368,f4365,f4366,f4374,f4367,f4376,f4369,f4373,f5634,f5635,f5647,f5645,f6604,f6664,f8720,f8723,f1102,f8721,f8803,f8806,f8804,f8808,f8809,f8811,f8813,f8815,f8819,f8821,f8825,f8827,f8829,f8831,f8835,f8837,f8841,f8843,f8847,f8851,f8857,f8863,f8873,f8875,f8885,f8894,f8905,f8910,f8916,f8921,f8923,f8929,f8934,f8940,f8945,f8947,f8949,f8951,f8953,f9233,f9234,f8965,f8967,f8972,f8977,f8985,f8995,f8997,f9003,f9005,f9007,f9012,f9014,f9019,f9021,f9026,f9028,f9033,f9035,f9041,f9046,f9054,f9060,f9062,f9076,f9078,f9080,f9089,f9091,f9100,f9102,f9111,f9113,f9122,f9124,f9126,f9132,f9137,f9142,f9144,f9149,f9151,f9157,f9162,f9168,f9173,f9178,f9180,f9182,f9187,f9189,f9191,f9195,f9199,f9201,f9203,f9207,f9209,f9213,f9235,f9236,f9237,f9238,f9239,f9240,f9241,f9242,f9243,f9244,f9245,f9246,f9247,f9248,f9249,f9250,f9251,f9252,f9253,f9254,f9255,f9256,f9257,f9258,f9259,f9260,f9261,f9262,f9263,f9264,f9265,f9266,f9267,f9268,f9269,f9270,f9271,f9272,f9273,f9274,f9275,f9276,f9277,f9278,f9279,f9280,f9281,f9282,f9283,f9284,f9285,f9286,f9287,f9288,f9289,f9290,f9291,f9292,f9293,f9294,f9295,f9296,f9322,f9324,f9323,f9326,f8805,f9327,f9776,f9854,f9852,f10344,f10345,f6662,f10497,f10510,f10514,f10518,f10523,f10529,f10533,f10540,f10548,f10551,f10555,f10571,f10574,f10576,f10578,f10580,f10583,f10586,f10589,f10592,f10594,f10597,f10611,f10624,f10699,f10713,f10734,f10736,f10749,f10758,f10772,f10849,f10854,f10993])).
fof(f10993,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f304,f10345])).
fof(f304,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X7,X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6)))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.integer.INT_MUL')).
fof(f10854,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10853,f1621])).
fof(f1621,plain,(
  ( ! [X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f304,f249])).
fof(f249,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X14] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.integer.INT_MUL_RZERO')).
fof(f10853,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))))) ) | $spl48),
  inference(forward_demodulation,[],[f10852,f2237])).
fof(f2237,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f1655,f245])).
fof(f245,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X14] : s('type.integer.int',X14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X14)))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.integer.INT_NEGNEG')).
fof(f1655,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(unit_resulting_resolution,[],[f248,f338])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2)))))),
  inference(nnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))) <=> s('type.integer.int',X2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X14,X15,X22] : (s('type.integer.int',X14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X22))) <=> s('type.integer.int',X22) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15))))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.int_arith.eq_move_left_left')).
fof(f248,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X14] : s('type.integer.int',X14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.integer.INT_ADD_RID')).
fof(f10852,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))) ) | $spl48),
  inference(forward_demodulation,[],[f10487,f246])).
fof(f246,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X20] : s('type.num.num',X20) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X20)))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.gcd.GCD_|48|L')).
fof(f10487,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f10185])).
fof(f10185,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))) )),
  inference(superposition,[],[f5706,f250])).
fof(f250,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X14] : s('type.integer.int',X14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X14)))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.integer.INT_MUL_LID')).
fof(f5706,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))) )),
  inference(unit_resulting_resolution,[],[f1623,f1711])).
fof(f1711,plain,(
  ( ! [X6,X4,X5,X3] : (s('type.integer.int',X5) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))))),s('type.integer.int',X6))) | s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X4)))))) )),
  inference(forward_demodulation,[],[f1672,f245])).
fof(f1672,plain,(
  ( ! [X6,X4,X5,X3] : (s('type.integer.int',X5) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))))),s('type.integer.int',X6))) | s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X4)))))))))) )),
  inference(superposition,[],[f338,f301])).
fof(f301,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X14,X15] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',X15)))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.integer.INT_MUL_CALCULATE|split|2')).
fof(f1623,plain,(
  ( ! [X2,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))) )),
  inference(superposition,[],[f236,f304])).
fof(f236,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X2,X3] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f124])).
fof(f124,plain,(
  ? [X0,X1] : ! [X2,X3] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ~! [X0,X1] : ? [X2,X3] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ? [X8,X9] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X6),s('type.num.num',X7))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ? [X8,X9] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X6),s('type.num.num',X7))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.int_arith.INT_LINEAR_GCD')).
fof(f10849,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10848,f249])).
fof(f10848,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl48),
  inference(forward_demodulation,[],[f10847,f2237])).
fof(f10847,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))) ) | $spl48),
  inference(forward_demodulation,[],[f10485,f246])).
fof(f10485,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f10104])).
fof(f10104,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))) )),
  inference(superposition,[],[f5703,f250])).
fof(f5703,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))) )),
  inference(unit_resulting_resolution,[],[f236,f1711])).
fof(f10772,plain,(
  $false | ($spl48 | $spl51)),
  inference(subsumption_resolution,[],[f10458,f1621])).
fof(f10458,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK16(sK11,sK10)),s('type.num.num','const.num.|48|_0'))))) | ($spl48 | $spl51)),
  inference(backward_demodulation,[],[f6662,f6664])).
fof(f10758,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10757,f1489])).
fof(f1489,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(backward_demodulation,[],[f1475,f1445])).
fof(f1445,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f301,f249])).
fof(f1475,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f407,f249])).
fof(f407,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(forward_demodulation,[],[f302,f301])).
fof(f302,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X14,X15] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X14))),s('type.integer.int',X15))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',X15)))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.integer.INT_MUL_CALCULATE|split|1')).
fof(f10757,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10756,f2237])).
fof(f10756,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10450,f246])).
fof(f10450,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f5254])).
fof(f5254,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))),
  inference(forward_demodulation,[],[f5253,f245])).
fof(f5253,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))))))),
  inference(unit_resulting_resolution,[],[f2284,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) )),
  inference(cnf_transformation,[],[f207])).
fof(f2284,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))))),
  inference(unit_resulting_resolution,[],[f2149,f339])).
fof(f2149,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))),
  inference(unit_resulting_resolution,[],[f2101,f339])).
fof(f2101,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))))),
  inference(unit_resulting_resolution,[],[f2100,f339])).
fof(f2100,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(forward_demodulation,[],[f2099,f245])).
fof(f2099,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))))))),
  inference(unit_resulting_resolution,[],[f2049,f339])).
fof(f2049,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))),
  inference(superposition,[],[f1745,f250])).
fof(f1745,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))) )),
  inference(superposition,[],[f1420,f407])).
fof(f1420,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))) )),
  inference(superposition,[],[f236,f250])).
fof(f10749,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10748])).
fof(f10748,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10747,f246])).
fof(f10747,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) | $spl48),
  inference(forward_demodulation,[],[f10746,f247])).
fof(f247,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X14] : s('type.integer.int',X14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14)))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.integer.INT_ADD_LID')).
fof(f10746,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))))) | $spl48),
  inference(forward_demodulation,[],[f10446,f1621])).
fof(f10446,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f4911])).
fof(f4911,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))) )),
  inference(unit_resulting_resolution,[],[f2214,f339])).
fof(f2214,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))))) )),
  inference(unit_resulting_resolution,[],[f2211,f339])).
fof(f2211,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK10)))))))) )),
  inference(superposition,[],[f2052,f304])).
fof(f2052,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))))) )),
  inference(forward_demodulation,[],[f2051,f245])).
fof(f2051,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))))))))) )),
  inference(forward_demodulation,[],[f2038,f301])).
fof(f2038,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))))))))) )),
  inference(unit_resulting_resolution,[],[f1745,f339])).
fof(f10736,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10735,f2237])).
fof(f10735,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10440,f246])).
fof(f10440,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f4137])).
fof(f4137,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))),
  inference(forward_demodulation,[],[f4136,f245])).
fof(f4136,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))))))),
  inference(unit_resulting_resolution,[],[f3985,f339])).
fof(f3985,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(forward_demodulation,[],[f3984,f245])).
fof(f3984,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))))))),
  inference(unit_resulting_resolution,[],[f1931,f339])).
fof(f1931,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))),
  inference(unit_resulting_resolution,[],[f1829,f339])).
fof(f1829,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))))),
  inference(unit_resulting_resolution,[],[f1748,f339])).
fof(f1748,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))))),
  inference(unit_resulting_resolution,[],[f1651,f339])).
fof(f1651,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))),
  inference(superposition,[],[f1419,f250])).
fof(f1419,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))) )),
  inference(superposition,[],[f236,f250])).
fof(f10734,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10733])).
fof(f10733,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) | $spl48),
  inference(forward_demodulation,[],[f10732,f246])).
fof(f10732,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10731,f247])).
fof(f10731,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10730,f1489])).
fof(f10730,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10439,f1621])).
fof(f10439,plain,(
  ( ! [X12] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f4020])).
fof(f4020,plain,(
  ( ! [X12] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X12),s('type.num.num',sK10))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))) )),
  inference(superposition,[],[f1745,f1626])).
fof(f1626,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))))) )),
  inference(superposition,[],[f301,f304])).
fof(f10713,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10712,f248])).
fof(f10712,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10431,f246])).
fof(f10431,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f3985])).
fof(f10699,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10698])).
fof(f10698,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10697,f246])).
fof(f10697,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) | $spl48),
  inference(forward_demodulation,[],[f10696,f247])).
fof(f10696,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))))) | $spl48),
  inference(forward_demodulation,[],[f10423,f249])).
fof(f10423,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f3316])).
fof(f3316,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))) )),
  inference(unit_resulting_resolution,[],[f2197,f339])).
fof(f2197,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))))) )),
  inference(unit_resulting_resolution,[],[f2052,f339])).
fof(f10624,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10623,f1489])).
fof(f10623,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10622,f1621])).
fof(f10622,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) ) | $spl48),
  inference(forward_demodulation,[],[f10621,f1655])).
fof(f10621,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))) ) | $spl48),
  inference(forward_demodulation,[],[f10394,f246])).
fof(f10394,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f2375])).
fof(f2375,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK10))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))))) )),
  inference(unit_resulting_resolution,[],[f1872,f339])).
fof(f1872,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK10)))))))))) )),
  inference(unit_resulting_resolution,[],[f1743,f339])).
fof(f1743,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))) )),
  inference(superposition,[],[f1420,f304])).
fof(f10611,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10610,f249])).
fof(f10610,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl48),
  inference(forward_demodulation,[],[f10609,f1489])).
fof(f10609,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl48),
  inference(forward_demodulation,[],[f10608,f1655])).
fof(f10608,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))) ) | $spl48),
  inference(forward_demodulation,[],[f10390,f246])).
fof(f10390,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f2325])).
fof(f2325,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))))) )),
  inference(unit_resulting_resolution,[],[f1793,f339])).
fof(f1793,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))))))) )),
  inference(forward_demodulation,[],[f1749,f301])).
fof(f1749,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))))))) )),
  inference(unit_resulting_resolution,[],[f1420,f339])).
fof(f10597,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10596,f248])).
fof(f10596,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10595,f246])).
fof(f10595,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10385,f1489])).
fof(f10385,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f2284])).
fof(f10594,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10593,f2188])).
fof(f2188,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f2187,f1578])).
fof(f1578,plain,(
  ( ! [X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f303,f248])).
fof(f303,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X7,X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6)))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.integer.INT_ADD')).
fof(f2187,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f2157,f303])).
fof(f2157,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f303,f1621])).
fof(f10593,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK11),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) ) | $spl48),
  inference(forward_demodulation,[],[f10384,f246])).
fof(f10384,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f2220])).
fof(f2220,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK10)))))))) )),
  inference(superposition,[],[f2211,f303])).
fof(f10592,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10591,f1621])).
fof(f10591,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))) ) | $spl48),
  inference(forward_demodulation,[],[f10590,f1655])).
fof(f10590,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))) ) | $spl48),
  inference(forward_demodulation,[],[f10383,f246])).
fof(f10383,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f2214])).
fof(f10589,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10588,f248])).
fof(f10588,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10587,f246])).
fof(f10587,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10382,f1621])).
fof(f10382,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f2211])).
fof(f10586,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10585,f249])).
fof(f10585,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl48),
  inference(forward_demodulation,[],[f10584,f1655])).
fof(f10584,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))) ) | $spl48),
  inference(forward_demodulation,[],[f10381,f246])).
fof(f10381,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11)))))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f2197])).
fof(f10583,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10582])).
fof(f10582,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10581,f246])).
fof(f10581,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) | $spl48),
  inference(forward_demodulation,[],[f10380,f247])).
fof(f10380,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f2149])).
fof(f10580,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10579,f1578])).
fof(f10579,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK11),s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10379,f246])).
fof(f10379,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))),s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f2102])).
fof(f2102,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num',sK10)))))),
  inference(superposition,[],[f2100,f303])).
fof(f10578,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10577,f1655])).
fof(f10577,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))))) | $spl48),
  inference(forward_demodulation,[],[f10378,f246])).
fof(f10378,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f2101])).
fof(f10576,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10575,f248])).
fof(f10575,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10377,f246])).
fof(f10377,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f2100])).
fof(f10574,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10573,f248])).
fof(f10573,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10572,f246])).
fof(f10572,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10376,f249])).
fof(f10376,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f2052])).
fof(f10571,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10570])).
fof(f10570,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) | $spl48),
  inference(forward_demodulation,[],[f10569,f246])).
fof(f10569,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10568,f247])).
fof(f10568,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10375,f1489])).
fof(f10375,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f2049])).
fof(f10555,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10554])).
fof(f10554,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10553,f246])).
fof(f10553,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) | $spl48),
  inference(forward_demodulation,[],[f10552,f247])).
fof(f10552,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))))) | $spl48),
  inference(forward_demodulation,[],[f10368,f1489])).
fof(f10368,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f1931])).
fof(f10551,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10550])).
fof(f10550,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) | $spl48),
  inference(forward_demodulation,[],[f10549,f246])).
fof(f10549,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10367,f2189])).
fof(f2189,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))),s('type.num.num',X7)))))) )),
  inference(forward_demodulation,[],[f2158,f247])).
fof(f2158,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))),s('type.num.num',X7)))))) )),
  inference(superposition,[],[f303,f1621])).
fof(f10367,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK11)))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f1877])).
fof(f1877,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK10))),s('type.num.num',sK11)))))) )),
  inference(superposition,[],[f1743,f303])).
fof(f10548,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10547,f248])).
fof(f10547,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10546,f246])).
fof(f10546,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10545,f1489])).
fof(f10545,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl48),
  inference(forward_demodulation,[],[f10366,f1621])).
fof(f10366,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f1872])).
fof(f10540,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10539,f1489])).
fof(f10539,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10538,f1655])).
fof(f10538,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))))) | $spl48),
  inference(forward_demodulation,[],[f10363,f246])).
fof(f10363,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f1829])).
fof(f10533,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10532,f248])).
fof(f10532,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10531,f246])).
fof(f10531,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10530,f249])).
fof(f10530,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl48),
  inference(forward_demodulation,[],[f10361,f1489])).
fof(f10361,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f1793])).
fof(f10529,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f10528,f248])).
fof(f10528,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10527,f246])).
fof(f10527,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(forward_demodulation,[],[f10360,f1489])).
fof(f10360,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f1748])).
fof(f10523,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10522])).
fof(f10522,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) | $spl48),
  inference(forward_demodulation,[],[f10521,f246])).
fof(f10521,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10520,f247])).
fof(f10520,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10519,f249])).
fof(f10519,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))) ) | $spl48),
  inference(forward_demodulation,[],[f10358,f1489])).
fof(f10358,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f1745])).
fof(f10518,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10517])).
fof(f10517,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) | $spl48),
  inference(forward_demodulation,[],[f10516,f246])).
fof(f10516,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10515,f247])).
fof(f10515,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10357,f1621])).
fof(f10357,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f1743])).
fof(f10514,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10513])).
fof(f10513,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) | $spl48),
  inference(forward_demodulation,[],[f10512,f246])).
fof(f10512,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10356,f1577])).
fof(f1577,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))))) )),
  inference(superposition,[],[f303,f247])).
fof(f10356,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f1653])).
fof(f1653,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))),
  inference(superposition,[],[f1651,f303])).
fof(f10510,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10509])).
fof(f10509,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) | $spl48),
  inference(forward_demodulation,[],[f10508,f246])).
fof(f10508,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10354,f247])).
fof(f10354,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(backward_demodulation,[],[f6662,f1651])).
fof(f10497,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f10496])).
fof(f10496,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) | $spl48),
  inference(forward_demodulation,[],[f10495,f246])).
fof(f10495,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10494,f247])).
fof(f10494,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11))))) | $spl48),
  inference(forward_demodulation,[],[f10348,f249])).
fof(f10348,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))) ) | $spl48),
  inference(backward_demodulation,[],[f6662,f1420])).
fof(f6662,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) | $spl48),
  inference(cnf_transformation,[],[f6662_D])).
fof(f6662_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f10345,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK16(sK11,sK10)),s('type.num.num',sK10)))))) ) | $spl68),
  inference(cnf_transformation,[],[f10345_D])).
fof(f10345_D,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK16(sK11,sK10)),s('type.num.num',sK10)))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f10344,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK16(sK11,sK10)),s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10)) )),
  inference(superposition,[],[f9855,f1943])).
fof(f1943,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK16(X1,X2)),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK17(X1,X2)),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(superposition,[],[f309,f300])).
fof(f300,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X20,X21] : s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X20),s('type.num.num',X21))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X21),s('type.num.num',X20)))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.gcd.GCD_SYM')).
fof(f309,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK16(X1,X0)),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK17(X1,X0)),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK16(X1,X0)),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK17(X1,X0)),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X2,X3] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(ennf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) => ? [X2,X3] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) => ? [X2,X3] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6,X7] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) => ? [X8,X9] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num',X7))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.gcd.LINEAR_GCD')).
fof(f9855,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',sK11))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))) )),
  inference(unit_resulting_resolution,[],[f4013,f1692])).
fof(f1692,plain,(
  ( ! [X4,X2,X3] : (s('type.integer.int',X4) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f338,f303])).
fof(f4013,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',sK11)))))))))) )),
  inference(superposition,[],[f1478,f1626])).
fof(f1478,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))))) )),
  inference(superposition,[],[f236,f407])).
fof(f9852,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'T_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f9852_D])).
fof(f9852_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'T_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f9854,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl67),
  inference(cnf_transformation,[],[f9854_D])).
fof(f9854_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f9776,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,f(s(bool,'T_0')))) ) | $spl46),
  inference(equality_factoring,[],[f8783])).
fof(f8783,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl46),
  inference(subsumption_resolution,[],[f8737,f426])).
fof(f426,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f406,f261])).
fof(f261,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.bool.EQ_CLAUSES|split|0')).
fof(f406,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f240,f237])).
fof(f237,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.int_arith.INT_LINEAR_GCD',aHL_TRUTH)).
fof(f240,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.bool.NOT_CLAUSES|split|1')).
fof(f8737,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl46),
  inference(superposition,[],[f5646,f499])).
fof(f499,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f252,f252])).
fof(f252,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.int_arith.INT_LINEAR_GCD',aHL_BOOL_CASES)).
fof(f5646,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl46),
  inference(cnf_transformation,[],[f5646_D])).
fof(f5646_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f9327,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl62),
  inference(cnf_transformation,[],[f9327_D])).
fof(f9327_D,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f8805,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) | $spl56),
  inference(cnf_transformation,[],[f8805_D])).
fof(f8805_D,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f9326,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f9320,f261])).
fof(f9320,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1150,f252])).
fof(f1150,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | sP1(X36,X35) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f906,f237])).
fof(f906,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | sP1(X36,X35) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f330,f499])).
fof(f330,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f200])).
fof(f200,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f9323,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl60),
  inference(cnf_transformation,[],[f9323_D])).
fof(f9323_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f9324,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f9316,f261])).
fof(f9316,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1150,f499])).
fof(f9322,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f9315,f261])).
fof(f9315,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1150,f499])).
fof(f9296,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8814,f8823])).
fof(f8823,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f8805,f8805,f308])).
fof(f308,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.bool.IMP_ANTISYM_AX')).
fof(f8814,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f8805,f265])).
fof(f265,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.bool.EQ_CLAUSES|split|2')).
fof(f9295,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8818,f237])).
fof(f8818,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f426,f8805,f308])).
fof(f9294,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8824,f237])).
fof(f8824,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f426,f8805,f308])).
fof(f9293,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8833,f237])).
fof(f8833,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f442,f8805,f326])).
fof(f326,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f201])).
fof(f442,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f406,f237,f330])).
fof(f9292,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8838,f237])).
fof(f8838,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f436,f8805,f326])).
fof(f436,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f406,f237,f327])).
fof(f327,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f201])).
fof(f9291,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8852,f237])).
fof(f8852,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f554,f406,f8805,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f209])).
fof(f209,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f554,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f406,f426,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f119,f154,f153,f152,f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f152,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f153,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f119,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X23,X9,X8] : ((p(s(bool,X8)) <=> s(bool,X9) = s(bool,X23)) <=> ((p(s(bool,X8)) | p(s(bool,X9)) | p(s(bool,X23))) & (p(s(bool,X8)) | ~p(s(bool,X23)) | ~p(s(bool,X9))) & (p(s(bool,X9)) | ~p(s(bool,X23)) | ~p(s(bool,X8))) & (p(s(bool,X23)) | ~p(s(bool,X9)) | ~p(s(bool,X8)))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.sat.sat_2|49|_prove_ltb')).
fof(f9290,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8871,f237])).
fof(f8871,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f743,f406,f8805,f343])).
fof(f743,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f417,f414,f412,f406,f345])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f412,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f406,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f153])).
fof(f414,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f352])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f215])).
fof(f215,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f152])).
fof(f417,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f151])).
fof(f9289,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8877,f8922])).
fof(f8922,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f8805,f348])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f213])).
fof(f8877,plain,(
  ~sP4('T_0','F_0','T_0') | $spl56),
  inference(unit_resulting_resolution,[],[f418,f415,f564,f8805,f344])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f564,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f237,f426,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f219])).
fof(f415,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f406,f353])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f216])).
fof(f418,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f237,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f218])).
fof(f9288,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8881,f8964])).
fof(f8964,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f426,f8805,f361])).
fof(f8881,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f418,f415,f411,f8805,f344])).
fof(f411,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f348])).
fof(f9287,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8883,f8946])).
fof(f8946,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f8805,f352])).
fof(f8883,plain,(
  ~sP3('T_0','F_0','T_0') | $spl56),
  inference(unit_resulting_resolution,[],[f418,f564,f411,f8805,f344])).
fof(f9286,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8884,f8952])).
fof(f8952,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f8805,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f218])).
fof(f8884,plain,(
  ~sP2('F_0','T_0','T_0') | $spl56),
  inference(unit_resulting_resolution,[],[f415,f564,f411,f8805,f344])).
fof(f9285,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8888,f8964])).
fof(f8888,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f418,f415,f412,f8805,f344])).
fof(f9284,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8890,f564])).
fof(f8890,plain,(
  sP5('T_0','T_0','F_0') | $spl56),
  inference(unit_resulting_resolution,[],[f418,f414,f412,f8805,f344])).
fof(f9283,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8892,f8946])).
fof(f8892,plain,(
  ~sP3('T_0','F_0','T_0') | $spl56),
  inference(unit_resulting_resolution,[],[f418,f564,f412,f8805,f344])).
fof(f9282,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8893,f8952])).
fof(f8893,plain,(
  ~sP2('F_0','T_0','T_0') | $spl56),
  inference(unit_resulting_resolution,[],[f415,f564,f412,f8805,f344])).
fof(f9281,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8901,f8805])).
fof(f8901,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl56),
  inference(unit_resulting_resolution,[],[f413,f406,f8805,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f213])).
fof(f413,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f406,f350])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f9280,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8909,f8922])).
fof(f8909,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f406,f8805,f8805,f347])).
fof(f9279,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8911,f8805])).
fof(f8911,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl56),
  inference(unit_resulting_resolution,[],[f412,f406,f8805,f347])).
fof(f9278,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8915,f8922])).
fof(f8915,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f406,f237,f8805,f347])).
fof(f9277,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8920,f8922])).
fof(f8920,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f406,f8805,f8805,f347])).
fof(f9276,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8925,f8805])).
fof(f8925,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl56),
  inference(unit_resulting_resolution,[],[f416,f406,f8805,f351])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f216])).
fof(f416,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f406,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f216])).
fof(f9275,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8933,f8946])).
fof(f8933,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f406,f8805,f8805,f351])).
fof(f9274,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8935,f8805])).
fof(f8935,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl56),
  inference(unit_resulting_resolution,[],[f415,f406,f8805,f351])).
fof(f9273,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8939,f8946])).
fof(f8939,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f406,f237,f8805,f351])).
fof(f9272,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8944,f8946])).
fof(f8944,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f406,f8805,f8805,f351])).
fof(f9271,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8960,f426])).
fof(f8960,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl56),
  inference(unit_resulting_resolution,[],[f553,f8805,f361])).
fof(f553,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f406,f426,f359])).
fof(f9270,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8961,f426])).
fof(f8961,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl56),
  inference(unit_resulting_resolution,[],[f554,f8805,f361])).
fof(f9269,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8971,f9190])).
fof(f9190,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f422,f8805,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f366])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP6(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f221])).
fof(f221,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1)) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f422,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f406,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f8971,plain,(
  ( ! [X0,X1] : (~sP7('F_0',X0,X1)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f406,f8805,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f9268,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8976,f9190])).
fof(f8976,plain,(
  ( ! [X0,X1] : (~sP7('F_0',X0,X1)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f406,f8805,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f222])).
fof(f9267,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8984,f9006])).
fof(f9006,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f8805,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f225])).
fof(f8984,plain,(
  ~sP6('F_0','T_0','F_0') | $spl56),
  inference(unit_resulting_resolution,[],[f451,f406,f8805,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f451,plain,(
  ( ! [X0] : (~sP7('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f237,f406,f364])).
fof(f9266,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f8992,f9006])).
fof(f8992,plain,(
  ~sP6('F_0','F_0','T_0') | $spl56),
  inference(unit_resulting_resolution,[],[f448,f406,f8805,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f222])).
fof(f448,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f406,f363])).
fof(f9265,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9002,f9006])).
fof(f9002,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f406,f406,f8805,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f225])).
fof(f9264,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9010,f237])).
fof(f9010,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f630,f8805,f375])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f121,f157,f156])).
fof(f121,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X23,X9,X8] : ((p(s(bool,X8)) <=> (p(s(bool,X9)) | p(s(bool,X23)))) <=> ((p(s(bool,X8)) | ~p(s(bool,X9))) & (p(s(bool,X8)) | ~p(s(bool,X23))) & (p(s(bool,X9)) | p(s(bool,X23)) | ~p(s(bool,X8)))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.sat.sat_25_prove_ltb')).
fof(f630,plain,(
  ~sP7('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f624,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f222])).
fof(f624,plain,(
  ~sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f406,f406,f237,f370])).
fof(f9263,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9015,f237])).
fof(f9015,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f448,f8805,f375])).
fof(f9262,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9016,f8805])).
fof(f9016,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl56),
  inference(unit_resulting_resolution,[],[f451,f8805,f375])).
fof(f9261,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9024,f237])).
fof(f9024,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f630,f8805,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f227])).
fof(f9260,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9029,f8805])).
fof(f9029,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl56),
  inference(unit_resulting_resolution,[],[f448,f8805,f376])).
fof(f9259,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9030,f237])).
fof(f9030,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f451,f8805,f376])).
fof(f9258,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9045,f9181])).
fof(f9181,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f425,f8805,f397])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f383])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f229])).
fof(f229,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f425,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f406,f390])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f232])).
fof(f232,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f231])).
fof(f231,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f9045,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f406,f8805,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f230])).
fof(f9257,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9048,f8805])).
fof(f9048,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl56),
  inference(unit_resulting_resolution,[],[f425,f465,f8805,f384])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f230])).
fof(f465,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f237,f406,f381])).
fof(f9256,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9049,f8805])).
fof(f9049,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl56),
  inference(unit_resulting_resolution,[],[f424,f465,f8805,f384])).
fof(f424,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f237,f389])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f233])).
fof(f9255,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9053,f9125])).
fof(f9125,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f8805,f389])).
fof(f9053,plain,(
  ~sP8('F_0','T_0','T_0') | $spl56),
  inference(unit_resulting_resolution,[],[f465,f237,f8805,f384])).
fof(f9254,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9059,f9125])).
fof(f9059,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f464,f8805,f8805,f384])).
fof(f464,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f406,f406,f380])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f230])).
fof(f9253,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9061,f9125])).
fof(f9061,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f465,f8805,f8805,f384])).
fof(f9252,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9072,f9125])).
fof(f9072,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f464,f8805,f8805,f384])).
fof(f9251,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9073,f9125])).
fof(f9073,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f465,f8805,f8805,f384])).
fof(f9250,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9086,f9125])).
fof(f9086,plain,(
  ~sP8('F_0','F_0','F_0') | $spl56),
  inference(unit_resulting_resolution,[],[f464,f406,f8805,f385])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f230])).
fof(f9249,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9097,f9125])).
fof(f9097,plain,(
  ~sP8('F_0','F_0','F_0') | $spl56),
  inference(unit_resulting_resolution,[],[f464,f406,f8805,f386])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f230])).
fof(f9248,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9109,f8805])).
fof(f9109,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl56),
  inference(unit_resulting_resolution,[],[f425,f406,f8805,f387])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f233])).
fof(f9247,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9112,f9125])).
fof(f9112,plain,(
  ( ! [X0,X1] : (~sP8(X0,'F_0',X1)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f8805,f406,f8805,f387])).
fof(f9246,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9119,f8805])).
fof(f9119,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl56),
  inference(unit_resulting_resolution,[],[f423,f406,f8805,f387])).
fof(f423,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f406,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f233])).
fof(f9245,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9121,f9125])).
fof(f9121,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f237,f406,f8805,f387])).
fof(f9244,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9123,f9125])).
fof(f9123,plain,(
  ( ! [X0,X1] : (~sP8(X0,'F_0',X1)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f8805,f406,f8805,f387])).
fof(f9243,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9140,f237])).
fof(f9140,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f666,f8805,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f146,f160,f159])).
fof(f146,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X23,X9,X8] : ((p(s(bool,X8)) <=> (p(s(bool,X9)) => p(s(bool,X23)))) <=> ((p(s(bool,X8)) | p(s(bool,X9))) & (p(s(bool,X8)) | ~p(s(bool,X23))) & (~p(s(bool,X9)) | p(s(bool,X23)) | ~p(s(bool,X8)))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.sat.sat_27_prove_ltb')).
fof(f666,plain,(
  ~sP9('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f660,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f230])).
fof(f660,plain,(
  ~sP8('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f237,f406,f237,f387])).
fof(f9242,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9145,f8805])).
fof(f9145,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl56),
  inference(unit_resulting_resolution,[],[f464,f8805,f393])).
fof(f9241,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9146,f237])).
fof(f9146,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f465,f8805,f393])).
fof(f9240,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9152,f237])).
fof(f9152,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f467,f406,f8805,f394])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f467,plain,(
  ( ! [X0] : (sP9('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f406,f237,f392])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f9239,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9164,f237])).
fof(f9164,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(unit_resulting_resolution,[],[f653,f406,f8805,f394])).
fof(f653,plain,(
  sP9('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f425,f237,f406,f386])).
fof(f9238,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9174,f423])).
fof(f9174,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'F_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f464,f8805,f397])).
fof(f9237,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9175,f424])).
fof(f9175,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f465,f8805,f397])).
fof(f9236,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9183,f420])).
fof(f420,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f371])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f9183,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl56),
  inference(unit_resulting_resolution,[],[f448,f8805,f398])).
fof(f9235,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f9184,f421])).
fof(f421,plain,(
  ( ! [X0,X1] : (sP6(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f237,f372])).
fof(f9184,plain,(
  ( ! [X0] : (~sP6('F_0','T_0',X0)) ) | $spl56),
  inference(unit_resulting_resolution,[],[f451,f8805,f398])).
fof(f9213,plain,(
  $false | ($spl26 | $spl56)),
  inference(unit_resulting_resolution,[],[f467,f406,f8805,f1138])).
fof(f9209,plain,(
  $false | ($spl24 | $spl56)),
  inference(unit_resulting_resolution,[],[f666,f8805,f1135])).
fof(f9207,plain,(
  $false | ($spl18 | $spl56)),
  inference(unit_resulting_resolution,[],[f425,f406,f8805,f1126])).
fof(f9203,plain,(
  $false | ($spl14 | $spl56)),
  inference(unit_resulting_resolution,[],[f630,f8805,f1120])).
fof(f9201,plain,(
  $false | ($spl12 | $spl56)),
  inference(unit_resulting_resolution,[],[f630,f8805,f1117])).
fof(f9199,plain,(
  $false | ($spl8 | $spl56)),
  inference(unit_resulting_resolution,[],[f416,f406,f8805,f1111])).
fof(f9195,plain,(
  $false | ($spl6 | $spl56)),
  inference(unit_resulting_resolution,[],[f413,f406,f8805,f1108])).
fof(f9191,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f448,f422,f8805,f398])).
fof(f9189,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f451,f421,f8805,f398])).
fof(f9187,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f448,f420,f8805,f398])).
fof(f9182,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f425,f8805,f397])).
fof(f9180,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f465,f424,f8805,f397])).
fof(f9178,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f423,f8805,f397])).
fof(f9173,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f467,f406,f8805,f8805,f394])).
fof(f9168,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f653,f406,f237,f8805,f394])).
fof(f9162,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f467,f406,f8805,f8805,f394])).
fof(f9157,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f467,f406,f237,f8805,f394])).
fof(f9151,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f8805,f8805,f393])).
fof(f9149,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f465,f237,f8805,f393])).
fof(f9144,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f8805,f8805,f393])).
fof(f9142,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f666,f237,f8805,f393])).
fof(f9137,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f406,f8805,f392])).
fof(f9132,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f406,f406,f8805,f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP9(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f235])).
fof(f9126,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f660,f8805,f389])).
fof(f9124,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f423,f8805,f406,f8805,f387])).
fof(f9122,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f423,f237,f406,f8805,f387])).
fof(f9113,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f423,f8805,f406,f8805,f387])).
fof(f9111,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f425,f237,f406,f8805,f387])).
fof(f9102,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f425,f406,f8805,f386])).
fof(f9100,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f423,f406,f8805,f386])).
fof(f9091,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f425,f406,f8805,f385])).
fof(f9089,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f423,f406,f8805,f385])).
fof(f9080,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f425,f8805,f8805,f384])).
fof(f9078,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f465,f424,f8805,f8805,f384])).
fof(f9076,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f464,f423,f8805,f8805,f384])).
fof(f9062,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f425,f465,f8805,f8805,f384])).
fof(f9060,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f425,f464,f8805,f8805,f384])).
fof(f9054,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f425,f465,f237,f8805,f384])).
fof(f9046,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f653,f406,f8805,f381])).
fof(f9041,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f617,f406,f406,f8805,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f227])).
fof(f617,plain,(
  sP7('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f422,f406,f406,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f9035,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f448,f8805,f8805,f376])).
fof(f9033,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f451,f237,f8805,f376])).
fof(f9028,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f448,f8805,f8805,f376])).
fof(f9026,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f630,f237,f8805,f376])).
fof(f9021,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f448,f8805,f8805,f375])).
fof(f9019,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f448,f237,f8805,f375])).
fof(f9014,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f448,f8805,f8805,f375])).
fof(f9012,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f630,f237,f8805,f375])).
fof(f9007,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f624,f8805,f372])).
fof(f9005,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f624,f8805,f371])).
fof(f9003,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f422,f406,f406,f8805,f370])).
fof(f8997,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f448,f422,f406,f8805,f368])).
fof(f8995,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f448,f420,f406,f8805,f368])).
fof(f8985,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f422,f451,f406,f8805,f367])).
fof(f8977,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f617,f406,f8805,f364])).
fof(f8972,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f617,f406,f8805,f363])).
fof(f8967,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f553,f426,f8805,f361])).
fof(f8965,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f554,f426,f8805,f361])).
fof(f9234,plain,(
  $false | $spl56),
  inference(trivial_inequality_removal,[],[f8957])).
fof(f8957,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl56),
  inference(unit_resulting_resolution,[],[f608,f8805,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f219])).
fof(f608,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f599,f340])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f599,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f406,f406,f406,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f218])).
fof(f9233,plain,(
  $false | $spl56),
  inference(trivial_inequality_removal,[],[f8956])).
fof(f8956,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl56),
  inference(unit_resulting_resolution,[],[f604,f8805,f360])).
fof(f604,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f594,f341])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f594,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f406,f237,f237,f351])).
fof(f8953,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f599,f8805,f358])).
fof(f8951,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f599,f8805,f357])).
fof(f8949,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f599,f8805,f356])).
fof(f8947,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f594,f8805,f352])).
fof(f8945,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f415,f406,f8805,f8805,f351])).
fof(f8940,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f415,f406,f237,f8805,f351])).
fof(f8934,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f415,f406,f8805,f8805,f351])).
fof(f8929,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f416,f406,f237,f8805,f351])).
fof(f8923,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f584,f8805,f348])).
fof(f584,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f406,f237,f237,f347])).
fof(f8921,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f412,f406,f8805,f8805,f347])).
fof(f8916,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f412,f406,f237,f8805,f347])).
fof(f8910,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f412,f406,f8805,f8805,f347])).
fof(f8905,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f413,f406,f237,f8805,f347])).
fof(f8894,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f418,f415,f564,f412,f8805,f344])).
fof(f8885,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f418,f415,f564,f411,f8805,f344])).
fof(f8875,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f554,f8805,f406,f8805,f343])).
fof(f8873,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f743,f237,f406,f8805,f343])).
fof(f8863,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f554,f406,f8805,f8805,f343])).
fof(f8857,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f554,f406,f237,f8805,f343])).
fof(f8851,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f433,f406,f8805,f330])).
fof(f433,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f406,f406,f325])).
fof(f325,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f201])).
fof(f8847,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f433,f406,f8805,f327])).
fof(f8843,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f436,f8805,f8805,f326])).
fof(f8841,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f436,f237,f8805,f326])).
fof(f8837,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f436,f8805,f8805,f326])).
fof(f8835,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f442,f237,f8805,f326])).
fof(f8831,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f447,f8805,f312])).
fof(f312,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f447,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f406,f406,f408])).
fof(f408,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f310,f406])).
fof(f310,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f194])).
fof(f8829,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f447,f8805,f311])).
fof(f311,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f194])).
fof(f8827,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f8805,f426,f8805,f308])).
fof(f8825,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f237,f426,f8805,f308])).
fof(f8821,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f8805,f426,f8805,f308])).
fof(f8819,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f237,f426,f8805,f308])).
fof(f8815,plain,(
  $false | ($spl46 | $spl56)),
  inference(unit_resulting_resolution,[],[f5646,f8805,f265])).
fof(f8813,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f426,f8805,f262])).
fof(f262,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f165])).
fof(f8811,plain,(
  $false | $spl56),
  inference(unit_resulting_resolution,[],[f406,f8805])).
fof(f8809,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl58),
  inference(cnf_transformation,[],[f8809_D])).
fof(f8809_D,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f8808,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f8801,f261])).
fof(f8801,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1149,f252])).
fof(f1149,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X32,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f905,f237])).
fof(f905,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X33)) | sP1(X32,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f327,f499])).
fof(f8804,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl54),
  inference(cnf_transformation,[],[f8804_D])).
fof(f8804_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8806,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f8797,f261])).
fof(f8797,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1149,f499])).
fof(f8803,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f8796,f261])).
fof(f8796,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1149,f499])).
fof(f8721,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~sP9(f(s(bool,'T_0')),X1,'T_0')) ) | $spl52),
  inference(cnf_transformation,[],[f8721_D])).
fof(f8721_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~sP9(f(s(bool,'T_0')),X1,'T_0')) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1102,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(cnf_transformation,[],[f1102_D])).
fof(f1102_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f8723,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP9(f(s(bool,'T_0')),X1,'T_0')) ) | $spl47),
  inference(condensation,[],[f8722])).
fof(f8722,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | ~sP9(f(s(bool,'T_0')),X18,'T_0') | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl47),
  inference(subsumption_resolution,[],[f8715,f545])).
fof(f545,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f542,f406])).
fof(f542,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f305,f252])).
fof(f305,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f139])).
fof(f8715,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | p(s(bool,X18)) | ~sP9(f(s(bool,'T_0')),X18,'T_0') | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl47),
  inference(superposition,[],[f6991,f499])).
fof(f6991,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | p(s(bool,X28)) | ~sP9(f(s(bool,'T_0')),X28,X27)) ) | $spl47),
  inference(subsumption_resolution,[],[f6947,f237])).
fof(f6947,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | p(s(bool,X28)) | ~sP9(f(s(bool,'T_0')),X28,X27)) ) | $spl47),
  inference(superposition,[],[f394,f5665])).
fof(f5665,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl47),
  inference(unit_resulting_resolution,[],[f406,f5647,f545])).
fof(f8720,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP9(f(s(bool,'T_0')),X1,'T_0')) ) | $spl47),
  inference(condensation,[],[f8719])).
fof(f8719,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | ~sP9(f(s(bool,'T_0')),X15,'T_0') | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl47),
  inference(subsumption_resolution,[],[f8714,f545])).
fof(f8714,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | p(s(bool,X15)) | ~sP9(f(s(bool,'T_0')),X15,'T_0') | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl47),
  inference(superposition,[],[f6991,f499])).
fof(f6664,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK16(sK11,sK10)),s('type.num.num',sK10))))) | $spl51),
  inference(cnf_transformation,[],[f6664_D])).
fof(f6664_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK16(sK11,sK10)),s('type.num.num',sK10))))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f6604,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK16(sK11,sK10)),s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10)),
  inference(superposition,[],[f4392,f1943])).
fof(f4392,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK11))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))) )),
  inference(superposition,[],[f4377,f303])).
fof(f4377,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))))) )),
  inference(unit_resulting_resolution,[],[f4014,f338])).
fof(f4014,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num',sK11)))))))))) )),
  inference(superposition,[],[f1652,f1626])).
fof(f1652,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK11)))))))))) )),
  inference(superposition,[],[f1419,f407])).
fof(f5645,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK22(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK22(X0,X1,bool,bool))))) ) | $spl44),
  inference(cnf_transformation,[],[f5645_D])).
fof(f5645_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK22(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK22(X0,X1,bool,bool))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5647,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl47),
  inference(cnf_transformation,[],[f5647_D])).
fof(f5647_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f5635,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK22(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK22(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1926,f499])).
fof(f1926,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK22(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK22(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f337,f252])).
fof(f337,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK22(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK22(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK22(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK22(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22])],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.int_arith.INT_LINEAR_GCD',aHL_EXT)).
fof(f5634,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK22(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK22(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1926,f499])).
fof(f4373,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl38),
  inference(cnf_transformation,[],[f4373_D])).
fof(f4373_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4369,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK22(X0,X1,bool,bool))) ) | $spl34),
  inference(cnf_transformation,[],[f4369_D])).
fof(f4369_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK22(X0,X1,bool,bool))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4376,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl43),
  inference(cnf_transformation,[],[f4376_D])).
fof(f4376_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f4367,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK22(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1921,f252])).
fof(f1921,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK22(X16,X17,X18,bool))) )),
  inference(superposition,[],[f337,f252])).
fof(f4374,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl40),
  inference(cnf_transformation,[],[f4374_D])).
fof(f4374_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4366,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK22(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1921,f499])).
fof(f4365,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK22(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1921,f499])).
fof(f4368,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl32),
  inference(cnf_transformation,[],[f4368_D])).
fof(f4368_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4371,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl37),
  inference(cnf_transformation,[],[f4371_D])).
fof(f4371_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f4364,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK22(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1921,f499])).
fof(f4363,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK22(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1921,f499])).
fof(f1209,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1200,f237])).
fof(f1200,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(backward_demodulation,[],[f1102,f406])).
fof(f1208,plain,(
  $false | $spl2),
  inference(global_subsumption,[],[f236,f237,f406,f244,f245,f246,f247,f248,f249,f250,f252,f262,f261,f266,f265,f291,f290,f299,f297,f300,f301,f407,f303,f304,f308,f305,f309,f312,f311,f408,f330,f327,f326,f325,f336,f335,f337,f339,f338,f346,f345,f344,f343,f342,f341,f340,f350,f349,f348,f347,f354,f353,f352,f351,f358,f357,f356,f355,f362,f361,f360,f359,f369,f368,f367,f398,f365,f364,f363,f373,f372,f371,f370,f377,f376,f375,f374,f386,f385,f384,f397,f382,f381,f380,f390,f389,f388,f387,f394,f393,f392,f391,f409,f410,f411,f412,f413,f414,f415,f416,f417,f418,f419,f420,f421,f422,f423,f424,f425,f426,f433,f434,f436,f442,f447,f448,f451,f454,f457,f464,f465,f467,f474,f499,f545,f554,f553,f564,f562,f584,f594,f599,f604,f608,f617,f624,f630,f653,f660,f666,f729,f743,f781,f863,f1100,f1104,f1107,f1110,f1113,f1116,f1119,f1122,f1125,f1128,f1131,f1134,f1137,f1140,f1143,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f1155,f1156,f1157,f1158,f1159,f1160,f1161,f1162,f1163,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1195,f1197,f1102,f1207])).
fof(f1207,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f1198,f1102])).
fof(f1198,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl2),
  inference(backward_demodulation,[],[f1102,f265])).
fof(f1163,plain,(
  ( ! [X88,X87,X89,X86] : (~sP6(X86,X87,X88) | sP7(X86,X87,X88) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f919,f237])).
fof(f919,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,'T_0')) | ~sP6(X86,X87,X88) | sP7(X86,X87,X88) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f398,f499])).
fof(f1162,plain,(
  ( ! [X85,X83,X84,X82] : (~sP8(X82,X83,X84) | sP9(X82,X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f918,f237])).
fof(f918,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'T_0')) | ~sP8(X82,X83,X84) | sP9(X82,X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f397,f499])).
fof(f1161,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X79)) | p(s(bool,X80)) | ~sP9(X78,X80,X79) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f917,f237])).
fof(f917,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,'T_0')) | ~p(s(bool,X79)) | p(s(bool,X80)) | ~sP9(X78,X80,X79) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f394,f499])).
fof(f1160,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X75)) | sP9(X74,X75,X76) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f916,f237])).
fof(f916,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X75)) | sP9(X74,X75,X76) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f393,f499])).
fof(f1159,plain,(
  ( ! [X70,X72,X71,X73] : (sP9(X70,X72,X71) | p(s(bool,X71)) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f915,f237])).
fof(f915,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | sP9(X70,X72,X71) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f392,f499])).
fof(f1158,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X67)) | sP9(X68,X67,X66) | p(s(bool,X68)) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f914,f237])).
fof(f914,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X67)) | sP9(X68,X67,X66) | p(s(bool,X68)) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f391,f499])).
fof(f1157,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | ~p(s(bool,X64)) | ~sP8(X64,X63,X62) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f913,f237])).
fof(f913,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X63)) | ~p(s(bool,X64)) | ~sP8(X64,X63,X62) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f387,f499])).
fof(f1156,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X59)) | p(s(bool,X60)) | ~sP7(X58,X60,X59) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f912,f237])).
fof(f912,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,'T_0')) | p(s(bool,X59)) | p(s(bool,X60)) | ~sP7(X58,X60,X59) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f377,f499])).
fof(f1155,plain,(
  ( ! [X57,X54,X56,X55] : (sP7(X54,X55,X56) | ~p(s(bool,X55)) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f911,f237])).
fof(f911,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X55)) | sP7(X54,X55,X56) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f376,f499])).
fof(f1154,plain,(
  ( ! [X52,X50,X53,X51] : (sP7(X50,X52,X51) | ~p(s(bool,X51)) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f910,f237])).
fof(f910,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X51)) | sP7(X50,X52,X51) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f375,f499])).
fof(f1153,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X47)) | p(s(bool,X48)) | ~sP6(X46,X47,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f909,f237])).
fof(f909,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | p(s(bool,X48)) | ~sP6(X46,X47,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f370,f499])).
fof(f1152,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | p(s(bool,X44)) | ~sP3(X43,X42,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f908,f237])).
fof(f908,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP3(X43,X42,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f351,f499])).
fof(f1151,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | ~sP4(X39,X38,X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f907,f237])).
fof(f907,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP4(X39,X38,X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f347,f499])).
fof(f1148,plain,(
  ( ! [X30,X31,X29] : (sP0(X29,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f904,f237])).
fof(f904,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | sP0(X29,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f312,f499])).
fof(f1147,plain,(
  ( ! [X28,X26,X27] : (sP0(X27,X26) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f903,f237])).
fof(f903,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | sP0(X27,X26) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f311,f499])).
fof(f1146,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f902,f237])).
fof(f902,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f308,f499])).
fof(f1145,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f898,f237])).
fof(f898,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f265,f499])).
fof(f863,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f261,f499])).
fof(f781,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f499,f499])).
fof(f729,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f419,f414,f411,f237,f344])).
fof(f562,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f237,f426,f361])).
fof(f474,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f237,f237,f393])).
fof(f457,plain,(
  ( ! [X0] : (sP7('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f237,f237,f376])).
fof(f454,plain,(
  ( ! [X0] : (sP7('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f237,f375])).
fof(f434,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f237,f237,f326])).
fof(f419,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f237,f358])).
fof(f410,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f237,f312])).
fof(f409,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f311])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f227])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f219])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f335,plain,(
  ( ! [X2,X0,X10,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK18(X2,X1,X0)))),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),sK19(X2,X1,X0)),s(X1,X8))))))) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK18(X2,X1,X0)))),s(X0,X10)))) | ! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),sK19(X2,X1,X0)),s(X1,X8))))))) & (! [X5] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK20(X5,X2,X1,X0)))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK20(X5,X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,sK21(X3,X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19,sK20,sK21])],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((? [X9] : ! [X10] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X9))),s(X0,X10)))) | ? [X7] : ! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X8))),s(X0,happ(s(fun(X1,X0),X7),s(X1,X8))))))) & (! [X5] : ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))) | ! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4))))))),
  inference(rectify,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((? [X3] : ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) | ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))))))) & (! [X5] : ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))) | ! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4))))))),
  inference(nnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : (! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) <=> ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ! [X2] : (! [X3] : ? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,X4)))) <=> ? [X5] : ! [X6] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X6))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X6)))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X18,X13] : ! [X16] : (! [X14] : ? [X15] : p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X13,fun(X18,bool)),X16),s(X13,X14))),s(X18,X15)))) <=> ? [X19] : ! [X14] : p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X13,fun(X18,bool)),X16),s(X13,X14))),s(X18,happ(s(fun(X13,X18),X19),s(X13,X14)))))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.bool.SKOLEM_THM')).
fof(f336,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,sK20(X5,X2,X1,X0)))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK20(X5,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X0,sK21(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f205])).
fof(f297,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK14(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK14(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X13] : ! [X16,X17] : ((p(s(bool,X16)) | ? [X14] : p(s(bool,happ(s(fun(X13,bool),X17),s(X13,X14))))) <=> ? [X14] : (p(s(bool,X16)) | p(s(bool,happ(s(fun(X13,bool),X17),s(X13,X14))))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f299,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f191])).
fof(f290,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK12(X1,X0)))))) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK12(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X13] : ! [X16] : (~? [X14] : p(s(bool,happ(s(fun(X13,bool),X16),s(X13,X14)))) <=> ! [X14] : ~p(s(bool,happ(s(fun(X13,bool),X16),s(X13,X14)))))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.bool.NOT_EXISTS_THM')).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X1,X0)))))) )),
  inference(cnf_transformation,[],[f184])).
fof(f266,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f244,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X20] : s('type.num.num',X20) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X20),s('type.num.num','const.num.|48|_0')))),
  file('thm.int_arith.INT_LINEAR_GCD','thm.gcd.GCD_|48|R')).
fof(f1144,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~sP6('T_0',X1,X2)) ) | $spl30),
  inference(cnf_transformation,[],[f1144_D])).
fof(f1144_D,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~sP6('T_0',X1,X2)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1197,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP6('T_0',X1,X2) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1196])).
fof(f1196,plain,(
  ( ! [X85,X83,X86,X84] : (~sP6('T_0',X84,X85) | sP7('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f946,f545])).
fof(f946,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~sP6('T_0',X84,X85) | sP7('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f398,f499])).
fof(f1141,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl28),
  inference(cnf_transformation,[],[f1141_D])).
fof(f1141_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1195,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1194])).
fof(f1194,plain,(
  ( ! [X80,X81,X79,X82] : (~sP8('T_0',X80,X81) | sP9('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f945,f545])).
fof(f945,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~sP8('T_0',X80,X81) | sP9('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f397,f499])).
fof(f1138,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f1138_D])).
fof(f1138_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1193,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1192])).
fof(f1192,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~sP9('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f944,f545])).
fof(f944,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | ~sP9('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f394,f499])).
fof(f1135,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f1135_D])).
fof(f1135_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1191,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1190])).
fof(f1190,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | sP9('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f943,f545])).
fof(f943,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | ~p(s(bool,X72)) | sP9('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f393,f499])).
fof(f1132,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f1132_D])).
fof(f1132_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1189,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1188])).
fof(f1188,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f942,f545])).
fof(f942,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f392,f499])).
fof(f1129,plain,(
  ( ! [X2,X1] : (sP9(X2,X1,'T_0') | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl20),
  inference(cnf_transformation,[],[f1129_D])).
fof(f1129_D,plain,(
  ( ! [X2,X1] : (sP9(X2,X1,'T_0') | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1187,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP9(X2,X1,'T_0') | p(s(bool,X2))) )),
  inference(condensation,[],[f1186])).
fof(f1186,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | sP9(X65,X64,'T_0') | p(s(bool,X65)) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f941,f545])).
fof(f941,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | sP9(X65,X64,'T_0') | p(s(bool,X65)) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f391,f499])).
fof(f1126,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f1126_D])).
fof(f1126_D,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1185,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f1184])).
fof(f1184,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X60)) | ~p(s(bool,X61)) | ~sP8(X61,X60,'T_0') | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f940,f545])).
fof(f940,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | p(s(bool,X60)) | ~p(s(bool,X61)) | ~sP8(X61,X60,'T_0') | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f387,f499])).
fof(f1123,plain,(
  ( ! [X2,X1] : (~sP7('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl16),
  inference(cnf_transformation,[],[f1123_D])).
fof(f1123_D,plain,(
  ( ! [X2,X1] : (~sP7('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1183,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f1182])).
fof(f1182,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP7('T_0',X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f939,f545])).
fof(f939,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP7('T_0',X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f377,f499])).
fof(f1120,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl14),
  inference(cnf_transformation,[],[f1120_D])).
fof(f1120_D,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1181,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1180])).
fof(f1180,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | sP7('T_0',X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f938,f545])).
fof(f938,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | sP7('T_0',X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f376,f499])).
fof(f1117,plain,(
  ( ! [X2,X1] : (sP7('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl12),
  inference(cnf_transformation,[],[f1117_D])).
fof(f1117_D,plain,(
  ( ! [X2,X1] : (sP7('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1179,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f1178])).
fof(f1178,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | sP7('T_0',X49,X48) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f937,f545])).
fof(f937,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | sP7('T_0',X49,X48) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f375,f499])).
fof(f1114,plain,(
  ( ! [X2,X1] : (~sP6('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl10),
  inference(cnf_transformation,[],[f1114_D])).
fof(f1114_D,plain,(
  ( ! [X2,X1] : (~sP6('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1177,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f1176])).
fof(f1176,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | ~sP6('T_0',X44,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f936,f545])).
fof(f936,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | p(s(bool,X44)) | p(s(bool,X45)) | ~sP6('T_0',X44,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f370,f499])).
fof(f1111,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl8),
  inference(cnf_transformation,[],[f1111_D])).
fof(f1111_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1175,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1174])).
fof(f1174,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f935,f545])).
fof(f935,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f351,f499])).
fof(f1108,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl6),
  inference(cnf_transformation,[],[f1108_D])).
fof(f1108_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1173,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1172])).
fof(f1172,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f934,f545])).
fof(f934,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f347,f499])).
fof(f1105,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl4),
  inference(cnf_transformation,[],[f1105_D])).
fof(f1105_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1171,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1170])).
fof(f1170,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f933,f545])).
fof(f933,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f330,f499])).
fof(f1101,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl0),
  inference(cnf_transformation,[],[f1101_D])).
fof(f1101_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1169,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1168])).
fof(f1168,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f932,f545])).
fof(f932,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f327,f499])).
fof(f1143,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP6('T_0',X1,X2) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1142])).
fof(f1142,plain,(
  ( ! [X85,X83,X86,X84] : (~sP6('T_0',X84,X85) | sP7('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f888,f545])).
fof(f888,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~sP6('T_0',X84,X85) | sP7('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f398,f499])).
fof(f1140,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1139])).
fof(f1139,plain,(
  ( ! [X80,X81,X79,X82] : (~sP8('T_0',X80,X81) | sP9('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f887,f545])).
fof(f887,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~sP8('T_0',X80,X81) | sP9('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f397,f499])).
fof(f1137,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1136])).
fof(f1136,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~sP9('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f886,f545])).
fof(f886,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | ~sP9('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f394,f499])).
fof(f1134,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1133])).
fof(f1133,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | sP9('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f885,f545])).
fof(f885,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | ~p(s(bool,X72)) | sP9('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f393,f499])).
fof(f1131,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1130])).
fof(f1130,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f884,f545])).
fof(f884,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f392,f499])).
fof(f1128,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP9(X2,X1,'T_0') | p(s(bool,X2))) )),
  inference(condensation,[],[f1127])).
fof(f1127,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | sP9(X65,X64,'T_0') | p(s(bool,X65)) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f883,f545])).
fof(f883,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | sP9(X65,X64,'T_0') | p(s(bool,X65)) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f391,f499])).
fof(f1125,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f1124])).
fof(f1124,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X60)) | ~p(s(bool,X61)) | ~sP8(X61,X60,'T_0') | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f882,f545])).
fof(f882,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | p(s(bool,X60)) | ~p(s(bool,X61)) | ~sP8(X61,X60,'T_0') | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f387,f499])).
fof(f1122,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f1121])).
fof(f1121,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP7('T_0',X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f881,f545])).
fof(f881,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP7('T_0',X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f377,f499])).
fof(f1119,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1118])).
fof(f1118,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | sP7('T_0',X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f880,f545])).
fof(f880,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | sP7('T_0',X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f376,f499])).
fof(f1116,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f1115])).
fof(f1115,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | sP7('T_0',X49,X48) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f879,f545])).
fof(f879,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | sP7('T_0',X49,X48) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f375,f499])).
fof(f1113,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f1112])).
fof(f1112,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | ~sP6('T_0',X44,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f878,f545])).
fof(f878,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | p(s(bool,X44)) | p(s(bool,X45)) | ~sP6('T_0',X44,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f370,f499])).
fof(f1110,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1109])).
fof(f1109,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f877,f545])).
fof(f877,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f351,f499])).
fof(f1107,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1106])).
fof(f1106,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f876,f545])).
fof(f876,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f347,f499])).
fof(f1104,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1103])).
fof(f1103,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f875,f545])).
fof(f875,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f330,f499])).
fof(f1100,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1099])).
fof(f1099,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f874,f545])).
fof(f874,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f327,f499])).
% SZS output end Proof for thm.int_arith
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 381
% Passive clauses: 1183
% Generated clauses: 10772
% Final active clauses: 239
% Final passive clauses: 636
% Input formulas: 51
% Initial clauses: 161
% 
% Duplicate literals: 187
% Trivial inequalities: 31
% Fw subsumption resolutions: 222
% Fw demodulations: 2374
% Bw demodulations: 157
% Condensations: 144
% Global subsumptions: 3
% 
% Simple tautologies: 60
% Equational tautologies: 77
% Forward subsumptions: 6240
% Fw demodulations to eq. taut.: 11
% 
% Unit resulting resolution: 1583
% Factoring: 1
% Forward superposition: 3749
% Backward superposition: 1850
% Self superposition: 80
% Equality factoring: 28
% Equality resolution: 18
% 
% Split clauses: 48
% Split components: 101
% SAT solver clauses: 1476
% SAT solver unit clauses: 580
% SAT solver binary clauses: 358
% 
% Sat splits: 48
% Sat splitting refutations: 192
% 
% Memory used [KB]: 14200
% Time elapsed: 0.476 s
% ------------------------------
% ------------------------------
% Success in time 1.794 s
