Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (81413)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 177
% Passive clauses: 2645
% Generated clauses: 18381
% Final active clauses: 165
% Final passive clauses: 2243
% Input formulas: 29
% Initial clauses: 82
% 
% Duplicate literals: 958
% Trivial inequalities: 255
% Fw subsumption resolutions: 2875
% Fw demodulations: 1941
% Bw demodulations: 12
% 
% Simple tautologies: 402
% Equational tautologies: 45
% Forward subsumptions: 6610
% Fw demodulations to eq. taut.: 14
% 
% Binary resolution: 656
% Forward superposition: 4684
% Backward superposition: 6350
% Self superposition: 174
% Equality factoring: 127
% Equality resolution: 48
% 
% Split clauses: 1245
% Split components: 2586
% SAT solver clauses: 1943
% SAT solver unit clauses: 33
% SAT solver binary clauses: 233
% 
% Sat splits: 1245
% 
% TWLsolver clauses: 1932
% TWLsolver calls for satisfiability: 239
% 
% Memory used [KB]: 10746
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (81426)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 146
% Passive clauses: 1507
% Generated clauses: 11286
% Final active clauses: 126
% Final passive clauses: 1266
% Input formulas: 29
% Initial clauses: 82
% 
% Duplicate literals: 512
% Trivial inequalities: 109
% Fw subsumption resolutions: 749
% Fw demodulations: 1060
% Bw demodulations: 98
% Condensations: 459
% Global subsumptions: 8
% 
% Simple tautologies: 149
% Equational tautologies: 24
% Forward subsumptions: 4155
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 195
% Forward superposition: 3804
% Backward superposition: 2445
% Self superposition: 79
% Equality factoring: 34
% Equality resolution: 38
% 
% Split clauses: 153
% Split components: 311
% SAT solver clauses: 2595
% SAT solver unit clauses: 86
% SAT solver binary clauses: 60
% SAT solver learnt clauses: 17
% SAT solver learnt literals: 1
% 
% Sat splits: 153
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 2562
% TWLsolver calls for satisfiability: 505
% 
% Memory used [KB]: 12665
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (81447)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 425
% Passive clauses: 3263
% Generated clauses: 149887
% Final active clauses: 121
% Final passive clauses: 871
% Input formulas: 29
% Initial clauses: 82
% Discarded non-redundant clauses: 362
% 
% Duplicate literals: 7470
% Trivial inequalities: 1393
% Fw subsumption resolutions: 16523
% Fw demodulations: 15619
% Bw demodulations: 1526
% Condensations: 17046
% Global subsumptions: 179
% 
% Simple tautologies: 1775
% Equational tautologies: 687
% Forward subsumptions: 85242
% Fw demodulations to eq. taut.: 316
% Bw demodulations to eq. taut.: 71
% 
% Unit resulting resolution: 3119
% Forward superposition: 39609
% Backward superposition: 37316
% Self superposition: 728
% Equality factoring: 347
% Equality resolution: 120
% 
% Split clauses: 336
% Split components: 681
% SAT solver clauses: 4361
% SAT solver unit clauses: 164
% SAT solver binary clauses: 234
% 
% Sat splits: 336
% Sat splitting refutations: 245
% 
% Memory used [KB]: 61022
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f15287,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f211,f212,f213,f214,f215,f269,f216,f218,f229,f267,f281,f336,f283,f284,f285,f338,f270,f337,f288,f298,f374,f361,f398,f401,f414,f413,f403,f419,f418,f421,f426,f407,f429,f428,f450,f451,f487,f453,f488,f484,f483,f492,f514,f515,f517,f518,f519,f520,f521,f560,f486,f564,f558,f557,f597,f606,f608,f609,f611,f613,f624,f671,f719,f726,f728,f729,f730,f733,f743,f782,f1042,f1045,f1046,f1044,f1068,f1070,f1071,f1072,f1075,f1085,f1121,f1123,f1124,f1126,f1128,f1139,f1181,f1183,f1184,f1185,f1188,f1198,f1113,f1232,f1234,f1235,f1236,f1239,f1249,f1283,f1285,f1286,f1288,f1290,f1301,f1332,f1339,f1341,f1342,f1343,f1346,f1356,f1277,f1440,f1442,f1443,f1444,f1447,f1457,f1485,f1690,f1691,f1753,f1751,f1750,f1755,f2633,f2632,f2631,f2638,f2673,f2672,f2671,f2712,f2713,f2714,f2715,f2761,f3010,f2764,f2765,f3011,f3016,f3017,f3019,f3022,f3025,f3026,f3028,f3029,f3031,f3038,f3040,f3043,f3049,f3051,f3055,f3056,f3057,f3059,f3062,f3063,f3065,f3066,f3076,f3077,f3083,f1754,f3093,f3095,f3096,f3098,f3100,f3111,f3144,f3170,f3172,f3173,f3174,f3177,f3187,f3219,f3276,f3278,f3279,f3280,f3283,f3293,f3326,f3333,f3335,f3336,f3337,f3340,f3350,f335,f3419,f3418,f3417,f3423,f3449,f3451,f3452,f3453,f3456,f3466,f3601,f3470,f3472,f3473,f3475,f3477,f3488,f3602,f3680,f3682,f3683,f3684,f3687,f3697,f3731,f3742,f3744,f3745,f3746,f3749,f3759,f3840,f3762,f3764,f3765,f3766,f3769,f3779,f3842,f3854,f3856,f3857,f3858,f3861,f3871,f3952,f3874,f3876,f3877,f3878,f3881,f3891,f3954,f4044,f4046,f4047,f4048,f4051,f4061,f4104,f4108,f4110,f4111,f4112,f4115,f4125,f4157,f4168,f4170,f4171,f4172,f4175,f4185,f4289,f4188,f4190,f4191,f4193,f4195,f4206,f4290,f4302,f4304,f4305,f4307,f4309,f4320,f4369,f4326,f4327,f4329,f4330,f4331,f4332,f4333,f4375,f4368,f4367,f4370,f4422,f4424,f4425,f4426,f4429,f4439,f4544,f4443,f4445,f4446,f4448,f4450,f4461,f4545,f4561,f4563,f4564,f4566,f4568,f4579,f4644,f4854,f4821,f4810,f4808,f4806,f4805,f4803,f5087,f411,f4865,f4867,f4868,f4869,f4872,f4882,f5088,f4966,f4967,f4970,f4972,f4975,f4976,f4977,f4979,f4980,f4981,f4984,f4986,f4987,f4988,f4989,f4991,f4995,f4996,f4997,f4999,f5029,f5094,f5030,f5093,f5031,f5096,f5033,f5097,f5108,f5110,f5111,f5113,f5115,f5126,f1174,f5443,f5445,f5446,f5448,f5450,f5461,f5477,f5479,f5480,f5482,f5484,f5495,f6303,f5819,f5820,f5823,f5825,f5828,f5829,f5830,f5832,f5833,f5834,f5837,f5839,f5840,f5841,f5842,f5844,f5848,f5849,f5850,f5852,f5882,f5883,f5884,f5886,f5913,f5914,f5916,f5918,f5919,f5920,f5921,f5934,f6385,f5937,f5938,f6080,f6081,f6084,f6086,f6089,f6090,f6091,f6093,f6094,f6095,f6098,f6100,f6101,f6102,f6103,f6105,f6109,f6110,f6111,f6113,f6143,f6144,f6145,f6372,f6147,f6373,f6174,f6175,f6177,f6179,f6180,f6181,f6182,f6382,f6195,f6409,f6198,f6199,f6384,f6427,f6429,f6430,f6431,f6434,f6444,f6487,f6489,f6490,f6492,f6494,f6505,f6554,f6556,f6557,f6559,f6561,f6572,f6834,f6706,f6707,f6710,f6712,f6715,f6716,f6717,f6719,f6720,f6721,f6724,f6726,f6727,f6728,f6729,f6731,f6735,f6736,f6737,f6739,f6769,f6843,f6770,f6842,f6771,f6845,f6773,f6846,f6889,f6877,f6918,f6833,f416,f423,f6927,f6929,f6930,f6931,f6934,f6944,f6481,f6988,f6990,f6991,f6992,f6995,f7005,f7047,f7010,f7011,f7051,f7013,f7055,f7046,f7045,f7050,f7110,f7109,f7113,f7112,f7117,f7118,f410,f7119,f7049,f7652,f7654,f7655,f7657,f7659,f7670,f7726,f7789,f7790,f7791,f7792,f7795,f7797,f7798,f7800,f7802,f7813,f7869,f7936,f7938,f7939,f7941,f7943,f7954,f8010,f8159,f8161,f8162,f8164,f8166,f8177,f8214,f8267,f8269,f8270,f8271,f8274,f8284,f8331,f8339,f8341,f8342,f8343,f8346,f8356,f8403,f8411,f8413,f8414,f8415,f8418,f8428,f8475,f8485,f8487,f8488,f8489,f8492,f8502,f8549,f8834,f8835,f8872,f8874,f8875,f8877,f8879,f8890,f4860,f9523,f9524,f9794,f10534,f10536,f10537,f10539,f10541,f10552,f10593,f10960,f10962,f10963,f10965,f10967,f10978,f11039,f11075,f11077,f11078,f11080,f11082,f11093,f11154,f11179,f11181,f11182,f11183,f11186,f11196,f11248,f11257,f11259,f11260,f11262,f11264,f11275,f11337,f11697,f11741,f11897,f11899,f11900,f11902,f11904,f11915,f11977,f11926,f11927,f11929,f11930,f11931,f11932,f11933,f11983,f11976,f11975,f11978,f12216,f12218,f12219,f12221,f12223,f12234,f12296,f12245,f12246,f12248,f12249,f12250,f12251,f12252,f12302,f12295,f12294,f12297,f12309,f12311,f12312,f12313,f12316,f12326,f12381,f12393,f12395,f12396,f12397,f12400,f12410,f12465,f12726,f12728,f12729,f12730,f12733,f12743,f12798,f12820,f12822,f12823,f12824,f12827,f12837,f12892,f13207,f13208,f13209,f13210,f13211,f13212,f13487,f13570,f9795,f14296,f14298,f14299,f14300,f14303,f14313,f14361,f14781,f14784,f14787,f14788,f14789,f14791,f14792,f14793,f14795,f14796,f14801,f493,f14802,f14806,f14807,f14809,f14811,f14812,f14810,f14818,f13488,f14820,f13571,f14821,f14822,f2637,f14827,f14828,f14830,f14831,f14782,f14840,f14842,f14843,f14844,f14847,f14857,f14905,f14912,f14914,f14915,f14916,f14919,f14929,f14785,f15187,f15189,f15190,f268,f15191,f272,f15194,f271,f15204,f15252,f273,f15286])).
fof(f15286,plain,(
  $false | ($spl26 | $spl29)),
  inference(trivial_inequality_removal,[],[f15279])).
fof(f15279,plain,(
  s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) != s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | ($spl26 | $spl29)),
  inference(superposition,[],[f15270,f423])).
fof(f15270,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) != s('type.sum.sum'(X0,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2))))) ) | $spl29),
  inference(equality_resolution,[],[f1491])).
fof(f1491,plain,(
  ( ! [X4,X5] : (s(sK2,X4) != s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)) | s('type.sum.sum'(X5,sK2),'const.sum.INR_1'(s(sK2,X4))) != s('type.sum.sum'(X5,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2))))) ) | $spl29),
  inference(superposition,[],[f428,f139])).
fof(f139,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,X3) | s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3] : ((s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) & (s(X1,X2) != s(X1,X3) | s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X7,X6] : ! [X11,X10] : (s('type.sum.sum'(X7,X6),'const.sum.INR_1'(s(X6,X10))) = s('type.sum.sum'(X7,X6),'const.sum.INR_1'(s(X6,X11))) <=> s(X6,X10) = s(X6,X11))),
  file('thm.pred_set.INJ_INR','thm.sum.INR_|49||49|')).
fof(f273,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl12),
  inference(cnf_transformation,[],[f273_D])).
fof(f273_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f15252,plain,(
  ( ! [X39,X37,X35,X33,X36,X34] : (sP0(X33,X34,X35,X36,X37) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(X33,sK5(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36))) | p(s(bool,X39)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X33,sK7(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36)))) ) | $spl218),
  inference(cnf_transformation,[],[f15252_D])).
fof(f15252_D,plain,(
  ( ! [X39,X37,X35,X33,X36,X34] : (sP0(X33,X34,X35,X36,X37) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(X33,sK5(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36))) | p(s(bool,X39)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X33,sK7(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f15204,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X148] : (sP0(X142,X143,X144,X145,X146) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X142,sK7(X146,X145,X144,X143,X142)),s(fun(X142,bool),X145))) | p(s(bool,X147)) | p(s(bool,X148)) | s(bool,X148) != s(bool,'const.bool.IN_2'(s(X142,sK5(X146,X145,X144,X143,X142)),s(fun(X142,bool),X145)))) )),
  inference(resolution,[],[f3869,f206])).
fof(f206,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f147])).
fof(f147,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X10,X15,X11,X16] : ((s(bool,X10) = s(bool,X15) & (p(s(bool,X15)) => s(bool,X11) = s(bool,X16))) => ((p(s(bool,X10)) => p(s(bool,X11))) <=> (p(s(bool,X15)) => p(s(bool,X16)))))),
  file('thm.pred_set.INJ_INR','thm.bool.IMP_CONG')).
fof(f3869,plain,(
  ( ! [X111,X109,X107,X110,X108] : (p(s(bool,'const.bool.IN_2'(s(X107,sK5(X111,X110,X109,X108,X107)),s(fun(X107,bool),X110)))) | sP0(X107,X108,X109,X110,X111) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X107,sK7(X111,X110,X109,X108,X107)),s(fun(X107,bool),X110)))) )),
  inference(resolution,[],[f130,f196])).
fof(f196,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.INJ_INR','thm.bool.EQ_CLAUSES|split|1')).
fof(f130,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X3,X4) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK5(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP0(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK5(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK5(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,sK6(X4,X3,X2,X1,X0)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK7(X4,X3,X2,X1,X0)))) & s(X0,sK6(X4,X3,X2,X1,X0)) != s(X0,sK7(X4,X3,X2,X1,X0))) | sP0(X0,X1,X2,X3,X4)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7])],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP0(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))) & s(X0,X5) != s(X0,X6)) | sP0(X0,X1,X2,X3,X4)))),
  inference(rectify,[],[f85])).
fof(f85,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP0(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6)) | sP0(X1,X0,X2,X3,X4)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP0(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6))) | sP0(X1,X0,X2,X3,X4)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X1,X0,X2,X3,X4] : (sP0(X1,X0,X2,X3,X4) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f271,plain,(
  ( ! [X6,X8] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,X8)) ) | $spl8),
  inference(cnf_transformation,[],[f271_D])).
fof(f271_D,plain,(
  ( ! [X6,X8] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,X8)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f15194,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89,X96] : (sP0(X89,X90,X91,X92,X93) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X89,sK7(X93,X92,X91,X90,X89)),s(fun(X89,bool),X92))) | p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,'const.bool.IN_2'(s(X89,sK5(X93,X92,X91,X90,X89)),s(fun(X89,bool),X92))) | s(bool,X95) != s(bool,X96)) )),
  inference(resolution,[],[f3869,f150])).
fof(f150,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f272,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | $spl10),
  inference(cnf_transformation,[],[f272_D])).
fof(f272_D,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f15191,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X65] : (sP0(X65,X66,X67,X68,X69) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X65,sK7(X69,X68,X67,X66,X65)),s(fun(X65,bool),X68))) | p(s(bool,X70)) | ~p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X72) != s(bool,'const.bool.IN_2'(s(X65,sK5(X69,X68,X67,X66,X65)),s(fun(X65,bool),X68)))) )),
  inference(resolution,[],[f3869,f149])).
fof(f149,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f268,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) | $spl2),
  inference(cnf_transformation,[],[f268_D])).
fof(f268_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f15190,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X63] : (sP0(X57,X58,X59,X60,X61) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X57,sK7(X61,X60,X59,X58,X57)),s(fun(X57,bool),X60))) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,'const.bool.IN_2'(s(X57,sK5(X61,X60,X59,X58,X57)),s(fun(X57,bool),X60)))) )),
  inference(resolution,[],[f3869,f148])).
fof(f148,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f15189,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (sP0(X49,X50,X51,X52,X53) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X49,sK7(X53,X52,X51,X50,X49)),s(fun(X49,bool),X52))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(X49,sK5(X53,X52,X51,X50,X49)),s(fun(X49,bool),X52))) | s(bool,X55) != s(bool,X56)) )),
  inference(resolution,[],[f3869,f146])).
fof(f146,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f15187,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X40] : (sP0(X33,X34,X35,X36,X37) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X33,sK7(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36))) | p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(X33,sK5(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36)))) )),
  inference(resolution,[],[f3869,f144])).
fof(f144,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f14785,plain,(
  ( ! [X44,X42] : (p(s(bool,X44)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X42),s(fun(sK2,bool),sK4))) | s(bool,X44) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X42))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl204),
  inference(cnf_transformation,[],[f14785_D])).
fof(f14785_D,plain,(
  ( ! [X44,X42] : (p(s(bool,X44)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X42),s(fun(sK2,bool),sK4))) | s(bool,X44) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X42))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f14929,plain,(
  ( ! [X57,X56,X55] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X55),s(fun(sK2,bool),sK4))) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X57) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X55))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f14769,f206])).
fof(f14769,plain,(
  ( ! [X177] : (p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X177))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X177),s(fun(sK2,bool),sK4)))) )),
  inference(trivial_inequality_removal,[],[f14616])).
fof(f14616,plain,(
  ( ! [X177] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X177))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X177),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f197,f8173])).
fof(f8173,plain,(
  ( ! [X66] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X66))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X66),s(fun(sK2,bool),sK4)))) )),
  inference(resolution,[],[f296,f192])).
fof(f192,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.INJ_INR','thm.bool.EQ_CLAUSES|split|3')).
fof(f296,plain,(
  ( ! [X53] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X53),s(fun(sK2,bool),sK4)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X53))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f209,f196])).
fof(f209,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X4))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK2,X4),s(fun(sK2,bool),sK4))))) )),
  inference(forward_demodulation,[],[f124,f138])).
fof(f138,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),happ(s(fun(X1,'type.sum.sum'(X0,X1)),'const.sum.INR_0'),s(X1,X2)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2] : s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),happ(s(fun(X1,'type.sum.sum'(X0,X1)),'const.sum.INR_0'),s(X1,X2)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ! [X2] : s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),happ(s(fun(X1,'type.sum.sum'(X0,X1)),'const.sum.INR_0'),s(X1,X2)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X7,X6] : ! [X18] : s('type.sum.sum'(X7,X6),'const.sum.INR_1'(s(X6,X18))) = s('type.sum.sum'(X7,X6),happ(s(fun(X6,'type.sum.sum'(X7,X6)),'const.sum.INR_0'),s(X6,X18)))),
  file('thm.pred_set.INJ_INR','ARITY_0_1_const.sum.INR')).
fof(f124,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,X4))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK2,X4),s(fun(sK2,bool),sK4))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X4),s(fun(sK2,bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,X4))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))) & ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f68])).
fof(f68,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X0,X1),happ(s(fun(X1,'type.sum.sum'(X0,X1)),'const.sum.INR_0'),s(X1,X4))),s(fun('type.sum.sum'(X0,X1),bool),X2))))) & ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,'type.sum.sum'(X0,X1)),'const.sum.INR_0'),s(fun(X1,bool),X3),s(fun('type.sum.sum'(X0,X1),bool),X2)))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ~! [X0,X1,X2,X3] : (! [X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X0,X1),happ(s(fun(X1,'type.sum.sum'(X0,X1)),'const.sum.INR_0'),s(X1,X4))),s(fun('type.sum.sum'(X0,X1),bool),X2))))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,'type.sum.sum'(X0,X1)),'const.sum.INR_0'),s(fun(X1,bool),X3),s(fun('type.sum.sum'(X0,X1),bool),X2)))))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ~! [X0,X1] : ! [X2,X3] : (! [X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X0,X1),happ(s(fun(X1,'type.sum.sum'(X0,X1)),'const.sum.INR_0'),s(X1,X4))),s(fun('type.sum.sum'(X0,X1),bool),X2))))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,'type.sum.sum'(X0,X1)),'const.sum.INR_0'),s(fun(X1,bool),X3),s(fun('type.sum.sum'(X0,X1),bool),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (! [X10] : (p(s(bool,'const.bool.IN_2'(s(X7,X10),s(fun(X7,bool),X9)))) => p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X6,X7),happ(s(fun(X7,'type.sum.sum'(X6,X7)),'const.sum.INR_0'),s(X7,X10))),s(fun('type.sum.sum'(X6,X7),bool),X8))))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,'type.sum.sum'(X6,X7)),'const.sum.INR_0'),s(fun(X7,bool),X9),s(fun('type.sum.sum'(X6,X7),bool),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (! [X10] : (p(s(bool,'const.bool.IN_2'(s(X7,X10),s(fun(X7,bool),X9)))) => p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X6,X7),happ(s(fun(X7,'type.sum.sum'(X6,X7)),'const.sum.INR_0'),s(X7,X10))),s(fun('type.sum.sum'(X6,X7),bool),X8))))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,'type.sum.sum'(X6,X7)),'const.sum.INR_0'),s(fun(X7,bool),X9),s(fun('type.sum.sum'(X6,X7),bool),X8)))))),
  file('thm.pred_set.INJ_INR','thm.pred_set.INJ_INR')).
fof(f197,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.INJ_INR','thm.bool.EQ_CLAUSES|split|0')).
fof(f14919,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X42),s(fun(sK2,bool),sK4))) | p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X42))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X44) != s(bool,X45)) )),
  inference(resolution,[],[f14769,f150])).
fof(f14916,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X30),s(fun(sK2,bool),sK4))) | p(s(bool,X31)) | ~p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X33) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X30))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f14769,f149])).
fof(f14915,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X26),s(fun(sK2,bool),sK4))) | p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X26))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f14769,f148])).
fof(f14914,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X22),s(fun(sK2,bool),sK4))) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X22))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X24) != s(bool,X25)) )),
  inference(resolution,[],[f14769,f146])).
fof(f14912,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X14),s(fun(sK2,bool),sK4))) | p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X14))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f14769,f144])).
fof(f14905,plain,(
  ( ! [X39,X37,X35,X33,X36,X34] : (sP0(X33,X34,X35,X36,X37) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(X33,sK6(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36))) | p(s(bool,X39)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X33,sK5(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36)))) ) | $spl216),
  inference(cnf_transformation,[],[f14905_D])).
fof(f14905_D,plain,(
  ( ! [X39,X37,X35,X33,X36,X34] : (sP0(X33,X34,X35,X36,X37) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(X33,sK6(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36))) | p(s(bool,X39)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X33,sK5(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f14857,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X148] : (sP0(X142,X143,X144,X145,X146) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X142,sK5(X146,X145,X144,X143,X142)),s(fun(X142,bool),X145))) | p(s(bool,X147)) | p(s(bool,X148)) | s(bool,X148) != s(bool,'const.bool.IN_2'(s(X142,sK6(X146,X145,X144,X143,X142)),s(fun(X142,bool),X145)))) )),
  inference(resolution,[],[f3777,f206])).
fof(f3777,plain,(
  ( ! [X111,X109,X107,X110,X108] : (p(s(bool,'const.bool.IN_2'(s(X107,sK6(X111,X110,X109,X108,X107)),s(fun(X107,bool),X110)))) | sP0(X107,X108,X109,X110,X111) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X107,sK5(X111,X110,X109,X108,X107)),s(fun(X107,bool),X110)))) )),
  inference(resolution,[],[f128,f196])).
fof(f128,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X3,X4) | p(s(bool,'const.bool.IN_2'(s(X0,sK6(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK5(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f14847,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89,X96] : (sP0(X89,X90,X91,X92,X93) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X89,sK5(X93,X92,X91,X90,X89)),s(fun(X89,bool),X92))) | p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,'const.bool.IN_2'(s(X89,sK6(X93,X92,X91,X90,X89)),s(fun(X89,bool),X92))) | s(bool,X95) != s(bool,X96)) )),
  inference(resolution,[],[f3777,f150])).
fof(f14844,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X65] : (sP0(X65,X66,X67,X68,X69) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X65,sK5(X69,X68,X67,X66,X65)),s(fun(X65,bool),X68))) | p(s(bool,X70)) | ~p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X72) != s(bool,'const.bool.IN_2'(s(X65,sK6(X69,X68,X67,X66,X65)),s(fun(X65,bool),X68)))) )),
  inference(resolution,[],[f3777,f149])).
fof(f14843,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X63] : (sP0(X57,X58,X59,X60,X61) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X57,sK5(X61,X60,X59,X58,X57)),s(fun(X57,bool),X60))) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,'const.bool.IN_2'(s(X57,sK6(X61,X60,X59,X58,X57)),s(fun(X57,bool),X60)))) )),
  inference(resolution,[],[f3777,f148])).
fof(f14842,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (sP0(X49,X50,X51,X52,X53) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X49,sK5(X53,X52,X51,X50,X49)),s(fun(X49,bool),X52))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(X49,sK6(X53,X52,X51,X50,X49)),s(fun(X49,bool),X52))) | s(bool,X55) != s(bool,X56)) )),
  inference(resolution,[],[f3777,f146])).
fof(f14840,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X40] : (sP0(X33,X34,X35,X36,X37) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X33,sK5(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36))) | p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(X33,sK6(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36)))) )),
  inference(resolution,[],[f3777,f144])).
fof(f14782,plain,(
  ( ! [X18] : (p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X18))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X18),s(fun(sK2,bool),sK4)))) ) | $spl202),
  inference(cnf_transformation,[],[f14782_D])).
fof(f14782_D,plain,(
  ( ! [X18] : (p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X18))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X18),s(fun(sK2,bool),sK4)))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f14831,plain,(
  ( ! [X461,X459,X460] : (s('type.sum.sum'(X459,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X459,bool),'const.sum.INR_1'(s(bool,X460))) | s(bool,X460) != s(bool,X461)) ) | $spl214),
  inference(cnf_transformation,[],[f14831_D])).
fof(f14831_D,plain,(
  ( ! [X461,X459,X460] : (s('type.sum.sum'(X459,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X459,bool),'const.sum.INR_1'(s(bool,X460))) | s(bool,X460) != s(bool,X461)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f14830,plain,(
  ( ! [X461,X459,X460,X458] : (s('type.sum.sum'(X459,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X459,bool),'const.sum.INR_1'(s(bool,X460))) | s(bool,X460) != s(bool,X461) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X458))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X458),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14709,f3698])).
fof(f3698,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,X64)) | s('type.sum.sum'(X62,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X62,bool),'const.sum.INR_1'(s(bool,X63))) | s(bool,X63) != s(bool,X64)) )),
  inference(resolution,[],[f3428,f230])).
fof(f230,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f125,f206])).
fof(f125,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))),
  inference(cnf_transformation,[],[f83])).
fof(f3428,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,X2)))) )),
  inference(superposition,[],[f3411,f194])).
fof(f194,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.INJ_INR','thm.bool.EQ_CLAUSES|split|2')).
fof(f3411,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0')))) )),
  inference(equality_resolution,[],[f1431])).
fof(f1431,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,X4)))) )),
  inference(forward_demodulation,[],[f1420,f225])).
fof(f225,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))),
  inference(resolution,[],[f125,f192])).
fof(f1420,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,X4))) != s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))))) )),
  inference(superposition,[],[f227,f139])).
fof(f227,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))),
  inference(resolution,[],[f125,f195])).
fof(f195,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f120])).
fof(f14709,plain,(
  ( ! [X461,X459,X460,X458] : (s('type.sum.sum'(X459,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X459,bool),'const.sum.INR_1'(s(bool,X460))) | s(bool,X460) != s(bool,X461) | ~p(s(bool,X461)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X458))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X458),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f7590,f8173])).
fof(f7590,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,X4))) != s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,X6))) | s(bool,X6) != s(bool,X7) | ~p(s(bool,X7)) | p(s(bool,X4))) )),
  inference(superposition,[],[f3382,f194])).
fof(f3382,plain,(
  ( ! [X4,X5,X3] : (s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,X4))) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f1431,f198])).
fof(f198,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f121])).
fof(f14828,plain,(
  ( ! [X445,X443,X444] : (s('type.sum.sum'(bool,X443),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X443),'const.sum.INL_1'(s(bool,X444))) | s(bool,X444) != s(bool,X445)) ) | $spl212),
  inference(cnf_transformation,[],[f14828_D])).
fof(f14828_D,plain,(
  ( ! [X445,X443,X444] : (s('type.sum.sum'(bool,X443),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X443),'const.sum.INL_1'(s(bool,X444))) | s(bool,X444) != s(bool,X445)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f14827,plain,(
  ( ! [X445,X443,X444,X442] : (s('type.sum.sum'(bool,X443),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X443),'const.sum.INL_1'(s(bool,X444))) | s(bool,X444) != s(bool,X445) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X442))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X442),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14705,f3294])).
fof(f3294,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,X64)) | s('type.sum.sum'(bool,X62),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X62),'const.sum.INL_1'(s(bool,X63))) | s(bool,X63) != s(bool,X64)) )),
  inference(resolution,[],[f3149,f230])).
fof(f3149,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s('type.sum.sum'(bool,X3),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X3),'const.sum.INL_1'(s(bool,X2)))) )),
  inference(superposition,[],[f2625,f194])).
fof(f2625,plain,(
  ( ! [X0] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))) )),
  inference(equality_resolution,[],[f1429])).
fof(f1429,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X0) | s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))) != s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0)))) )),
  inference(forward_demodulation,[],[f1418,f225])).
fof(f1418,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X0) | s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0))) != s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))))) )),
  inference(superposition,[],[f227,f141])).
fof(f141,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,X3) | s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) != s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3)))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3] : ((s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) != s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) & (s(X1,X2) != s(X1,X3) | s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) = s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3)))))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3] : (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) = s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) = s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7] : ! [X11,X10] : (s('type.sum.sum'(X7,X6),'const.sum.INL_1'(s(X7,X10))) = s('type.sum.sum'(X7,X6),'const.sum.INL_1'(s(X7,X11))) <=> s(X7,X10) = s(X7,X11))),
  file('thm.pred_set.INJ_INR','thm.sum.INR_INL_|49||49||split|0')).
fof(f14705,plain,(
  ( ! [X445,X443,X444,X442] : (s('type.sum.sum'(bool,X443),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X443),'const.sum.INL_1'(s(bool,X444))) | s(bool,X444) != s(bool,X445) | ~p(s(bool,X445)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X442))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X442),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f7527,f8173])).
fof(f7527,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.sum.sum'(bool,X5),'const.sum.INL_1'(s(bool,X4))) != s('type.sum.sum'(bool,X5),'const.sum.INL_1'(s(bool,X6))) | s(bool,X6) != s(bool,X7) | ~p(s(bool,X7)) | p(s(bool,X4))) )),
  inference(superposition,[],[f2590,f194])).
fof(f2590,plain,(
  ( ! [X4,X5,X3] : (s('type.sum.sum'(bool,X5),'const.sum.INL_1'(s(bool,'F_0'))) != s('type.sum.sum'(bool,X5),'const.sum.INL_1'(s(bool,X4))) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f1429,f198])).
fof(f2637,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,X12)) ) | $spl54),
  inference(cnf_transformation,[],[f2637_D])).
fof(f2637_D,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,X12)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f14822,plain,(
  ( ! [X408,X405,X407] : (s(bool,X407) != s(bool,X408) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X405),s(fun(sK2,bool),sK4))) | s(bool,X407) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X405))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X408)) ) | $spl210),
  inference(cnf_transformation,[],[f14822_D])).
fof(f14822_D,plain,(
  ( ! [X408,X405,X407] : (s(bool,X407) != s(bool,X408) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X405),s(fun(sK2,bool),sK4))) | s(bool,X407) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X405))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X408)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f14821,plain,(
  ( ! [X408,X405,X407,X406] : (s(bool,X407) != s(bool,X408) | s(bool,'T_0') = s(bool,X408) | s(bool,X407) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X405))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,X406) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X405),s(fun(sK2,bool),sK4)))) ) | $spl34),
  inference(subsumption_resolution,[],[f14691,f205])).
fof(f205,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.INJ_INR',aHL_BOOL_CASES)).
fof(f14691,plain,(
  ( ! [X408,X405,X407,X406] : (s(bool,'T_0') != s(bool,X406) | s(bool,X407) != s(bool,X408) | s(bool,'T_0') = s(bool,X408) | s(bool,X407) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X405))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,X406) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X405),s(fun(sK2,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f4322,f8173])).
fof(f4322,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X89) != s(bool,X91) | s(bool,X88) != s(bool,X90) | s(bool,'T_0') = s(bool,X90) | s(bool,X88) != s(bool,X89) | s(bool,'F_0') = s(bool,X91)) ) | $spl34),
  inference(resolution,[],[f1758,f620])).
fof(f620,plain,(
  ( ! [X52,X53] : (p(s(bool,X52)) | s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53)) )),
  inference(resolution,[],[f230,f192])).
fof(f1758,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X29)) | s(bool,X28) != s(bool,X29) | s(bool,X27) != s(bool,X28) | s(bool,'T_0') = s(bool,X27)) ) | $spl34),
  inference(subsumption_resolution,[],[f1716,f162])).
fof(f162,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.INJ_INR','thm.bool.IMP_CLAUSES|split|4')).
fof(f1716,plain,(
  ( ! [X28,X29,X27] : (s(bool,X27) != s(bool,X28) | s(bool,X28) != s(bool,X29) | ~p(s(bool,X29)) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X27)) ) | $spl34),
  inference(superposition,[],[f671,f205])).
fof(f13571,plain,(
  ( ! [X61,X60] : (s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61)))) ) | $spl198),
  inference(cnf_transformation,[],[f13571_D])).
fof(f13571_D,plain,(
  ( ! [X61,X60] : (s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61)))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f14820,plain,(
  ( ! [X377,X379,X378] : (s('type.sum.sum'(X378,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X378,bool),'const.sum.INR_1'(s(bool,X379))) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X377))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X377),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14682,f3428])).
fof(f14682,plain,(
  ( ! [X377,X379,X378] : (s('type.sum.sum'(X378,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X378,bool),'const.sum.INR_1'(s(bool,X379))) | ~p(s(bool,X379)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X377))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X377),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f3628,f8173])).
fof(f3628,plain,(
  ( ! [X4,X5,X3] : (s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X3))) != s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))) | ~p(s(bool,X5)) | p(s(bool,X3))) )),
  inference(superposition,[],[f3421,f194])).
fof(f3421,plain,(
  ( ! [X6,X7] : (s('type.sum.sum'(X7,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X7,bool),'const.sum.INR_1'(s(bool,X6))) | ~p(s(bool,X6))) )),
  inference(trivial_inequality_removal,[],[f3394])).
fof(f3394,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'T_0') | s('type.sum.sum'(X7,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X7,bool),'const.sum.INR_1'(s(bool,X6))) | ~p(s(bool,X6))) )),
  inference(superposition,[],[f1431,f198])).
fof(f13488,plain,(
  ( ! [X61,X60] : (s('type.sum.sum'(bool,X60),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X60),'const.sum.INL_1'(s(bool,X61)))) ) | $spl196),
  inference(cnf_transformation,[],[f13488_D])).
fof(f13488_D,plain,(
  ( ! [X61,X60] : (s('type.sum.sum'(bool,X60),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X60),'const.sum.INL_1'(s(bool,X61)))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f14818,plain,(
  ( ! [X344,X346,X345] : (s('type.sum.sum'(bool,X345),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X345),'const.sum.INL_1'(s(bool,X346))) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X344))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X344),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14670,f3149])).
fof(f14670,plain,(
  ( ! [X344,X346,X345] : (s('type.sum.sum'(bool,X345),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X345),'const.sum.INL_1'(s(bool,X346))) | ~p(s(bool,X346)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X344))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X344),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f3224,f8173])).
fof(f3224,plain,(
  ( ! [X4,X5,X3] : (s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X3))) != s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))) | ~p(s(bool,X5)) | p(s(bool,X3))) )),
  inference(superposition,[],[f2634,f194])).
fof(f2634,plain,(
  ( ! [X6,X7] : (s('type.sum.sum'(bool,X7),'const.sum.INL_1'(s(bool,'F_0'))) != s('type.sum.sum'(bool,X7),'const.sum.INL_1'(s(bool,X6))) | ~p(s(bool,X6))) )),
  inference(trivial_inequality_removal,[],[f2604])).
fof(f2604,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'T_0') | s('type.sum.sum'(bool,X7),'const.sum.INL_1'(s(bool,'F_0'))) != s('type.sum.sum'(bool,X7),'const.sum.INL_1'(s(bool,X6))) | ~p(s(bool,X6))) )),
  inference(superposition,[],[f1429,f196])).
fof(f14810,plain,(
  ( ! [X256,X258] : (s(bool,X258) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X256))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X256),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X258)) ) | $spl208),
  inference(cnf_transformation,[],[f14810_D])).
fof(f14810_D,plain,(
  ( ! [X256,X258] : (s(bool,X258) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X256))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X256),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X258)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f14812,plain,(
  ( ! [X280,X279,X281,X278] : (s('type.sum.sum'(X279,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X279,bool),'const.sum.INR_1'(s(bool,X280))) | s(bool,'T_0') = s(bool,X281) | s(bool,X281) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X278))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X278),s(fun(sK2,bool),sK4)))) ) | $spl39),
  inference(subsumption_resolution,[],[f14652,f3428])).
fof(f14652,plain,(
  ( ! [X280,X279,X281,X278] : (s('type.sum.sum'(X279,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X279,bool),'const.sum.INR_1'(s(bool,X280))) | s(bool,'T_0') = s(bool,X281) | ~p(s(bool,X280)) | s(bool,X281) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X278))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X278),s(fun(sK2,bool),sK4)))) ) | $spl39),
  inference(superposition,[],[f1602,f8173])).
fof(f1602,plain,(
  ( ! [X45,X43,X46,X44] : (s('type.sum.sum'(X46,bool),'const.sum.INR_1'(s(bool,X43))) != s('type.sum.sum'(X46,bool),'const.sum.INR_1'(s(bool,X44))) | s(bool,'T_0') = s(bool,X45) | ~p(s(bool,X43)) | s(bool,X44) != s(bool,X45)) ) | $spl39),
  inference(superposition,[],[f1583,f139])).
fof(f1583,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X14))) ) | $spl39),
  inference(subsumption_resolution,[],[f1561,f1044])).
fof(f1561,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X14))) )),
  inference(superposition,[],[f741,f198])).
fof(f741,plain,(
  ( ! [X54,X55] : (s(bool,X54) != s(bool,X55) | s(bool,'F_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X54)) )),
  inference(resolution,[],[f620,f196])).
fof(f14811,plain,(
  ( ! [X271,X273,X270,X272] : (s('type.sum.sum'(bool,X271),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X271),'const.sum.INL_1'(s(bool,X272))) | s(bool,'T_0') = s(bool,X273) | s(bool,X273) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X270))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X270),s(fun(sK2,bool),sK4)))) ) | $spl39),
  inference(subsumption_resolution,[],[f14650,f3149])).
fof(f14650,plain,(
  ( ! [X271,X273,X270,X272] : (s('type.sum.sum'(bool,X271),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X271),'const.sum.INL_1'(s(bool,X272))) | s(bool,'T_0') = s(bool,X273) | ~p(s(bool,X272)) | s(bool,X273) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X270))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X270),s(fun(sK2,bool),sK4)))) ) | $spl39),
  inference(superposition,[],[f1600,f8173])).
fof(f1600,plain,(
  ( ! [X37,X35,X38,X36] : (s('type.sum.sum'(bool,X38),'const.sum.INL_1'(s(bool,X35))) != s('type.sum.sum'(bool,X38),'const.sum.INL_1'(s(bool,X36))) | s(bool,'T_0') = s(bool,X37) | ~p(s(bool,X35)) | s(bool,X36) != s(bool,X37)) ) | $spl39),
  inference(superposition,[],[f1583,f141])).
fof(f14809,plain,(
  ( ! [X257,X256,X258] : (s(bool,'F_0') = s(bool,X257) | s(bool,X258) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X256))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X258) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X256),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14645,f205])).
fof(f14645,plain,(
  ( ! [X257,X256,X258] : (s(bool,'T_0') != s(bool,X257) | s(bool,'F_0') = s(bool,X257) | s(bool,X258) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X256))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X258) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X256),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f1495,f8173])).
fof(f1495,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f739,f205])).
fof(f739,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') != s(bool,X50) | s(bool,'F_0') = s(bool,X51) | s(bool,X50) != s(bool,X51)) )),
  inference(resolution,[],[f620,f191])).
fof(f191,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f118])).
fof(f14807,plain,(
  ( ! [X213,X214] : (s(bool,'T_0') != s(bool,X213) | s(bool,X213) != s(bool,X214)) ) | $spl206),
  inference(cnf_transformation,[],[f14807_D])).
fof(f14807_D,plain,(
  ( ! [X213,X214] : (s(bool,'T_0') != s(bool,X213) | s(bool,X213) != s(bool,X214)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f14806,plain,(
  ( ! [X212,X213,X214] : (s(bool,'T_0') != s(bool,X213) | s(bool,X213) != s(bool,X214) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X212))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X212),s(fun(sK2,bool),sK4)))) ) | $spl34),
  inference(subsumption_resolution,[],[f14630,f622])).
fof(f622,plain,(
  ( ! [X57,X56] : (s(bool,'T_0') != s(bool,X57) | s(bool,X56) != s(bool,X57) | p(s(bool,X56))) )),
  inference(resolution,[],[f230,f195])).
fof(f14630,plain,(
  ( ! [X212,X213,X214] : (s(bool,'T_0') != s(bool,X213) | s(bool,X213) != s(bool,X214) | ~p(s(bool,X214)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X212))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X212),s(fun(sK2,bool),sK4)))) ) | $spl34),
  inference(superposition,[],[f671,f8173])).
fof(f14802,plain,(
  ( ! [X196,X197,X195] : (p(s(bool,X196)) | p(s(bool,X197)) | s(bool,X197) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X195))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X195),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14624,f207])).
fof(f207,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f159])).
fof(f159,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X7] : ! [X10] : (s(X7,X10) = s(X7,X10) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.INJ_INR','thm.bool.REFL_CLAUSE')).
fof(f14624,plain,(
  ( ! [X196,X197,X195] : (~p(s(bool,'T_0')) | p(s(bool,X196)) | p(s(bool,X197)) | s(bool,X197) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X195))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X195),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f206,f8173])).
fof(f493,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,X19)) ) | $spl32),
  inference(cnf_transformation,[],[f493_D])).
fof(f493_D,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,X19)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f14801,plain,(
  ( ! [X187,X188,X186] : (s(bool,'T_0') != s(bool,X187) | p(s(bool,X188)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X186))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X186),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14621,f195])).
fof(f14621,plain,(
  ( ! [X187,X188,X186] : (s(bool,'T_0') != s(bool,X187) | p(s(bool,X188)) | ~p(s(bool,X187)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X186))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X186),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f206,f8173])).
fof(f14796,plain,(
  ( ! [X152,X151,X153,X150] : (p(s(bool,X151)) | ~p(s(bool,X152)) | s(bool,X151) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X150))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X152) != s(bool,X153) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X150),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14593,f207])).
fof(f14593,plain,(
  ( ! [X152,X151,X153,X150] : (~p(s(bool,'T_0')) | p(s(bool,X151)) | ~p(s(bool,X152)) | s(bool,X151) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X150))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X152) != s(bool,X153) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X150),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f150,f8173])).
fof(f14795,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,'T_0') != s(bool,X135) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X134))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | ~p(s(bool,X136)) | s(bool,X136) != s(bool,X137) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X134),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14589,f195])).
fof(f14589,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,'T_0') != s(bool,X135) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X134))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | ~p(s(bool,X136)) | ~p(s(bool,X135)) | s(bool,X136) != s(bool,X137) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X134),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f150,f8173])).
fof(f14793,plain,(
  ( ! [X121,X118,X120,X119] : (p(s(bool,X119)) | ~p(s(bool,X120)) | p(s(bool,X121)) | s(bool,X121) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X118))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X118),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14585,f207])).
fof(f14585,plain,(
  ( ! [X121,X118,X120,X119] : (~p(s(bool,'T_0')) | p(s(bool,X119)) | ~p(s(bool,X120)) | p(s(bool,X121)) | s(bool,X121) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X118))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X118),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f149,f8173])).
fof(f14792,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'T_0') != s(bool,X107) | p(s(bool,X108)) | ~p(s(bool,X109)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X106))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X106),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14582,f195])).
fof(f14582,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'T_0') != s(bool,X107) | p(s(bool,X108)) | ~p(s(bool,X107)) | ~p(s(bool,X109)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X106))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X106),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f149,f8173])).
fof(f14791,plain,(
  ( ! [X101,X99,X100,X98] : (p(s(bool,X99)) | p(s(bool,X100)) | s(bool,X99) != s(bool,X101) | s(bool,X100) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X98))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X98),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14580,f207])).
fof(f14580,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,'T_0')) | p(s(bool,X99)) | p(s(bool,X100)) | s(bool,X99) != s(bool,X101) | s(bool,X100) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X98))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X98),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f148,f8173])).
fof(f14789,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') != s(bool,X79) | p(s(bool,X80)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X78))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,X80) != s(bool,X81) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X78),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14575,f195])).
fof(f14575,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') != s(bool,X79) | p(s(bool,X80)) | ~p(s(bool,X79)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X78))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,X80) != s(bool,X81) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X78),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f148,f8173])).
fof(f14788,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X71) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X70))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X72) != s(bool,X73) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X70),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14573,f207])).
fof(f14573,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X71) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X70))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X72) != s(bool,X73) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X70),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f146,f8173])).
fof(f14787,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'T_0') != s(bool,X59) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X58))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,X61) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X58),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14570,f195])).
fof(f14570,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'T_0') != s(bool,X59) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X58))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | ~p(s(bool,X60)) | ~p(s(bool,X59)) | s(bool,X60) != s(bool,X61) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X58),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f146,f8173])).
fof(f14784,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X42))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X42),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14566,f207])).
fof(f14566,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X42))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X42),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f144,f8173])).
fof(f14781,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') != s(bool,X19) | p(s(bool,X20)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X18))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X18),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f14560,f195])).
fof(f14560,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') != s(bool,X19) | p(s(bool,X20)) | ~p(s(bool,X19)) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X18))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,X18),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f144,f8173])).
fof(f14361,plain,(
  ( ! [X39,X37,X35,X33,X36,X34] : (sP0(X33,X34,X35,X36,X37) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(X33,sK5(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36))) | p(s(bool,X39)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X33,sK6(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36)))) ) | $spl200),
  inference(cnf_transformation,[],[f14361_D])).
fof(f14361_D,plain,(
  ( ! [X39,X37,X35,X33,X36,X34] : (sP0(X33,X34,X35,X36,X37) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(X33,sK5(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36))) | p(s(bool,X39)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X33,sK6(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36)))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f14313,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X148] : (sP0(X142,X143,X144,X145,X146) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X142,sK6(X146,X145,X144,X143,X142)),s(fun(X142,bool),X145))) | p(s(bool,X147)) | p(s(bool,X148)) | s(bool,X148) != s(bool,'const.bool.IN_2'(s(X142,sK5(X146,X145,X144,X143,X142)),s(fun(X142,bool),X145)))) )),
  inference(resolution,[],[f3757,f206])).
fof(f3757,plain,(
  ( ! [X111,X109,X107,X110,X108] : (p(s(bool,'const.bool.IN_2'(s(X107,sK5(X111,X110,X109,X108,X107)),s(fun(X107,bool),X110)))) | sP0(X107,X108,X109,X110,X111) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X107,sK6(X111,X110,X109,X108,X107)),s(fun(X107,bool),X110)))) )),
  inference(resolution,[],[f128,f196])).
fof(f14303,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89,X96] : (sP0(X89,X90,X91,X92,X93) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X89,sK6(X93,X92,X91,X90,X89)),s(fun(X89,bool),X92))) | p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,'const.bool.IN_2'(s(X89,sK5(X93,X92,X91,X90,X89)),s(fun(X89,bool),X92))) | s(bool,X95) != s(bool,X96)) )),
  inference(resolution,[],[f3757,f150])).
fof(f14300,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X65] : (sP0(X65,X66,X67,X68,X69) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X65,sK6(X69,X68,X67,X66,X65)),s(fun(X65,bool),X68))) | p(s(bool,X70)) | ~p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X72) != s(bool,'const.bool.IN_2'(s(X65,sK5(X69,X68,X67,X66,X65)),s(fun(X65,bool),X68)))) )),
  inference(resolution,[],[f3757,f149])).
fof(f14299,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X63] : (sP0(X57,X58,X59,X60,X61) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X57,sK6(X61,X60,X59,X58,X57)),s(fun(X57,bool),X60))) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,'const.bool.IN_2'(s(X57,sK5(X61,X60,X59,X58,X57)),s(fun(X57,bool),X60)))) )),
  inference(resolution,[],[f3757,f148])).
fof(f14298,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (sP0(X49,X50,X51,X52,X53) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X49,sK6(X53,X52,X51,X50,X49)),s(fun(X49,bool),X52))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(X49,sK5(X53,X52,X51,X50,X49)),s(fun(X49,bool),X52))) | s(bool,X55) != s(bool,X56)) )),
  inference(resolution,[],[f3757,f146])).
fof(f14296,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X40] : (sP0(X33,X34,X35,X36,X37) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X33,sK6(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36))) | p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(X33,sK5(X37,X36,X35,X34,X33)),s(fun(X33,bool),X36)))) )),
  inference(resolution,[],[f3757,f144])).
fof(f9795,plain,(
  ( ! [X138] : (s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) | $spl162),
  inference(cnf_transformation,[],[f9795_D])).
fof(f9795_D,plain,(
  ( ! [X138] : (s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f13570,plain,(
  ( ! [X61,X62,X60] : (s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61))) | s(bool,'T_0') = s(bool,X62) | s(bool,X62) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | ($spl20 | $spl39)),
  inference(subsumption_resolution,[],[f13539,f3428])).
fof(f13539,plain,(
  ( ! [X61,X62,X60] : (s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61))) | s(bool,'T_0') = s(bool,X62) | ~p(s(bool,X61)) | s(bool,X62) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | ($spl20 | $spl39)),
  inference(superposition,[],[f1602,f8216])).
fof(f8216,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | $spl20),
  inference(subsumption_resolution,[],[f8200,f207])).
fof(f8200,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | $spl20),
  inference(superposition,[],[f296,f1354])).
fof(f1354,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | $spl20),
  inference(resolution,[],[f413,f196])).
fof(f13487,plain,(
  ( ! [X61,X62,X60] : (s('type.sum.sum'(bool,X60),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X60),'const.sum.INL_1'(s(bool,X61))) | s(bool,'T_0') = s(bool,X62) | s(bool,X62) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | ($spl20 | $spl39)),
  inference(subsumption_resolution,[],[f13456,f3149])).
fof(f13456,plain,(
  ( ! [X61,X62,X60] : (s('type.sum.sum'(bool,X60),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X60),'const.sum.INL_1'(s(bool,X61))) | s(bool,'T_0') = s(bool,X62) | ~p(s(bool,X61)) | s(bool,X62) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | ($spl20 | $spl39)),
  inference(superposition,[],[f1600,f8216])).
fof(f13212,plain,(
  ( ! [X61,X62] : (s(bool,X61) != s(bool,X62) | s(bool,X61) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X62)) ) | $spl194),
  inference(cnf_transformation,[],[f13212_D])).
fof(f13212_D,plain,(
  ( ! [X61,X62] : (s(bool,X61) != s(bool,X62) | s(bool,X61) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X62)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f13211,plain,(
  ( ! [X61,X62,X60] : (s(bool,X61) != s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,X61) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,X60)) ) | ($spl20 | $spl34)),
  inference(subsumption_resolution,[],[f13166,f205])).
fof(f13166,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') != s(bool,X60) | s(bool,X61) != s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,X61) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'F_0') = s(bool,X60)) ) | ($spl20 | $spl34)),
  inference(superposition,[],[f4322,f8216])).
fof(f13210,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X53) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X53)) ) | $spl192),
  inference(cnf_transformation,[],[f13210_D])).
fof(f13210_D,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X53) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X53)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f13209,plain,(
  ( ! [X52,X53,X51] : (s(bool,X52) != s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'F_0') = s(bool,X51)) ) | ($spl20 | $spl34)),
  inference(subsumption_resolution,[],[f13163,f205])).
fof(f13163,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,X52) != s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'F_0') = s(bool,X51)) ) | ($spl20 | $spl34)),
  inference(superposition,[],[f4322,f1354])).
fof(f13208,plain,(
  ( ! [X50,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X50)) ) | $spl190),
  inference(cnf_transformation,[],[f13208_D])).
fof(f13208_D,plain,(
  ( ! [X50,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X50)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f13207,plain,(
  ( ! [X50,X48,X49] : (s(bool,X49) != s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,X49) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'F_0') = s(bool,X48)) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f13162,f205])).
fof(f13162,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') != s(bool,X48) | s(bool,X49) != s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,X49) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'F_0') = s(bool,X48)) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f4322,f1455])).
fof(f1455,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | $spl24),
  inference(resolution,[],[f418,f196])).
fof(f12892,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3) | s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X3)))) ) | $spl188),
  inference(cnf_transformation,[],[f12892_D])).
fof(f12892_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3) | s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X3)))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f12837,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,X92) != s(bool,X93) | s('type.sum.sum'(X94,bool),'const.sum.INR_1'(s(bool,X93))) != s('type.sum.sum'(X94,bool),'const.sum.INR_1'(s(bool,X95))) | s(bool,'F_0') = s(bool,X95) | p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X92) != s(bool,X97)) )),
  inference(resolution,[],[f11089,f206])).
fof(f11089,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X80)) | s(bool,X80) != s(bool,X81) | s('type.sum.sum'(X82,bool),'const.sum.INR_1'(s(bool,X81))) != s('type.sum.sum'(X82,bool),'const.sum.INR_1'(s(bool,X83))) | s(bool,'F_0') = s(bool,X83)) )),
  inference(resolution,[],[f665,f192])).
fof(f665,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X42) != s(bool,X44) | s('type.sum.sum'(X45,bool),'const.sum.INR_1'(s(bool,X42))) != s('type.sum.sum'(X45,bool),'const.sum.INR_1'(s(bool,X43)))) )),
  inference(superposition,[],[f230,f139])).
fof(f12827,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s('type.sum.sum'(X51,bool),'const.sum.INR_1'(s(bool,X50))) != s('type.sum.sum'(X51,bool),'const.sum.INR_1'(s(bool,X52))) | s(bool,'F_0') = s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X49) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f11089,f150])).
fof(f12824,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s('type.sum.sum'(X30,bool),'const.sum.INR_1'(s(bool,X29))) != s('type.sum.sum'(X30,bool),'const.sum.INR_1'(s(bool,X31))) | s(bool,'F_0') = s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X28) != s(bool,X34)) )),
  inference(resolution,[],[f11089,f149])).
fof(f12823,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s('type.sum.sum'(X23,bool),'const.sum.INR_1'(s(bool,X22))) != s('type.sum.sum'(X23,bool),'const.sum.INR_1'(s(bool,X24))) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X21) != s(bool,X26)) )),
  inference(resolution,[],[f11089,f148])).
fof(f12822,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s('type.sum.sum'(X16,bool),'const.sum.INR_1'(s(bool,X15))) != s('type.sum.sum'(X16,bool),'const.sum.INR_1'(s(bool,X17))) | s(bool,'F_0') = s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X14) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f11089,f146])).
fof(f12820,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X0) != s(bool,X5)) )),
  inference(resolution,[],[f11089,f144])).
fof(f12798,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3) | s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X3)))) ) | $spl186),
  inference(cnf_transformation,[],[f12798_D])).
fof(f12798_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3) | s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X3)))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f12743,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,X92) != s(bool,X93) | s('type.sum.sum'(bool,X94),'const.sum.INL_1'(s(bool,X93))) != s('type.sum.sum'(bool,X94),'const.sum.INL_1'(s(bool,X95))) | s(bool,'F_0') = s(bool,X95) | p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X92) != s(bool,X97)) )),
  inference(resolution,[],[f10974,f206])).
fof(f10974,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X80)) | s(bool,X80) != s(bool,X81) | s('type.sum.sum'(bool,X82),'const.sum.INL_1'(s(bool,X81))) != s('type.sum.sum'(bool,X82),'const.sum.INL_1'(s(bool,X83))) | s(bool,'F_0') = s(bool,X83)) )),
  inference(resolution,[],[f663,f192])).
fof(f663,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X34) != s(bool,X36) | s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X34))) != s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X35)))) )),
  inference(superposition,[],[f230,f141])).
fof(f12733,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s('type.sum.sum'(bool,X51),'const.sum.INL_1'(s(bool,X50))) != s('type.sum.sum'(bool,X51),'const.sum.INL_1'(s(bool,X52))) | s(bool,'F_0') = s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X49) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f10974,f150])).
fof(f12730,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s('type.sum.sum'(bool,X30),'const.sum.INL_1'(s(bool,X29))) != s('type.sum.sum'(bool,X30),'const.sum.INL_1'(s(bool,X31))) | s(bool,'F_0') = s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X28) != s(bool,X34)) )),
  inference(resolution,[],[f10974,f149])).
fof(f12729,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s('type.sum.sum'(bool,X23),'const.sum.INL_1'(s(bool,X22))) != s('type.sum.sum'(bool,X23),'const.sum.INL_1'(s(bool,X24))) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X21) != s(bool,X26)) )),
  inference(resolution,[],[f10974,f148])).
fof(f12728,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s('type.sum.sum'(bool,X16),'const.sum.INL_1'(s(bool,X15))) != s('type.sum.sum'(bool,X16),'const.sum.INL_1'(s(bool,X17))) | s(bool,'F_0') = s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X14) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f10974,f146])).
fof(f12726,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X0) != s(bool,X5)) )),
  inference(resolution,[],[f10974,f144])).
fof(f12465,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X0))) != s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X1)) ) | $spl184),
  inference(cnf_transformation,[],[f12465_D])).
fof(f12465_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X0))) != s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f12410,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,X92) != s(bool,X93) | s(bool,'F_0') = s(bool,X93) | s('type.sum.sum'(X94,bool),'const.sum.INR_1'(s(bool,X92))) != s('type.sum.sum'(X94,bool),'const.sum.INR_1'(s(bool,X95))) | p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X95) != s(bool,X97)) )),
  inference(resolution,[],[f776,f206])).
fof(f776,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X43)) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | s('type.sum.sum'(X45,bool),'const.sum.INR_1'(s(bool,X42))) != s('type.sum.sum'(X45,bool),'const.sum.INR_1'(s(bool,X43)))) )),
  inference(superposition,[],[f620,f139])).
fof(f12400,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,'F_0') = s(bool,X50) | s('type.sum.sum'(X51,bool),'const.sum.INR_1'(s(bool,X49))) != s('type.sum.sum'(X51,bool),'const.sum.INR_1'(s(bool,X52))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f776,f150])).
fof(f12397,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,'F_0') = s(bool,X29) | s('type.sum.sum'(X30,bool),'const.sum.INR_1'(s(bool,X28))) != s('type.sum.sum'(X30,bool),'const.sum.INR_1'(s(bool,X31))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X31) != s(bool,X34)) )),
  inference(resolution,[],[f776,f149])).
fof(f12396,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X22) | s('type.sum.sum'(X23,bool),'const.sum.INR_1'(s(bool,X21))) != s('type.sum.sum'(X23,bool),'const.sum.INR_1'(s(bool,X24))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f776,f148])).
fof(f12395,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,'F_0') = s(bool,X15) | s('type.sum.sum'(X16,bool),'const.sum.INR_1'(s(bool,X14))) != s('type.sum.sum'(X16,bool),'const.sum.INR_1'(s(bool,X17))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f776,f146])).
fof(f12393,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X0))) != s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X3))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X3) != s(bool,X5)) )),
  inference(resolution,[],[f776,f144])).
fof(f12381,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X0))) != s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X1)) ) | $spl182),
  inference(cnf_transformation,[],[f12381_D])).
fof(f12381_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X0))) != s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f12326,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,X92) != s(bool,X93) | s(bool,'F_0') = s(bool,X93) | s('type.sum.sum'(bool,X94),'const.sum.INL_1'(s(bool,X92))) != s('type.sum.sum'(bool,X94),'const.sum.INL_1'(s(bool,X95))) | p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X95) != s(bool,X97)) )),
  inference(resolution,[],[f774,f206])).
fof(f774,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X34))) != s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X35)))) )),
  inference(superposition,[],[f620,f141])).
fof(f12316,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,'F_0') = s(bool,X50) | s('type.sum.sum'(bool,X51),'const.sum.INL_1'(s(bool,X49))) != s('type.sum.sum'(bool,X51),'const.sum.INL_1'(s(bool,X52))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f774,f150])).
fof(f12313,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,'F_0') = s(bool,X29) | s('type.sum.sum'(bool,X30),'const.sum.INL_1'(s(bool,X28))) != s('type.sum.sum'(bool,X30),'const.sum.INL_1'(s(bool,X31))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X31) != s(bool,X34)) )),
  inference(resolution,[],[f774,f149])).
fof(f12312,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X22) | s('type.sum.sum'(bool,X23),'const.sum.INL_1'(s(bool,X21))) != s('type.sum.sum'(bool,X23),'const.sum.INL_1'(s(bool,X24))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f774,f148])).
fof(f12311,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,'F_0') = s(bool,X15) | s('type.sum.sum'(bool,X16),'const.sum.INL_1'(s(bool,X14))) != s('type.sum.sum'(bool,X16),'const.sum.INL_1'(s(bool,X17))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f774,f146])).
fof(f12309,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X0))) != s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X3))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X3) != s(bool,X5)) )),
  inference(resolution,[],[f774,f144])).
fof(f12297,plain,(
  ( ! [X251,X249,X252,X250,X248,X247] : (sP0(X247,X248,X249,X250,X251) | s('type.sum.sum'(X252,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X252,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X247,X248),X249),s(fun(X247,bool),X250),s(fun(X248,bool),X251)))))) ) | $spl180),
  inference(cnf_transformation,[],[f12297_D])).
fof(f12297_D,plain,(
  ( ! [X251,X249,X252,X250,X248,X247] : (sP0(X247,X248,X249,X250,X251) | s('type.sum.sum'(X252,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X252,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X247,X248),X249),s(fun(X247,bool),X250),s(fun(X248,bool),X251)))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f12294,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (~p(s(bool,X56)) | sP0(X57,X58,X59,X60,X61) | s('type.sum.sum'(X62,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X62,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X57,X58),X59),s(fun(X57,bool),X60),s(fun(X58,bool),X61)))))) )),
  inference(duplicate_literal_removal,[],[f12276])).
fof(f12276,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (~p(s(bool,X56)) | sP0(X57,X58,X59,X60,X61) | s('type.sum.sum'(X62,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X62,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X57,X58),X59),s(fun(X57,bool),X60),s(fun(X58,bool),X61))))) | ~p(s(bool,X56))) )),
  inference(superposition,[],[f3141,f196])).
fof(f3141,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X49] : (~p(s(bool,X49)) | sP0(X44,X45,X46,X47,X48) | s('type.sum.sum'(X50,bool),'const.sum.INR_1'(s(bool,X49))) != s('type.sum.sum'(X50,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X44,X45),X46),s(fun(X44,bool),X47),s(fun(X45,bool),X48)))))) )),
  inference(superposition,[],[f136,f139])).
fof(f136,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X1,X0,X2,X3,X4) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | sP0(X1,X0,X2,X3,X4)) & (~sP0(X1,X0,X2,X3,X4) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> sP0(X1,X0,X2,X3,X4))),
  inference(definition_folding,[],[f70,f81])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) | (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X7] : ! [X17,X9,X8] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X17),s(fun(X7,bool),X9),s(fun(X6,bool),X8)))) <=> (! [X10] : (p(s(bool,'const.bool.IN_2'(s(X7,X10),s(fun(X7,bool),X9)))) => p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(X7,X6),X17),s(X7,X10))),s(fun(X6,bool),X8))))) & ! [X10,X11] : ((p(s(bool,'const.bool.IN_2'(s(X7,X10),s(fun(X7,bool),X9)))) & p(s(bool,'const.bool.IN_2'(s(X7,X11),s(fun(X7,bool),X9))))) => (s(X6,happ(s(fun(X7,X6),X17),s(X7,X10))) = s(X6,happ(s(fun(X7,X6),X17),s(X7,X11))) => s(X7,X10) = s(X7,X11)))))),
  file('thm.pred_set.INJ_INR','thm.pred_set.INJ_DEF')).
fof(f12295,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,X49)) | sP0(X50,X51,X52,X53,X54) | s('type.sum.sum'(X55,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X55,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X50,X51),X52),s(fun(X50,bool),X53),s(fun(X51,bool),X54)))))) )),
  inference(duplicate_literal_removal,[],[f12275])).
fof(f12275,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,X49)) | sP0(X50,X51,X52,X53,X54) | s('type.sum.sum'(X55,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X55,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X50,X51),X52),s(fun(X50,bool),X53),s(fun(X51,bool),X54))))) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f3141,f198])).
fof(f12302,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,X42)) | sP0(X43,X44,X45,X46,X47) | s('type.sum.sum'(X48,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X48,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X43,X44),X45),s(fun(X43,bool),X46),s(fun(X44,bool),X47)))))) )),
  inference(subsumption_resolution,[],[f12274,f191])).
fof(f12274,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,X42)) | sP0(X43,X44,X45,X46,X47) | s('type.sum.sum'(X48,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X48,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X43,X44),X45),s(fun(X43,bool),X46),s(fun(X44,bool),X47))))) | s(bool,'F_0') = s(bool,X42)) )),
  inference(superposition,[],[f3141,f205])).
fof(f12252,plain,(
  ( ! [X300,X296,X298,X295,X301,X297,X299] : (sP0(X295,X296,X297,X298,X299) | s('type.sum.sum'(X300,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X300,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X295,X296),X297),s(fun(X295,bool),X298),s(fun(X296,bool),X299))))) | ~p(s(bool,X301))) )),
  inference(resolution,[],[f3141,f189])).
fof(f189,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8] : ((p(s(bool,'T_0')) & p(s(bool,X8))) <=> p(s(bool,X8)))),
  file('thm.pred_set.INJ_INR','thm.bool.AND_CLAUSES|split|0')).
fof(f12251,plain,(
  ( ! [X292,X294,X288,X290,X293,X289,X291] : (sP0(X288,X289,X290,X291,X292) | s('type.sum.sum'(X293,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X293,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X288,X289),X290),s(fun(X288,bool),X291),s(fun(X289,bool),X292))))) | ~p(s(bool,X294))) )),
  inference(resolution,[],[f3141,f181])).
fof(f181,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X8] : ((p(s(bool,X8)) & p(s(bool,'T_0'))) <=> p(s(bool,X8)))),
  file('thm.pred_set.INJ_INR','thm.bool.AND_CLAUSES|split|1')).
fof(f12250,plain,(
  ( ! [X282,X285,X287,X281,X283,X284,X286] : (sP0(X281,X282,X283,X284,X285) | s('type.sum.sum'(X286,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X286,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X281,X282),X283),s(fun(X281,bool),X284),s(fun(X282,bool),X285))))) | p(s(bool,X287))) )),
  inference(resolution,[],[f3141,f173])).
fof(f173,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8] : ((p(s(bool,'T_0')) => p(s(bool,X8))) <=> p(s(bool,X8)))),
  file('thm.pred_set.INJ_INR','thm.bool.IMP_CLAUSES|split|0')).
fof(f12249,plain,(
  ( ! [X280,X277,X279,X275,X276,X278,X274] : (sP0(X274,X275,X276,X277,X278) | s('type.sum.sum'(X279,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X279,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X274,X275),X276),s(fun(X274,bool),X277),s(fun(X275,bool),X278))))) | p(s(bool,X280))) )),
  inference(resolution,[],[f3141,f170])).
fof(f170,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.INJ_INR','thm.bool.IMP_CLAUSES|split|1')).
fof(f12248,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272] : (sP0(X267,X268,X269,X270,X271) | s('type.sum.sum'(X272,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X272,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X267,X268),X269),s(fun(X267,bool),X270),s(fun(X268,bool),X271))))) | ~p(s(bool,X273))) )),
  inference(resolution,[],[f3141,f168])).
fof(f168,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : ((p(s(bool,'F_0')) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.INJ_INR','thm.bool.IMP_CLAUSES|split|2')).
fof(f12246,plain,(
  ( ! [X255,X257,X259,X260,X254,X256,X258] : (sP0(X254,X255,X256,X257,X258) | s('type.sum.sum'(X259,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X259,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X254,X255),X256),s(fun(X254,bool),X257),s(fun(X255,bool),X258))))) | ~p(s(bool,X260))) )),
  inference(resolution,[],[f3141,f165])).
fof(f165,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.INJ_INR','thm.bool.IMP_CLAUSES|split|3')).
fof(f12245,plain,(
  ( ! [X253,X251,X249,X252,X250,X248,X247] : (sP0(X247,X248,X249,X250,X251) | s('type.sum.sum'(X252,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X252,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X247,X248),X249),s(fun(X247,bool),X250),s(fun(X248,bool),X251))))) | p(s(bool,X253))) )),
  inference(resolution,[],[f3141,f164])).
fof(f164,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f101])).
fof(f12296,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (sP0(X0,X1,X2,X3,X4) | s(bool,X6) != s(bool,X8) | ~p(s(bool,X8)) | s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,X6))) != s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) ) | $spl178),
  inference(cnf_transformation,[],[f12296_D])).
fof(f12296_D,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (sP0(X0,X1,X2,X3,X4) | s(bool,X6) != s(bool,X8) | ~p(s(bool,X8)) | s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,X6))) != s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f12234,plain,(
  ( ! [X158,X156,X154,X152,X151,X157,X155,X153,X150] : (sP0(X150,X151,X152,X153,X154) | s('type.sum.sum'(X155,bool),'const.sum.INR_1'(s(bool,X156))) != s('type.sum.sum'(X155,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X150,X151),X152),s(fun(X150,bool),X153),s(fun(X151,bool),X154))))) | p(s(bool,X157)) | ~p(s(bool,X158)) | s(bool,X156) != s(bool,X158)) )),
  inference(resolution,[],[f3141,f206])).
fof(f12223,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X79,X77,X75,X73] : (sP0(X70,X71,X72,X73,X74) | s('type.sum.sum'(X75,bool),'const.sum.INR_1'(s(bool,X76))) != s('type.sum.sum'(X75,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X70,X71),X72),s(fun(X70,bool),X73),s(fun(X71,bool),X74))))) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X76) != s(bool,X78) | s(bool,X77) != s(bool,X79)) )),
  inference(resolution,[],[f3141,f150])).
fof(f12221,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51] : (sP0(X50,X51,X52,X53,X54) | s('type.sum.sum'(X55,bool),'const.sum.INR_1'(s(bool,X56))) != s('type.sum.sum'(X55,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X50,X51),X52),s(fun(X50,bool),X53),s(fun(X51,bool),X54))))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X56) != s(bool,X58)) )),
  inference(resolution,[],[f3141,f149])).
fof(f12219,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (sP0(X30,X31,X32,X33,X34) | s('type.sum.sum'(X35,bool),'const.sum.INR_1'(s(bool,X36))) != s('type.sum.sum'(X35,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X30,X31),X32),s(fun(X30,bool),X33),s(fun(X31,bool),X34))))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X36) != s(bool,X38)) )),
  inference(resolution,[],[f3141,f148])).
fof(f12218,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (sP0(X20,X21,X22,X23,X24) | s('type.sum.sum'(X25,bool),'const.sum.INR_1'(s(bool,X26))) != s('type.sum.sum'(X25,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X20,X21),X22),s(fun(X20,bool),X23),s(fun(X21,bool),X24))))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f3141,f146])).
fof(f12216,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (sP0(X0,X1,X2,X3,X4) | s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,X6))) != s('type.sum.sum'(X5,bool),'const.sum.INR_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f3141,f144])).
fof(f11978,plain,(
  ( ! [X251,X249,X252,X250,X248,X247] : (sP0(X247,X248,X249,X250,X251) | s('type.sum.sum'(bool,X252),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X252),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X247,X248),X249),s(fun(X247,bool),X250),s(fun(X248,bool),X251)))))) ) | $spl176),
  inference(cnf_transformation,[],[f11978_D])).
fof(f11978_D,plain,(
  ( ! [X251,X249,X252,X250,X248,X247] : (sP0(X247,X248,X249,X250,X251) | s('type.sum.sum'(bool,X252),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X252),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X247,X248),X249),s(fun(X247,bool),X250),s(fun(X248,bool),X251)))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f11975,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (~p(s(bool,X56)) | sP0(X57,X58,X59,X60,X61) | s('type.sum.sum'(bool,X62),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X62),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X57,X58),X59),s(fun(X57,bool),X60),s(fun(X58,bool),X61)))))) )),
  inference(duplicate_literal_removal,[],[f11957])).
fof(f11957,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (~p(s(bool,X56)) | sP0(X57,X58,X59,X60,X61) | s('type.sum.sum'(bool,X62),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X62),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X57,X58),X59),s(fun(X57,bool),X60),s(fun(X58,bool),X61))))) | ~p(s(bool,X56))) )),
  inference(superposition,[],[f3139,f196])).
fof(f3139,plain,(
  ( ! [X30,X35,X33,X31,X36,X34,X32] : (~p(s(bool,X35)) | sP0(X30,X31,X32,X33,X34) | s('type.sum.sum'(bool,X36),'const.sum.INL_1'(s(bool,X35))) != s('type.sum.sum'(bool,X36),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X30,X31),X32),s(fun(X30,bool),X33),s(fun(X31,bool),X34)))))) )),
  inference(superposition,[],[f136,f141])).
fof(f11976,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,X49)) | sP0(X50,X51,X52,X53,X54) | s('type.sum.sum'(bool,X55),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X55),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X50,X51),X52),s(fun(X50,bool),X53),s(fun(X51,bool),X54)))))) )),
  inference(duplicate_literal_removal,[],[f11956])).
fof(f11956,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,X49)) | sP0(X50,X51,X52,X53,X54) | s('type.sum.sum'(bool,X55),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X55),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X50,X51),X52),s(fun(X50,bool),X53),s(fun(X51,bool),X54))))) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f3139,f198])).
fof(f11983,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,X42)) | sP0(X43,X44,X45,X46,X47) | s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X43,X44),X45),s(fun(X43,bool),X46),s(fun(X44,bool),X47)))))) )),
  inference(subsumption_resolution,[],[f11955,f191])).
fof(f11955,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,X42)) | sP0(X43,X44,X45,X46,X47) | s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X43,X44),X45),s(fun(X43,bool),X46),s(fun(X44,bool),X47))))) | s(bool,'F_0') = s(bool,X42)) )),
  inference(superposition,[],[f3139,f205])).
fof(f11933,plain,(
  ( ! [X300,X296,X298,X295,X301,X297,X299] : (sP0(X295,X296,X297,X298,X299) | s('type.sum.sum'(bool,X300),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X300),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X295,X296),X297),s(fun(X295,bool),X298),s(fun(X296,bool),X299))))) | ~p(s(bool,X301))) )),
  inference(resolution,[],[f3139,f189])).
fof(f11932,plain,(
  ( ! [X292,X294,X288,X290,X293,X289,X291] : (sP0(X288,X289,X290,X291,X292) | s('type.sum.sum'(bool,X293),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X293),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X288,X289),X290),s(fun(X288,bool),X291),s(fun(X289,bool),X292))))) | ~p(s(bool,X294))) )),
  inference(resolution,[],[f3139,f181])).
fof(f11931,plain,(
  ( ! [X282,X285,X287,X281,X283,X284,X286] : (sP0(X281,X282,X283,X284,X285) | s('type.sum.sum'(bool,X286),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X286),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X281,X282),X283),s(fun(X281,bool),X284),s(fun(X282,bool),X285))))) | p(s(bool,X287))) )),
  inference(resolution,[],[f3139,f173])).
fof(f11930,plain,(
  ( ! [X280,X277,X279,X275,X276,X278,X274] : (sP0(X274,X275,X276,X277,X278) | s('type.sum.sum'(bool,X279),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X279),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X274,X275),X276),s(fun(X274,bool),X277),s(fun(X275,bool),X278))))) | p(s(bool,X280))) )),
  inference(resolution,[],[f3139,f170])).
fof(f11929,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272] : (sP0(X267,X268,X269,X270,X271) | s('type.sum.sum'(bool,X272),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X272),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X267,X268),X269),s(fun(X267,bool),X270),s(fun(X268,bool),X271))))) | ~p(s(bool,X273))) )),
  inference(resolution,[],[f3139,f168])).
fof(f11927,plain,(
  ( ! [X255,X257,X259,X260,X254,X256,X258] : (sP0(X254,X255,X256,X257,X258) | s('type.sum.sum'(bool,X259),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X259),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X254,X255),X256),s(fun(X254,bool),X257),s(fun(X255,bool),X258))))) | ~p(s(bool,X260))) )),
  inference(resolution,[],[f3139,f165])).
fof(f11926,plain,(
  ( ! [X253,X251,X249,X252,X250,X248,X247] : (sP0(X247,X248,X249,X250,X251) | s('type.sum.sum'(bool,X252),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X252),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X247,X248),X249),s(fun(X247,bool),X250),s(fun(X248,bool),X251))))) | p(s(bool,X253))) )),
  inference(resolution,[],[f3139,f164])).
fof(f11977,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (sP0(X0,X1,X2,X3,X4) | s(bool,X6) != s(bool,X8) | ~p(s(bool,X8)) | s('type.sum.sum'(bool,X5),'const.sum.INL_1'(s(bool,X6))) != s('type.sum.sum'(bool,X5),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) ) | $spl174),
  inference(cnf_transformation,[],[f11977_D])).
fof(f11977_D,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (sP0(X0,X1,X2,X3,X4) | s(bool,X6) != s(bool,X8) | ~p(s(bool,X8)) | s('type.sum.sum'(bool,X5),'const.sum.INL_1'(s(bool,X6))) != s('type.sum.sum'(bool,X5),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f11915,plain,(
  ( ! [X158,X156,X154,X152,X151,X157,X155,X153,X150] : (sP0(X150,X151,X152,X153,X154) | s('type.sum.sum'(bool,X155),'const.sum.INL_1'(s(bool,X156))) != s('type.sum.sum'(bool,X155),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X150,X151),X152),s(fun(X150,bool),X153),s(fun(X151,bool),X154))))) | p(s(bool,X157)) | ~p(s(bool,X158)) | s(bool,X156) != s(bool,X158)) )),
  inference(resolution,[],[f3139,f206])).
fof(f11904,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X79,X77,X75,X73] : (sP0(X70,X71,X72,X73,X74) | s('type.sum.sum'(bool,X75),'const.sum.INL_1'(s(bool,X76))) != s('type.sum.sum'(bool,X75),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X70,X71),X72),s(fun(X70,bool),X73),s(fun(X71,bool),X74))))) | ~p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X76) != s(bool,X78) | s(bool,X77) != s(bool,X79)) )),
  inference(resolution,[],[f3139,f150])).
fof(f11902,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51] : (sP0(X50,X51,X52,X53,X54) | s('type.sum.sum'(bool,X55),'const.sum.INL_1'(s(bool,X56))) != s('type.sum.sum'(bool,X55),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X50,X51),X52),s(fun(X50,bool),X53),s(fun(X51,bool),X54))))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X56) != s(bool,X58)) )),
  inference(resolution,[],[f3139,f149])).
fof(f11900,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (sP0(X30,X31,X32,X33,X34) | s('type.sum.sum'(bool,X35),'const.sum.INL_1'(s(bool,X36))) != s('type.sum.sum'(bool,X35),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X30,X31),X32),s(fun(X30,bool),X33),s(fun(X31,bool),X34))))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X36) != s(bool,X38)) )),
  inference(resolution,[],[f3139,f148])).
fof(f11899,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (sP0(X20,X21,X22,X23,X24) | s('type.sum.sum'(bool,X25),'const.sum.INL_1'(s(bool,X26))) != s('type.sum.sum'(bool,X25),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X20,X21),X22),s(fun(X20,bool),X23),s(fun(X21,bool),X24))))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f3139,f146])).
fof(f11897,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (sP0(X0,X1,X2,X3,X4) | s('type.sum.sum'(bool,X5),'const.sum.INL_1'(s(bool,X6))) != s('type.sum.sum'(bool,X5),'const.sum.INL_1'(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f3139,f144])).
fof(f11741,plain,(
  $false | ($spl23 | $spl120)),
  inference(subsumption_resolution,[],[f11701,f11669])).
fof(f11669,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) | ($spl23 | $spl120)),
  inference(subsumption_resolution,[],[f11668,f194])).
fof(f11668,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X12))) ) | ($spl23 | $spl120)),
  inference(forward_demodulation,[],[f6481,f8886])).
fof(f8886,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | $spl23),
  inference(resolution,[],[f416,f192])).
fof(f11701,plain,(
  ( ! [X71] : (~p(s(bool,X71))) ) | ($spl23 | $spl120)),
  inference(resolution,[],[f11669,f162])).
fof(f11697,plain,(
  $false | ($spl23 | $spl120)),
  inference(resolution,[],[f11669,f199])).
fof(f199,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.INJ_INR',aHL_FALSITY)).
fof(f11337,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl172),
  inference(cnf_transformation,[],[f11337_D])).
fof(f11337_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f11275,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (s(bool,X96) != s(bool,X97) | s(bool,X96) != s(bool,X98) | s(bool,'T_0') = s(bool,X99) | s(bool,X98) != s(bool,X99) | p(s(bool,X100)) | ~p(s(bool,X101)) | s(bool,X97) != s(bool,X101)) ) | $spl88),
  inference(resolution,[],[f4369,f206])).
fof(f11264,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,'T_0') = s(bool,X52) | s(bool,X51) != s(bool,X52) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl88),
  inference(resolution,[],[f4369,f150])).
fof(f11262,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,X35) != s(bool,X36) | s(bool,X35) != s(bool,X37) | s(bool,'T_0') = s(bool,X38) | s(bool,X37) != s(bool,X38) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X36) != s(bool,X40)) ) | $spl88),
  inference(resolution,[],[f4369,f149])).
fof(f11260,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl88),
  inference(resolution,[],[f4369,f148])).
fof(f11259,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,'T_0') = s(bool,X17) | s(bool,X16) != s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X19) | s(bool,X18) != s(bool,X20)) ) | $spl88),
  inference(resolution,[],[f4369,f146])).
fof(f11257,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl88),
  inference(resolution,[],[f4369,f144])).
fof(f11248,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2)) ) | $spl170),
  inference(cnf_transformation,[],[f11248_D])).
fof(f11248_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X3) != s(bool,X5) | p(s(bool,X5)) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f11196,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,X92) != s(bool,X93) | s(bool,X93) != s(bool,X94) | s(bool,'F_0') = s(bool,X92) | s(bool,X94) != s(bool,X95) | p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X95) != s(bool,X97)) )),
  inference(resolution,[],[f4126,f206])).
fof(f4126,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X83)) | s(bool,X80) != s(bool,X81) | s(bool,X81) != s(bool,X82) | s(bool,'F_0') = s(bool,X80) | s(bool,X82) != s(bool,X83)) )),
  inference(resolution,[],[f744,f230])).
fof(f744,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,X64)) | s(bool,'F_0') = s(bool,X63) | s(bool,X62) != s(bool,X63) | s(bool,X62) != s(bool,X64)) )),
  inference(resolution,[],[f620,f230])).
fof(f11186,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X50) != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X51) != s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f4126,f150])).
fof(f11183,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,X29) != s(bool,X30) | s(bool,'F_0') = s(bool,X28) | s(bool,X30) != s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X31) != s(bool,X34)) )),
  inference(resolution,[],[f4126,f149])).
fof(f11182,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X21) | s(bool,X23) != s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f4126,f148])).
fof(f11181,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X14) | s(bool,X16) != s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f4126,f146])).
fof(f11179,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X3) != s(bool,X5)) )),
  inference(resolution,[],[f4126,f144])).
fof(f11154,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X3) != s(bool,X5) | ~p(s(bool,X5)) | s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X3))) | s(bool,X0) != s(bool,X1)) ) | $spl168),
  inference(cnf_transformation,[],[f11154_D])).
fof(f11154_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X3) != s(bool,X5) | ~p(s(bool,X5)) | s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X3))) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f11093,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (p(s(bool,X96)) | s(bool,X96) != s(bool,X97) | s('type.sum.sum'(X98,bool),'const.sum.INR_1'(s(bool,X97))) != s('type.sum.sum'(X98,bool),'const.sum.INR_1'(s(bool,X99))) | p(s(bool,X100)) | ~p(s(bool,X101)) | s(bool,X99) != s(bool,X101)) )),
  inference(resolution,[],[f665,f206])).
fof(f11082,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,X49)) | s(bool,X49) != s(bool,X50) | s('type.sum.sum'(X51,bool),'const.sum.INR_1'(s(bool,X50))) != s('type.sum.sum'(X51,bool),'const.sum.INR_1'(s(bool,X52))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X54) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f665,f150])).
fof(f11080,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,X35)) | s(bool,X35) != s(bool,X36) | s('type.sum.sum'(X37,bool),'const.sum.INR_1'(s(bool,X36))) != s('type.sum.sum'(X37,bool),'const.sum.INR_1'(s(bool,X38))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) )),
  inference(resolution,[],[f665,f149])).
fof(f11078,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,X21) != s(bool,X22) | s('type.sum.sum'(X23,bool),'const.sum.INR_1'(s(bool,X22))) != s('type.sum.sum'(X23,bool),'const.sum.INR_1'(s(bool,X24))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f665,f148])).
fof(f11077,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,X14)) | s(bool,X14) != s(bool,X15) | s('type.sum.sum'(X16,bool),'const.sum.INR_1'(s(bool,X15))) != s('type.sum.sum'(X16,bool),'const.sum.INR_1'(s(bool,X17))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f665,f146])).
fof(f11075,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X3))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X3) != s(bool,X5)) )),
  inference(resolution,[],[f665,f144])).
fof(f11039,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X3) != s(bool,X5) | ~p(s(bool,X5)) | s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X3))) | s(bool,X0) != s(bool,X1)) ) | $spl166),
  inference(cnf_transformation,[],[f11039_D])).
fof(f11039_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X3) != s(bool,X5) | ~p(s(bool,X5)) | s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X3))) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f10978,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (p(s(bool,X96)) | s(bool,X96) != s(bool,X97) | s('type.sum.sum'(bool,X98),'const.sum.INL_1'(s(bool,X97))) != s('type.sum.sum'(bool,X98),'const.sum.INL_1'(s(bool,X99))) | p(s(bool,X100)) | ~p(s(bool,X101)) | s(bool,X99) != s(bool,X101)) )),
  inference(resolution,[],[f663,f206])).
fof(f10967,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,X49)) | s(bool,X49) != s(bool,X50) | s('type.sum.sum'(bool,X51),'const.sum.INL_1'(s(bool,X50))) != s('type.sum.sum'(bool,X51),'const.sum.INL_1'(s(bool,X52))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X54) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f663,f150])).
fof(f10965,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,X35)) | s(bool,X35) != s(bool,X36) | s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X36))) != s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X38))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X38) != s(bool,X40)) )),
  inference(resolution,[],[f663,f149])).
fof(f10963,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,X21) != s(bool,X22) | s('type.sum.sum'(bool,X23),'const.sum.INL_1'(s(bool,X22))) != s('type.sum.sum'(bool,X23),'const.sum.INL_1'(s(bool,X24))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f663,f148])).
fof(f10962,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,X14)) | s(bool,X14) != s(bool,X15) | s('type.sum.sum'(bool,X16),'const.sum.INL_1'(s(bool,X15))) != s('type.sum.sum'(bool,X16),'const.sum.INL_1'(s(bool,X17))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f663,f146])).
fof(f10960,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X3))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X3) != s(bool,X5)) )),
  inference(resolution,[],[f663,f144])).
fof(f10593,plain,(
  ( ! [X37,X38,X40] : (p(s(bool,X37)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK2,X38),s(fun(sK2,bool),sK4))) | s(bool,X37) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X38))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl164),
  inference(cnf_transformation,[],[f10593_D])).
fof(f10593_D,plain,(
  ( ! [X37,X38,X40] : (p(s(bool,X37)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK2,X38),s(fun(sK2,bool),sK4))) | s(bool,X37) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X38))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f10552,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,X97)) | s(bool,X97) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X98))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X100) != s(bool,'const.bool.IN_2'(s(sK2,X98),s(fun(sK2,bool),sK4)))) ) | $spl14),
  inference(resolution,[],[f335,f206])).
fof(f10541,plain,(
  ( ! [X76,X74,X72,X75,X73] : (p(s(bool,X72)) | s(bool,X72) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X73))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | ~p(s(bool,X74)) | ~p(s(bool,X75)) | s(bool,X75) != s(bool,'const.bool.IN_2'(s(sK2,X73),s(fun(sK2,bool),sK4))) | s(bool,X74) != s(bool,X76)) ) | $spl14),
  inference(resolution,[],[f335,f150])).
fof(f10539,plain,(
  ( ! [X66,X64,X62,X65,X63] : (p(s(bool,X62)) | s(bool,X62) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X63))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | p(s(bool,X64)) | ~p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.bool.IN_2'(s(sK2,X63),s(fun(sK2,bool),sK4)))) ) | $spl14),
  inference(resolution,[],[f335,f149])).
fof(f10537,plain,(
  ( ! [X54,X52,X56,X55,X53] : (p(s(bool,X52)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X53))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,X56) | s(bool,X55) != s(bool,'const.bool.IN_2'(s(sK2,X53),s(fun(sK2,bool),sK4)))) ) | $spl14),
  inference(resolution,[],[f335,f148])).
fof(f10536,plain,(
  ( ! [X47,X50,X48,X51,X49] : (p(s(bool,X47)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X48))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.bool.IN_2'(s(sK2,X48),s(fun(sK2,bool),sK4))) | s(bool,X49) != s(bool,X51)) ) | $spl14),
  inference(resolution,[],[f335,f146])).
fof(f10534,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,X37)) | s(bool,X37) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X38))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK2,X38),s(fun(sK2,bool),sK4)))) ) | $spl14),
  inference(resolution,[],[f335,f144])).
fof(f9794,plain,(
  ( ! [X138,X137] : (s(bool,'F_0') = s(bool,X137) | s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) | $spl20),
  inference(subsumption_resolution,[],[f9670,f205])).
fof(f9670,plain,(
  ( ! [X138,X137] : (s(bool,'T_0') != s(bool,X137) | s(bool,'F_0') = s(bool,X137) | s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) | $spl20),
  inference(superposition,[],[f1495,f8216])).
fof(f9524,plain,(
  ( ! [X138] : (s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) | $spl160),
  inference(cnf_transformation,[],[f9524_D])).
fof(f9524_D,plain,(
  ( ! [X138] : (s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f9523,plain,(
  ( ! [X138,X137] : (s(bool,'F_0') = s(bool,X137) | s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) | $spl24),
  inference(subsumption_resolution,[],[f9399,f205])).
fof(f9399,plain,(
  ( ! [X138,X137] : (s(bool,'T_0') != s(bool,X137) | s(bool,'F_0') = s(bool,X137) | s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) | $spl24),
  inference(superposition,[],[f1495,f8215])).
fof(f8215,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | $spl24),
  inference(subsumption_resolution,[],[f8199,f207])).
fof(f8199,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | $spl24),
  inference(superposition,[],[f296,f1455])).
fof(f4860,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl98),
  inference(cnf_transformation,[],[f4860_D])).
fof(f4860_D,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f8890,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f206])).
fof(f8879,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X21) != s(bool,X23)) ) | $spl23),
  inference(resolution,[],[f416,f150])).
fof(f8877,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f149])).
fof(f8875,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f148])).
fof(f8874,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X6) != s(bool,X8)) ) | $spl23),
  inference(resolution,[],[f416,f146])).
fof(f8872,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f144])).
fof(f8835,plain,(
  ( ! [X138] : (s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) | $spl158),
  inference(cnf_transformation,[],[f8835_D])).
fof(f8835_D,plain,(
  ( ! [X138] : (s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f8834,plain,(
  ( ! [X138,X137] : (s(bool,'F_0') = s(bool,X137) | s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) | $spl22),
  inference(subsumption_resolution,[],[f8710,f205])).
fof(f8710,plain,(
  ( ! [X138,X137] : (s(bool,'T_0') != s(bool,X137) | s(bool,'F_0') = s(bool,X137) | s(bool,X138) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,'T_0') = s(bool,X138)) ) | $spl22),
  inference(superposition,[],[f1495,f6942])).
fof(f6942,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | $spl22),
  inference(resolution,[],[f415,f196])).
fof(f415,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | $spl22),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f8549,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X2))) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X1)) ) | $spl156),
  inference(cnf_transformation,[],[f8549_D])).
fof(f8549_D,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X2))) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f8502,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s('type.sum.sum'(X75,bool),'const.sum.INR_1'(s(bool,X76))) != s('type.sum.sum'(X75,bool),'const.sum.INR_1'(s(bool,X77))) | s(bool,'F_0') = s(bool,X76) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X77) != s(bool,X79)) )),
  inference(resolution,[],[f4794,f206])).
fof(f4794,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.sum.sum'(X19,bool),'const.sum.INR_1'(s(bool,X18))) != s('type.sum.sum'(X19,bool),'const.sum.INR_1'(s(bool,X20))) | s(bool,'F_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f4757,f207])).
fof(f4757,plain,(
  ( ! [X19,X20,X18] : (s('type.sum.sum'(X19,bool),'const.sum.INR_1'(s(bool,X18))) != s('type.sum.sum'(X19,bool),'const.sum.INR_1'(s(bool,X20))) | ~p(s(bool,'T_0')) | p(s(bool,X20)) | s(bool,'F_0') = s(bool,X18)) )),
  inference(superposition,[],[f3628,f205])).
fof(f8492,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.sum.sum'(X42,bool),'const.sum.INR_1'(s(bool,X43))) != s('type.sum.sum'(X42,bool),'const.sum.INR_1'(s(bool,X44))) | s(bool,'F_0') = s(bool,X43) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f4794,f150])).
fof(f8489,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.sum.sum'(X24,bool),'const.sum.INR_1'(s(bool,X25))) != s('type.sum.sum'(X24,bool),'const.sum.INR_1'(s(bool,X26))) | s(bool,'F_0') = s(bool,X25) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) )),
  inference(resolution,[],[f4794,f149])).
fof(f8488,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.sum.sum'(X18,bool),'const.sum.INR_1'(s(bool,X19))) != s('type.sum.sum'(X18,bool),'const.sum.INR_1'(s(bool,X20))) | s(bool,'F_0') = s(bool,X19) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4794,f148])).
fof(f8487,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.sum.sum'(X12,bool),'const.sum.INR_1'(s(bool,X13))) != s('type.sum.sum'(X12,bool),'const.sum.INR_1'(s(bool,X14))) | s(bool,'F_0') = s(bool,X13) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f4794,f146])).
fof(f8485,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f4794,f144])).
fof(f8475,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X2))) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X1)) ) | $spl154),
  inference(cnf_transformation,[],[f8475_D])).
fof(f8475_D,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X2))) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f8428,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s('type.sum.sum'(bool,X75),'const.sum.INL_1'(s(bool,X76))) != s('type.sum.sum'(bool,X75),'const.sum.INL_1'(s(bool,X77))) | s(bool,'F_0') = s(bool,X76) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X77) != s(bool,X79)) )),
  inference(resolution,[],[f4717,f206])).
fof(f4717,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.sum.sum'(bool,X19),'const.sum.INL_1'(s(bool,X18))) != s('type.sum.sum'(bool,X19),'const.sum.INL_1'(s(bool,X20))) | s(bool,'F_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f4680,f207])).
fof(f4680,plain,(
  ( ! [X19,X20,X18] : (s('type.sum.sum'(bool,X19),'const.sum.INL_1'(s(bool,X18))) != s('type.sum.sum'(bool,X19),'const.sum.INL_1'(s(bool,X20))) | ~p(s(bool,'T_0')) | p(s(bool,X20)) | s(bool,'F_0') = s(bool,X18)) )),
  inference(superposition,[],[f3224,f205])).
fof(f8418,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.sum.sum'(bool,X42),'const.sum.INL_1'(s(bool,X43))) != s('type.sum.sum'(bool,X42),'const.sum.INL_1'(s(bool,X44))) | s(bool,'F_0') = s(bool,X43) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f4717,f150])).
fof(f8415,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.sum.sum'(bool,X24),'const.sum.INL_1'(s(bool,X25))) != s('type.sum.sum'(bool,X24),'const.sum.INL_1'(s(bool,X26))) | s(bool,'F_0') = s(bool,X25) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) )),
  inference(resolution,[],[f4717,f149])).
fof(f8414,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.sum.sum'(bool,X18),'const.sum.INL_1'(s(bool,X19))) != s('type.sum.sum'(bool,X18),'const.sum.INL_1'(s(bool,X20))) | s(bool,'F_0') = s(bool,X19) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4717,f148])).
fof(f8413,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.sum.sum'(bool,X12),'const.sum.INL_1'(s(bool,X13))) != s('type.sum.sum'(bool,X12),'const.sum.INL_1'(s(bool,X14))) | s(bool,'F_0') = s(bool,X13) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f4717,f146])).
fof(f8411,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f4717,f144])).
fof(f8403,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,X1) != s(bool,X2)) ) | $spl152),
  inference(cnf_transformation,[],[f8403_D])).
fof(f8403_D,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,X1) != s(bool,X2)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f8356,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s('type.sum.sum'(X75,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X75,bool),'const.sum.INR_1'(s(bool,X76))) | s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X77) != s(bool,X79)) )),
  inference(resolution,[],[f3698,f206])).
fof(f8346,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.sum.sum'(X42,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X42,bool),'const.sum.INR_1'(s(bool,X43))) | s(bool,X43) != s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f3698,f150])).
fof(f8343,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.sum.sum'(X24,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X24,bool),'const.sum.INR_1'(s(bool,X25))) | s(bool,X25) != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) )),
  inference(resolution,[],[f3698,f149])).
fof(f8342,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.sum.sum'(X18,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X18,bool),'const.sum.INR_1'(s(bool,X19))) | s(bool,X19) != s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f3698,f148])).
fof(f8341,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.sum.sum'(X12,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X12,bool),'const.sum.INR_1'(s(bool,X13))) | s(bool,X13) != s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3698,f146])).
fof(f8339,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) | s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f3698,f144])).
fof(f8331,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,X1) != s(bool,X2)) ) | $spl150),
  inference(cnf_transformation,[],[f8331_D])).
fof(f8331_D,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,X1) != s(bool,X2)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f8284,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s('type.sum.sum'(bool,X75),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X75),'const.sum.INL_1'(s(bool,X76))) | s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X77) != s(bool,X79)) )),
  inference(resolution,[],[f3294,f206])).
fof(f8274,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.sum.sum'(bool,X42),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X42),'const.sum.INL_1'(s(bool,X43))) | s(bool,X43) != s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f3294,f150])).
fof(f8271,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.sum.sum'(bool,X24),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X24),'const.sum.INL_1'(s(bool,X25))) | s(bool,X25) != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) )),
  inference(resolution,[],[f3294,f149])).
fof(f8270,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.sum.sum'(bool,X18),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X18),'const.sum.INL_1'(s(bool,X19))) | s(bool,X19) != s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f3294,f148])).
fof(f8269,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.sum.sum'(bool,X12),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X12),'const.sum.INL_1'(s(bool,X13))) | s(bool,X13) != s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3294,f146])).
fof(f8267,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) | s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f3294,f144])).
fof(f8214,plain,(
  ( ! [X30,X28] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X28))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(sK2,X28),s(fun(sK2,bool),sK4)))) ) | $spl148),
  inference(cnf_transformation,[],[f8214_D])).
fof(f8214_D,plain,(
  ( ! [X30,X28] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X28))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(sK2,X28),s(fun(sK2,bool),sK4)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f8177,plain,(
  ( ! [X70,X72,X71] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X70))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X72) != s(bool,'const.bool.IN_2'(s(sK2,X70),s(fun(sK2,bool),sK4)))) )),
  inference(resolution,[],[f296,f206])).
fof(f8166,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X56))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(sK2,X56),s(fun(sK2,bool),sK4))) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f296,f150])).
fof(f8164,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X48))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | p(s(bool,X49)) | ~p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.bool.IN_2'(s(sK2,X48),s(fun(sK2,bool),sK4)))) )),
  inference(resolution,[],[f296,f149])).
fof(f8162,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X40))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(sK2,X40),s(fun(sK2,bool),sK4)))) )),
  inference(resolution,[],[f296,f148])).
fof(f8161,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X36))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(sK2,X36),s(fun(sK2,bool),sK4))) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f296,f146])).
fof(f8159,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X28))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(sK2,X28),s(fun(sK2,bool),sK4)))) )),
  inference(resolution,[],[f296,f144])).
fof(f8010,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X2))) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X1)) ) | $spl146),
  inference(cnf_transformation,[],[f8010_D])).
fof(f8010_D,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X2))) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f7954,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,X79))) != s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,X80))) | s(bool,'T_0') = s(bool,X79) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X80) != s(bool,X82)) )),
  inference(resolution,[],[f3627,f206])).
fof(f3627,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X0))) != s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f3421,f205])).
fof(f7943,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.sum.sum'(X42,bool),'const.sum.INR_1'(s(bool,X43))) != s('type.sum.sum'(X42,bool),'const.sum.INR_1'(s(bool,X44))) | s(bool,'T_0') = s(bool,X43) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f3627,f150])).
fof(f7941,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s('type.sum.sum'(X30,bool),'const.sum.INR_1'(s(bool,X31))) != s('type.sum.sum'(X30,bool),'const.sum.INR_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,X31) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f3627,f149])).
fof(f7939,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.sum.sum'(X18,bool),'const.sum.INR_1'(s(bool,X19))) != s('type.sum.sum'(X18,bool),'const.sum.INR_1'(s(bool,X20))) | s(bool,'T_0') = s(bool,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f3627,f148])).
fof(f7938,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.sum.sum'(X12,bool),'const.sum.INR_1'(s(bool,X13))) != s('type.sum.sum'(X12,bool),'const.sum.INR_1'(s(bool,X14))) | s(bool,'T_0') = s(bool,X13) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f3627,f146])).
fof(f7936,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X1) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f3627,f144])).
fof(f7869,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X2))) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X1)) ) | $spl144),
  inference(cnf_transformation,[],[f7869_D])).
fof(f7869_D,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X2))) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f7813,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,X79))) != s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,X80))) | s(bool,'T_0') = s(bool,X79) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X80) != s(bool,X82)) )),
  inference(resolution,[],[f3223,f206])).
fof(f3223,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0))) != s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2634,f205])).
fof(f7802,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.sum.sum'(bool,X42),'const.sum.INL_1'(s(bool,X43))) != s('type.sum.sum'(bool,X42),'const.sum.INL_1'(s(bool,X44))) | s(bool,'T_0') = s(bool,X43) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f3223,f150])).
fof(f7800,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s('type.sum.sum'(bool,X30),'const.sum.INL_1'(s(bool,X31))) != s('type.sum.sum'(bool,X30),'const.sum.INL_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,X31) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f3223,f149])).
fof(f7798,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.sum.sum'(bool,X18),'const.sum.INL_1'(s(bool,X19))) != s('type.sum.sum'(bool,X18),'const.sum.INL_1'(s(bool,X20))) | s(bool,'T_0') = s(bool,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f3223,f148])).
fof(f7797,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.sum.sum'(bool,X12),'const.sum.INL_1'(s(bool,X13))) != s('type.sum.sum'(bool,X12),'const.sum.INL_1'(s(bool,X14))) | s(bool,'T_0') = s(bool,X13) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f3223,f146])).
fof(f7795,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X1) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f3223,f144])).
fof(f7792,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X39)) ) | $spl142),
  inference(cnf_transformation,[],[f7792_D])).
fof(f7792_D,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X39)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f7791,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') = s(bool,X38) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X39)) ) | $spl20),
  inference(subsumption_resolution,[],[f7748,f205])).
fof(f7748,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X39)) ) | $spl20),
  inference(superposition,[],[f1495,f1354])).
fof(f7790,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X37)) ) | $spl140),
  inference(cnf_transformation,[],[f7790_D])).
fof(f7790_D,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X37)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f7789,plain,(
  ( ! [X37,X36] : (s(bool,'F_0') = s(bool,X36) | s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X37)) ) | $spl24),
  inference(subsumption_resolution,[],[f7747,f205])).
fof(f7747,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X37)) ) | $spl24),
  inference(superposition,[],[f1495,f1455])).
fof(f7726,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,'const.pred_set.INJ_3'(s(fun(X1,X2),X3),s(fun(X1,bool),X4),s(fun(X2,bool),X5))) | s(bool,X0) != s(bool,X7) | ~p(s(bool,X7)) | sP0(X1,X2,X3,X4,X5)) ) | $spl138),
  inference(cnf_transformation,[],[f7726_D])).
fof(f7726_D,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,'const.pred_set.INJ_3'(s(fun(X1,X2),X3),s(fun(X1,bool),X4),s(fun(X2,bool),X5))) | s(bool,X0) != s(bool,X7) | ~p(s(bool,X7)) | sP0(X1,X2,X3,X4,X5)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f7670,plain,(
  ( ! [X134,X132,X138,X136,X135,X133,X139,X137] : (s(bool,X132) != s(bool,'const.pred_set.INJ_3'(s(fun(X133,X134),X135),s(fun(X133,bool),X136),s(fun(X134,bool),X137))) | sP0(X133,X134,X135,X136,X137) | p(s(bool,X138)) | ~p(s(bool,X139)) | s(bool,X132) != s(bool,X139)) ) | $spl58),
  inference(resolution,[],[f3144,f206])).
fof(f7659,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (s(bool,X63) != s(bool,'const.pred_set.INJ_3'(s(fun(X64,X65),X66),s(fun(X64,bool),X67),s(fun(X65,bool),X68))) | sP0(X64,X65,X66,X67,X68) | ~p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X63) != s(bool,X70) | s(bool,X69) != s(bool,X71)) ) | $spl58),
  inference(resolution,[],[f3144,f150])).
fof(f7657,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X53,X51,X49] : (s(bool,X45) != s(bool,'const.pred_set.INJ_3'(s(fun(X46,X47),X48),s(fun(X46,bool),X49),s(fun(X47,bool),X50))) | sP0(X46,X47,X48,X49,X50) | p(s(bool,X51)) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X45) != s(bool,X52)) ) | $spl58),
  inference(resolution,[],[f3144,f149])).
fof(f7655,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (s(bool,X27) != s(bool,'const.pred_set.INJ_3'(s(fun(X28,X29),X30),s(fun(X28,bool),X31),s(fun(X29,bool),X32))) | sP0(X28,X29,X30,X31,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X27) != s(bool,X34)) ) | $spl58),
  inference(resolution,[],[f3144,f148])).
fof(f7654,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,X18) != s(bool,'const.pred_set.INJ_3'(s(fun(X19,X20),X21),s(fun(X19,bool),X22),s(fun(X20,bool),X23))) | sP0(X19,X20,X21,X22,X23) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X18) != s(bool,X25) | s(bool,X24) != s(bool,X26)) ) | $spl58),
  inference(resolution,[],[f3144,f146])).
fof(f7652,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,X0) != s(bool,'const.pred_set.INJ_3'(s(fun(X1,X2),X3),s(fun(X1,bool),X4),s(fun(X2,bool),X5))) | sP0(X1,X2,X3,X4,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X0) != s(bool,X7)) ) | $spl58),
  inference(resolution,[],[f3144,f144])).
fof(f7049,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | $spl134),
  inference(cnf_transformation,[],[f7049_D])).
fof(f7049_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f7119,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl42),
  inference(subsumption_resolution,[],[f7100,f196])).
fof(f7100,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,X4) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl42),
  inference(superposition,[],[f1174,f205])).
fof(f410,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) | $spl19),
  inference(cnf_transformation,[],[f410_D])).
fof(f410_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f7118,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))) ) | $spl42),
  inference(subsumption_resolution,[],[f7099,f196])).
fof(f7099,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))) ) | $spl42),
  inference(superposition,[],[f1174,f198])).
fof(f7117,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))) ) | $spl42),
  inference(subsumption_resolution,[],[f7098,f196])).
fof(f7098,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2)) | ~p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))) ) | $spl42),
  inference(superposition,[],[f1174,f196])).
fof(f7112,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | $spl137),
  inference(cnf_transformation,[],[f7112_D])).
fof(f7112_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f7113,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | ~p(s(bool,X4))) ) | $spl42),
  inference(subsumption_resolution,[],[f7069,f191])).
fof(f7069,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X4)) ) | $spl42),
  inference(superposition,[],[f1174,f205])).
fof(f7109,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | ~p(s(bool,X3))) ) | $spl42),
  inference(duplicate_literal_removal,[],[f7068])).
fof(f7068,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | ~p(s(bool,X3)) | ~p(s(bool,X3))) ) | $spl42),
  inference(superposition,[],[f1174,f198])).
fof(f7110,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | ~p(s(bool,X2))) ) | $spl42),
  inference(duplicate_literal_removal,[],[f7067])).
fof(f7067,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | ~p(s(bool,X2)) | ~p(s(bool,X2))) ) | $spl42),
  inference(superposition,[],[f1174,f196])).
fof(f7050,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | $spl135),
  inference(cnf_transformation,[],[f7050_D])).
fof(f7050_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f7045,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl40),
  inference(duplicate_literal_removal,[],[f7033])).
fof(f7033,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X11))) ) | $spl40),
  inference(superposition,[],[f1113,f192])).
fof(f7046,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl40),
  inference(duplicate_literal_removal,[],[f7032])).
fof(f7032,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X10))) ) | $spl40),
  inference(superposition,[],[f1113,f194])).
fof(f7055,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl40),
  inference(subsumption_resolution,[],[f7031,f195])).
fof(f7031,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,'T_0') = s(bool,X9)) ) | $spl40),
  inference(superposition,[],[f1113,f205])).
fof(f7013,plain,(
  ( ! [X53] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | ~p(s(bool,X53))) ) | $spl40),
  inference(resolution,[],[f1113,f162])).
fof(f7051,plain,(
  ( ! [X52] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X52))) ) | $spl40),
  inference(subsumption_resolution,[],[f7012,f207])).
fof(f7012,plain,(
  ( ! [X52] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X52)) | ~p(s(bool,'T_0'))) ) | $spl40),
  inference(resolution,[],[f1113,f169])).
fof(f169,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f103])).
fof(f7011,plain,(
  ( ! [X51] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X51))) ) | $spl40),
  inference(resolution,[],[f1113,f178])).
fof(f178,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8] : ((p(s(bool,'F_0')) & p(s(bool,X8))) <=> p(s(bool,'F_0')))),
  file('thm.pred_set.INJ_INR','thm.bool.AND_CLAUSES|split|2')).
fof(f7010,plain,(
  ( ! [X50] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X50))) ) | $spl40),
  inference(resolution,[],[f1113,f186])).
fof(f186,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : ((p(s(bool,X8)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.pred_set.INJ_INR','thm.bool.AND_CLAUSES|split|3')).
fof(f7047,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl132),
  inference(cnf_transformation,[],[f7047_D])).
fof(f7047_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f7005,plain,(
  ( ! [X43,X41,X42] : (s(bool,X41) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X41) != s(bool,X43)) ) | $spl40),
  inference(resolution,[],[f1113,f206])).
fof(f6995,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl40),
  inference(resolution,[],[f1113,f150])).
fof(f6992,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl40),
  inference(resolution,[],[f1113,f149])).
fof(f6991,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl40),
  inference(resolution,[],[f1113,f148])).
fof(f6990,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl40),
  inference(resolution,[],[f1113,f146])).
fof(f6988,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl40),
  inference(resolution,[],[f1113,f144])).
fof(f6481,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl120),
  inference(cnf_transformation,[],[f6481_D])).
fof(f6481_D,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f6944,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl22),
  inference(resolution,[],[f415,f206])).
fof(f6934,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X33) != s(bool,X34)) ) | $spl22),
  inference(resolution,[],[f415,f150])).
fof(f6931,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl22),
  inference(resolution,[],[f415,f149])).
fof(f6930,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl22),
  inference(resolution,[],[f415,f148])).
fof(f6929,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X18) != s(bool,X19)) ) | $spl22),
  inference(resolution,[],[f415,f146])).
fof(f6927,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl22),
  inference(resolution,[],[f415,f144])).
fof(f423,plain,(
  s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | $spl26),
  inference(cnf_transformation,[],[f423_D])).
fof(f423_D,plain,(
  s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f416,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | $spl23),
  inference(cnf_transformation,[],[f416_D])).
fof(f416_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f6833,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2))))),
  inference(forward_demodulation,[],[f6832,f138])).
fof(f6832,plain,(
  s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))),
  inference(forward_demodulation,[],[f6831,f138])).
fof(f6831,plain,(
  s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))),
  inference(forward_demodulation,[],[f6549,f138])).
fof(f6549,plain,(
  s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))),
  inference(resolution,[],[f133,f210])).
fof(f210,plain,(
  ~sP0(sK2,'type.sum.sum'(sK1,sK2),'const.sum.INR_0',sK4,sK3)),
  inference(resolution,[],[f125,f137])).
fof(f137,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | ~sP0(X1,X0,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f88])).
fof(f133,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X3,X4) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK6(X4,X3,X2,X1,X0)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK7(X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK5(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f6918,plain,(
  $false | ($spl19 | $spl40)),
  inference(subsumption_resolution,[],[f6881,f6855])).
fof(f6855,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) | ($spl19 | $spl40)),
  inference(subsumption_resolution,[],[f6854,f194])).
fof(f6854,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X12))) ) | ($spl19 | $spl40)),
  inference(forward_demodulation,[],[f1113,f5122])).
fof(f5122,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | $spl19),
  inference(resolution,[],[f410,f192])).
fof(f6881,plain,(
  ( ! [X47] : (~p(s(bool,X47))) ) | ($spl19 | $spl40)),
  inference(resolution,[],[f6855,f162])).
fof(f6877,plain,(
  $false | ($spl19 | $spl40)),
  inference(resolution,[],[f6855,f199])).
fof(f6889,plain,(
  $false | ($spl19 | $spl23 | $spl40)),
  inference(resolution,[],[f6855,f416])).
fof(f6846,plain,(
  ( ! [X269,X271,X267,X268,X266] : (~p(s(bool,happ(s(fun(X266,bool),X267),s(X266,sK6(X268,X269,X267,bool,X266))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X266,bool),X267),s(X266,sK5(X268,X269,X267,bool,X266)))),s(fun(bool,bool),X268)))) | sP0(X266,bool,X267,X269,X268) | s(bool,X271) != s(bool,happ(s(fun(X266,bool),X267),s(X266,sK7(X268,X269,X267,bool,X266)))) | p(s(bool,X271))) ) | $spl130),
  inference(cnf_transformation,[],[f6846_D])).
fof(f6846_D,plain,(
  ( ! [X269,X271,X267,X268,X266] : (~p(s(bool,happ(s(fun(X266,bool),X267),s(X266,sK6(X268,X269,X267,bool,X266))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X266,bool),X267),s(X266,sK5(X268,X269,X267,bool,X266)))),s(fun(bool,bool),X268)))) | sP0(X266,bool,X267,X269,X268) | s(bool,X271) != s(bool,happ(s(fun(X266,bool),X267),s(X266,sK7(X268,X269,X267,bool,X266)))) | p(s(bool,X271))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f6773,plain,(
  ( ! [X606,X602,X605,X607,X603,X604] : (~p(s(bool,happ(s(fun(X602,bool),X603),s(X602,sK6(X604,X605,X603,bool,X602))))) | p(s(bool,X606)) | p(s(bool,X607)) | s(bool,X607) != s(bool,happ(s(fun(X602,bool),X603),s(X602,sK7(X604,X605,X603,bool,X602)))) | sP0(X602,bool,X603,X605,X604) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X602,bool),X603),s(X602,sK5(X604,X605,X603,bool,X602)))),s(fun(bool,bool),X604))))) )),
  inference(superposition,[],[f206,f133])).
fof(f6845,plain,(
  ( ! [X255,X253,X257,X254,X252] : (p(s(bool,happ(s(fun(X252,bool),X253),s(X252,sK6(X254,X255,X253,bool,X252))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X252,bool),X253),s(X252,sK5(X254,X255,X253,bool,X252)))),s(fun(bool,bool),X254)))) | sP0(X252,bool,X253,X255,X254) | s(bool,X257) != s(bool,happ(s(fun(X252,bool),X253),s(X252,sK7(X254,X255,X253,bool,X252)))) | ~p(s(bool,X257))) ) | $spl128),
  inference(cnf_transformation,[],[f6845_D])).
fof(f6845_D,plain,(
  ( ! [X255,X253,X257,X254,X252] : (p(s(bool,happ(s(fun(X252,bool),X253),s(X252,sK6(X254,X255,X253,bool,X252))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X252,bool),X253),s(X252,sK5(X254,X255,X253,bool,X252)))),s(fun(bool,bool),X254)))) | sP0(X252,bool,X253,X255,X254) | s(bool,X257) != s(bool,happ(s(fun(X252,bool),X253),s(X252,sK7(X254,X255,X253,bool,X252)))) | ~p(s(bool,X257))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f6771,plain,(
  ( ! [X592,X594,X590,X593,X595,X591] : (p(s(bool,happ(s(fun(X590,bool),X591),s(X590,sK6(X592,X593,X591,bool,X590))))) | p(s(bool,X594)) | ~p(s(bool,X595)) | s(bool,X595) != s(bool,happ(s(fun(X590,bool),X591),s(X590,sK7(X592,X593,X591,bool,X590)))) | sP0(X590,bool,X591,X593,X592) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X590,bool),X591),s(X590,sK5(X592,X593,X591,bool,X590)))),s(fun(bool,bool),X592))))) )),
  inference(superposition,[],[f206,f133])).
fof(f6842,plain,(
  ( ! [X227,X225,X228,X226,X224] : (s(bool,X228) != s(bool,happ(s(fun(X224,bool),X225),s(X224,sK6(X226,X227,X225,bool,X224)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X224,bool),X225),s(X224,sK5(X226,X227,X225,bool,X224)))),s(fun(bool,bool),X226)))) | sP0(X224,bool,X225,X227,X226) | p(s(bool,happ(s(fun(X224,bool),X225),s(X224,sK7(X226,X227,X225,bool,X224))))) | ~p(s(bool,X228))) ) | $spl124),
  inference(cnf_transformation,[],[f6842_D])).
fof(f6842_D,plain,(
  ( ! [X227,X225,X228,X226,X224] : (s(bool,X228) != s(bool,happ(s(fun(X224,bool),X225),s(X224,sK6(X226,X227,X225,bool,X224)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X224,bool),X225),s(X224,sK5(X226,X227,X225,bool,X224)))),s(fun(bool,bool),X226)))) | sP0(X224,bool,X225,X227,X226) | p(s(bool,happ(s(fun(X224,bool),X225),s(X224,sK7(X226,X227,X225,bool,X224))))) | ~p(s(bool,X228))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f6770,plain,(
  ( ! [X588,X584,X586,X589,X585,X587] : (s(bool,X588) != s(bool,happ(s(fun(X584,bool),X585),s(X584,sK6(X586,X587,X585,bool,X584)))) | p(s(bool,X589)) | ~p(s(bool,X588)) | p(s(bool,happ(s(fun(X584,bool),X585),s(X584,sK7(X586,X587,X585,bool,X584))))) | sP0(X584,bool,X585,X587,X586) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X584,bool),X585),s(X584,sK5(X586,X587,X585,bool,X584)))),s(fun(bool,bool),X586))))) )),
  inference(superposition,[],[f206,f133])).
fof(f6843,plain,(
  ( ! [X231,X235,X233,X234,X232] : (s(bool,X235) != s(bool,happ(s(fun(X231,bool),X232),s(X231,sK6(X233,X234,X232,bool,X231)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X231,bool),X232),s(X231,sK5(X233,X234,X232,bool,X231)))),s(fun(bool,bool),X233)))) | sP0(X231,bool,X232,X234,X233) | p(s(bool,X235)) | ~p(s(bool,happ(s(fun(X231,bool),X232),s(X231,sK7(X233,X234,X232,bool,X231)))))) ) | $spl126),
  inference(cnf_transformation,[],[f6843_D])).
fof(f6843_D,plain,(
  ( ! [X231,X235,X233,X234,X232] : (s(bool,X235) != s(bool,happ(s(fun(X231,bool),X232),s(X231,sK6(X233,X234,X232,bool,X231)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X231,bool),X232),s(X231,sK5(X233,X234,X232,bool,X231)))),s(fun(bool,bool),X233)))) | sP0(X231,bool,X232,X234,X233) | p(s(bool,X235)) | ~p(s(bool,happ(s(fun(X231,bool),X232),s(X231,sK7(X233,X234,X232,bool,X231)))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f6769,plain,(
  ( ! [X580,X582,X578,X581,X583,X579] : (s(bool,X582) != s(bool,happ(s(fun(X578,bool),X579),s(X578,sK6(X580,X581,X579,bool,X578)))) | p(s(bool,X583)) | ~p(s(bool,happ(s(fun(X578,bool),X579),s(X578,sK7(X580,X581,X579,bool,X578))))) | p(s(bool,X582)) | sP0(X578,bool,X579,X581,X580) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X578,bool),X579),s(X578,sK5(X580,X581,X579,bool,X578)))),s(fun(bool,bool),X580))))) )),
  inference(superposition,[],[f206,f133])).
fof(f6739,plain,(
  ( ! [X455,X461,X457,X459,X460,X456,X458] : (~p(s(bool,happ(s(fun(X455,bool),X456),s(X455,sK6(X457,X458,X456,bool,X455))))) | p(s(bool,X459)) | ~p(s(bool,X460)) | s(bool,X459) != s(bool,happ(s(fun(X455,bool),X456),s(X455,sK7(X457,X458,X456,bool,X455)))) | s(bool,X460) != s(bool,X461) | sP0(X455,bool,X456,X458,X457) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X455,bool),X456),s(X455,sK5(X457,X458,X456,bool,X455)))),s(fun(bool,bool),X457))))) )),
  inference(superposition,[],[f150,f133])).
fof(f6737,plain,(
  ( ! [X445,X447,X441,X443,X444,X446,X442] : (p(s(bool,happ(s(fun(X441,bool),X442),s(X441,sK6(X443,X444,X442,bool,X441))))) | ~p(s(bool,X445)) | ~p(s(bool,X446)) | s(bool,X446) != s(bool,happ(s(fun(X441,bool),X442),s(X441,sK7(X443,X444,X442,bool,X441)))) | s(bool,X445) != s(bool,X447) | sP0(X441,bool,X442,X444,X443) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X441,bool),X442),s(X441,sK5(X443,X444,X442,bool,X441)))),s(fun(bool,bool),X443))))) )),
  inference(superposition,[],[f150,f133])).
fof(f6736,plain,(
  ( ! [X436,X438,X434,X440,X437,X439,X435] : (s(bool,X438) != s(bool,happ(s(fun(X434,bool),X435),s(X434,sK6(X436,X437,X435,bool,X434)))) | p(s(bool,X438)) | ~p(s(bool,X439)) | ~p(s(bool,happ(s(fun(X434,bool),X435),s(X434,sK7(X436,X437,X435,bool,X434))))) | s(bool,X439) != s(bool,X440) | sP0(X434,bool,X435,X437,X436) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X434,bool),X435),s(X434,sK5(X436,X437,X435,bool,X434)))),s(fun(bool,bool),X436))))) )),
  inference(superposition,[],[f150,f133])).
fof(f6735,plain,(
  ( ! [X428,X430,X432,X429,X431,X427,X433] : (s(bool,X431) != s(bool,happ(s(fun(X427,bool),X428),s(X427,sK6(X429,X430,X428,bool,X427)))) | p(s(bool,happ(s(fun(X427,bool),X428),s(X427,sK7(X429,X430,X428,bool,X427))))) | ~p(s(bool,X432)) | ~p(s(bool,X431)) | s(bool,X432) != s(bool,X433) | sP0(X427,bool,X428,X430,X429) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X427,bool),X428),s(X427,sK5(X429,X430,X428,bool,X427)))),s(fun(bool,bool),X429))))) )),
  inference(superposition,[],[f150,f133])).
fof(f6731,plain,(
  ( ! [X399,X405,X401,X403,X404,X400,X402] : (~p(s(bool,happ(s(fun(X399,bool),X400),s(X399,sK6(X401,X402,X400,bool,X399))))) | p(s(bool,X403)) | ~p(s(bool,X404)) | p(s(bool,X405)) | s(bool,X405) != s(bool,happ(s(fun(X399,bool),X400),s(X399,sK7(X401,X402,X400,bool,X399)))) | sP0(X399,bool,X400,X402,X401) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X399,bool),X400),s(X399,sK5(X401,X402,X400,bool,X399)))),s(fun(bool,bool),X401))))) )),
  inference(superposition,[],[f149,f133])).
fof(f6729,plain,(
  ( ! [X389,X391,X385,X387,X388,X390,X386] : (p(s(bool,happ(s(fun(X385,bool),X386),s(X385,sK6(X387,X388,X386,bool,X385))))) | p(s(bool,X389)) | ~p(s(bool,X390)) | ~p(s(bool,X391)) | s(bool,X390) != s(bool,happ(s(fun(X385,bool),X386),s(X385,sK7(X387,X388,X386,bool,X385)))) | sP0(X385,bool,X386,X388,X387) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X385,bool),X386),s(X385,sK5(X387,X388,X386,bool,X385)))),s(fun(bool,bool),X387))))) )),
  inference(superposition,[],[f149,f133])).
fof(f6728,plain,(
  ( ! [X383,X379,X380,X382,X378,X384,X381] : (s(bool,X382) != s(bool,happ(s(fun(X378,bool),X379),s(X378,sK6(X380,X381,X379,bool,X378)))) | p(s(bool,X383)) | ~p(s(bool,X382)) | ~p(s(bool,X384)) | p(s(bool,happ(s(fun(X378,bool),X379),s(X378,sK7(X380,X381,X379,bool,X378))))) | sP0(X378,bool,X379,X381,X380) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X378,bool),X379),s(X378,sK5(X380,X381,X379,bool,X378)))),s(fun(bool,bool),X380))))) )),
  inference(superposition,[],[f149,f133])).
fof(f6727,plain,(
  ( ! [X377,X372,X374,X376,X373,X375,X371] : (s(bool,X375) != s(bool,happ(s(fun(X371,bool),X372),s(X371,sK6(X373,X374,X372,bool,X371)))) | p(s(bool,X376)) | ~p(s(bool,happ(s(fun(X371,bool),X372),s(X371,sK7(X373,X374,X372,bool,X371))))) | ~p(s(bool,X377)) | p(s(bool,X375)) | sP0(X371,bool,X372,X374,X373) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X371,bool),X372),s(X371,sK5(X373,X374,X372,bool,X371)))),s(fun(bool,bool),X373))))) )),
  inference(superposition,[],[f149,f133])).
fof(f6726,plain,(
  ( ! [X364,X366,X368,X370,X365,X367,X369] : (~p(s(bool,happ(s(fun(X364,bool),X365),s(X364,sK6(X366,X367,X365,bool,X364))))) | p(s(bool,X368)) | p(s(bool,X369)) | s(bool,X368) != s(bool,X370) | s(bool,X369) != s(bool,happ(s(fun(X364,bool),X365),s(X364,sK7(X366,X367,X365,bool,X364)))) | sP0(X364,bool,X365,X367,X366) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X364,bool),X365),s(X364,sK5(X366,X367,X365,bool,X364)))),s(fun(bool,bool),X366))))) )),
  inference(superposition,[],[f148,f133])).
fof(f6724,plain,(
  ( ! [X350,X356,X352,X354,X351,X353,X355] : (p(s(bool,happ(s(fun(X350,bool),X351),s(X350,sK6(X352,X353,X351,bool,X350))))) | p(s(bool,X354)) | ~p(s(bool,X355)) | s(bool,X354) != s(bool,X356) | s(bool,X355) != s(bool,happ(s(fun(X350,bool),X351),s(X350,sK7(X352,X353,X351,bool,X350)))) | sP0(X350,bool,X351,X353,X352) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X350,bool),X351),s(X350,sK5(X352,X353,X351,bool,X350)))),s(fun(bool,bool),X352))))) )),
  inference(superposition,[],[f148,f133])).
fof(f6721,plain,(
  ( ! [X333,X335,X329,X331,X332,X334,X330] : (s(bool,X333) != s(bool,happ(s(fun(X329,bool),X330),s(X329,sK6(X331,X332,X330,bool,X329)))) | p(s(bool,X334)) | ~p(s(bool,X333)) | p(s(bool,happ(s(fun(X329,bool),X330),s(X329,sK7(X331,X332,X330,bool,X329))))) | s(bool,X334) != s(bool,X335) | sP0(X329,bool,X330,X332,X331) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X329,bool),X330),s(X329,sK5(X331,X332,X330,bool,X329)))),s(fun(bool,bool),X331))))) )),
  inference(superposition,[],[f148,f133])).
fof(f6720,plain,(
  ( ! [X325,X327,X323,X324,X326,X322,X328] : (s(bool,X326) != s(bool,happ(s(fun(X322,bool),X323),s(X322,sK6(X324,X325,X323,bool,X322)))) | p(s(bool,X327)) | ~p(s(bool,happ(s(fun(X322,bool),X323),s(X322,sK7(X324,X325,X323,bool,X322))))) | p(s(bool,X326)) | s(bool,X327) != s(bool,X328) | sP0(X322,bool,X323,X325,X324) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X322,bool),X323),s(X322,sK5(X324,X325,X323,bool,X322)))),s(fun(bool,bool),X324))))) )),
  inference(superposition,[],[f148,f133])).
fof(f6719,plain,(
  ( ! [X315,X321,X316,X318,X320,X317,X319] : (~p(s(bool,happ(s(fun(X315,bool),X316),s(X315,sK6(X317,X318,X316,bool,X315))))) | p(s(bool,X319)) | ~p(s(bool,X320)) | s(bool,X319) != s(bool,happ(s(fun(X315,bool),X316),s(X315,sK7(X317,X318,X316,bool,X315)))) | s(bool,X320) != s(bool,X321) | sP0(X315,bool,X316,X318,X317) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X315,bool),X316),s(X315,sK5(X317,X318,X316,bool,X315)))),s(fun(bool,bool),X317))))) )),
  inference(superposition,[],[f146,f133])).
fof(f6717,plain,(
  ( ! [X302,X304,X306,X301,X303,X305,X307] : (p(s(bool,happ(s(fun(X301,bool),X302),s(X301,sK6(X303,X304,X302,bool,X301))))) | ~p(s(bool,X305)) | ~p(s(bool,X306)) | s(bool,X306) != s(bool,happ(s(fun(X301,bool),X302),s(X301,sK7(X303,X304,X302,bool,X301)))) | s(bool,X305) != s(bool,X307) | sP0(X301,bool,X302,X304,X303) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X301,bool),X302),s(X301,sK5(X303,X304,X302,bool,X301)))),s(fun(bool,bool),X303))))) )),
  inference(superposition,[],[f146,f133])).
fof(f6716,plain,(
  ( ! [X294,X300,X296,X298,X295,X297,X299] : (s(bool,X298) != s(bool,happ(s(fun(X294,bool),X295),s(X294,sK6(X296,X297,X295,bool,X294)))) | p(s(bool,happ(s(fun(X294,bool),X295),s(X294,sK7(X296,X297,X295,bool,X294))))) | ~p(s(bool,X299)) | ~p(s(bool,X298)) | s(bool,X299) != s(bool,X300) | sP0(X294,bool,X295,X297,X296) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X294,bool),X295),s(X294,sK5(X296,X297,X295,bool,X294)))),s(fun(bool,bool),X296))))) )),
  inference(superposition,[],[f146,f133])).
fof(f6715,plain,(
  ( ! [X292,X288,X290,X287,X293,X289,X291] : (s(bool,X291) != s(bool,happ(s(fun(X287,bool),X288),s(X287,sK6(X289,X290,X288,bool,X287)))) | p(s(bool,X291)) | ~p(s(bool,X292)) | ~p(s(bool,happ(s(fun(X287,bool),X288),s(X287,sK7(X289,X290,X288,bool,X287))))) | s(bool,X292) != s(bool,X293) | sP0(X287,bool,X288,X290,X289) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X287,bool),X288),s(X287,sK5(X289,X290,X288,bool,X287)))),s(fun(bool,bool),X289))))) )),
  inference(superposition,[],[f146,f133])).
fof(f6712,plain,(
  ( ! [X269,X271,X267,X268,X270,X266,X272] : (~p(s(bool,happ(s(fun(X266,bool),X267),s(X266,sK6(X268,X269,X267,bool,X266))))) | p(s(bool,X270)) | p(s(bool,X271)) | s(bool,X270) != s(bool,X272) | s(bool,X271) != s(bool,happ(s(fun(X266,bool),X267),s(X266,sK7(X268,X269,X267,bool,X266)))) | sP0(X266,bool,X267,X269,X268) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X266,bool),X267),s(X266,sK5(X268,X269,X267,bool,X266)))),s(fun(bool,bool),X268))))) )),
  inference(superposition,[],[f144,f133])).
fof(f6710,plain,(
  ( ! [X255,X253,X257,X254,X252,X256,X258] : (p(s(bool,happ(s(fun(X252,bool),X253),s(X252,sK6(X254,X255,X253,bool,X252))))) | p(s(bool,X256)) | ~p(s(bool,X257)) | s(bool,X256) != s(bool,X258) | s(bool,X257) != s(bool,happ(s(fun(X252,bool),X253),s(X252,sK7(X254,X255,X253,bool,X252)))) | sP0(X252,bool,X253,X255,X254) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X252,bool),X253),s(X252,sK5(X254,X255,X253,bool,X252)))),s(fun(bool,bool),X254))))) )),
  inference(superposition,[],[f144,f133])).
fof(f6707,plain,(
  ( ! [X231,X237,X235,X233,X236,X234,X232] : (s(bool,X235) != s(bool,happ(s(fun(X231,bool),X232),s(X231,sK6(X233,X234,X232,bool,X231)))) | p(s(bool,X236)) | ~p(s(bool,happ(s(fun(X231,bool),X232),s(X231,sK7(X233,X234,X232,bool,X231))))) | p(s(bool,X235)) | s(bool,X236) != s(bool,X237) | sP0(X231,bool,X232,X234,X233) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X231,bool),X232),s(X231,sK5(X233,X234,X232,bool,X231)))),s(fun(bool,bool),X233))))) )),
  inference(superposition,[],[f144,f133])).
fof(f6706,plain,(
  ( ! [X229,X227,X225,X230,X228,X226,X224] : (s(bool,X228) != s(bool,happ(s(fun(X224,bool),X225),s(X224,sK6(X226,X227,X225,bool,X224)))) | p(s(bool,X229)) | ~p(s(bool,X228)) | p(s(bool,happ(s(fun(X224,bool),X225),s(X224,sK7(X226,X227,X225,bool,X224))))) | s(bool,X229) != s(bool,X230) | sP0(X224,bool,X225,X227,X226) | ~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X224,bool),X225),s(X224,sK5(X226,X227,X225,bool,X224)))),s(fun(bool,bool),X226))))) )),
  inference(superposition,[],[f144,f133])).
fof(f6834,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (sP0(X6,X7,X8,X9,X10) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10))) | ~p(s(bool,X12)) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK6(X10,X9,X8,X7,X6)))) = s(X7,happ(s(fun(X6,X7),X8),s(X6,sK7(X10,X9,X8,X7,X6))))) ) | $spl122),
  inference(cnf_transformation,[],[f6834_D])).
fof(f6834_D,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (sP0(X6,X7,X8,X9,X10) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10))) | ~p(s(bool,X12)) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK6(X10,X9,X8,X7,X6)))) = s(X7,happ(s(fun(X6,X7),X8),s(X6,sK7(X10,X9,X8,X7,X6))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f6572,plain,(
  ( ! [X125,X123,X121,X126,X124,X122,X120] : (sP0(X120,X121,X122,X123,X124) | s(X121,happ(s(fun(X120,X121),X122),s(X120,sK6(X124,X123,X122,X121,X120)))) = s(X121,happ(s(fun(X120,X121),X122),s(X120,sK7(X124,X123,X122,X121,X120)))) | p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,'const.bool.IN_2'(s(X121,happ(s(fun(X120,X121),X122),s(X120,sK5(X124,X123,X122,X121,X120)))),s(fun(X121,bool),X124)))) )),
  inference(resolution,[],[f133,f206])).
fof(f6561,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (sP0(X62,X63,X64,X65,X66) | s(X63,happ(s(fun(X62,X63),X64),s(X62,sK6(X66,X65,X64,X63,X62)))) = s(X63,happ(s(fun(X62,X63),X64),s(X62,sK7(X66,X65,X64,X63,X62)))) | ~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X68) != s(bool,'const.bool.IN_2'(s(X63,happ(s(fun(X62,X63),X64),s(X62,sK5(X66,X65,X64,X63,X62)))),s(fun(X63,bool),X66))) | s(bool,X67) != s(bool,X69)) )),
  inference(resolution,[],[f133,f150])).
fof(f6559,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (sP0(X46,X47,X48,X49,X50) | s(X47,happ(s(fun(X46,X47),X48),s(X46,sK6(X50,X49,X48,X47,X46)))) = s(X47,happ(s(fun(X46,X47),X48),s(X46,sK7(X50,X49,X48,X47,X46)))) | p(s(bool,X51)) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(X47,happ(s(fun(X46,X47),X48),s(X46,sK5(X50,X49,X48,X47,X46)))),s(fun(X47,bool),X50)))) )),
  inference(resolution,[],[f133,f149])).
fof(f6557,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (sP0(X30,X31,X32,X33,X34) | s(X31,happ(s(fun(X30,X31),X32),s(X30,sK6(X34,X33,X32,X31,X30)))) = s(X31,happ(s(fun(X30,X31),X32),s(X30,sK7(X34,X33,X32,X31,X30)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(X31,happ(s(fun(X30,X31),X32),s(X30,sK5(X34,X33,X32,X31,X30)))),s(fun(X31,bool),X34)))) )),
  inference(resolution,[],[f133,f148])).
fof(f6556,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (sP0(X22,X23,X24,X25,X26) | s(X23,happ(s(fun(X22,X23),X24),s(X22,sK6(X26,X25,X24,X23,X22)))) = s(X23,happ(s(fun(X22,X23),X24),s(X22,sK7(X26,X25,X24,X23,X22)))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(X23,happ(s(fun(X22,X23),X24),s(X22,sK5(X26,X25,X24,X23,X22)))),s(fun(X23,bool),X26))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f133,f146])).
fof(f6554,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (sP0(X6,X7,X8,X9,X10) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK6(X10,X9,X8,X7,X6)))) = s(X7,happ(s(fun(X6,X7),X8),s(X6,sK7(X10,X9,X8,X7,X6)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10)))) )),
  inference(resolution,[],[f133,f144])).
fof(f6505,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f206])).
fof(f6494,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X21) != s(bool,X23)) ) | $spl23),
  inference(resolution,[],[f416,f150])).
fof(f6492,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f149])).
fof(f6490,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f148])).
fof(f6489,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X6) != s(bool,X8)) ) | $spl23),
  inference(resolution,[],[f416,f146])).
fof(f6487,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f144])).
fof(f6444,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl22),
  inference(resolution,[],[f415,f206])).
fof(f6434,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X33) != s(bool,X34)) ) | $spl22),
  inference(resolution,[],[f415,f150])).
fof(f6431,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl22),
  inference(resolution,[],[f415,f149])).
fof(f6430,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl22),
  inference(resolution,[],[f415,f148])).
fof(f6429,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X18) != s(bool,X19)) ) | $spl22),
  inference(resolution,[],[f415,f146])).
fof(f6427,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl22),
  inference(resolution,[],[f415,f144])).
fof(f6384,plain,(
  ( ! [X1246,X1248,X1243,X1247,X1245] : (~p(s(bool,X1243)) | ~sP0(bool,X1245,X1246,X1247,X1248) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1247)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1243),s(fun(bool,bool),X1247)))) | s(X1245,happ(s(fun(bool,X1245),X1246),s(bool,'F_0'))) != s(X1245,happ(s(fun(bool,X1245),X1246),s(bool,X1243)))) ) | $spl118),
  inference(cnf_transformation,[],[f6384_D])).
fof(f6384_D,plain,(
  ( ! [X1246,X1248,X1243,X1247,X1245] : (~p(s(bool,X1243)) | ~sP0(bool,X1245,X1246,X1247,X1248) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1247)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1243),s(fun(bool,bool),X1247)))) | s(X1245,happ(s(fun(bool,X1245),X1246),s(bool,'F_0'))) != s(X1245,happ(s(fun(bool,X1245),X1246),s(bool,X1243)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f6199,plain,(
  ( ! [X1267,X1271,X1269,X1266,X1270,X1268] : (~p(s(bool,X1266)) | p(s(bool,X1267)) | s(X1268,happ(s(fun(bool,X1268),X1269),s(bool,'F_0'))) != s(X1268,happ(s(fun(bool,X1268),X1269),s(bool,X1266))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1270)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1266),s(fun(bool,bool),X1270)))) | ~sP0(bool,X1268,X1269,X1270,X1271)) )),
  inference(superposition,[],[f186,f127])).
fof(f127,plain,(
  ( ! [X4,X2,X0,X8,X3,X1,X9] : (s(X0,X8) = s(X0,X9) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~sP0(X0,X1,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f87])).
fof(f6198,plain,(
  ( ! [X1265,X1262,X1260,X1264,X1263,X1261] : (~p(s(bool,X1260)) | p(s(bool,X1261)) | s(X1262,happ(s(fun(bool,X1262),X1263),s(bool,'F_0'))) != s(X1262,happ(s(fun(bool,X1262),X1263),s(bool,X1260))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1264)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1260),s(fun(bool,bool),X1264)))) | ~sP0(bool,X1262,X1263,X1264,X1265)) )),
  inference(superposition,[],[f178,f127])).
fof(f6409,plain,(
  ( ! [X1258,X1256,X1254,X1259,X1257,X1255] : (~p(s(bool,X1254)) | p(s(bool,X1255)) | s(X1256,happ(s(fun(bool,X1256),X1257),s(bool,'F_0'))) != s(X1256,happ(s(fun(bool,X1256),X1257),s(bool,X1254))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1258)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1254),s(fun(bool,bool),X1258)))) | ~sP0(bool,X1256,X1257,X1258,X1259)) )),
  inference(subsumption_resolution,[],[f6197,f207])).
fof(f6197,plain,(
  ( ! [X1258,X1256,X1254,X1259,X1257,X1255] : (~p(s(bool,X1254)) | p(s(bool,X1255)) | ~p(s(bool,'T_0')) | s(X1256,happ(s(fun(bool,X1256),X1257),s(bool,'F_0'))) != s(X1256,happ(s(fun(bool,X1256),X1257),s(bool,X1254))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1258)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1254),s(fun(bool,bool),X1258)))) | ~sP0(bool,X1256,X1257,X1258,X1259)) )),
  inference(superposition,[],[f169,f127])).
fof(f6195,plain,(
  ( ! [X1246,X1244,X1248,X1243,X1247,X1245] : (~p(s(bool,X1243)) | ~p(s(bool,X1244)) | s(X1245,happ(s(fun(bool,X1245),X1246),s(bool,'F_0'))) != s(X1245,happ(s(fun(bool,X1245),X1246),s(bool,X1243))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1247)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1243),s(fun(bool,bool),X1247)))) | ~sP0(bool,X1245,X1246,X1247,X1248)) )),
  inference(superposition,[],[f162,f127])).
fof(f6382,plain,(
  ( ! [X1112,X1110,X1115,X1113,X1114] : (p(s(bool,X1110)) | ~sP0(bool,X1112,X1113,X1114,X1115) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1114)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1110),s(fun(bool,bool),X1114)))) | s(X1112,happ(s(fun(bool,X1112),X1113),s(bool,'T_0'))) != s(X1112,happ(s(fun(bool,X1112),X1113),s(bool,X1110)))) ) | $spl116),
  inference(cnf_transformation,[],[f6382_D])).
fof(f6382_D,plain,(
  ( ! [X1112,X1110,X1115,X1113,X1114] : (p(s(bool,X1110)) | ~sP0(bool,X1112,X1113,X1114,X1115) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1114)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1110),s(fun(bool,bool),X1114)))) | s(X1112,happ(s(fun(bool,X1112),X1113),s(bool,'T_0'))) != s(X1112,happ(s(fun(bool,X1112),X1113),s(bool,X1110)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f6182,plain,(
  ( ! [X1161,X1159,X1157,X1162,X1160,X1158] : (p(s(bool,X1157)) | ~p(s(bool,X1158)) | s(X1159,happ(s(fun(bool,X1159),X1160),s(bool,'T_0'))) != s(X1159,happ(s(fun(bool,X1159),X1160),s(bool,X1157))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1161)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1157),s(fun(bool,bool),X1161)))) | ~sP0(bool,X1159,X1160,X1161,X1162)) )),
  inference(superposition,[],[f189,f127])).
fof(f6181,plain,(
  ( ! [X1151,X1155,X1153,X1154,X1152,X1156] : (p(s(bool,X1151)) | ~p(s(bool,X1152)) | s(X1153,happ(s(fun(bool,X1153),X1154),s(bool,'T_0'))) != s(X1153,happ(s(fun(bool,X1153),X1154),s(bool,X1151))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1155)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1151),s(fun(bool,bool),X1155)))) | ~sP0(bool,X1153,X1154,X1155,X1156)) )),
  inference(superposition,[],[f181,f127])).
fof(f6180,plain,(
  ( ! [X1149,X1146,X1150,X1148,X1147,X1145] : (p(s(bool,X1145)) | p(s(bool,X1146)) | s(X1147,happ(s(fun(bool,X1147),X1148),s(bool,'T_0'))) != s(X1147,happ(s(fun(bool,X1147),X1148),s(bool,X1145))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1149)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1145),s(fun(bool,bool),X1149)))) | ~sP0(bool,X1147,X1148,X1149,X1150)) )),
  inference(superposition,[],[f173,f127])).
fof(f6179,plain,(
  ( ! [X1139,X1143,X1141,X1144,X1142,X1140] : (p(s(bool,X1139)) | p(s(bool,X1140)) | s(X1141,happ(s(fun(bool,X1141),X1142),s(bool,'T_0'))) != s(X1141,happ(s(fun(bool,X1141),X1142),s(bool,X1139))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1143)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1139),s(fun(bool,bool),X1143)))) | ~sP0(bool,X1141,X1142,X1143,X1144)) )),
  inference(superposition,[],[f170,f127])).
fof(f6177,plain,(
  ( ! [X1130,X1128,X1132,X1131,X1129,X1127] : (p(s(bool,X1127)) | ~p(s(bool,X1128)) | s(X1129,happ(s(fun(bool,X1129),X1130),s(bool,'T_0'))) != s(X1129,happ(s(fun(bool,X1129),X1130),s(bool,X1127))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1131)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1127),s(fun(bool,bool),X1131)))) | ~sP0(bool,X1129,X1130,X1131,X1132)) )),
  inference(superposition,[],[f168,f127])).
fof(f6175,plain,(
  ( ! [X1118,X1116,X1120,X1119,X1117,X1121] : (p(s(bool,X1116)) | ~p(s(bool,X1117)) | s(X1118,happ(s(fun(bool,X1118),X1119),s(bool,'T_0'))) != s(X1118,happ(s(fun(bool,X1118),X1119),s(bool,X1116))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1120)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1116),s(fun(bool,bool),X1120)))) | ~sP0(bool,X1118,X1119,X1120,X1121)) )),
  inference(superposition,[],[f165,f127])).
fof(f6174,plain,(
  ( ! [X1112,X1110,X1115,X1113,X1111,X1114] : (p(s(bool,X1110)) | p(s(bool,X1111)) | s(X1112,happ(s(fun(bool,X1112),X1113),s(bool,'T_0'))) != s(X1112,happ(s(fun(bool,X1112),X1113),s(bool,X1110))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1114)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1110),s(fun(bool,bool),X1114)))) | ~sP0(bool,X1112,X1113,X1114,X1115)) )),
  inference(superposition,[],[f164,f127])).
fof(f6373,plain,(
  ( ! [X414,X410,X416,X415,X409,X411,X417] : (s(bool,X410) != s(bool,X411) | ~sP0(bool,X414,X415,X416,X417) | ~p(s(bool,'const.bool.IN_2'(s(bool,X409),s(fun(bool,bool),X416)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X410),s(fun(bool,bool),X416)))) | s(X414,happ(s(fun(bool,X414),X415),s(bool,X409))) != s(X414,happ(s(fun(bool,X414),X415),s(bool,X410))) | p(s(bool,X411)) | ~p(s(bool,X409))) ) | $spl114),
  inference(cnf_transformation,[],[f6373_D])).
fof(f6373_D,plain,(
  ( ! [X414,X410,X416,X415,X409,X411,X417] : (s(bool,X410) != s(bool,X411) | ~sP0(bool,X414,X415,X416,X417) | ~p(s(bool,'const.bool.IN_2'(s(bool,X409),s(fun(bool,bool),X416)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X410),s(fun(bool,bool),X416)))) | s(X414,happ(s(fun(bool,X414),X415),s(bool,X409))) != s(X414,happ(s(fun(bool,X414),X415),s(bool,X410))) | p(s(bool,X411)) | ~p(s(bool,X409))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f6147,plain,(
  ( ! [X912,X918,X916,X914,X915,X913,X919,X917] : (~p(s(bool,X913)) | p(s(bool,X914)) | p(s(bool,X915)) | s(bool,X912) != s(bool,X915) | s(X916,happ(s(fun(bool,X916),X917),s(bool,X912))) != s(X916,happ(s(fun(bool,X916),X917),s(bool,X913))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X912),s(fun(bool,bool),X918)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X913),s(fun(bool,bool),X918)))) | ~sP0(bool,X916,X917,X918,X919)) )),
  inference(superposition,[],[f206,f127])).
fof(f6372,plain,(
  ( ! [X408,X405,X407,X401,X406,X400,X402] : (s(bool,X401) != s(bool,X402) | ~sP0(bool,X405,X406,X407,X408) | ~p(s(bool,'const.bool.IN_2'(s(bool,X400),s(fun(bool,bool),X407)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X401),s(fun(bool,bool),X407)))) | s(X405,happ(s(fun(bool,X405),X406),s(bool,X400))) != s(X405,happ(s(fun(bool,X405),X406),s(bool,X401))) | p(s(bool,X400)) | ~p(s(bool,X402))) ) | $spl112),
  inference(cnf_transformation,[],[f6372_D])).
fof(f6372_D,plain,(
  ( ! [X408,X405,X407,X401,X406,X400,X402] : (s(bool,X401) != s(bool,X402) | ~sP0(bool,X405,X406,X407,X408) | ~p(s(bool,'const.bool.IN_2'(s(bool,X400),s(fun(bool,bool),X407)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X401),s(fun(bool,bool),X407)))) | s(X405,happ(s(fun(bool,X405),X406),s(bool,X400))) != s(X405,happ(s(fun(bool,X405),X406),s(bool,X401))) | p(s(bool,X400)) | ~p(s(bool,X402))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f6145,plain,(
  ( ! [X899,X897,X903,X901,X898,X896,X902,X900] : (p(s(bool,X897)) | p(s(bool,X898)) | ~p(s(bool,X899)) | s(bool,X896) != s(bool,X899) | s(X900,happ(s(fun(bool,X900),X901),s(bool,X896))) != s(X900,happ(s(fun(bool,X900),X901),s(bool,X897))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X896),s(fun(bool,bool),X902)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X897),s(fun(bool,bool),X902)))) | ~sP0(bool,X900,X901,X902,X903)) )),
  inference(superposition,[],[f206,f127])).
fof(f6144,plain,(
  ( ! [X895,X893,X890,X888,X894,X892,X891,X889] : (s(bool,X889) != s(bool,X890) | p(s(bool,X891)) | ~p(s(bool,X890)) | p(s(bool,X888)) | s(X892,happ(s(fun(bool,X892),X893),s(bool,X888))) != s(X892,happ(s(fun(bool,X892),X893),s(bool,X889))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X888),s(fun(bool,bool),X894)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X889),s(fun(bool,bool),X894)))) | ~sP0(bool,X892,X893,X894,X895)) )),
  inference(superposition,[],[f206,f127])).
fof(f6143,plain,(
  ( ! [X883,X881,X887,X885,X882,X880,X886,X884] : (s(bool,X881) != s(bool,X882) | p(s(bool,X883)) | ~p(s(bool,X880)) | p(s(bool,X882)) | s(X884,happ(s(fun(bool,X884),X885),s(bool,X880))) != s(X884,happ(s(fun(bool,X884),X885),s(bool,X881))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X880),s(fun(bool,bool),X886)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X881),s(fun(bool,bool),X886)))) | ~sP0(bool,X884,X885,X886,X887)) )),
  inference(superposition,[],[f206,f127])).
fof(f6113,plain,(
  ( ! [X697,X699,X705,X700,X702,X704,X698,X701,X703] : (~p(s(bool,X698)) | p(s(bool,X699)) | ~p(s(bool,X700)) | s(bool,X697) != s(bool,X699) | s(bool,X700) != s(bool,X701) | s(X702,happ(s(fun(bool,X702),X703),s(bool,X697))) != s(X702,happ(s(fun(bool,X702),X703),s(bool,X698))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X697),s(fun(bool,bool),X704)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X698),s(fun(bool,bool),X704)))) | ~sP0(bool,X702,X703,X704,X705)) )),
  inference(superposition,[],[f150,f127])).
fof(f6111,plain,(
  ( ! [X685,X687,X681,X683,X679,X684,X686,X680,X682] : (p(s(bool,X680)) | ~p(s(bool,X681)) | ~p(s(bool,X682)) | s(bool,X679) != s(bool,X682) | s(bool,X681) != s(bool,X683) | s(X684,happ(s(fun(bool,X684),X685),s(bool,X679))) != s(X684,happ(s(fun(bool,X684),X685),s(bool,X680))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X679),s(fun(bool,bool),X686)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X680),s(fun(bool,bool),X686)))) | ~sP0(bool,X684,X685,X686,X687)) )),
  inference(superposition,[],[f150,f127])).
fof(f6110,plain,(
  ( ! [X670,X677,X673,X675,X671,X676,X678,X672,X674] : (s(bool,X671) != s(bool,X672) | p(s(bool,X672)) | ~p(s(bool,X673)) | ~p(s(bool,X670)) | s(bool,X673) != s(bool,X674) | s(X675,happ(s(fun(bool,X675),X676),s(bool,X670))) != s(X675,happ(s(fun(bool,X675),X676),s(bool,X671))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X670),s(fun(bool,bool),X677)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X671),s(fun(bool,bool),X677)))) | ~sP0(bool,X675,X676,X677,X678)) )),
  inference(superposition,[],[f150,f127])).
fof(f6109,plain,(
  ( ! [X666,X668,X662,X664,X667,X661,X663,X665,X669] : (s(bool,X662) != s(bool,X663) | p(s(bool,X661)) | ~p(s(bool,X664)) | ~p(s(bool,X663)) | s(bool,X664) != s(bool,X665) | s(X666,happ(s(fun(bool,X666),X667),s(bool,X661))) != s(X666,happ(s(fun(bool,X666),X667),s(bool,X662))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X661),s(fun(bool,bool),X668)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X662),s(fun(bool,bool),X668)))) | ~sP0(bool,X666,X667,X668,X669)) )),
  inference(superposition,[],[f150,f127])).
fof(f6105,plain,(
  ( ! [X633,X629,X631,X625,X627,X632,X628,X630,X626] : (~p(s(bool,X626)) | p(s(bool,X627)) | ~p(s(bool,X628)) | p(s(bool,X629)) | s(bool,X625) != s(bool,X629) | s(X630,happ(s(fun(bool,X630),X631),s(bool,X625))) != s(X630,happ(s(fun(bool,X630),X631),s(bool,X626))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X625),s(fun(bool,bool),X632)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X626),s(fun(bool,bool),X632)))) | ~sP0(bool,X630,X631,X632,X633)) )),
  inference(superposition,[],[f149,f127])).
fof(f6103,plain,(
  ( ! [X613,X615,X609,X611,X607,X612,X614,X608,X610] : (p(s(bool,X608)) | p(s(bool,X609)) | ~p(s(bool,X610)) | ~p(s(bool,X611)) | s(bool,X607) != s(bool,X610) | s(X612,happ(s(fun(bool,X612),X613),s(bool,X607))) != s(X612,happ(s(fun(bool,X612),X613),s(bool,X608))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X607),s(fun(bool,bool),X614)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X608),s(fun(bool,bool),X614)))) | ~sP0(bool,X612,X613,X614,X615)) )),
  inference(superposition,[],[f149,f127])).
fof(f6102,plain,(
  ( ! [X606,X600,X602,X598,X605,X601,X603,X599,X604] : (s(bool,X599) != s(bool,X600) | p(s(bool,X601)) | ~p(s(bool,X600)) | ~p(s(bool,X602)) | p(s(bool,X598)) | s(X603,happ(s(fun(bool,X603),X604),s(bool,X598))) != s(X603,happ(s(fun(bool,X603),X604),s(bool,X599))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X598),s(fun(bool,bool),X605)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X599),s(fun(bool,bool),X605)))) | ~sP0(bool,X603,X604,X605,X606)) )),
  inference(superposition,[],[f149,f127])).
fof(f6101,plain,(
  ( ! [X596,X592,X594,X590,X597,X593,X595,X589,X591] : (s(bool,X590) != s(bool,X591) | p(s(bool,X592)) | ~p(s(bool,X589)) | ~p(s(bool,X593)) | p(s(bool,X591)) | s(X594,happ(s(fun(bool,X594),X595),s(bool,X589))) != s(X594,happ(s(fun(bool,X594),X595),s(bool,X590))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X589),s(fun(bool,bool),X596)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X590),s(fun(bool,bool),X596)))) | ~sP0(bool,X594,X595,X596,X597)) )),
  inference(superposition,[],[f149,f127])).
fof(f6100,plain,(
  ( ! [X588,X584,X586,X580,X582,X585,X587,X581,X583] : (~p(s(bool,X581)) | p(s(bool,X582)) | p(s(bool,X583)) | s(bool,X582) != s(bool,X584) | s(bool,X580) != s(bool,X583) | s(X585,happ(s(fun(bool,X585),X586),s(bool,X580))) != s(X585,happ(s(fun(bool,X585),X586),s(bool,X581))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X580),s(fun(bool,bool),X587)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X581),s(fun(bool,bool),X587)))) | ~sP0(bool,X585,X586,X587,X588)) )),
  inference(superposition,[],[f148,f127])).
fof(f6098,plain,(
  ( ! [X569,X565,X567,X563,X568,X570,X564,X566,X562] : (p(s(bool,X563)) | p(s(bool,X564)) | ~p(s(bool,X565)) | s(bool,X564) != s(bool,X566) | s(bool,X562) != s(bool,X565) | s(X567,happ(s(fun(bool,X567),X568),s(bool,X562))) != s(X567,happ(s(fun(bool,X567),X568),s(bool,X563))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X562),s(fun(bool,bool),X569)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X563),s(fun(bool,bool),X569)))) | ~sP0(bool,X567,X568,X569,X570)) )),
  inference(superposition,[],[f148,f127])).
fof(f6095,plain,(
  ( ! [X540,X542,X536,X538,X541,X543,X537,X539,X535] : (s(bool,X536) != s(bool,X537) | p(s(bool,X538)) | ~p(s(bool,X537)) | p(s(bool,X535)) | s(bool,X538) != s(bool,X539) | s(X540,happ(s(fun(bool,X540),X541),s(bool,X535))) != s(X540,happ(s(fun(bool,X540),X541),s(bool,X536))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X535),s(fun(bool,bool),X542)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X536),s(fun(bool,bool),X542)))) | ~sP0(bool,X540,X541,X542,X543)) )),
  inference(superposition,[],[f148,f127])).
fof(f6094,plain,(
  ( ! [X532,X534,X528,X530,X526,X533,X529,X531,X527] : (s(bool,X527) != s(bool,X528) | p(s(bool,X529)) | ~p(s(bool,X526)) | p(s(bool,X528)) | s(bool,X529) != s(bool,X530) | s(X531,happ(s(fun(bool,X531),X532),s(bool,X526))) != s(X531,happ(s(fun(bool,X531),X532),s(bool,X527))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X526),s(fun(bool,bool),X533)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X527),s(fun(bool,bool),X533)))) | ~sP0(bool,X531,X532,X533,X534)) )),
  inference(superposition,[],[f148,f127])).
fof(f6093,plain,(
  ( ! [X524,X520,X522,X518,X525,X521,X523,X517,X519] : (~p(s(bool,X518)) | p(s(bool,X519)) | ~p(s(bool,X520)) | s(bool,X517) != s(bool,X519) | s(bool,X520) != s(bool,X521) | s(X522,happ(s(fun(bool,X522),X523),s(bool,X517))) != s(X522,happ(s(fun(bool,X522),X523),s(bool,X518))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X517),s(fun(bool,bool),X524)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X518),s(fun(bool,bool),X524)))) | ~sP0(bool,X522,X523,X524,X525)) )),
  inference(superposition,[],[f146,f127])).
fof(f6091,plain,(
  ( ! [X499,X505,X507,X500,X502,X504,X506,X501,X503] : (p(s(bool,X500)) | ~p(s(bool,X501)) | ~p(s(bool,X502)) | s(bool,X499) != s(bool,X502) | s(bool,X501) != s(bool,X503) | s(X504,happ(s(fun(bool,X504),X505),s(bool,X499))) != s(X504,happ(s(fun(bool,X504),X505),s(bool,X500))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X499),s(fun(bool,bool),X506)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X500),s(fun(bool,bool),X506)))) | ~sP0(bool,X504,X505,X506,X507)) )),
  inference(superposition,[],[f146,f127])).
fof(f6090,plain,(
  ( ! [X492,X494,X490,X496,X498,X493,X495,X491,X497] : (s(bool,X491) != s(bool,X492) | p(s(bool,X490)) | ~p(s(bool,X493)) | ~p(s(bool,X492)) | s(bool,X493) != s(bool,X494) | s(X495,happ(s(fun(bool,X495),X496),s(bool,X490))) != s(X495,happ(s(fun(bool,X495),X496),s(bool,X491))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X490),s(fun(bool,bool),X497)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X491),s(fun(bool,bool),X497)))) | ~sP0(bool,X495,X496,X497,X498)) )),
  inference(superposition,[],[f146,f127])).
fof(f6089,plain,(
  ( ! [X484,X486,X482,X488,X485,X487,X481,X483,X489] : (s(bool,X482) != s(bool,X483) | p(s(bool,X483)) | ~p(s(bool,X484)) | ~p(s(bool,X481)) | s(bool,X484) != s(bool,X485) | s(X486,happ(s(fun(bool,X486),X487),s(bool,X481))) != s(X486,happ(s(fun(bool,X486),X487),s(bool,X482))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X481),s(fun(bool,bool),X488)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X482),s(fun(bool,bool),X488)))) | ~sP0(bool,X486,X487,X488,X489)) )),
  inference(superposition,[],[f146,f127])).
fof(f6086,plain,(
  ( ! [X455,X461,X457,X459,X454,X460,X462,X456,X458] : (~p(s(bool,X455)) | p(s(bool,X456)) | p(s(bool,X457)) | s(bool,X456) != s(bool,X458) | s(bool,X454) != s(bool,X457) | s(X459,happ(s(fun(bool,X459),X460),s(bool,X454))) != s(X459,happ(s(fun(bool,X459),X460),s(bool,X455))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X454),s(fun(bool,bool),X461)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X455),s(fun(bool,bool),X461)))) | ~sP0(bool,X459,X460,X461,X462)) )),
  inference(superposition,[],[f144,f127])).
fof(f6084,plain,(
  ( ! [X441,X443,X436,X438,X444,X440,X442,X437,X439] : (p(s(bool,X437)) | p(s(bool,X438)) | ~p(s(bool,X439)) | s(bool,X438) != s(bool,X440) | s(bool,X436) != s(bool,X439) | s(X441,happ(s(fun(bool,X441),X442),s(bool,X436))) != s(X441,happ(s(fun(bool,X441),X442),s(bool,X437))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X436),s(fun(bool,bool),X443)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X437),s(fun(bool,bool),X443)))) | ~sP0(bool,X441,X442,X443,X444)) )),
  inference(superposition,[],[f144,f127])).
fof(f6081,plain,(
  ( ! [X412,X414,X410,X416,X413,X415,X409,X411,X417] : (s(bool,X410) != s(bool,X411) | p(s(bool,X412)) | ~p(s(bool,X409)) | p(s(bool,X411)) | s(bool,X412) != s(bool,X413) | s(X414,happ(s(fun(bool,X414),X415),s(bool,X409))) != s(X414,happ(s(fun(bool,X414),X415),s(bool,X410))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X409),s(fun(bool,bool),X416)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X410),s(fun(bool,bool),X416)))) | ~sP0(bool,X414,X415,X416,X417)) )),
  inference(superposition,[],[f144,f127])).
fof(f6080,plain,(
  ( ! [X408,X405,X407,X401,X403,X404,X406,X400,X402] : (s(bool,X401) != s(bool,X402) | p(s(bool,X403)) | ~p(s(bool,X402)) | p(s(bool,X400)) | s(bool,X403) != s(bool,X404) | s(X405,happ(s(fun(bool,X405),X406),s(bool,X400))) != s(X405,happ(s(fun(bool,X405),X406),s(bool,X401))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X400),s(fun(bool,bool),X407)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X401),s(fun(bool,bool),X407)))) | ~sP0(bool,X405,X406,X407,X408)) )),
  inference(superposition,[],[f144,f127])).
fof(f5938,plain,(
  ( ! [X1267,X1271,X1269,X1266,X1270,X1268] : (~p(s(bool,X1266)) | p(s(bool,X1267)) | s(X1268,happ(s(fun(bool,X1268),X1269),s(bool,'F_0'))) != s(X1268,happ(s(fun(bool,X1268),X1269),s(bool,X1266))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1266),s(fun(bool,bool),X1270)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1270)))) | ~sP0(bool,X1268,X1269,X1270,X1271)) )),
  inference(superposition,[],[f186,f127])).
fof(f5937,plain,(
  ( ! [X1265,X1262,X1260,X1264,X1263,X1261] : (~p(s(bool,X1260)) | p(s(bool,X1261)) | s(X1262,happ(s(fun(bool,X1262),X1263),s(bool,'F_0'))) != s(X1262,happ(s(fun(bool,X1262),X1263),s(bool,X1260))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1260),s(fun(bool,bool),X1264)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1264)))) | ~sP0(bool,X1262,X1263,X1264,X1265)) )),
  inference(superposition,[],[f178,f127])).
fof(f6385,plain,(
  ( ! [X1258,X1256,X1254,X1259,X1257,X1255] : (~p(s(bool,X1254)) | p(s(bool,X1255)) | s(X1256,happ(s(fun(bool,X1256),X1257),s(bool,'F_0'))) != s(X1256,happ(s(fun(bool,X1256),X1257),s(bool,X1254))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1254),s(fun(bool,bool),X1258)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1258)))) | ~sP0(bool,X1256,X1257,X1258,X1259)) )),
  inference(subsumption_resolution,[],[f5936,f207])).
fof(f5936,plain,(
  ( ! [X1258,X1256,X1254,X1259,X1257,X1255] : (~p(s(bool,X1254)) | p(s(bool,X1255)) | ~p(s(bool,'T_0')) | s(X1256,happ(s(fun(bool,X1256),X1257),s(bool,'F_0'))) != s(X1256,happ(s(fun(bool,X1256),X1257),s(bool,X1254))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1254),s(fun(bool,bool),X1258)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1258)))) | ~sP0(bool,X1256,X1257,X1258,X1259)) )),
  inference(superposition,[],[f169,f127])).
fof(f5934,plain,(
  ( ! [X1246,X1244,X1248,X1243,X1247,X1245] : (~p(s(bool,X1243)) | ~p(s(bool,X1244)) | s(X1245,happ(s(fun(bool,X1245),X1246),s(bool,'F_0'))) != s(X1245,happ(s(fun(bool,X1245),X1246),s(bool,X1243))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1243),s(fun(bool,bool),X1247)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1247)))) | ~sP0(bool,X1245,X1246,X1247,X1248)) )),
  inference(superposition,[],[f162,f127])).
fof(f5921,plain,(
  ( ! [X1161,X1159,X1157,X1162,X1160,X1158] : (p(s(bool,X1157)) | ~p(s(bool,X1158)) | s(X1159,happ(s(fun(bool,X1159),X1160),s(bool,'T_0'))) != s(X1159,happ(s(fun(bool,X1159),X1160),s(bool,X1157))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1157),s(fun(bool,bool),X1161)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1161)))) | ~sP0(bool,X1159,X1160,X1161,X1162)) )),
  inference(superposition,[],[f189,f127])).
fof(f5920,plain,(
  ( ! [X1151,X1155,X1153,X1154,X1152,X1156] : (p(s(bool,X1151)) | ~p(s(bool,X1152)) | s(X1153,happ(s(fun(bool,X1153),X1154),s(bool,'T_0'))) != s(X1153,happ(s(fun(bool,X1153),X1154),s(bool,X1151))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1151),s(fun(bool,bool),X1155)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1155)))) | ~sP0(bool,X1153,X1154,X1155,X1156)) )),
  inference(superposition,[],[f181,f127])).
fof(f5919,plain,(
  ( ! [X1149,X1146,X1150,X1148,X1147,X1145] : (p(s(bool,X1145)) | p(s(bool,X1146)) | s(X1147,happ(s(fun(bool,X1147),X1148),s(bool,'T_0'))) != s(X1147,happ(s(fun(bool,X1147),X1148),s(bool,X1145))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1145),s(fun(bool,bool),X1149)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1149)))) | ~sP0(bool,X1147,X1148,X1149,X1150)) )),
  inference(superposition,[],[f173,f127])).
fof(f5918,plain,(
  ( ! [X1139,X1143,X1141,X1144,X1142,X1140] : (p(s(bool,X1139)) | p(s(bool,X1140)) | s(X1141,happ(s(fun(bool,X1141),X1142),s(bool,'T_0'))) != s(X1141,happ(s(fun(bool,X1141),X1142),s(bool,X1139))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1139),s(fun(bool,bool),X1143)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1143)))) | ~sP0(bool,X1141,X1142,X1143,X1144)) )),
  inference(superposition,[],[f170,f127])).
fof(f5916,plain,(
  ( ! [X1130,X1128,X1132,X1131,X1129,X1127] : (p(s(bool,X1127)) | ~p(s(bool,X1128)) | s(X1129,happ(s(fun(bool,X1129),X1130),s(bool,'T_0'))) != s(X1129,happ(s(fun(bool,X1129),X1130),s(bool,X1127))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1127),s(fun(bool,bool),X1131)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1131)))) | ~sP0(bool,X1129,X1130,X1131,X1132)) )),
  inference(superposition,[],[f168,f127])).
fof(f5914,plain,(
  ( ! [X1118,X1116,X1120,X1119,X1117,X1121] : (p(s(bool,X1116)) | ~p(s(bool,X1117)) | s(X1118,happ(s(fun(bool,X1118),X1119),s(bool,'T_0'))) != s(X1118,happ(s(fun(bool,X1118),X1119),s(bool,X1116))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1116),s(fun(bool,bool),X1120)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1120)))) | ~sP0(bool,X1118,X1119,X1120,X1121)) )),
  inference(superposition,[],[f165,f127])).
fof(f5913,plain,(
  ( ! [X1112,X1110,X1115,X1113,X1111,X1114] : (p(s(bool,X1110)) | p(s(bool,X1111)) | s(X1112,happ(s(fun(bool,X1112),X1113),s(bool,'T_0'))) != s(X1112,happ(s(fun(bool,X1112),X1113),s(bool,X1110))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1110),s(fun(bool,bool),X1114)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1114)))) | ~sP0(bool,X1112,X1113,X1114,X1115)) )),
  inference(superposition,[],[f164,f127])).
fof(f5886,plain,(
  ( ! [X912,X918,X916,X914,X915,X913,X919,X917] : (~p(s(bool,X913)) | p(s(bool,X914)) | p(s(bool,X915)) | s(bool,X912) != s(bool,X915) | s(X916,happ(s(fun(bool,X916),X917),s(bool,X912))) != s(X916,happ(s(fun(bool,X916),X917),s(bool,X913))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X913),s(fun(bool,bool),X918)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X912),s(fun(bool,bool),X918)))) | ~sP0(bool,X916,X917,X918,X919)) )),
  inference(superposition,[],[f206,f127])).
fof(f5884,plain,(
  ( ! [X899,X897,X903,X901,X898,X896,X902,X900] : (p(s(bool,X897)) | p(s(bool,X898)) | ~p(s(bool,X899)) | s(bool,X896) != s(bool,X899) | s(X900,happ(s(fun(bool,X900),X901),s(bool,X896))) != s(X900,happ(s(fun(bool,X900),X901),s(bool,X897))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X897),s(fun(bool,bool),X902)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X896),s(fun(bool,bool),X902)))) | ~sP0(bool,X900,X901,X902,X903)) )),
  inference(superposition,[],[f206,f127])).
fof(f5883,plain,(
  ( ! [X895,X893,X890,X888,X894,X892,X891,X889] : (s(bool,X889) != s(bool,X890) | p(s(bool,X891)) | ~p(s(bool,X890)) | p(s(bool,X888)) | s(X892,happ(s(fun(bool,X892),X893),s(bool,X888))) != s(X892,happ(s(fun(bool,X892),X893),s(bool,X889))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X889),s(fun(bool,bool),X894)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X888),s(fun(bool,bool),X894)))) | ~sP0(bool,X892,X893,X894,X895)) )),
  inference(superposition,[],[f206,f127])).
fof(f5882,plain,(
  ( ! [X883,X881,X887,X885,X882,X880,X886,X884] : (s(bool,X881) != s(bool,X882) | p(s(bool,X883)) | ~p(s(bool,X880)) | p(s(bool,X882)) | s(X884,happ(s(fun(bool,X884),X885),s(bool,X880))) != s(X884,happ(s(fun(bool,X884),X885),s(bool,X881))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X881),s(fun(bool,bool),X886)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X880),s(fun(bool,bool),X886)))) | ~sP0(bool,X884,X885,X886,X887)) )),
  inference(superposition,[],[f206,f127])).
fof(f5852,plain,(
  ( ! [X697,X699,X705,X700,X702,X704,X698,X701,X703] : (~p(s(bool,X698)) | p(s(bool,X699)) | ~p(s(bool,X700)) | s(bool,X697) != s(bool,X699) | s(bool,X700) != s(bool,X701) | s(X702,happ(s(fun(bool,X702),X703),s(bool,X697))) != s(X702,happ(s(fun(bool,X702),X703),s(bool,X698))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X698),s(fun(bool,bool),X704)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X697),s(fun(bool,bool),X704)))) | ~sP0(bool,X702,X703,X704,X705)) )),
  inference(superposition,[],[f150,f127])).
fof(f5850,plain,(
  ( ! [X685,X687,X681,X683,X679,X684,X686,X680,X682] : (p(s(bool,X680)) | ~p(s(bool,X681)) | ~p(s(bool,X682)) | s(bool,X679) != s(bool,X682) | s(bool,X681) != s(bool,X683) | s(X684,happ(s(fun(bool,X684),X685),s(bool,X679))) != s(X684,happ(s(fun(bool,X684),X685),s(bool,X680))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X680),s(fun(bool,bool),X686)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X679),s(fun(bool,bool),X686)))) | ~sP0(bool,X684,X685,X686,X687)) )),
  inference(superposition,[],[f150,f127])).
fof(f5849,plain,(
  ( ! [X670,X677,X673,X675,X671,X676,X678,X672,X674] : (s(bool,X671) != s(bool,X672) | p(s(bool,X672)) | ~p(s(bool,X673)) | ~p(s(bool,X670)) | s(bool,X673) != s(bool,X674) | s(X675,happ(s(fun(bool,X675),X676),s(bool,X670))) != s(X675,happ(s(fun(bool,X675),X676),s(bool,X671))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X671),s(fun(bool,bool),X677)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X670),s(fun(bool,bool),X677)))) | ~sP0(bool,X675,X676,X677,X678)) )),
  inference(superposition,[],[f150,f127])).
fof(f5848,plain,(
  ( ! [X666,X668,X662,X664,X667,X661,X663,X665,X669] : (s(bool,X662) != s(bool,X663) | p(s(bool,X661)) | ~p(s(bool,X664)) | ~p(s(bool,X663)) | s(bool,X664) != s(bool,X665) | s(X666,happ(s(fun(bool,X666),X667),s(bool,X661))) != s(X666,happ(s(fun(bool,X666),X667),s(bool,X662))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X662),s(fun(bool,bool),X668)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X661),s(fun(bool,bool),X668)))) | ~sP0(bool,X666,X667,X668,X669)) )),
  inference(superposition,[],[f150,f127])).
fof(f5844,plain,(
  ( ! [X633,X629,X631,X625,X627,X632,X628,X630,X626] : (~p(s(bool,X626)) | p(s(bool,X627)) | ~p(s(bool,X628)) | p(s(bool,X629)) | s(bool,X625) != s(bool,X629) | s(X630,happ(s(fun(bool,X630),X631),s(bool,X625))) != s(X630,happ(s(fun(bool,X630),X631),s(bool,X626))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X626),s(fun(bool,bool),X632)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X625),s(fun(bool,bool),X632)))) | ~sP0(bool,X630,X631,X632,X633)) )),
  inference(superposition,[],[f149,f127])).
fof(f5842,plain,(
  ( ! [X613,X615,X609,X611,X607,X612,X614,X608,X610] : (p(s(bool,X608)) | p(s(bool,X609)) | ~p(s(bool,X610)) | ~p(s(bool,X611)) | s(bool,X607) != s(bool,X610) | s(X612,happ(s(fun(bool,X612),X613),s(bool,X607))) != s(X612,happ(s(fun(bool,X612),X613),s(bool,X608))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X608),s(fun(bool,bool),X614)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X607),s(fun(bool,bool),X614)))) | ~sP0(bool,X612,X613,X614,X615)) )),
  inference(superposition,[],[f149,f127])).
fof(f5841,plain,(
  ( ! [X606,X600,X602,X598,X605,X601,X603,X599,X604] : (s(bool,X599) != s(bool,X600) | p(s(bool,X601)) | ~p(s(bool,X600)) | ~p(s(bool,X602)) | p(s(bool,X598)) | s(X603,happ(s(fun(bool,X603),X604),s(bool,X598))) != s(X603,happ(s(fun(bool,X603),X604),s(bool,X599))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X599),s(fun(bool,bool),X605)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X598),s(fun(bool,bool),X605)))) | ~sP0(bool,X603,X604,X605,X606)) )),
  inference(superposition,[],[f149,f127])).
fof(f5840,plain,(
  ( ! [X596,X592,X594,X590,X597,X593,X595,X589,X591] : (s(bool,X590) != s(bool,X591) | p(s(bool,X592)) | ~p(s(bool,X589)) | ~p(s(bool,X593)) | p(s(bool,X591)) | s(X594,happ(s(fun(bool,X594),X595),s(bool,X589))) != s(X594,happ(s(fun(bool,X594),X595),s(bool,X590))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X590),s(fun(bool,bool),X596)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X589),s(fun(bool,bool),X596)))) | ~sP0(bool,X594,X595,X596,X597)) )),
  inference(superposition,[],[f149,f127])).
fof(f5839,plain,(
  ( ! [X588,X584,X586,X580,X582,X585,X587,X581,X583] : (~p(s(bool,X581)) | p(s(bool,X582)) | p(s(bool,X583)) | s(bool,X582) != s(bool,X584) | s(bool,X580) != s(bool,X583) | s(X585,happ(s(fun(bool,X585),X586),s(bool,X580))) != s(X585,happ(s(fun(bool,X585),X586),s(bool,X581))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X581),s(fun(bool,bool),X587)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X580),s(fun(bool,bool),X587)))) | ~sP0(bool,X585,X586,X587,X588)) )),
  inference(superposition,[],[f148,f127])).
fof(f5837,plain,(
  ( ! [X569,X565,X567,X563,X568,X570,X564,X566,X562] : (p(s(bool,X563)) | p(s(bool,X564)) | ~p(s(bool,X565)) | s(bool,X564) != s(bool,X566) | s(bool,X562) != s(bool,X565) | s(X567,happ(s(fun(bool,X567),X568),s(bool,X562))) != s(X567,happ(s(fun(bool,X567),X568),s(bool,X563))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X563),s(fun(bool,bool),X569)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X562),s(fun(bool,bool),X569)))) | ~sP0(bool,X567,X568,X569,X570)) )),
  inference(superposition,[],[f148,f127])).
fof(f5834,plain,(
  ( ! [X540,X542,X536,X538,X541,X543,X537,X539,X535] : (s(bool,X536) != s(bool,X537) | p(s(bool,X538)) | ~p(s(bool,X537)) | p(s(bool,X535)) | s(bool,X538) != s(bool,X539) | s(X540,happ(s(fun(bool,X540),X541),s(bool,X535))) != s(X540,happ(s(fun(bool,X540),X541),s(bool,X536))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X536),s(fun(bool,bool),X542)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X535),s(fun(bool,bool),X542)))) | ~sP0(bool,X540,X541,X542,X543)) )),
  inference(superposition,[],[f148,f127])).
fof(f5833,plain,(
  ( ! [X532,X534,X528,X530,X526,X533,X529,X531,X527] : (s(bool,X527) != s(bool,X528) | p(s(bool,X529)) | ~p(s(bool,X526)) | p(s(bool,X528)) | s(bool,X529) != s(bool,X530) | s(X531,happ(s(fun(bool,X531),X532),s(bool,X526))) != s(X531,happ(s(fun(bool,X531),X532),s(bool,X527))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X527),s(fun(bool,bool),X533)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X526),s(fun(bool,bool),X533)))) | ~sP0(bool,X531,X532,X533,X534)) )),
  inference(superposition,[],[f148,f127])).
fof(f5832,plain,(
  ( ! [X524,X520,X522,X518,X525,X521,X523,X517,X519] : (~p(s(bool,X518)) | p(s(bool,X519)) | ~p(s(bool,X520)) | s(bool,X517) != s(bool,X519) | s(bool,X520) != s(bool,X521) | s(X522,happ(s(fun(bool,X522),X523),s(bool,X517))) != s(X522,happ(s(fun(bool,X522),X523),s(bool,X518))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X518),s(fun(bool,bool),X524)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X517),s(fun(bool,bool),X524)))) | ~sP0(bool,X522,X523,X524,X525)) )),
  inference(superposition,[],[f146,f127])).
fof(f5830,plain,(
  ( ! [X499,X505,X507,X500,X502,X504,X506,X501,X503] : (p(s(bool,X500)) | ~p(s(bool,X501)) | ~p(s(bool,X502)) | s(bool,X499) != s(bool,X502) | s(bool,X501) != s(bool,X503) | s(X504,happ(s(fun(bool,X504),X505),s(bool,X499))) != s(X504,happ(s(fun(bool,X504),X505),s(bool,X500))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X500),s(fun(bool,bool),X506)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X499),s(fun(bool,bool),X506)))) | ~sP0(bool,X504,X505,X506,X507)) )),
  inference(superposition,[],[f146,f127])).
fof(f5829,plain,(
  ( ! [X492,X494,X490,X496,X498,X493,X495,X491,X497] : (s(bool,X491) != s(bool,X492) | p(s(bool,X490)) | ~p(s(bool,X493)) | ~p(s(bool,X492)) | s(bool,X493) != s(bool,X494) | s(X495,happ(s(fun(bool,X495),X496),s(bool,X490))) != s(X495,happ(s(fun(bool,X495),X496),s(bool,X491))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X491),s(fun(bool,bool),X497)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X490),s(fun(bool,bool),X497)))) | ~sP0(bool,X495,X496,X497,X498)) )),
  inference(superposition,[],[f146,f127])).
fof(f5828,plain,(
  ( ! [X484,X486,X482,X488,X485,X487,X481,X483,X489] : (s(bool,X482) != s(bool,X483) | p(s(bool,X483)) | ~p(s(bool,X484)) | ~p(s(bool,X481)) | s(bool,X484) != s(bool,X485) | s(X486,happ(s(fun(bool,X486),X487),s(bool,X481))) != s(X486,happ(s(fun(bool,X486),X487),s(bool,X482))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X482),s(fun(bool,bool),X488)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X481),s(fun(bool,bool),X488)))) | ~sP0(bool,X486,X487,X488,X489)) )),
  inference(superposition,[],[f146,f127])).
fof(f5825,plain,(
  ( ! [X455,X461,X457,X459,X454,X460,X462,X456,X458] : (~p(s(bool,X455)) | p(s(bool,X456)) | p(s(bool,X457)) | s(bool,X456) != s(bool,X458) | s(bool,X454) != s(bool,X457) | s(X459,happ(s(fun(bool,X459),X460),s(bool,X454))) != s(X459,happ(s(fun(bool,X459),X460),s(bool,X455))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X455),s(fun(bool,bool),X461)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X454),s(fun(bool,bool),X461)))) | ~sP0(bool,X459,X460,X461,X462)) )),
  inference(superposition,[],[f144,f127])).
fof(f5823,plain,(
  ( ! [X441,X443,X436,X438,X444,X440,X442,X437,X439] : (p(s(bool,X437)) | p(s(bool,X438)) | ~p(s(bool,X439)) | s(bool,X438) != s(bool,X440) | s(bool,X436) != s(bool,X439) | s(X441,happ(s(fun(bool,X441),X442),s(bool,X436))) != s(X441,happ(s(fun(bool,X441),X442),s(bool,X437))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X437),s(fun(bool,bool),X443)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X436),s(fun(bool,bool),X443)))) | ~sP0(bool,X441,X442,X443,X444)) )),
  inference(superposition,[],[f144,f127])).
fof(f5820,plain,(
  ( ! [X412,X414,X410,X416,X413,X415,X409,X411,X417] : (s(bool,X410) != s(bool,X411) | p(s(bool,X412)) | ~p(s(bool,X409)) | p(s(bool,X411)) | s(bool,X412) != s(bool,X413) | s(X414,happ(s(fun(bool,X414),X415),s(bool,X409))) != s(X414,happ(s(fun(bool,X414),X415),s(bool,X410))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X410),s(fun(bool,bool),X416)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X409),s(fun(bool,bool),X416)))) | ~sP0(bool,X414,X415,X416,X417)) )),
  inference(superposition,[],[f144,f127])).
fof(f5819,plain,(
  ( ! [X408,X405,X407,X401,X403,X404,X406,X400,X402] : (s(bool,X401) != s(bool,X402) | p(s(bool,X403)) | ~p(s(bool,X402)) | p(s(bool,X400)) | s(bool,X403) != s(bool,X404) | s(X405,happ(s(fun(bool,X405),X406),s(bool,X400))) != s(X405,happ(s(fun(bool,X405),X406),s(bool,X401))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X401),s(fun(bool,bool),X407)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X400),s(fun(bool,bool),X407)))) | ~sP0(bool,X405,X406,X407,X408)) )),
  inference(superposition,[],[f144,f127])).
fof(f6303,plain,(
  ( ! [X92,X90,X88,X87,X95,X93,X91,X89] : (s(X87,X88) = s(X87,X89) | s(bool,X95) != s(bool,'const.bool.IN_2'(s(X87,X89),s(fun(X87,bool),X92))) | ~p(s(bool,X95)) | ~sP0(X87,X90,X91,X92,X93) | ~p(s(bool,'const.bool.IN_2'(s(X87,X88),s(fun(X87,bool),X92)))) | s(X90,happ(s(fun(X87,X90),X91),s(X87,X88))) != s(X90,happ(s(fun(X87,X90),X91),s(X87,X89)))) ) | $spl110),
  inference(cnf_transformation,[],[f6303_D])).
fof(f6303_D,plain,(
  ( ! [X92,X90,X88,X87,X95,X93,X91,X89] : (s(X87,X88) = s(X87,X89) | s(bool,X95) != s(bool,'const.bool.IN_2'(s(X87,X89),s(fun(X87,bool),X92))) | ~p(s(bool,X95)) | ~sP0(X87,X90,X91,X92,X93) | ~p(s(bool,'const.bool.IN_2'(s(X87,X88),s(fun(X87,bool),X92)))) | s(X90,happ(s(fun(X87,X90),X91),s(X87,X88))) != s(X90,happ(s(fun(X87,X90),X91),s(X87,X89)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f5495,plain,(
  ( ! [X243,X241,X239,X237,X244,X242,X240,X238,X245] : (s(X237,X238) = s(X237,X239) | s(X240,happ(s(fun(X237,X240),X241),s(X237,X238))) != s(X240,happ(s(fun(X237,X240),X241),s(X237,X239))) | ~p(s(bool,'const.bool.IN_2'(s(X237,X239),s(fun(X237,bool),X242)))) | ~sP0(X237,X240,X241,X242,X243) | p(s(bool,X244)) | ~p(s(bool,X245)) | s(bool,X245) != s(bool,'const.bool.IN_2'(s(X237,X238),s(fun(X237,bool),X242)))) )),
  inference(resolution,[],[f127,f206])).
fof(f5484,plain,(
  ( ! [X158,X165,X163,X161,X159,X157,X166,X164,X162,X160] : (s(X157,X158) = s(X157,X159) | s(X160,happ(s(fun(X157,X160),X161),s(X157,X158))) != s(X160,happ(s(fun(X157,X160),X161),s(X157,X159))) | ~p(s(bool,'const.bool.IN_2'(s(X157,X159),s(fun(X157,bool),X162)))) | ~sP0(X157,X160,X161,X162,X163) | ~p(s(bool,X164)) | ~p(s(bool,X165)) | s(bool,X165) != s(bool,'const.bool.IN_2'(s(X157,X158),s(fun(X157,bool),X162))) | s(bool,X164) != s(bool,X166)) )),
  inference(resolution,[],[f127,f150])).
fof(f5482,plain,(
  ( ! [X146,X144,X142,X140,X138,X145,X143,X141,X139,X137] : (s(X137,X138) = s(X137,X139) | s(X140,happ(s(fun(X137,X140),X141),s(X137,X138))) != s(X140,happ(s(fun(X137,X140),X141),s(X137,X139))) | ~p(s(bool,'const.bool.IN_2'(s(X137,X139),s(fun(X137,bool),X142)))) | ~sP0(X137,X140,X141,X142,X143) | p(s(bool,X144)) | ~p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X145) != s(bool,'const.bool.IN_2'(s(X137,X138),s(fun(X137,bool),X142)))) )),
  inference(resolution,[],[f127,f149])).
fof(f5480,plain,(
  ( ! [X125,X123,X121,X118,X126,X124,X122,X120,X119,X117] : (s(X117,X118) = s(X117,X119) | s(X120,happ(s(fun(X117,X120),X121),s(X117,X118))) != s(X120,happ(s(fun(X117,X120),X121),s(X117,X119))) | ~p(s(bool,'const.bool.IN_2'(s(X117,X119),s(fun(X117,bool),X122)))) | ~sP0(X117,X120,X121,X122,X123) | p(s(bool,X124)) | ~p(s(bool,X125)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,'const.bool.IN_2'(s(X117,X118),s(fun(X117,bool),X122)))) )),
  inference(resolution,[],[f127,f148])).
fof(f5479,plain,(
  ( ! [X111,X109,X107,X116,X114,X112,X110,X108,X115,X113] : (s(X107,X108) = s(X107,X109) | s(X110,happ(s(fun(X107,X110),X111),s(X107,X108))) != s(X110,happ(s(fun(X107,X110),X111),s(X107,X109))) | ~p(s(bool,'const.bool.IN_2'(s(X107,X109),s(fun(X107,bool),X112)))) | ~sP0(X107,X110,X111,X112,X113) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X115) != s(bool,'const.bool.IN_2'(s(X107,X108),s(fun(X107,bool),X112))) | s(bool,X114) != s(bool,X116)) )),
  inference(resolution,[],[f127,f146])).
fof(f5477,plain,(
  ( ! [X94,X92,X90,X88,X87,X95,X93,X91,X89,X96] : (s(X87,X88) = s(X87,X89) | s(X90,happ(s(fun(X87,X90),X91),s(X87,X88))) != s(X90,happ(s(fun(X87,X90),X91),s(X87,X89))) | ~p(s(bool,'const.bool.IN_2'(s(X87,X89),s(fun(X87,bool),X92)))) | ~sP0(X87,X90,X91,X92,X93) | p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,X96) | s(bool,X95) != s(bool,'const.bool.IN_2'(s(X87,X88),s(fun(X87,bool),X92)))) )),
  inference(resolution,[],[f127,f144])).
fof(f5461,plain,(
  ( ! [X243,X241,X239,X237,X244,X242,X240,X238,X245] : (s(X237,X238) = s(X237,X239) | s(X240,happ(s(fun(X237,X240),X241),s(X237,X238))) != s(X240,happ(s(fun(X237,X240),X241),s(X237,X239))) | ~p(s(bool,'const.bool.IN_2'(s(X237,X238),s(fun(X237,bool),X242)))) | ~sP0(X237,X240,X241,X242,X243) | p(s(bool,X244)) | ~p(s(bool,X245)) | s(bool,X245) != s(bool,'const.bool.IN_2'(s(X237,X239),s(fun(X237,bool),X242)))) )),
  inference(resolution,[],[f127,f206])).
fof(f5450,plain,(
  ( ! [X158,X165,X163,X161,X159,X157,X166,X164,X162,X160] : (s(X157,X158) = s(X157,X159) | s(X160,happ(s(fun(X157,X160),X161),s(X157,X158))) != s(X160,happ(s(fun(X157,X160),X161),s(X157,X159))) | ~p(s(bool,'const.bool.IN_2'(s(X157,X158),s(fun(X157,bool),X162)))) | ~sP0(X157,X160,X161,X162,X163) | ~p(s(bool,X164)) | ~p(s(bool,X165)) | s(bool,X165) != s(bool,'const.bool.IN_2'(s(X157,X159),s(fun(X157,bool),X162))) | s(bool,X164) != s(bool,X166)) )),
  inference(resolution,[],[f127,f150])).
fof(f5448,plain,(
  ( ! [X146,X144,X142,X140,X138,X145,X143,X141,X139,X137] : (s(X137,X138) = s(X137,X139) | s(X140,happ(s(fun(X137,X140),X141),s(X137,X138))) != s(X140,happ(s(fun(X137,X140),X141),s(X137,X139))) | ~p(s(bool,'const.bool.IN_2'(s(X137,X138),s(fun(X137,bool),X142)))) | ~sP0(X137,X140,X141,X142,X143) | p(s(bool,X144)) | ~p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X145) != s(bool,'const.bool.IN_2'(s(X137,X139),s(fun(X137,bool),X142)))) )),
  inference(resolution,[],[f127,f149])).
fof(f5446,plain,(
  ( ! [X125,X123,X121,X118,X126,X124,X122,X120,X119,X117] : (s(X117,X118) = s(X117,X119) | s(X120,happ(s(fun(X117,X120),X121),s(X117,X118))) != s(X120,happ(s(fun(X117,X120),X121),s(X117,X119))) | ~p(s(bool,'const.bool.IN_2'(s(X117,X118),s(fun(X117,bool),X122)))) | ~sP0(X117,X120,X121,X122,X123) | p(s(bool,X124)) | ~p(s(bool,X125)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,'const.bool.IN_2'(s(X117,X119),s(fun(X117,bool),X122)))) )),
  inference(resolution,[],[f127,f148])).
fof(f5445,plain,(
  ( ! [X111,X109,X107,X116,X114,X112,X110,X108,X115,X113] : (s(X107,X108) = s(X107,X109) | s(X110,happ(s(fun(X107,X110),X111),s(X107,X108))) != s(X110,happ(s(fun(X107,X110),X111),s(X107,X109))) | ~p(s(bool,'const.bool.IN_2'(s(X107,X108),s(fun(X107,bool),X112)))) | ~sP0(X107,X110,X111,X112,X113) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X115) != s(bool,'const.bool.IN_2'(s(X107,X109),s(fun(X107,bool),X112))) | s(bool,X114) != s(bool,X116)) )),
  inference(resolution,[],[f127,f146])).
fof(f5443,plain,(
  ( ! [X94,X92,X90,X88,X87,X95,X93,X91,X89,X96] : (s(X87,X88) = s(X87,X89) | s(X90,happ(s(fun(X87,X90),X91),s(X87,X88))) != s(X90,happ(s(fun(X87,X90),X91),s(X87,X89))) | ~p(s(bool,'const.bool.IN_2'(s(X87,X88),s(fun(X87,bool),X92)))) | ~sP0(X87,X90,X91,X92,X93) | p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X94) != s(bool,X96) | s(bool,X95) != s(bool,'const.bool.IN_2'(s(X87,X89),s(fun(X87,bool),X92)))) )),
  inference(resolution,[],[f127,f144])).
fof(f1174,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | ~p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f1174_D])).
fof(f1174_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | ~p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f5126,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl19),
  inference(resolution,[],[f410,f206])).
fof(f5115,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X21) != s(bool,X23)) ) | $spl19),
  inference(resolution,[],[f410,f150])).
fof(f5113,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl19),
  inference(resolution,[],[f410,f149])).
fof(f5111,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl19),
  inference(resolution,[],[f410,f148])).
fof(f5110,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X6) != s(bool,X8)) ) | $spl19),
  inference(resolution,[],[f410,f146])).
fof(f5108,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl19),
  inference(resolution,[],[f410,f144])).
fof(f5097,plain,(
  ( ! [X187,X185,X190,X188,X186] : (~p(s(bool,happ(s(fun(X185,bool),X186),s(X185,sK6(X187,X188,X186,bool,X185))))) | p(s(bool,'const.bool.IN_2'(s(X185,sK5(X187,X188,X186,bool,X185)),s(fun(X185,bool),X188)))) | sP0(X185,bool,X186,X188,X187) | s(bool,X190) != s(bool,happ(s(fun(X185,bool),X186),s(X185,sK7(X187,X188,X186,bool,X185)))) | p(s(bool,X190))) ) | $spl108),
  inference(cnf_transformation,[],[f5097_D])).
fof(f5097_D,plain,(
  ( ! [X187,X185,X190,X188,X186] : (~p(s(bool,happ(s(fun(X185,bool),X186),s(X185,sK6(X187,X188,X186,bool,X185))))) | p(s(bool,'const.bool.IN_2'(s(X185,sK5(X187,X188,X186,bool,X185)),s(fun(X185,bool),X188)))) | sP0(X185,bool,X186,X188,X187) | s(bool,X190) != s(bool,happ(s(fun(X185,bool),X186),s(X185,sK7(X187,X188,X186,bool,X185)))) | p(s(bool,X190))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f5033,plain,(
  ( ! [X524,X526,X522,X525,X521,X523] : (~p(s(bool,happ(s(fun(X521,bool),X522),s(X521,sK6(X523,X524,X522,bool,X521))))) | p(s(bool,X525)) | p(s(bool,X526)) | s(bool,X526) != s(bool,happ(s(fun(X521,bool),X522),s(X521,sK7(X523,X524,X522,bool,X521)))) | sP0(X521,bool,X522,X524,X523) | p(s(bool,'const.bool.IN_2'(s(X521,sK5(X523,X524,X522,bool,X521)),s(fun(X521,bool),X524))))) )),
  inference(superposition,[],[f206,f132])).
fof(f132,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X3,X4) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK6(X4,X3,X2,X1,X0)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK7(X4,X3,X2,X1,X0)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK5(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f5096,plain,(
  ( ! [X173,X171,X176,X174,X172] : (p(s(bool,happ(s(fun(X171,bool),X172),s(X171,sK6(X173,X174,X172,bool,X171))))) | p(s(bool,'const.bool.IN_2'(s(X171,sK5(X173,X174,X172,bool,X171)),s(fun(X171,bool),X174)))) | sP0(X171,bool,X172,X174,X173) | s(bool,X176) != s(bool,happ(s(fun(X171,bool),X172),s(X171,sK7(X173,X174,X172,bool,X171)))) | ~p(s(bool,X176))) ) | $spl106),
  inference(cnf_transformation,[],[f5096_D])).
fof(f5096_D,plain,(
  ( ! [X173,X171,X176,X174,X172] : (p(s(bool,happ(s(fun(X171,bool),X172),s(X171,sK6(X173,X174,X172,bool,X171))))) | p(s(bool,'const.bool.IN_2'(s(X171,sK5(X173,X174,X172,bool,X171)),s(fun(X171,bool),X174)))) | sP0(X171,bool,X172,X174,X173) | s(bool,X176) != s(bool,happ(s(fun(X171,bool),X172),s(X171,sK7(X173,X174,X172,bool,X171)))) | ~p(s(bool,X176))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f5031,plain,(
  ( ! [X509,X511,X512,X514,X510,X513] : (p(s(bool,happ(s(fun(X509,bool),X510),s(X509,sK6(X511,X512,X510,bool,X509))))) | p(s(bool,X513)) | ~p(s(bool,X514)) | s(bool,X514) != s(bool,happ(s(fun(X509,bool),X510),s(X509,sK7(X511,X512,X510,bool,X509)))) | sP0(X509,bool,X510,X512,X511) | p(s(bool,'const.bool.IN_2'(s(X509,sK5(X511,X512,X510,bool,X509)),s(fun(X509,bool),X512))))) )),
  inference(superposition,[],[f206,f132])).
fof(f5093,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,X147) != s(bool,happ(s(fun(X143,bool),X144),s(X143,sK6(X145,X146,X144,bool,X143)))) | p(s(bool,'const.bool.IN_2'(s(X143,sK5(X145,X146,X144,bool,X143)),s(fun(X143,bool),X146)))) | sP0(X143,bool,X144,X146,X145) | p(s(bool,happ(s(fun(X143,bool),X144),s(X143,sK7(X145,X146,X144,bool,X143))))) | ~p(s(bool,X147))) ) | $spl102),
  inference(cnf_transformation,[],[f5093_D])).
fof(f5093_D,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,X147) != s(bool,happ(s(fun(X143,bool),X144),s(X143,sK6(X145,X146,X144,bool,X143)))) | p(s(bool,'const.bool.IN_2'(s(X143,sK5(X145,X146,X144,bool,X143)),s(fun(X143,bool),X146)))) | sP0(X143,bool,X144,X146,X145) | p(s(bool,happ(s(fun(X143,bool),X144),s(X143,sK7(X145,X146,X144,bool,X143))))) | ~p(s(bool,X147))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f5030,plain,(
  ( ! [X505,X507,X508,X504,X506,X503] : (s(bool,X507) != s(bool,happ(s(fun(X503,bool),X504),s(X503,sK6(X505,X506,X504,bool,X503)))) | p(s(bool,X508)) | ~p(s(bool,X507)) | p(s(bool,happ(s(fun(X503,bool),X504),s(X503,sK7(X505,X506,X504,bool,X503))))) | sP0(X503,bool,X504,X506,X505) | p(s(bool,'const.bool.IN_2'(s(X503,sK5(X505,X506,X504,bool,X503)),s(fun(X503,bool),X506))))) )),
  inference(superposition,[],[f206,f132])).
fof(f5094,plain,(
  ( ! [X154,X152,X151,X153,X150] : (s(bool,X154) != s(bool,happ(s(fun(X150,bool),X151),s(X150,sK6(X152,X153,X151,bool,X150)))) | p(s(bool,'const.bool.IN_2'(s(X150,sK5(X152,X153,X151,bool,X150)),s(fun(X150,bool),X153)))) | sP0(X150,bool,X151,X153,X152) | p(s(bool,X154)) | ~p(s(bool,happ(s(fun(X150,bool),X151),s(X150,sK7(X152,X153,X151,bool,X150)))))) ) | $spl104),
  inference(cnf_transformation,[],[f5094_D])).
fof(f5094_D,plain,(
  ( ! [X154,X152,X151,X153,X150] : (s(bool,X154) != s(bool,happ(s(fun(X150,bool),X151),s(X150,sK6(X152,X153,X151,bool,X150)))) | p(s(bool,'const.bool.IN_2'(s(X150,sK5(X152,X153,X151,bool,X150)),s(fun(X150,bool),X153)))) | sP0(X150,bool,X151,X153,X152) | p(s(bool,X154)) | ~p(s(bool,happ(s(fun(X150,bool),X151),s(X150,sK7(X152,X153,X151,bool,X150)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f5029,plain,(
  ( ! [X499,X500,X502,X498,X501,X497] : (s(bool,X501) != s(bool,happ(s(fun(X497,bool),X498),s(X497,sK6(X499,X500,X498,bool,X497)))) | p(s(bool,X502)) | ~p(s(bool,happ(s(fun(X497,bool),X498),s(X497,sK7(X499,X500,X498,bool,X497))))) | p(s(bool,X501)) | sP0(X497,bool,X498,X500,X499) | p(s(bool,'const.bool.IN_2'(s(X497,sK5(X499,X500,X498,bool,X497)),s(fun(X497,bool),X500))))) )),
  inference(superposition,[],[f206,f132])).
fof(f4999,plain,(
  ( ! [X377,X379,X374,X380,X376,X378,X375] : (~p(s(bool,happ(s(fun(X374,bool),X375),s(X374,sK6(X376,X377,X375,bool,X374))))) | p(s(bool,X378)) | ~p(s(bool,X379)) | s(bool,X378) != s(bool,happ(s(fun(X374,bool),X375),s(X374,sK7(X376,X377,X375,bool,X374)))) | s(bool,X379) != s(bool,X380) | sP0(X374,bool,X375,X377,X376) | p(s(bool,'const.bool.IN_2'(s(X374,sK5(X376,X377,X375,bool,X374)),s(fun(X374,bool),X377))))) )),
  inference(superposition,[],[f150,f132])).
fof(f4997,plain,(
  ( ! [X364,X366,X360,X362,X365,X361,X363] : (p(s(bool,happ(s(fun(X360,bool),X361),s(X360,sK6(X362,X363,X361,bool,X360))))) | ~p(s(bool,X364)) | ~p(s(bool,X365)) | s(bool,X365) != s(bool,happ(s(fun(X360,bool),X361),s(X360,sK7(X362,X363,X361,bool,X360)))) | s(bool,X364) != s(bool,X366) | sP0(X360,bool,X361,X363,X362) | p(s(bool,'const.bool.IN_2'(s(X360,sK5(X362,X363,X361,bool,X360)),s(fun(X360,bool),X363))))) )),
  inference(superposition,[],[f150,f132])).
fof(f4996,plain,(
  ( ! [X356,X358,X354,X357,X359,X353,X355] : (s(bool,X357) != s(bool,happ(s(fun(X353,bool),X354),s(X353,sK6(X355,X356,X354,bool,X353)))) | p(s(bool,X357)) | ~p(s(bool,X358)) | ~p(s(bool,happ(s(fun(X353,bool),X354),s(X353,sK7(X355,X356,X354,bool,X353))))) | s(bool,X358) != s(bool,X359) | sP0(X353,bool,X354,X356,X355) | p(s(bool,'const.bool.IN_2'(s(X353,sK5(X355,X356,X354,bool,X353)),s(fun(X353,bool),X356))))) )),
  inference(superposition,[],[f150,f132])).
fof(f4995,plain,(
  ( ! [X350,X346,X352,X349,X351,X347,X348] : (s(bool,X350) != s(bool,happ(s(fun(X346,bool),X347),s(X346,sK6(X348,X349,X347,bool,X346)))) | p(s(bool,happ(s(fun(X346,bool),X347),s(X346,sK7(X348,X349,X347,bool,X346))))) | ~p(s(bool,X351)) | ~p(s(bool,X350)) | s(bool,X351) != s(bool,X352) | sP0(X346,bool,X347,X349,X348) | p(s(bool,'const.bool.IN_2'(s(X346,sK5(X348,X349,X347,bool,X346)),s(fun(X346,bool),X349))))) )),
  inference(superposition,[],[f150,f132])).
fof(f4991,plain,(
  ( ! [X321,X323,X318,X324,X320,X322,X319] : (~p(s(bool,happ(s(fun(X318,bool),X319),s(X318,sK6(X320,X321,X319,bool,X318))))) | p(s(bool,X322)) | ~p(s(bool,X323)) | p(s(bool,X324)) | s(bool,X324) != s(bool,happ(s(fun(X318,bool),X319),s(X318,sK7(X320,X321,X319,bool,X318)))) | sP0(X318,bool,X319,X321,X320) | p(s(bool,'const.bool.IN_2'(s(X318,sK5(X320,X321,X319,bool,X318)),s(fun(X318,bool),X321))))) )),
  inference(superposition,[],[f149,f132])).
fof(f4989,plain,(
  ( ! [X308,X310,X304,X306,X309,X305,X307] : (p(s(bool,happ(s(fun(X304,bool),X305),s(X304,sK6(X306,X307,X305,bool,X304))))) | p(s(bool,X308)) | ~p(s(bool,X309)) | ~p(s(bool,X310)) | s(bool,X309) != s(bool,happ(s(fun(X304,bool),X305),s(X304,sK7(X306,X307,X305,bool,X304)))) | sP0(X304,bool,X305,X307,X306) | p(s(bool,'const.bool.IN_2'(s(X304,sK5(X306,X307,X305,bool,X304)),s(fun(X304,bool),X307))))) )),
  inference(superposition,[],[f149,f132])).
fof(f4988,plain,(
  ( ! [X300,X302,X298,X301,X303,X297,X299] : (s(bool,X301) != s(bool,happ(s(fun(X297,bool),X298),s(X297,sK6(X299,X300,X298,bool,X297)))) | p(s(bool,X302)) | ~p(s(bool,X301)) | ~p(s(bool,X303)) | p(s(bool,happ(s(fun(X297,bool),X298),s(X297,sK7(X299,X300,X298,bool,X297))))) | sP0(X297,bool,X298,X300,X299) | p(s(bool,'const.bool.IN_2'(s(X297,sK5(X299,X300,X298,bool,X297)),s(fun(X297,bool),X300))))) )),
  inference(superposition,[],[f149,f132])).
fof(f4987,plain,(
  ( ! [X292,X294,X290,X296,X293,X295,X291] : (s(bool,X294) != s(bool,happ(s(fun(X290,bool),X291),s(X290,sK6(X292,X293,X291,bool,X290)))) | p(s(bool,X295)) | ~p(s(bool,happ(s(fun(X290,bool),X291),s(X290,sK7(X292,X293,X291,bool,X290))))) | ~p(s(bool,X296)) | p(s(bool,X294)) | sP0(X290,bool,X291,X293,X292) | p(s(bool,'const.bool.IN_2'(s(X290,sK5(X292,X293,X291,bool,X290)),s(fun(X290,bool),X293))))) )),
  inference(superposition,[],[f149,f132])).
fof(f4986,plain,(
  ( ! [X288,X285,X287,X283,X289,X284,X286] : (~p(s(bool,happ(s(fun(X283,bool),X284),s(X283,sK6(X285,X286,X284,bool,X283))))) | p(s(bool,X287)) | p(s(bool,X288)) | s(bool,X287) != s(bool,X289) | s(bool,X288) != s(bool,happ(s(fun(X283,bool),X284),s(X283,sK7(X285,X286,X284,bool,X283)))) | sP0(X283,bool,X284,X286,X285) | p(s(bool,'const.bool.IN_2'(s(X283,sK5(X285,X286,X284,bool,X283)),s(fun(X283,bool),X286))))) )),
  inference(superposition,[],[f148,f132])).
fof(f4984,plain,(
  ( ! [X269,X271,X273,X275,X270,X272,X274] : (p(s(bool,happ(s(fun(X269,bool),X270),s(X269,sK6(X271,X272,X270,bool,X269))))) | p(s(bool,X273)) | ~p(s(bool,X274)) | s(bool,X273) != s(bool,X275) | s(bool,X274) != s(bool,happ(s(fun(X269,bool),X270),s(X269,sK7(X271,X272,X270,bool,X269)))) | sP0(X269,bool,X270,X272,X271) | p(s(bool,'const.bool.IN_2'(s(X269,sK5(X271,X272,X270,bool,X269)),s(fun(X269,bool),X272))))) )),
  inference(superposition,[],[f148,f132])).
fof(f4981,plain,(
  ( ! [X253,X251,X249,X254,X252,X250,X248] : (s(bool,X252) != s(bool,happ(s(fun(X248,bool),X249),s(X248,sK6(X250,X251,X249,bool,X248)))) | p(s(bool,X253)) | ~p(s(bool,X252)) | p(s(bool,happ(s(fun(X248,bool),X249),s(X248,sK7(X250,X251,X249,bool,X248))))) | s(bool,X253) != s(bool,X254) | sP0(X248,bool,X249,X251,X250) | p(s(bool,'const.bool.IN_2'(s(X248,sK5(X250,X251,X249,bool,X248)),s(fun(X248,bool),X251))))) )),
  inference(superposition,[],[f148,f132])).
fof(f4980,plain,(
  ( ! [X243,X241,X246,X244,X242,X247,X245] : (s(bool,X245) != s(bool,happ(s(fun(X241,bool),X242),s(X241,sK6(X243,X244,X242,bool,X241)))) | p(s(bool,X246)) | ~p(s(bool,happ(s(fun(X241,bool),X242),s(X241,sK7(X243,X244,X242,bool,X241))))) | p(s(bool,X245)) | s(bool,X246) != s(bool,X247) | sP0(X241,bool,X242,X244,X243) | p(s(bool,'const.bool.IN_2'(s(X241,sK5(X243,X244,X242,bool,X241)),s(fun(X241,bool),X244))))) )),
  inference(superposition,[],[f148,f132])).
fof(f4979,plain,(
  ( ! [X239,X237,X235,X240,X238,X236,X234] : (~p(s(bool,happ(s(fun(X234,bool),X235),s(X234,sK6(X236,X237,X235,bool,X234))))) | p(s(bool,X238)) | ~p(s(bool,X239)) | s(bool,X238) != s(bool,happ(s(fun(X234,bool),X235),s(X234,sK7(X236,X237,X235,bool,X234)))) | s(bool,X239) != s(bool,X240) | sP0(X234,bool,X235,X237,X236) | p(s(bool,'const.bool.IN_2'(s(X234,sK5(X236,X237,X235,bool,X234)),s(fun(X234,bool),X237))))) )),
  inference(superposition,[],[f146,f132])).
fof(f4977,plain,(
  ( ! [X222,X220,X225,X223,X221,X226,X224] : (p(s(bool,happ(s(fun(X220,bool),X221),s(X220,sK6(X222,X223,X221,bool,X220))))) | ~p(s(bool,X224)) | ~p(s(bool,X225)) | s(bool,X225) != s(bool,happ(s(fun(X220,bool),X221),s(X220,sK7(X222,X223,X221,bool,X220)))) | s(bool,X224) != s(bool,X226) | sP0(X220,bool,X221,X223,X222) | p(s(bool,'const.bool.IN_2'(s(X220,sK5(X222,X223,X221,bool,X220)),s(fun(X220,bool),X223))))) )),
  inference(superposition,[],[f146,f132])).
fof(f4976,plain,(
  ( ! [X218,X216,X215,X213,X219,X217,X214] : (s(bool,X217) != s(bool,happ(s(fun(X213,bool),X214),s(X213,sK6(X215,X216,X214,bool,X213)))) | p(s(bool,happ(s(fun(X213,bool),X214),s(X213,sK7(X215,X216,X214,bool,X213))))) | ~p(s(bool,X218)) | ~p(s(bool,X217)) | s(bool,X218) != s(bool,X219) | sP0(X213,bool,X214,X216,X215) | p(s(bool,'const.bool.IN_2'(s(X213,sK5(X215,X216,X214,bool,X213)),s(fun(X213,bool),X216))))) )),
  inference(superposition,[],[f146,f132])).
fof(f4975,plain,(
  ( ! [X212,X210,X208,X206,X211,X209,X207] : (s(bool,X210) != s(bool,happ(s(fun(X206,bool),X207),s(X206,sK6(X208,X209,X207,bool,X206)))) | p(s(bool,X210)) | ~p(s(bool,X211)) | ~p(s(bool,happ(s(fun(X206,bool),X207),s(X206,sK7(X208,X209,X207,bool,X206))))) | s(bool,X211) != s(bool,X212) | sP0(X206,bool,X207,X209,X208) | p(s(bool,'const.bool.IN_2'(s(X206,sK5(X208,X209,X207,bool,X206)),s(fun(X206,bool),X209))))) )),
  inference(superposition,[],[f146,f132])).
fof(f4972,plain,(
  ( ! [X191,X189,X187,X185,X190,X188,X186] : (~p(s(bool,happ(s(fun(X185,bool),X186),s(X185,sK6(X187,X188,X186,bool,X185))))) | p(s(bool,X189)) | p(s(bool,X190)) | s(bool,X189) != s(bool,X191) | s(bool,X190) != s(bool,happ(s(fun(X185,bool),X186),s(X185,sK7(X187,X188,X186,bool,X185)))) | sP0(X185,bool,X186,X188,X187) | p(s(bool,'const.bool.IN_2'(s(X185,sK5(X187,X188,X186,bool,X185)),s(fun(X185,bool),X188))))) )),
  inference(superposition,[],[f144,f132])).
fof(f4970,plain,(
  ( ! [X177,X175,X173,X171,X176,X174,X172] : (p(s(bool,happ(s(fun(X171,bool),X172),s(X171,sK6(X173,X174,X172,bool,X171))))) | p(s(bool,X175)) | ~p(s(bool,X176)) | s(bool,X175) != s(bool,X177) | s(bool,X176) != s(bool,happ(s(fun(X171,bool),X172),s(X171,sK7(X173,X174,X172,bool,X171)))) | sP0(X171,bool,X172,X174,X173) | p(s(bool,'const.bool.IN_2'(s(X171,sK5(X173,X174,X172,bool,X171)),s(fun(X171,bool),X174))))) )),
  inference(superposition,[],[f144,f132])).
fof(f4967,plain,(
  ( ! [X156,X154,X152,X151,X155,X153,X150] : (s(bool,X154) != s(bool,happ(s(fun(X150,bool),X151),s(X150,sK6(X152,X153,X151,bool,X150)))) | p(s(bool,X155)) | ~p(s(bool,happ(s(fun(X150,bool),X151),s(X150,sK7(X152,X153,X151,bool,X150))))) | p(s(bool,X154)) | s(bool,X155) != s(bool,X156) | sP0(X150,bool,X151,X153,X152) | p(s(bool,'const.bool.IN_2'(s(X150,sK5(X152,X153,X151,bool,X150)),s(fun(X150,bool),X153))))) )),
  inference(superposition,[],[f144,f132])).
fof(f4966,plain,(
  ( ! [X146,X144,X149,X147,X145,X143,X148] : (s(bool,X147) != s(bool,happ(s(fun(X143,bool),X144),s(X143,sK6(X145,X146,X144,bool,X143)))) | p(s(bool,X148)) | ~p(s(bool,X147)) | p(s(bool,happ(s(fun(X143,bool),X144),s(X143,sK7(X145,X146,X144,bool,X143))))) | s(bool,X148) != s(bool,X149) | sP0(X143,bool,X144,X146,X145) | p(s(bool,'const.bool.IN_2'(s(X143,sK5(X145,X146,X144,bool,X143)),s(fun(X143,bool),X146))))) )),
  inference(superposition,[],[f144,f132])).
fof(f5088,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | s(X9,happ(s(fun(X8,X9),X10),s(X8,sK6(X12,X11,X10,X9,X8)))) = s(X9,happ(s(fun(X8,X9),X10),s(X8,sK7(X12,X11,X10,X9,X8))))) ) | $spl100),
  inference(cnf_transformation,[],[f5088_D])).
fof(f5088_D,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | s(X9,happ(s(fun(X8,X9),X10),s(X8,sK6(X12,X11,X10,X9,X8)))) = s(X9,happ(s(fun(X8,X9),X10),s(X8,sK7(X12,X11,X10,X9,X8))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f4882,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (sP0(X117,X118,X119,X120,X121) | s(X118,happ(s(fun(X117,X118),X119),s(X117,sK6(X121,X120,X119,X118,X117)))) = s(X118,happ(s(fun(X117,X118),X119),s(X117,sK7(X121,X120,X119,X118,X117)))) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X123) != s(bool,'const.bool.IN_2'(s(X117,sK5(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) )),
  inference(resolution,[],[f132,f206])).
fof(f4872,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (sP0(X64,X65,X66,X67,X68) | s(X65,happ(s(fun(X64,X65),X66),s(X64,sK6(X68,X67,X66,X65,X64)))) = s(X65,happ(s(fun(X64,X65),X66),s(X64,sK7(X68,X67,X66,X65,X64)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.bool.IN_2'(s(X64,sK5(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67))) | s(bool,X70) != s(bool,X71)) )),
  inference(resolution,[],[f132,f150])).
fof(f4869,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (sP0(X40,X41,X42,X43,X44) | s(X41,happ(s(fun(X40,X41),X42),s(X40,sK6(X44,X43,X42,X41,X40)))) = s(X41,happ(s(fun(X40,X41),X42),s(X40,sK7(X44,X43,X42,X41,X40)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(X40,sK5(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) )),
  inference(resolution,[],[f132,f149])).
fof(f4868,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (sP0(X32,X33,X34,X35,X36) | s(X33,happ(s(fun(X32,X33),X34),s(X32,sK6(X36,X35,X34,X33,X32)))) = s(X33,happ(s(fun(X32,X33),X34),s(X32,sK7(X36,X35,X34,X33,X32)))) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(X32,sK5(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) )),
  inference(resolution,[],[f132,f148])).
fof(f4867,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (sP0(X24,X25,X26,X27,X28) | s(X25,happ(s(fun(X24,X25),X26),s(X24,sK6(X28,X27,X26,X25,X24)))) = s(X25,happ(s(fun(X24,X25),X26),s(X24,sK7(X28,X27,X26,X25,X24)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(X24,sK5(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f132,f146])).
fof(f4865,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(X9,happ(s(fun(X8,X9),X10),s(X8,sK6(X12,X11,X10,X9,X8)))) = s(X9,happ(s(fun(X8,X9),X10),s(X8,sK7(X12,X11,X10,X9,X8)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) )),
  inference(resolution,[],[f132,f144])).
fof(f411,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) | $spl18),
  inference(cnf_transformation,[],[f411_D])).
fof(f411_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5087,plain,(
  s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))),
  inference(forward_demodulation,[],[f5086,f138])).
fof(f5086,plain,(
  s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))),
  inference(forward_demodulation,[],[f4861,f138])).
fof(f4861,plain,(
  s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))),
  inference(resolution,[],[f132,f210])).
fof(f4803,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f144])).
fof(f4805,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X6) != s(bool,X8)) ) | $spl23),
  inference(resolution,[],[f416,f146])).
fof(f4806,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f148])).
fof(f4808,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f149])).
fof(f4810,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X21) != s(bool,X23)) ) | $spl23),
  inference(resolution,[],[f416,f150])).
fof(f4821,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl23),
  inference(resolution,[],[f416,f206])).
fof(f4854,plain,(
  $false | ($spl18 | $spl23)),
  inference(subsumption_resolution,[],[f4853,f201])).
fof(f201,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INJ_INR',aHL_TRUTH)).
fof(f4853,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl23)),
  inference(forward_demodulation,[],[f4802,f1196])).
fof(f1196,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | $spl18),
  inference(resolution,[],[f411,f196])).
fof(f4802,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) | $spl23),
  inference(resolution,[],[f416,f209])).
fof(f4644,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (sP0(X6,X7,X8,X9,X10) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10))) | ~p(s(bool,X12)) | s(X6,sK6(X10,X9,X8,X7,X6)) != s(X6,sK7(X10,X9,X8,X7,X6))) ) | $spl96),
  inference(cnf_transformation,[],[f4644_D])).
fof(f4644_D,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (sP0(X6,X7,X8,X9,X10) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10))) | ~p(s(bool,X12)) | s(X6,sK6(X10,X9,X8,X7,X6)) != s(X6,sK7(X10,X9,X8,X7,X6))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f4579,plain,(
  ( ! [X125,X123,X121,X126,X124,X122,X120] : (sP0(X120,X121,X122,X123,X124) | s(X120,sK6(X124,X123,X122,X121,X120)) != s(X120,sK7(X124,X123,X122,X121,X120)) | p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,'const.bool.IN_2'(s(X121,happ(s(fun(X120,X121),X122),s(X120,sK5(X124,X123,X122,X121,X120)))),s(fun(X121,bool),X124)))) )),
  inference(resolution,[],[f135,f206])).
fof(f135,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X3,X4) | s(X0,sK6(X4,X3,X2,X1,X0)) != s(X0,sK7(X4,X3,X2,X1,X0)) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK5(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f4568,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (sP0(X62,X63,X64,X65,X66) | s(X62,sK6(X66,X65,X64,X63,X62)) != s(X62,sK7(X66,X65,X64,X63,X62)) | ~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X68) != s(bool,'const.bool.IN_2'(s(X63,happ(s(fun(X62,X63),X64),s(X62,sK5(X66,X65,X64,X63,X62)))),s(fun(X63,bool),X66))) | s(bool,X67) != s(bool,X69)) )),
  inference(resolution,[],[f135,f150])).
fof(f4566,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (sP0(X46,X47,X48,X49,X50) | s(X46,sK6(X50,X49,X48,X47,X46)) != s(X46,sK7(X50,X49,X48,X47,X46)) | p(s(bool,X51)) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(X47,happ(s(fun(X46,X47),X48),s(X46,sK5(X50,X49,X48,X47,X46)))),s(fun(X47,bool),X50)))) )),
  inference(resolution,[],[f135,f149])).
fof(f4564,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (sP0(X30,X31,X32,X33,X34) | s(X30,sK6(X34,X33,X32,X31,X30)) != s(X30,sK7(X34,X33,X32,X31,X30)) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(X31,happ(s(fun(X30,X31),X32),s(X30,sK5(X34,X33,X32,X31,X30)))),s(fun(X31,bool),X34)))) )),
  inference(resolution,[],[f135,f148])).
fof(f4563,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (sP0(X22,X23,X24,X25,X26) | s(X22,sK6(X26,X25,X24,X23,X22)) != s(X22,sK7(X26,X25,X24,X23,X22)) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(X23,happ(s(fun(X22,X23),X24),s(X22,sK5(X26,X25,X24,X23,X22)))),s(fun(X23,bool),X26))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f135,f146])).
fof(f4561,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (sP0(X6,X7,X8,X9,X10) | s(X6,sK6(X10,X9,X8,X7,X6)) != s(X6,sK7(X10,X9,X8,X7,X6)) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10)))) )),
  inference(resolution,[],[f135,f144])).
fof(f4545,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (sP0(X6,X7,X8,X9,X10) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10))) | ~p(s(bool,X12)) | p(s(bool,'const.bool.IN_2'(s(X6,sK7(X10,X9,X8,X7,X6)),s(fun(X6,bool),X9))))) ) | $spl94),
  inference(cnf_transformation,[],[f4545_D])).
fof(f4545_D,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (sP0(X6,X7,X8,X9,X10) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10))) | ~p(s(bool,X12)) | p(s(bool,'const.bool.IN_2'(s(X6,sK7(X10,X9,X8,X7,X6)),s(fun(X6,bool),X9))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f4461,plain,(
  ( ! [X125,X123,X121,X126,X124,X122,X120] : (sP0(X120,X121,X122,X123,X124) | p(s(bool,'const.bool.IN_2'(s(X120,sK7(X124,X123,X122,X121,X120)),s(fun(X120,bool),X123)))) | p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,'const.bool.IN_2'(s(X121,happ(s(fun(X120,X121),X122),s(X120,sK5(X124,X123,X122,X121,X120)))),s(fun(X121,bool),X124)))) )),
  inference(resolution,[],[f131,f206])).
fof(f131,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X3,X4) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK5(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f4450,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (sP0(X62,X63,X64,X65,X66) | p(s(bool,'const.bool.IN_2'(s(X62,sK7(X66,X65,X64,X63,X62)),s(fun(X62,bool),X65)))) | ~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X68) != s(bool,'const.bool.IN_2'(s(X63,happ(s(fun(X62,X63),X64),s(X62,sK5(X66,X65,X64,X63,X62)))),s(fun(X63,bool),X66))) | s(bool,X67) != s(bool,X69)) )),
  inference(resolution,[],[f131,f150])).
fof(f4448,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (sP0(X46,X47,X48,X49,X50) | p(s(bool,'const.bool.IN_2'(s(X46,sK7(X50,X49,X48,X47,X46)),s(fun(X46,bool),X49)))) | p(s(bool,X51)) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(X47,happ(s(fun(X46,X47),X48),s(X46,sK5(X50,X49,X48,X47,X46)))),s(fun(X47,bool),X50)))) )),
  inference(resolution,[],[f131,f149])).
fof(f4446,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (sP0(X30,X31,X32,X33,X34) | p(s(bool,'const.bool.IN_2'(s(X30,sK7(X34,X33,X32,X31,X30)),s(fun(X30,bool),X33)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(X31,happ(s(fun(X30,X31),X32),s(X30,sK5(X34,X33,X32,X31,X30)))),s(fun(X31,bool),X34)))) )),
  inference(resolution,[],[f131,f148])).
fof(f4445,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (sP0(X22,X23,X24,X25,X26) | p(s(bool,'const.bool.IN_2'(s(X22,sK7(X26,X25,X24,X23,X22)),s(fun(X22,bool),X25)))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(X23,happ(s(fun(X22,X23),X24),s(X22,sK5(X26,X25,X24,X23,X22)))),s(fun(X23,bool),X26))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f131,f146])).
fof(f4443,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (sP0(X6,X7,X8,X9,X10) | p(s(bool,'const.bool.IN_2'(s(X6,sK7(X10,X9,X8,X7,X6)),s(fun(X6,bool),X9)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10)))) )),
  inference(resolution,[],[f131,f144])).
fof(f4544,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK7(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | ~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X8,X9),X10),s(X8,sK5(X12,X11,X10,X9,X8)))),s(fun(X9,bool),X12))))) ) | $spl92),
  inference(cnf_transformation,[],[f4544_D])).
fof(f4544_D,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK7(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | ~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X8,X9),X10),s(X8,sK5(X12,X11,X10,X9,X8)))),s(fun(X9,bool),X12))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4439,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (sP0(X117,X118,X119,X120,X121) | ~p(s(bool,'const.bool.IN_2'(s(X118,happ(s(fun(X117,X118),X119),s(X117,sK5(X121,X120,X119,X118,X117)))),s(fun(X118,bool),X121)))) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X123) != s(bool,'const.bool.IN_2'(s(X117,sK7(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) )),
  inference(resolution,[],[f131,f206])).
fof(f4429,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (sP0(X64,X65,X66,X67,X68) | ~p(s(bool,'const.bool.IN_2'(s(X65,happ(s(fun(X64,X65),X66),s(X64,sK5(X68,X67,X66,X65,X64)))),s(fun(X65,bool),X68)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.bool.IN_2'(s(X64,sK7(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67))) | s(bool,X70) != s(bool,X71)) )),
  inference(resolution,[],[f131,f150])).
fof(f4426,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (sP0(X40,X41,X42,X43,X44) | ~p(s(bool,'const.bool.IN_2'(s(X41,happ(s(fun(X40,X41),X42),s(X40,sK5(X44,X43,X42,X41,X40)))),s(fun(X41,bool),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(X40,sK7(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) )),
  inference(resolution,[],[f131,f149])).
fof(f4425,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (sP0(X32,X33,X34,X35,X36) | ~p(s(bool,'const.bool.IN_2'(s(X33,happ(s(fun(X32,X33),X34),s(X32,sK5(X36,X35,X34,X33,X32)))),s(fun(X33,bool),X36)))) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(X32,sK7(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) )),
  inference(resolution,[],[f131,f148])).
fof(f4424,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (sP0(X24,X25,X26,X27,X28) | ~p(s(bool,'const.bool.IN_2'(s(X25,happ(s(fun(X24,X25),X26),s(X24,sK5(X28,X27,X26,X25,X24)))),s(fun(X25,bool),X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(X24,sK7(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f131,f146])).
fof(f4422,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (sP0(X8,X9,X10,X11,X12) | ~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X8,X9),X10),s(X8,sK5(X12,X11,X10,X9,X8)))),s(fun(X9,bool),X12)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK7(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) )),
  inference(resolution,[],[f131,f144])).
fof(f4370,plain,(
  ( ! [X105,X106] : (s(bool,'T_0') != s(bool,X105) | s(bool,'T_0') = s(bool,X106) | s(bool,X105) != s(bool,X106)) ) | $spl90),
  inference(cnf_transformation,[],[f4370_D])).
fof(f4370_D,plain,(
  ( ! [X105,X106] : (s(bool,'T_0') != s(bool,X105) | s(bool,'T_0') = s(bool,X106) | s(bool,X105) != s(bool,X106)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f4367,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | s(bool,'T_0') != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'T_0') = s(bool,X26)) ) | $spl34),
  inference(duplicate_literal_removal,[],[f4355])).
fof(f4355,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | s(bool,'T_0') != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | ~p(s(bool,X24))) ) | $spl34),
  inference(superposition,[],[f1758,f196])).
fof(f4368,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | s(bool,'T_0') != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23)) ) | $spl34),
  inference(duplicate_literal_removal,[],[f4354])).
fof(f4354,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | s(bool,'T_0') != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | ~p(s(bool,X21))) ) | $spl34),
  inference(superposition,[],[f1758,f198])).
fof(f4375,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | $spl34),
  inference(subsumption_resolution,[],[f4353,f191])).
fof(f4353,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,'F_0') = s(bool,X18)) ) | $spl34),
  inference(superposition,[],[f1758,f205])).
fof(f4333,plain,(
  ( ! [X127,X125,X126] : (s(bool,'T_0') != s(bool,X125) | s(bool,X125) != s(bool,X126) | s(bool,'T_0') = s(bool,X126) | ~p(s(bool,X127))) ) | $spl34),
  inference(resolution,[],[f1758,f189])).
fof(f4332,plain,(
  ( ! [X123,X124,X122] : (s(bool,'T_0') != s(bool,X122) | s(bool,X122) != s(bool,X123) | s(bool,'T_0') = s(bool,X123) | ~p(s(bool,X124))) ) | $spl34),
  inference(resolution,[],[f1758,f181])).
fof(f4331,plain,(
  ( ! [X121,X120,X119] : (s(bool,'T_0') != s(bool,X119) | s(bool,X119) != s(bool,X120) | s(bool,'T_0') = s(bool,X120) | p(s(bool,X121))) ) | $spl34),
  inference(resolution,[],[f1758,f173])).
fof(f4330,plain,(
  ( ! [X118,X116,X117] : (s(bool,'T_0') != s(bool,X116) | s(bool,X116) != s(bool,X117) | s(bool,'T_0') = s(bool,X117) | p(s(bool,X118))) ) | $spl34),
  inference(resolution,[],[f1758,f170])).
fof(f4329,plain,(
  ( ! [X114,X115,X113] : (s(bool,'T_0') != s(bool,X113) | s(bool,X113) != s(bool,X114) | s(bool,'T_0') = s(bool,X114) | ~p(s(bool,X115))) ) | $spl34),
  inference(resolution,[],[f1758,f168])).
fof(f4327,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,X108) | s(bool,X108) != s(bool,X109) | s(bool,'T_0') = s(bool,X109) | ~p(s(bool,X110))) ) | $spl34),
  inference(resolution,[],[f1758,f165])).
fof(f4326,plain,(
  ( ! [X107,X105,X106] : (s(bool,'T_0') != s(bool,X105) | s(bool,X105) != s(bool,X106) | s(bool,'T_0') = s(bool,X106) | p(s(bool,X107))) ) | $spl34),
  inference(resolution,[],[f1758,f164])).
fof(f4369,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl88),
  inference(cnf_transformation,[],[f4369_D])).
fof(f4369_D,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f4320,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,X78) != s(bool,X79) | s(bool,X78) != s(bool,X80) | s(bool,'T_0') = s(bool,X80) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X79) != s(bool,X82)) ) | $spl34),
  inference(resolution,[],[f1758,f206])).
fof(f4309,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'T_0') = s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl34),
  inference(resolution,[],[f1758,f150])).
fof(f4307,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X31) != s(bool,X34)) ) | $spl34),
  inference(resolution,[],[f1758,f149])).
fof(f4305,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl34),
  inference(resolution,[],[f1758,f148])).
fof(f4304,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'T_0') = s(bool,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl34),
  inference(resolution,[],[f1758,f146])).
fof(f4302,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl34),
  inference(resolution,[],[f1758,f144])).
fof(f4290,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (sP0(X6,X7,X8,X9,X10) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10))) | ~p(s(bool,X12)) | p(s(bool,'const.bool.IN_2'(s(X6,sK6(X10,X9,X8,X7,X6)),s(fun(X6,bool),X9))))) ) | $spl86),
  inference(cnf_transformation,[],[f4290_D])).
fof(f4290_D,plain,(
  ( ! [X6,X12,X10,X8,X7,X9] : (sP0(X6,X7,X8,X9,X10) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10))) | ~p(s(bool,X12)) | p(s(bool,'const.bool.IN_2'(s(X6,sK6(X10,X9,X8,X7,X6)),s(fun(X6,bool),X9))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f4206,plain,(
  ( ! [X125,X123,X121,X126,X124,X122,X120] : (sP0(X120,X121,X122,X123,X124) | p(s(bool,'const.bool.IN_2'(s(X120,sK6(X124,X123,X122,X121,X120)),s(fun(X120,bool),X123)))) | p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,'const.bool.IN_2'(s(X121,happ(s(fun(X120,X121),X122),s(X120,sK5(X124,X123,X122,X121,X120)))),s(fun(X121,bool),X124)))) )),
  inference(resolution,[],[f129,f206])).
fof(f129,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X3,X4) | p(s(bool,'const.bool.IN_2'(s(X0,sK6(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK5(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f4195,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (sP0(X62,X63,X64,X65,X66) | p(s(bool,'const.bool.IN_2'(s(X62,sK6(X66,X65,X64,X63,X62)),s(fun(X62,bool),X65)))) | ~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X68) != s(bool,'const.bool.IN_2'(s(X63,happ(s(fun(X62,X63),X64),s(X62,sK5(X66,X65,X64,X63,X62)))),s(fun(X63,bool),X66))) | s(bool,X67) != s(bool,X69)) )),
  inference(resolution,[],[f129,f150])).
fof(f4193,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (sP0(X46,X47,X48,X49,X50) | p(s(bool,'const.bool.IN_2'(s(X46,sK6(X50,X49,X48,X47,X46)),s(fun(X46,bool),X49)))) | p(s(bool,X51)) | ~p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(X47,happ(s(fun(X46,X47),X48),s(X46,sK5(X50,X49,X48,X47,X46)))),s(fun(X47,bool),X50)))) )),
  inference(resolution,[],[f129,f149])).
fof(f4191,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (sP0(X30,X31,X32,X33,X34) | p(s(bool,'const.bool.IN_2'(s(X30,sK6(X34,X33,X32,X31,X30)),s(fun(X30,bool),X33)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(X31,happ(s(fun(X30,X31),X32),s(X30,sK5(X34,X33,X32,X31,X30)))),s(fun(X31,bool),X34)))) )),
  inference(resolution,[],[f129,f148])).
fof(f4190,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (sP0(X22,X23,X24,X25,X26) | p(s(bool,'const.bool.IN_2'(s(X22,sK6(X26,X25,X24,X23,X22)),s(fun(X22,bool),X25)))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(X23,happ(s(fun(X22,X23),X24),s(X22,sK5(X26,X25,X24,X23,X22)))),s(fun(X23,bool),X26))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f129,f146])).
fof(f4188,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (sP0(X6,X7,X8,X9,X10) | p(s(bool,'const.bool.IN_2'(s(X6,sK6(X10,X9,X8,X7,X6)),s(fun(X6,bool),X9)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,sK5(X10,X9,X8,X7,X6)))),s(fun(X7,bool),X10)))) )),
  inference(resolution,[],[f129,f144])).
fof(f4289,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK6(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | ~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X8,X9),X10),s(X8,sK5(X12,X11,X10,X9,X8)))),s(fun(X9,bool),X12))))) ) | $spl84),
  inference(cnf_transformation,[],[f4289_D])).
fof(f4289_D,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK6(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | ~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X8,X9),X10),s(X8,sK5(X12,X11,X10,X9,X8)))),s(fun(X9,bool),X12))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f4185,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (sP0(X117,X118,X119,X120,X121) | ~p(s(bool,'const.bool.IN_2'(s(X118,happ(s(fun(X117,X118),X119),s(X117,sK5(X121,X120,X119,X118,X117)))),s(fun(X118,bool),X121)))) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X123) != s(bool,'const.bool.IN_2'(s(X117,sK6(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) )),
  inference(resolution,[],[f129,f206])).
fof(f4175,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (sP0(X64,X65,X66,X67,X68) | ~p(s(bool,'const.bool.IN_2'(s(X65,happ(s(fun(X64,X65),X66),s(X64,sK5(X68,X67,X66,X65,X64)))),s(fun(X65,bool),X68)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.bool.IN_2'(s(X64,sK6(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67))) | s(bool,X70) != s(bool,X71)) )),
  inference(resolution,[],[f129,f150])).
fof(f4172,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (sP0(X40,X41,X42,X43,X44) | ~p(s(bool,'const.bool.IN_2'(s(X41,happ(s(fun(X40,X41),X42),s(X40,sK5(X44,X43,X42,X41,X40)))),s(fun(X41,bool),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(X40,sK6(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) )),
  inference(resolution,[],[f129,f149])).
fof(f4171,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (sP0(X32,X33,X34,X35,X36) | ~p(s(bool,'const.bool.IN_2'(s(X33,happ(s(fun(X32,X33),X34),s(X32,sK5(X36,X35,X34,X33,X32)))),s(fun(X33,bool),X36)))) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(X32,sK6(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) )),
  inference(resolution,[],[f129,f148])).
fof(f4170,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (sP0(X24,X25,X26,X27,X28) | ~p(s(bool,'const.bool.IN_2'(s(X25,happ(s(fun(X24,X25),X26),s(X24,sK5(X28,X27,X26,X25,X24)))),s(fun(X25,bool),X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(X24,sK6(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f129,f146])).
fof(f4168,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (sP0(X8,X9,X10,X11,X12) | ~p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X8,X9),X10),s(X8,sK5(X12,X11,X10,X9,X8)))),s(fun(X9,bool),X12)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK6(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) )),
  inference(resolution,[],[f129,f144])).
fof(f4157,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1)) ) | $spl82),
  inference(cnf_transformation,[],[f4157_D])).
fof(f4157_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X4) | p(s(bool,X4)) | s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f4125,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,'F_0') = s(bool,X75) | s(bool,X75) != s(bool,X76) | s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X77) != s(bool,X79)) )),
  inference(resolution,[],[f744,f206])).
fof(f4115,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'F_0') = s(bool,X42) | s(bool,X42) != s(bool,X43) | s(bool,X43) != s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f744,f150])).
fof(f4112,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'F_0') = s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X26) != s(bool,X29)) )),
  inference(resolution,[],[f744,f149])).
fof(f4111,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,X18) | s(bool,X18) != s(bool,X19) | s(bool,X19) != s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f744,f148])).
fof(f4110,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'F_0') = s(bool,X12) | s(bool,X12) != s(bool,X13) | s(bool,X13) != s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f744,f146])).
fof(f4108,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f744,f144])).
fof(f4104,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | s(X8,sK6(X12,X11,X10,X9,X8)) != s(X8,sK7(X12,X11,X10,X9,X8))) ) | $spl80),
  inference(cnf_transformation,[],[f4104_D])).
fof(f4104_D,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | s(X8,sK6(X12,X11,X10,X9,X8)) != s(X8,sK7(X12,X11,X10,X9,X8))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f4061,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (sP0(X117,X118,X119,X120,X121) | s(X117,sK6(X121,X120,X119,X118,X117)) != s(X117,sK7(X121,X120,X119,X118,X117)) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X123) != s(bool,'const.bool.IN_2'(s(X117,sK5(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) )),
  inference(resolution,[],[f134,f206])).
fof(f134,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,X1,X2,X3,X4) | s(X0,sK6(X4,X3,X2,X1,X0)) != s(X0,sK7(X4,X3,X2,X1,X0)) | p(s(bool,'const.bool.IN_2'(s(X0,sK5(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f4051,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (sP0(X64,X65,X66,X67,X68) | s(X64,sK6(X68,X67,X66,X65,X64)) != s(X64,sK7(X68,X67,X66,X65,X64)) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.bool.IN_2'(s(X64,sK5(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67))) | s(bool,X70) != s(bool,X71)) )),
  inference(resolution,[],[f134,f150])).
fof(f4048,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (sP0(X40,X41,X42,X43,X44) | s(X40,sK6(X44,X43,X42,X41,X40)) != s(X40,sK7(X44,X43,X42,X41,X40)) | p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(X40,sK5(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) )),
  inference(resolution,[],[f134,f149])).
fof(f4047,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (sP0(X32,X33,X34,X35,X36) | s(X32,sK6(X36,X35,X34,X33,X32)) != s(X32,sK7(X36,X35,X34,X33,X32)) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(X32,sK5(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) )),
  inference(resolution,[],[f134,f148])).
fof(f4046,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (sP0(X24,X25,X26,X27,X28) | s(X24,sK6(X28,X27,X26,X25,X24)) != s(X24,sK7(X28,X27,X26,X25,X24)) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(X24,sK5(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f134,f146])).
fof(f4044,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(X8,sK6(X12,X11,X10,X9,X8)) != s(X8,sK7(X12,X11,X10,X9,X8)) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) )),
  inference(resolution,[],[f134,f144])).
fof(f3954,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | p(s(bool,'const.bool.IN_2'(s(X8,sK7(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))))) ) | $spl78),
  inference(cnf_transformation,[],[f3954_D])).
fof(f3954_D,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | p(s(bool,'const.bool.IN_2'(s(X8,sK7(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f3891,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (sP0(X117,X118,X119,X120,X121) | p(s(bool,'const.bool.IN_2'(s(X117,sK7(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X123) != s(bool,'const.bool.IN_2'(s(X117,sK5(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) )),
  inference(resolution,[],[f130,f206])).
fof(f3881,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (sP0(X64,X65,X66,X67,X68) | p(s(bool,'const.bool.IN_2'(s(X64,sK7(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.bool.IN_2'(s(X64,sK5(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67))) | s(bool,X70) != s(bool,X71)) )),
  inference(resolution,[],[f130,f150])).
fof(f3878,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (sP0(X40,X41,X42,X43,X44) | p(s(bool,'const.bool.IN_2'(s(X40,sK7(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(X40,sK5(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) )),
  inference(resolution,[],[f130,f149])).
fof(f3877,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (sP0(X32,X33,X34,X35,X36) | p(s(bool,'const.bool.IN_2'(s(X32,sK7(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(X32,sK5(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) )),
  inference(resolution,[],[f130,f148])).
fof(f3876,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (sP0(X24,X25,X26,X27,X28) | p(s(bool,'const.bool.IN_2'(s(X24,sK7(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(X24,sK5(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f130,f146])).
fof(f3874,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (sP0(X8,X9,X10,X11,X12) | p(s(bool,'const.bool.IN_2'(s(X8,sK7(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) )),
  inference(resolution,[],[f130,f144])).
fof(f3952,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK7(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | p(s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))))) ) | $spl76),
  inference(cnf_transformation,[],[f3952_D])).
fof(f3952_D,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK7(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | p(s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f3871,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (sP0(X117,X118,X119,X120,X121) | p(s(bool,'const.bool.IN_2'(s(X117,sK5(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X123) != s(bool,'const.bool.IN_2'(s(X117,sK7(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) )),
  inference(resolution,[],[f130,f206])).
fof(f3861,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (sP0(X64,X65,X66,X67,X68) | p(s(bool,'const.bool.IN_2'(s(X64,sK5(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.bool.IN_2'(s(X64,sK7(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67))) | s(bool,X70) != s(bool,X71)) )),
  inference(resolution,[],[f130,f150])).
fof(f3858,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (sP0(X40,X41,X42,X43,X44) | p(s(bool,'const.bool.IN_2'(s(X40,sK5(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(X40,sK7(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) )),
  inference(resolution,[],[f130,f149])).
fof(f3857,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (sP0(X32,X33,X34,X35,X36) | p(s(bool,'const.bool.IN_2'(s(X32,sK5(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(X32,sK7(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) )),
  inference(resolution,[],[f130,f148])).
fof(f3856,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (sP0(X24,X25,X26,X27,X28) | p(s(bool,'const.bool.IN_2'(s(X24,sK5(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(X24,sK7(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f130,f146])).
fof(f3854,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (sP0(X8,X9,X10,X11,X12) | p(s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK7(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) )),
  inference(resolution,[],[f130,f144])).
fof(f3842,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | p(s(bool,'const.bool.IN_2'(s(X8,sK6(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))))) ) | $spl74),
  inference(cnf_transformation,[],[f3842_D])).
fof(f3842_D,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | p(s(bool,'const.bool.IN_2'(s(X8,sK6(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3779,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (sP0(X117,X118,X119,X120,X121) | p(s(bool,'const.bool.IN_2'(s(X117,sK6(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X123) != s(bool,'const.bool.IN_2'(s(X117,sK5(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) )),
  inference(resolution,[],[f128,f206])).
fof(f3769,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (sP0(X64,X65,X66,X67,X68) | p(s(bool,'const.bool.IN_2'(s(X64,sK6(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.bool.IN_2'(s(X64,sK5(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67))) | s(bool,X70) != s(bool,X71)) )),
  inference(resolution,[],[f128,f150])).
fof(f3766,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (sP0(X40,X41,X42,X43,X44) | p(s(bool,'const.bool.IN_2'(s(X40,sK6(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(X40,sK5(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) )),
  inference(resolution,[],[f128,f149])).
fof(f3765,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (sP0(X32,X33,X34,X35,X36) | p(s(bool,'const.bool.IN_2'(s(X32,sK6(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(X32,sK5(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) )),
  inference(resolution,[],[f128,f148])).
fof(f3764,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (sP0(X24,X25,X26,X27,X28) | p(s(bool,'const.bool.IN_2'(s(X24,sK6(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(X24,sK5(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f128,f146])).
fof(f3762,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (sP0(X8,X9,X10,X11,X12) | p(s(bool,'const.bool.IN_2'(s(X8,sK6(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) )),
  inference(resolution,[],[f128,f144])).
fof(f3840,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK6(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | p(s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))))) ) | $spl72),
  inference(cnf_transformation,[],[f3840_D])).
fof(f3840_D,plain,(
  ( ! [X14,X12,X10,X8,X11,X9] : (sP0(X8,X9,X10,X11,X12) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK6(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))) | p(s(bool,X14)) | p(s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3759,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (sP0(X117,X118,X119,X120,X121) | p(s(bool,'const.bool.IN_2'(s(X117,sK5(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X123) != s(bool,'const.bool.IN_2'(s(X117,sK6(X121,X120,X119,X118,X117)),s(fun(X117,bool),X120)))) )),
  inference(resolution,[],[f128,f206])).
fof(f3749,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (sP0(X64,X65,X66,X67,X68) | p(s(bool,'const.bool.IN_2'(s(X64,sK5(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.bool.IN_2'(s(X64,sK6(X68,X67,X66,X65,X64)),s(fun(X64,bool),X67))) | s(bool,X70) != s(bool,X71)) )),
  inference(resolution,[],[f128,f150])).
fof(f3746,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (sP0(X40,X41,X42,X43,X44) | p(s(bool,'const.bool.IN_2'(s(X40,sK5(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(X40,sK6(X44,X43,X42,X41,X40)),s(fun(X40,bool),X43)))) )),
  inference(resolution,[],[f128,f149])).
fof(f3745,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (sP0(X32,X33,X34,X35,X36) | p(s(bool,'const.bool.IN_2'(s(X32,sK5(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(X32,sK6(X36,X35,X34,X33,X32)),s(fun(X32,bool),X35)))) )),
  inference(resolution,[],[f128,f148])).
fof(f3744,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (sP0(X24,X25,X26,X27,X28) | p(s(bool,'const.bool.IN_2'(s(X24,sK5(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(X24,sK6(X28,X27,X26,X25,X24)),s(fun(X24,bool),X27))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f128,f146])).
fof(f3742,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (sP0(X8,X9,X10,X11,X12) | p(s(bool,'const.bool.IN_2'(s(X8,sK5(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(X8,sK6(X12,X11,X10,X9,X8)),s(fun(X8,bool),X11)))) )),
  inference(resolution,[],[f128,f144])).
fof(f3731,plain,(
  ( ! [X0,X3,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) | s(bool,X1) != s(bool,X3) | p(s(bool,X3))) ) | $spl70),
  inference(cnf_transformation,[],[f3731_D])).
fof(f3731_D,plain,(
  ( ! [X0,X3,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) | s(bool,X1) != s(bool,X3) | p(s(bool,X3))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3697,plain,(
  ( ! [X61,X59,X60,X58] : (s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,X59))) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X59) != s(bool,X61)) )),
  inference(resolution,[],[f3428,f206])).
fof(f3687,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.sum.sum'(X35,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X35,bool),'const.sum.INR_1'(s(bool,X36))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f3428,f150])).
fof(f3684,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.sum.sum'(X20,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X20,bool),'const.sum.INR_1'(s(bool,X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f3428,f149])).
fof(f3683,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.sum.sum'(X15,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X15,bool),'const.sum.INR_1'(s(bool,X16))) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f3428,f148])).
fof(f3682,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.sum.sum'(X10,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X10,bool),'const.sum.INR_1'(s(bool,X11))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f3428,f146])).
fof(f3680,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f3428,f144])).
fof(f3602,plain,(
  ( ! [X14,X17,X15,X13,X20,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(X13,happ(s(fun(X14,X13),X15),s(X14,X16))),s(fun(X13,bool),X17)))) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(X14,X16),s(fun(X14,bool),X18))) | ~p(s(bool,X20)) | ~sP0(X14,X13,X15,X18,X17)) ) | $spl68),
  inference(cnf_transformation,[],[f3602_D])).
fof(f3602_D,plain,(
  ( ! [X14,X17,X15,X13,X20,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(X13,happ(s(fun(X14,X13),X15),s(X14,X16))),s(fun(X13,bool),X17)))) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(X14,X16),s(fun(X14,bool),X18))) | ~p(s(bool,X20)) | ~sP0(X14,X13,X15,X18,X17)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3488,plain,(
  ( ! [X146,X152,X151,X149,X147,X145,X150,X148] : (p(s(bool,'const.bool.IN_2'(s(X145,happ(s(fun(X146,X145),X147),s(X146,X148))),s(fun(X145,bool),X149)))) | ~sP0(X146,X145,X147,X150,X149) | p(s(bool,X151)) | ~p(s(bool,X152)) | s(bool,X152) != s(bool,'const.bool.IN_2'(s(X146,X148),s(fun(X146,bool),X150)))) )),
  inference(resolution,[],[f126,f206])).
fof(f126,plain,(
  ( ! [X4,X2,X0,X10,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | ~sP0(X0,X1,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f87])).
fof(f3477,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X84,X82] : (p(s(bool,'const.bool.IN_2'(s(X76,happ(s(fun(X77,X76),X78),s(X77,X79))),s(fun(X76,bool),X80)))) | ~sP0(X77,X76,X78,X81,X80) | ~p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X83) != s(bool,'const.bool.IN_2'(s(X77,X79),s(fun(X77,bool),X81))) | s(bool,X82) != s(bool,X84)) )),
  inference(resolution,[],[f126,f150])).
fof(f3475,plain,(
  ( ! [X61,X59,X66,X64,X62,X60,X58,X65,X63] : (p(s(bool,'const.bool.IN_2'(s(X58,happ(s(fun(X59,X58),X60),s(X59,X61))),s(fun(X58,bool),X62)))) | ~sP0(X59,X58,X60,X63,X62) | p(s(bool,X64)) | ~p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.bool.IN_2'(s(X59,X61),s(fun(X59,bool),X63)))) )),
  inference(resolution,[],[f126,f149])).
fof(f3473,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(X40,happ(s(fun(X41,X40),X42),s(X41,X43))),s(fun(X40,bool),X44)))) | ~sP0(X41,X40,X42,X45,X44) | p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(X41,X43),s(fun(X41,bool),X45)))) )),
  inference(resolution,[],[f126,f148])).
fof(f3472,plain,(
  ( ! [X39,X37,X35,X33,X31,X38,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(X31,happ(s(fun(X32,X31),X33),s(X32,X34))),s(fun(X31,bool),X35)))) | ~sP0(X32,X31,X33,X36,X35) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(X32,X34),s(fun(X32,bool),X36))) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f126,f146])).
fof(f3470,plain,(
  ( ! [X14,X21,X19,X17,X15,X13,X20,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(X13,happ(s(fun(X14,X13),X15),s(X14,X16))),s(fun(X13,bool),X17)))) | ~sP0(X14,X13,X15,X18,X17) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(X14,X16),s(fun(X14,bool),X18)))) )),
  inference(resolution,[],[f126,f144])).
fof(f3601,plain,(
  ( ! [X14,X12,X10,X13,X11,X9,X16] : (~p(s(bool,'const.bool.IN_2'(s(X9,X10),s(fun(X9,bool),X11)))) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(X12,happ(s(fun(X9,X12),X13),s(X9,X10))),s(fun(X12,bool),X14))) | p(s(bool,X16)) | ~sP0(X9,X12,X13,X11,X14)) ) | $spl66),
  inference(cnf_transformation,[],[f3601_D])).
fof(f3601_D,plain,(
  ( ! [X14,X12,X10,X13,X11,X9,X16] : (~p(s(bool,'const.bool.IN_2'(s(X9,X10),s(fun(X9,bool),X11)))) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(X12,happ(s(fun(X9,X12),X13),s(X9,X10))),s(fun(X12,bool),X14))) | p(s(bool,X16)) | ~sP0(X9,X12,X13,X11,X14)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3466,plain,(
  ( ! [X142,X140,X138,X136,X135,X141,X139,X137] : (~p(s(bool,'const.bool.IN_2'(s(X135,X136),s(fun(X135,bool),X137)))) | ~sP0(X135,X138,X139,X137,X140) | p(s(bool,X141)) | p(s(bool,X142)) | s(bool,X142) != s(bool,'const.bool.IN_2'(s(X138,happ(s(fun(X135,X138),X139),s(X135,X136))),s(fun(X138,bool),X140)))) )),
  inference(resolution,[],[f126,f206])).
fof(f3456,plain,(
  ( ! [X80,X78,X76,X74,X72,X79,X77,X75,X73] : (~p(s(bool,'const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),X74)))) | ~sP0(X72,X75,X76,X74,X77) | p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,X78) != s(bool,'const.bool.IN_2'(s(X75,happ(s(fun(X72,X75),X76),s(X72,X73))),s(fun(X75,bool),X77))) | s(bool,X79) != s(bool,X80)) )),
  inference(resolution,[],[f126,f150])).
fof(f3453,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X53,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(X45,X46),s(fun(X45,bool),X47)))) | ~sP0(X45,X48,X49,X47,X50) | p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X53) != s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X45,X48),X49),s(X45,X46))),s(fun(X48,bool),X50)))) )),
  inference(resolution,[],[f126,f149])).
fof(f3452,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X44,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | ~sP0(X36,X39,X40,X38,X41) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X42) != s(bool,X44) | s(bool,X43) != s(bool,'const.bool.IN_2'(s(X39,happ(s(fun(X36,X39),X40),s(X36,X37))),s(fun(X39,bool),X41)))) )),
  inference(resolution,[],[f126,f148])).
fof(f3451,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X27,X28),s(fun(X27,bool),X29)))) | ~sP0(X27,X30,X31,X29,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.bool.IN_2'(s(X30,happ(s(fun(X27,X30),X31),s(X27,X28))),s(fun(X30,bool),X32))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f126,f146])).
fof(f3449,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (~p(s(bool,'const.bool.IN_2'(s(X9,X10),s(fun(X9,bool),X11)))) | ~sP0(X9,X12,X13,X11,X14) | p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(X12,happ(s(fun(X9,X12),X13),s(X9,X10))),s(fun(X12,bool),X14)))) )),
  inference(resolution,[],[f126,f144])).
fof(f3423,plain,(
  ( ! [X13] : (s('type.sum.sum'(X13,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X13,bool),'const.sum.INR_1'(s(bool,'F_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f3423_D])).
fof(f3423_D,plain,(
  ( ! [X13] : (s('type.sum.sum'(X13,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X13,bool),'const.sum.INR_1'(s(bool,'F_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3417,plain,(
  ( ! [X17,X16] : (s('type.sum.sum'(X17,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X17,bool),'const.sum.INR_1'(s(bool,'F_0'))) | ~p(s(bool,X16))) )),
  inference(trivial_inequality_removal,[],[f3399])).
fof(f3399,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X16) | s('type.sum.sum'(X17,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X17,bool),'const.sum.INR_1'(s(bool,'F_0'))) | ~p(s(bool,X16))) )),
  inference(superposition,[],[f1431,f196])).
fof(f3418,plain,(
  ( ! [X14,X15] : (s('type.sum.sum'(X15,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X15,bool),'const.sum.INR_1'(s(bool,'F_0'))) | ~p(s(bool,X14))) )),
  inference(trivial_inequality_removal,[],[f3398])).
fof(f3398,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X14) | s('type.sum.sum'(X15,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X15,bool),'const.sum.INR_1'(s(bool,'F_0'))) | ~p(s(bool,X14))) )),
  inference(superposition,[],[f1431,f198])).
fof(f3419,plain,(
  ( ! [X12,X13] : (s('type.sum.sum'(X13,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X13,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(trivial_inequality_removal,[],[f3397])).
fof(f3397,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X12) | s('type.sum.sum'(X13,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X13,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f1431,f205])).
fof(f335,plain,(
  ( ! [X14,X16] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X14),s(fun(sK2,bool),sK4)))) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X14))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl14),
  inference(cnf_transformation,[],[f335_D])).
fof(f335_D,plain,(
  ( ! [X14,X16] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X14),s(fun(sK2,bool),sK4)))) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X14))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3350,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X41),s(fun(sK2,bool),sK4)))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X41))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f209,f206])).
fof(f3340,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X28),s(fun(sK2,bool),sK4)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X28))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f209,f150])).
fof(f3337,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X16),s(fun(sK2,bool),sK4)))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X16))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f209,f149])).
fof(f3336,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X12),s(fun(sK2,bool),sK4)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X12))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f209,f148])).
fof(f3335,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X8),s(fun(sK2,bool),sK4)))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X8))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f209,f146])).
fof(f3333,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X0),s(fun(sK2,bool),sK4)))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X0))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f209,f144])).
fof(f3326,plain,(
  ( ! [X0,X3,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) | s(bool,X1) != s(bool,X3) | p(s(bool,X3))) ) | $spl62),
  inference(cnf_transformation,[],[f3326_D])).
fof(f3326_D,plain,(
  ( ! [X0,X3,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) | s(bool,X1) != s(bool,X3) | p(s(bool,X3))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3293,plain,(
  ( ! [X61,X59,X60,X58] : (s('type.sum.sum'(bool,X58),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X58),'const.sum.INL_1'(s(bool,X59))) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X59) != s(bool,X61)) )),
  inference(resolution,[],[f3149,f206])).
fof(f3283,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.sum.sum'(bool,X35),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X35),'const.sum.INL_1'(s(bool,X36))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f3149,f150])).
fof(f3280,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.sum.sum'(bool,X20),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X20),'const.sum.INL_1'(s(bool,X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f3149,f149])).
fof(f3279,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.sum.sum'(bool,X15),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X15),'const.sum.INL_1'(s(bool,X16))) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f3149,f148])).
fof(f3278,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.sum.sum'(bool,X10),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X10),'const.sum.INL_1'(s(bool,X11))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f3149,f146])).
fof(f3276,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f3149,f144])).
fof(f3219,plain,(
  ( ! [X6,X8,X7,X5,X11,X9] : (~sP0(X5,X6,X7,X8,X9) | s(bool,X11) != s(bool,'const.pred_set.INJ_3'(s(fun(X5,X6),X7),s(fun(X5,bool),X8),s(fun(X6,bool),X9))) | p(s(bool,X11))) ) | $spl60),
  inference(cnf_transformation,[],[f3219_D])).
fof(f3219_D,plain,(
  ( ! [X6,X8,X7,X5,X11,X9] : (~sP0(X5,X6,X7,X8,X9) | s(bool,X11) != s(bool,'const.pred_set.INJ_3'(s(fun(X5,X6),X7),s(fun(X5,bool),X8),s(fun(X6,bool),X9))) | p(s(bool,X11))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f3187,plain,(
  ( ! [X118,X116,X114,X120,X119,X117,X115] : (~sP0(X114,X115,X116,X117,X118) | p(s(bool,X119)) | p(s(bool,X120)) | s(bool,X120) != s(bool,'const.pred_set.INJ_3'(s(fun(X114,X115),X116),s(fun(X114,bool),X117),s(fun(X115,bool),X118)))) )),
  inference(resolution,[],[f137,f206])).
fof(f3177,plain,(
  ( ! [X61,X68,X66,X64,X62,X67,X65,X63] : (~sP0(X61,X62,X63,X64,X65) | p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X66) != s(bool,'const.pred_set.INJ_3'(s(fun(X61,X62),X63),s(fun(X61,bool),X64),s(fun(X62,bool),X65))) | s(bool,X67) != s(bool,X68)) )),
  inference(resolution,[],[f137,f150])).
fof(f3174,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (~sP0(X37,X38,X39,X40,X41) | p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X44) != s(bool,'const.pred_set.INJ_3'(s(fun(X37,X38),X39),s(fun(X37,bool),X40),s(fun(X38,bool),X41)))) )),
  inference(resolution,[],[f137,f149])).
fof(f3173,plain,(
  ( ! [X30,X35,X33,X31,X29,X36,X34,X32] : (~sP0(X29,X30,X31,X32,X33) | p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.pred_set.INJ_3'(s(fun(X29,X30),X31),s(fun(X29,bool),X32),s(fun(X30,bool),X33)))) )),
  inference(resolution,[],[f137,f148])).
fof(f3172,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (~sP0(X21,X22,X23,X24,X25) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.pred_set.INJ_3'(s(fun(X21,X22),X23),s(fun(X21,bool),X24),s(fun(X22,bool),X25))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f137,f146])).
fof(f3170,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X11,X9] : (~sP0(X5,X6,X7,X8,X9) | p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.pred_set.INJ_3'(s(fun(X5,X6),X7),s(fun(X5,bool),X8),s(fun(X6,bool),X9)))) )),
  inference(resolution,[],[f137,f144])).
fof(f3144,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))) | sP0(X0,X1,X2,X3,X4)) ) | $spl58),
  inference(cnf_transformation,[],[f3144_D])).
fof(f3144_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))) | sP0(X0,X1,X2,X3,X4)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f3111,plain,(
  ( ! [X118,X116,X114,X120,X119,X117,X115] : (sP0(X114,X115,X116,X117,X118) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X120) != s(bool,'const.pred_set.INJ_3'(s(fun(X114,X115),X116),s(fun(X114,bool),X117),s(fun(X115,bool),X118)))) )),
  inference(resolution,[],[f136,f206])).
fof(f3100,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (sP0(X56,X57,X58,X59,X60) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,'const.pred_set.INJ_3'(s(fun(X56,X57),X58),s(fun(X56,bool),X59),s(fun(X57,bool),X60))) | s(bool,X61) != s(bool,X63)) )),
  inference(resolution,[],[f136,f150])).
fof(f3098,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (sP0(X40,X41,X42,X43,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.pred_set.INJ_3'(s(fun(X40,X41),X42),s(fun(X40,bool),X43),s(fun(X41,bool),X44)))) )),
  inference(resolution,[],[f136,f149])).
fof(f3096,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (sP0(X24,X25,X26,X27,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.pred_set.INJ_3'(s(fun(X24,X25),X26),s(fun(X24,bool),X27),s(fun(X25,bool),X28)))) )),
  inference(resolution,[],[f136,f148])).
fof(f3095,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (sP0(X16,X17,X18,X19,X20) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.INJ_3'(s(fun(X16,X17),X18),s(fun(X16,bool),X19),s(fun(X17,bool),X20))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f136,f146])).
fof(f3093,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP0(X0,X1,X2,X3,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) )),
  inference(resolution,[],[f136,f144])).
fof(f1754,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3)) ) | $spl50),
  inference(cnf_transformation,[],[f1754_D])).
fof(f1754_D,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f3083,plain,(
  ( ! [X200,X201] : (s(bool,'T_0') = s(bool,X201) | ~p(s(bool,X200))) ) | ($spl39 | $spl54)),
  inference(subsumption_resolution,[],[f2943,f205])).
fof(f2943,plain,(
  ( ! [X200,X201] : (s(bool,'F_0') != s(bool,X201) | s(bool,'T_0') = s(bool,X201) | ~p(s(bool,X200))) ) | ($spl39 | $spl54)),
  inference(superposition,[],[f1583,f2637])).
fof(f3077,plain,(
  ( ! [X173,X171,X172] : (p(s(bool,X172)) | ~p(s(bool,X173)) | s(bool,X171) != s(bool,X173)) ) | $spl54),
  inference(subsumption_resolution,[],[f2931,f178])).
fof(f2931,plain,(
  ( ! [X173,X171,X172] : (p(s(bool,'F_0')) | p(s(bool,X172)) | ~p(s(bool,X173)) | s(bool,X171) != s(bool,X173)) ) | $spl54),
  inference(superposition,[],[f206,f2637])).
fof(f3076,plain,(
  ( ! [X167,X165,X166] : (p(s(bool,X167)) | ~p(s(bool,X165)) | p(s(bool,X166))) ) | $spl54),
  inference(subsumption_resolution,[],[f2929,f192])).
fof(f2929,plain,(
  ( ! [X167,X165,X166] : (s(bool,'F_0') != s(bool,X166) | p(s(bool,X167)) | ~p(s(bool,X165)) | p(s(bool,X166))) ) | $spl54),
  inference(superposition,[],[f206,f2637])).
fof(f3066,plain,(
  ( ! [X127,X125,X126,X124] : (~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,X127)) ) | $spl54),
  inference(subsumption_resolution,[],[f2900,f162])).
fof(f2900,plain,(
  ( ! [X127,X125,X126,X124] : (p(s(bool,'F_0')) | ~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,X127)) ) | $spl54),
  inference(superposition,[],[f150,f2637])).
fof(f3065,plain,(
  ( ! [X123,X121,X122,X120] : (p(s(bool,X121)) | ~p(s(bool,X122)) | ~p(s(bool,X120)) | s(bool,X122) != s(bool,X123)) ) | $spl54),
  inference(subsumption_resolution,[],[f2899,f192])).
fof(f2899,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'F_0') != s(bool,X121) | p(s(bool,X121)) | ~p(s(bool,X122)) | ~p(s(bool,X120)) | s(bool,X122) != s(bool,X123)) ) | $spl54),
  inference(superposition,[],[f150,f2637])).
fof(f3063,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,X93)) | ~p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X92) != s(bool,X94)) ) | $spl54),
  inference(subsumption_resolution,[],[f2892,f178])).
fof(f2892,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'F_0')) | p(s(bool,X93)) | ~p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X92) != s(bool,X94)) ) | $spl54),
  inference(superposition,[],[f149,f2637])).
fof(f3062,plain,(
  ( ! [X87,X85,X86,X84] : (p(s(bool,X86)) | ~p(s(bool,X84)) | ~p(s(bool,X87)) | p(s(bool,X85))) ) | $spl54),
  inference(subsumption_resolution,[],[f2890,f192])).
fof(f2890,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'F_0') != s(bool,X85) | p(s(bool,X86)) | ~p(s(bool,X84)) | ~p(s(bool,X87)) | p(s(bool,X85))) ) | $spl54),
  inference(superposition,[],[f149,f2637])).
fof(f3059,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X73) != s(bool,X75) | s(bool,X72) != s(bool,X74)) ) | $spl54),
  inference(subsumption_resolution,[],[f2887,f178])).
fof(f2887,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,'F_0')) | p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X73) != s(bool,X75) | s(bool,X72) != s(bool,X74)) ) | $spl54),
  inference(superposition,[],[f148,f2637])).
fof(f3057,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X58)) | ~p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X58) != s(bool,X59)) ) | $spl54),
  inference(subsumption_resolution,[],[f2883,f192])).
fof(f2883,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'F_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X58) != s(bool,X59)) ) | $spl54),
  inference(superposition,[],[f148,f2637])).
fof(f3056,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl54),
  inference(subsumption_resolution,[],[f2880,f162])).
fof(f2880,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,'F_0')) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl54),
  inference(superposition,[],[f146,f2637])).
fof(f3055,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X36)) | s(bool,X38) != s(bool,X39)) ) | $spl54),
  inference(subsumption_resolution,[],[f2878,f192])).
fof(f2878,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X36)) | s(bool,X38) != s(bool,X39)) ) | $spl54),
  inference(superposition,[],[f146,f2637])).
fof(f3051,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl54),
  inference(subsumption_resolution,[],[f2873,f178])).
fof(f2873,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'F_0')) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl54),
  inference(superposition,[],[f144,f2637])).
fof(f3049,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X6)) | ~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,X7)) ) | $spl54),
  inference(subsumption_resolution,[],[f2870,f192])).
fof(f2870,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,X7)) ) | $spl54),
  inference(superposition,[],[f144,f2637])).
fof(f3043,plain,(
  ( ! [X212,X211] : (s(bool,'T_0') != s(bool,X211) | p(s(bool,X212))) ) | $spl54),
  inference(subsumption_resolution,[],[f2842,f192])).
fof(f2842,plain,(
  ( ! [X212,X211] : (s(bool,'T_0') != s(bool,X211) | s(bool,'F_0') != s(bool,X212) | p(s(bool,X212))) ) | $spl54),
  inference(superposition,[],[f622,f2637])).
fof(f3040,plain,(
  ( ! [X204,X205] : (~p(s(bool,X204)) | p(s(bool,X205))) ) | $spl54),
  inference(subsumption_resolution,[],[f3039,f2637])).
fof(f3039,plain,(
  ( ! [X206,X204,X205] : (~p(s(bool,X204)) | p(s(bool,X205)) | s(bool,'F_0') != s(bool,X206)) ) | $spl54),
  inference(subsumption_resolution,[],[f2839,f191])).
fof(f2839,plain,(
  ( ! [X206,X204,X205] : (~p(s(bool,X204)) | p(s(bool,X205)) | p(s(bool,X206)) | s(bool,'F_0') != s(bool,X206)) ) | $spl54),
  inference(superposition,[],[f206,f2637])).
fof(f3038,plain,(
  ( ! [X196,X197,X195] : (s(bool,X195) != s(bool,X196) | p(s(bool,X197)) | ~p(s(bool,X196))) ) | $spl54),
  inference(subsumption_resolution,[],[f2836,f178])).
fof(f2836,plain,(
  ( ! [X196,X197,X195] : (s(bool,X195) != s(bool,X196) | p(s(bool,X197)) | ~p(s(bool,X196)) | p(s(bool,'F_0'))) ) | $spl54),
  inference(superposition,[],[f206,f2637])).
fof(f3031,plain,(
  ( ! [X161,X159,X162] : (~p(s(bool,X159)) | ~p(s(bool,X161)) | s(bool,X161) != s(bool,X162)) ) | $spl54),
  inference(subsumption_resolution,[],[f3030,f2637])).
fof(f3030,plain,(
  ( ! [X161,X159,X162,X160] : (~p(s(bool,X159)) | ~p(s(bool,X161)) | s(bool,'F_0') != s(bool,X160) | s(bool,X161) != s(bool,X162)) ) | $spl54),
  inference(subsumption_resolution,[],[f2808,f191])).
fof(f2808,plain,(
  ( ! [X161,X159,X162,X160] : (~p(s(bool,X159)) | p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,'F_0') != s(bool,X160) | s(bool,X161) != s(bool,X162)) ) | $spl54),
  inference(superposition,[],[f150,f2637])).
fof(f3029,plain,(
  ( ! [X146,X144,X145,X143] : (s(bool,X143) != s(bool,X144) | ~p(s(bool,X145)) | ~p(s(bool,X144)) | s(bool,X145) != s(bool,X146)) ) | $spl54),
  inference(subsumption_resolution,[],[f2804,f162])).
fof(f2804,plain,(
  ( ! [X146,X144,X145,X143] : (s(bool,X143) != s(bool,X144) | p(s(bool,'F_0')) | ~p(s(bool,X145)) | ~p(s(bool,X144)) | s(bool,X145) != s(bool,X146)) ) | $spl54),
  inference(superposition,[],[f150,f2637])).
fof(f3028,plain,(
  ( ! [X127,X128,X129] : (~p(s(bool,X127)) | p(s(bool,X128)) | ~p(s(bool,X129))) ) | $spl54),
  inference(subsumption_resolution,[],[f3027,f2637])).
fof(f3027,plain,(
  ( ! [X127,X130,X128,X129] : (~p(s(bool,X127)) | p(s(bool,X128)) | ~p(s(bool,X129)) | s(bool,'F_0') != s(bool,X130)) ) | $spl54),
  inference(subsumption_resolution,[],[f2800,f191])).
fof(f2800,plain,(
  ( ! [X127,X130,X128,X129] : (~p(s(bool,X127)) | p(s(bool,X128)) | ~p(s(bool,X129)) | p(s(bool,X130)) | s(bool,'F_0') != s(bool,X130)) ) | $spl54),
  inference(superposition,[],[f149,f2637])).
fof(f3026,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,X115) != s(bool,X116) | p(s(bool,X117)) | ~p(s(bool,X116)) | ~p(s(bool,X118))) ) | $spl54),
  inference(subsumption_resolution,[],[f2797,f178])).
fof(f2797,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,X115) != s(bool,X116) | p(s(bool,X117)) | ~p(s(bool,X116)) | ~p(s(bool,X118)) | p(s(bool,'F_0'))) ) | $spl54),
  inference(superposition,[],[f149,f2637])).
fof(f3025,plain,(
  ( ! [X107,X110,X108] : (~p(s(bool,X107)) | p(s(bool,X108)) | s(bool,X108) != s(bool,X110)) ) | $spl54),
  inference(subsumption_resolution,[],[f3024,f2637])).
fof(f3024,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,X107)) | p(s(bool,X108)) | s(bool,X108) != s(bool,X110) | s(bool,'F_0') != s(bool,X109)) ) | $spl54),
  inference(subsumption_resolution,[],[f2795,f191])).
fof(f2795,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,X107)) | p(s(bool,X108)) | p(s(bool,X109)) | s(bool,X108) != s(bool,X110) | s(bool,'F_0') != s(bool,X109)) ) | $spl54),
  inference(superposition,[],[f148,f2637])).
fof(f3022,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,X87) != s(bool,X88) | p(s(bool,X89)) | ~p(s(bool,X88)) | s(bool,X89) != s(bool,X90)) ) | $spl54),
  inference(subsumption_resolution,[],[f2790,f178])).
fof(f2790,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,X87) != s(bool,X88) | p(s(bool,X89)) | ~p(s(bool,X88)) | p(s(bool,'F_0')) | s(bool,X89) != s(bool,X90)) ) | $spl54),
  inference(superposition,[],[f148,f2637])).
fof(f3019,plain,(
  ( ! [X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,X82)) ) | $spl54),
  inference(subsumption_resolution,[],[f3018,f2637])).
fof(f3018,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X81)) | s(bool,'F_0') != s(bool,X80) | s(bool,X81) != s(bool,X82)) ) | $spl54),
  inference(subsumption_resolution,[],[f2788,f191])).
fof(f2788,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,'F_0') != s(bool,X80) | s(bool,X81) != s(bool,X82)) ) | $spl54),
  inference(superposition,[],[f146,f2637])).
fof(f3017,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X67) != s(bool,X68) | ~p(s(bool,X69)) | ~p(s(bool,X68)) | s(bool,X69) != s(bool,X70)) ) | $spl54),
  inference(subsumption_resolution,[],[f2785,f162])).
fof(f2785,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X67) != s(bool,X68) | p(s(bool,'F_0')) | ~p(s(bool,X69)) | ~p(s(bool,X68)) | s(bool,X69) != s(bool,X70)) ) | $spl54),
  inference(superposition,[],[f146,f2637])).
fof(f3016,plain,(
  ( ! [X54,X52,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,X54)) ) | $spl54),
  inference(subsumption_resolution,[],[f3015,f2637])).
fof(f3015,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,X54) | s(bool,'F_0') != s(bool,X53)) ) | $spl54),
  inference(subsumption_resolution,[],[f2781,f191])).
fof(f2781,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X52) != s(bool,X54) | s(bool,'F_0') != s(bool,X53)) ) | $spl54),
  inference(superposition,[],[f144,f2637])).
fof(f3011,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X27) != s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X28)) | s(bool,X29) != s(bool,X30)) ) | $spl54),
  inference(subsumption_resolution,[],[f2775,f178])).
fof(f2775,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X27) != s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X28)) | p(s(bool,'F_0')) | s(bool,X29) != s(bool,X30)) ) | $spl54),
  inference(superposition,[],[f144,f2637])).
fof(f2765,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8))) ) | $spl54),
  inference(superposition,[],[f186,f2637])).
fof(f2764,plain,(
  ( ! [X6,X5] : (~p(s(bool,X5)) | p(s(bool,X6))) ) | $spl54),
  inference(superposition,[],[f178,f2637])).
fof(f3010,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | p(s(bool,X4))) ) | $spl54),
  inference(subsumption_resolution,[],[f2763,f207])).
fof(f2763,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,'T_0'))) ) | $spl54),
  inference(superposition,[],[f169,f2637])).
fof(f2761,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl54),
  inference(superposition,[],[f162,f2637])).
fof(f2715,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f2668,f272])).
fof(f2668,plain,(
  ( ! [X46] : (p(s(bool,X46))) ) | $spl10),
  inference(resolution,[],[f272,f173])).
fof(f2714,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f2667,f272])).
fof(f2667,plain,(
  ( ! [X45] : (p(s(bool,X45))) ) | $spl10),
  inference(resolution,[],[f272,f170])).
fof(f2713,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f2665,f272])).
fof(f2665,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f272,f167])).
fof(f167,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f103])).
fof(f2712,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f2663,f272])).
fof(f2663,plain,(
  ( ! [X42] : (p(s(bool,X42))) ) | $spl10),
  inference(resolution,[],[f272,f164])).
fof(f2671,plain,(
  $false | $spl10),
  inference(resolution,[],[f272,f200])).
fof(f200,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INJ_INR','thm.bool.TRUTH')).
fof(f2672,plain,(
  $false | $spl10),
  inference(resolution,[],[f272,f201])).
fof(f2673,plain,(
  $false | $spl10),
  inference(resolution,[],[f272,f207])).
fof(f2638,plain,(
  ( ! [X13] : (s('type.sum.sum'(bool,X13),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X13),'const.sum.INL_1'(s(bool,'F_0')))) ) | $spl56),
  inference(cnf_transformation,[],[f2638_D])).
fof(f2638_D,plain,(
  ( ! [X13] : (s('type.sum.sum'(bool,X13),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X13),'const.sum.INL_1'(s(bool,'F_0')))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2631,plain,(
  ( ! [X17,X16] : (s('type.sum.sum'(bool,X17),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X17),'const.sum.INL_1'(s(bool,'F_0'))) | ~p(s(bool,X16))) )),
  inference(trivial_inequality_removal,[],[f2609])).
fof(f2609,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X16) | s('type.sum.sum'(bool,X17),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X17),'const.sum.INL_1'(s(bool,'F_0'))) | ~p(s(bool,X16))) )),
  inference(superposition,[],[f1429,f196])).
fof(f2632,plain,(
  ( ! [X14,X15] : (s('type.sum.sum'(bool,X15),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X15),'const.sum.INL_1'(s(bool,'F_0'))) | ~p(s(bool,X14))) )),
  inference(trivial_inequality_removal,[],[f2608])).
fof(f2608,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X14) | s('type.sum.sum'(bool,X15),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X15),'const.sum.INL_1'(s(bool,'F_0'))) | ~p(s(bool,X14))) )),
  inference(superposition,[],[f1429,f198])).
fof(f2633,plain,(
  ( ! [X12,X13] : (s('type.sum.sum'(bool,X13),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X13),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(trivial_inequality_removal,[],[f2607])).
fof(f2607,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X12) | s('type.sum.sum'(bool,X13),'const.sum.INL_1'(s(bool,'T_0'))) != s('type.sum.sum'(bool,X13),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f1429,f205])).
fof(f1755,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5)) ) | $spl52),
  inference(cnf_transformation,[],[f1755_D])).
fof(f1755_D,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1750,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,X16) != s(bool,X17) | ~p(s(bool,X17)) | p(s(bool,X15))) ) | $spl34),
  inference(duplicate_literal_removal,[],[f1712])).
fof(f1712,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,X16) != s(bool,X17) | ~p(s(bool,X17)) | p(s(bool,X15)) | p(s(bool,X15))) ) | $spl34),
  inference(superposition,[],[f671,f192])).
fof(f1751,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X13) != s(bool,X14) | ~p(s(bool,X14)) | p(s(bool,X12))) ) | $spl34),
  inference(duplicate_literal_removal,[],[f1711])).
fof(f1711,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X13) != s(bool,X14) | ~p(s(bool,X14)) | p(s(bool,X12)) | p(s(bool,X12))) ) | $spl34),
  inference(superposition,[],[f671,f194])).
fof(f1753,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5) | ~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X3)) ) | $spl34),
  inference(subsumption_resolution,[],[f1708,f196])).
fof(f1708,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5) | ~p(s(bool,X5)) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X3)) ) | $spl34),
  inference(superposition,[],[f671,f205])).
fof(f1691,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f1644])).
fof(f1644,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl4),
  inference(superposition,[],[f269,f225])).
fof(f1690,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f269])).
fof(f1485,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl48),
  inference(cnf_transformation,[],[f1485_D])).
fof(f1485_D,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1457,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl24),
  inference(resolution,[],[f418,f206])).
fof(f1447,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X33) != s(bool,X34)) ) | $spl24),
  inference(resolution,[],[f418,f150])).
fof(f1444,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl24),
  inference(resolution,[],[f418,f149])).
fof(f1443,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl24),
  inference(resolution,[],[f418,f148])).
fof(f1442,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X18) != s(bool,X19)) ) | $spl24),
  inference(resolution,[],[f418,f146])).
fof(f1440,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl24),
  inference(resolution,[],[f418,f144])).
fof(f1277,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl44),
  inference(cnf_transformation,[],[f1277_D])).
fof(f1277_D,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1356,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl20),
  inference(resolution,[],[f413,f206])).
fof(f1346,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X33) != s(bool,X34)) ) | $spl20),
  inference(resolution,[],[f413,f150])).
fof(f1343,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl20),
  inference(resolution,[],[f413,f149])).
fof(f1342,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl20),
  inference(resolution,[],[f413,f148])).
fof(f1341,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X18) != s(bool,X19)) ) | $spl20),
  inference(resolution,[],[f413,f146])).
fof(f1339,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl20),
  inference(resolution,[],[f413,f144])).
fof(f1332,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl46),
  inference(cnf_transformation,[],[f1332_D])).
fof(f1332_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1301,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl21),
  inference(resolution,[],[f412,f206])).
fof(f412,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) | $spl21),
  inference(cnf_transformation,[],[f412_D])).
fof(f412_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1290,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X21) != s(bool,X23)) ) | $spl21),
  inference(resolution,[],[f412,f150])).
fof(f1288,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl21),
  inference(resolution,[],[f412,f149])).
fof(f1286,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl21),
  inference(resolution,[],[f412,f148])).
fof(f1285,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X6) != s(bool,X8)) ) | $spl21),
  inference(resolution,[],[f412,f146])).
fof(f1283,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl21),
  inference(resolution,[],[f412,f144])).
fof(f1249,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl20),
  inference(resolution,[],[f413,f206])).
fof(f1239,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X33) != s(bool,X34)) ) | $spl20),
  inference(resolution,[],[f413,f150])).
fof(f1236,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl20),
  inference(resolution,[],[f413,f149])).
fof(f1235,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl20),
  inference(resolution,[],[f413,f148])).
fof(f1234,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X18) != s(bool,X19)) ) | $spl20),
  inference(resolution,[],[f413,f146])).
fof(f1232,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl20),
  inference(resolution,[],[f413,f144])).
fof(f1113,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl40),
  inference(cnf_transformation,[],[f1113_D])).
fof(f1113_D,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1198,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl18),
  inference(resolution,[],[f411,f206])).
fof(f1188,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X33) != s(bool,X34)) ) | $spl18),
  inference(resolution,[],[f411,f150])).
fof(f1185,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl18),
  inference(resolution,[],[f411,f149])).
fof(f1184,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl18),
  inference(resolution,[],[f411,f148])).
fof(f1183,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X18) != s(bool,X19)) ) | $spl18),
  inference(resolution,[],[f411,f146])).
fof(f1181,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl18),
  inference(resolution,[],[f411,f144])).
fof(f1139,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl19),
  inference(resolution,[],[f410,f206])).
fof(f1128,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X21) != s(bool,X23)) ) | $spl19),
  inference(resolution,[],[f410,f150])).
fof(f1126,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl19),
  inference(resolution,[],[f410,f149])).
fof(f1124,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl19),
  inference(resolution,[],[f410,f148])).
fof(f1123,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X6) != s(bool,X8)) ) | $spl19),
  inference(resolution,[],[f410,f146])).
fof(f1121,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl19),
  inference(resolution,[],[f410,f144])).
fof(f1085,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl18),
  inference(resolution,[],[f411,f206])).
fof(f1075,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X33) != s(bool,X34)) ) | $spl18),
  inference(resolution,[],[f411,f150])).
fof(f1072,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl18),
  inference(resolution,[],[f411,f149])).
fof(f1071,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl18),
  inference(resolution,[],[f411,f148])).
fof(f1070,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))) | s(bool,X18) != s(bool,X19)) ) | $spl18),
  inference(resolution,[],[f411,f146])).
fof(f1068,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) ) | $spl18),
  inference(resolution,[],[f411,f144])).
fof(f1044,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl39),
  inference(cnf_transformation,[],[f1044_D])).
fof(f1044_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f1046,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X10) != s(bool,X11) | p(s(bool,X10))) )),
  inference(subsumption_resolution,[],[f1019,f230])).
fof(f1019,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X10) != s(bool,X11) | p(s(bool,X11)) | p(s(bool,X10))) )),
  inference(superposition,[],[f622,f192])).
fof(f1045,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X8) != s(bool,X9) | p(s(bool,X8))) )),
  inference(subsumption_resolution,[],[f1018,f230])).
fof(f1018,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X8) != s(bool,X9) | p(s(bool,X9)) | p(s(bool,X8))) )),
  inference(superposition,[],[f622,f194])).
fof(f1042,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X2) != s(bool,X3) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f1015,f622])).
fof(f1015,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f622,f205])).
fof(f782,plain,(
  ( ! [X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X1)) ) | $spl36),
  inference(cnf_transformation,[],[f782_D])).
fof(f782_D,plain,(
  ( ! [X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f743,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,X58) != s(bool,X59) | s(bool,'F_0') = s(bool,X59) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X58) != s(bool,X61)) )),
  inference(resolution,[],[f620,f206])).
fof(f733,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f620,f150])).
fof(f730,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X20) != s(bool,X24)) )),
  inference(resolution,[],[f620,f149])).
fof(f729,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) )),
  inference(resolution,[],[f620,f148])).
fof(f728,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f620,f146])).
fof(f726,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) )),
  inference(resolution,[],[f620,f144])).
fof(f719,plain,(
  $false | $spl32),
  inference(equality_resolution,[],[f493])).
fof(f671,plain,(
  ( ! [X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | p(s(bool,X0))) ) | $spl34),
  inference(cnf_transformation,[],[f671_D])).
fof(f671_D,plain,(
  ( ! [X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | p(s(bool,X0))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f624,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X60)) | s(bool,X60) != s(bool,X61) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X61) != s(bool,X63)) )),
  inference(resolution,[],[f230,f206])).
fof(f613,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,X35)) | s(bool,X35) != s(bool,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f230,f150])).
fof(f611,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,X25)) | s(bool,X25) != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) )),
  inference(resolution,[],[f230,f149])).
fof(f609,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,X15) != s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f230,f148])).
fof(f608,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X10)) | s(bool,X10) != s(bool,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f230,f146])).
fof(f606,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f230,f144])).
fof(f597,plain,(
  $false | $spl30),
  inference(equality_resolution,[],[f486])).
fof(f557,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,'T_0') != s(bool,X23)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f548])).
fof(f548,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X22))) ) | $spl8),
  inference(superposition,[],[f271,f196])).
fof(f558,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') != s(bool,X21)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f547])).
fof(f547,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') != s(bool,X21) | ~p(s(bool,X20))) ) | $spl8),
  inference(superposition,[],[f271,f198])).
fof(f564,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19)) ) | $spl8),
  inference(subsumption_resolution,[],[f546,f271])).
fof(f546,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') = s(bool,X18)) ) | $spl8),
  inference(superposition,[],[f271,f205])).
fof(f486,plain,(
  ( ! [X65] : (s(bool,'F_0') != s(bool,X65)) ) | $spl30),
  inference(cnf_transformation,[],[f486_D])).
fof(f486_D,plain,(
  ( ! [X65] : (s(bool,'F_0') != s(bool,X65)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f560,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1)) ) | $spl8),
  inference(subsumption_resolution,[],[f537,f271])).
fof(f537,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f271,f205])).
fof(f521,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') != s(bool,X81) | ~p(s(bool,X82))) ) | $spl8),
  inference(resolution,[],[f271,f189])).
fof(f520,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') != s(bool,X79) | ~p(s(bool,X80))) ) | $spl8),
  inference(resolution,[],[f271,f181])).
fof(f519,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,X77) | p(s(bool,X78))) ) | $spl8),
  inference(resolution,[],[f271,f173])).
fof(f518,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') != s(bool,X75) | p(s(bool,X76))) ) | $spl8),
  inference(resolution,[],[f271,f170])).
fof(f517,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') != s(bool,X73) | ~p(s(bool,X74))) ) | $spl8),
  inference(resolution,[],[f271,f168])).
fof(f515,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,X70) | ~p(s(bool,X71))) ) | $spl8),
  inference(resolution,[],[f271,f165])).
fof(f514,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') != s(bool,X68) | p(s(bool,X69))) ) | $spl8),
  inference(resolution,[],[f271,f164])).
fof(f492,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19)) ) | $spl2),
  inference(subsumption_resolution,[],[f472,f268])).
fof(f472,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') = s(bool,X18)) ) | $spl2),
  inference(superposition,[],[f268,f205])).
fof(f483,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f465])).
fof(f465,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X4))) ) | $spl2),
  inference(superposition,[],[f268,f192])).
fof(f484,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f464])).
fof(f464,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3) | p(s(bool,X2))) ) | $spl2),
  inference(superposition,[],[f268,f194])).
fof(f488,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f463,f268])).
fof(f463,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(superposition,[],[f268,f205])).
fof(f453,plain,(
  ( ! [X72,X71] : (s(bool,'F_0') != s(bool,X71) | ~p(s(bool,X72))) ) | $spl2),
  inference(resolution,[],[f268,f162])).
fof(f487,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,X69) | p(s(bool,X70))) ) | $spl2),
  inference(subsumption_resolution,[],[f452,f207])).
fof(f452,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,'T_0'))) ) | $spl2),
  inference(resolution,[],[f268,f169])).
fof(f451,plain,(
  ( ! [X68,X67] : (s(bool,'F_0') != s(bool,X67) | p(s(bool,X68))) ) | $spl2),
  inference(resolution,[],[f268,f178])).
fof(f450,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X66))) ) | $spl2),
  inference(resolution,[],[f268,f186])).
fof(f428,plain,(
  s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)) != s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)) | $spl29),
  inference(cnf_transformation,[],[f428_D])).
fof(f428_D,plain,(
  s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)) != s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f429,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)) != s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2))),
  inference(forward_demodulation,[],[f408,f138])).
fof(f408,plain,(
  s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)) != s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)) | ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))),
  inference(resolution,[],[f210,f135])).
fof(f407,plain,(
  s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)) != s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)) | p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))),
  inference(resolution,[],[f210,f134])).
fof(f426,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2))))),
  inference(forward_demodulation,[],[f425,f138])).
fof(f425,plain,(
  s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))),
  inference(forward_demodulation,[],[f424,f138])).
fof(f424,plain,(
  s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))),
  inference(forward_demodulation,[],[f406,f138])).
fof(f406,plain,(
  s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))),
  inference(resolution,[],[f210,f133])).
fof(f421,plain,(
  s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))),
  inference(forward_demodulation,[],[f420,f138])).
fof(f420,plain,(
  s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))),
  inference(forward_demodulation,[],[f405,f138])).
fof(f405,plain,(
  s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) = s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))),
  inference(resolution,[],[f210,f132])).
fof(f418,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) | $spl24),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f419,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))),
  inference(forward_demodulation,[],[f404,f138])).
fof(f404,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) | ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))),
  inference(resolution,[],[f210,f131])).
fof(f403,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK7(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))),
  inference(resolution,[],[f210,f130])).
fof(f413,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) | $spl20),
  inference(cnf_transformation,[],[f413_D])).
fof(f413_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f414,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))),
  inference(forward_demodulation,[],[f402,f138])).
fof(f402,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) | ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),happ(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))))),
  inference(resolution,[],[f210,f129])).
fof(f401,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,sK6(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(sK2,sK5(sK3,sK4,'const.sum.INR_0','type.sum.sum'(sK1,sK2),sK2)),s(fun(sK2,bool),sK4))))),
  inference(resolution,[],[f210,f128])).
fof(f398,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f365,f273])).
fof(f365,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) | $spl12),
  inference(resolution,[],[f273,f162])).
fof(f361,plain,(
  $false | $spl12),
  inference(resolution,[],[f273,f199])).
fof(f374,plain,(
  $false | $spl12),
  inference(resolution,[],[f273,f125])).
fof(f298,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X55),s(fun(sK2,bool),sK4)))) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X57) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X55))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f209,f206])).
fof(f288,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X42),s(fun(sK2,bool),sK4)))) | p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X42))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X44) != s(bool,X45)) )),
  inference(resolution,[],[f209,f150])).
fof(f337,plain,(
  ( ! [X21,X18] : (s(bool,X21) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X18))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl16),
  inference(cnf_transformation,[],[f337_D])).
fof(f337_D,plain,(
  ( ! [X21,X18] : (s(bool,X21) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X18))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f270,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) | $spl6),
  inference(cnf_transformation,[],[f270_D])).
fof(f270_D,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f338,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X40) | s(bool,X41) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X38))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f287,f206])).
fof(f287,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X38),s(fun(sK2,bool),sK4)))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X40) | s(bool,X41) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X38))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f209,f150])).
fof(f285,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X30),s(fun(sK2,bool),sK4)))) | p(s(bool,X31)) | ~p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X33) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X30))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f209,f149])).
fof(f284,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X26),s(fun(sK2,bool),sK4)))) | p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X26))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f209,f148])).
fof(f283,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X22),s(fun(sK2,bool),sK4)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X22))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X24) != s(bool,X25)) )),
  inference(resolution,[],[f209,f146])).
fof(f336,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X18))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f282,f206])).
fof(f282,plain,(
  ( ! [X21,X19,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X18),s(fun(sK2,bool),sK4)))) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X18))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f209,f146])).
fof(f281,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.IN_2'(s(sK2,X14),s(fun(sK2,bool),sK4)))) | p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s('type.sum.sum'(sK1,sK2),'const.sum.INR_1'(s(sK2,X14))),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f209,f144])).
fof(f267,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl0),
  inference(cnf_transformation,[],[f267_D])).
fof(f267_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f229,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f125,f206])).
fof(f218,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f125,f150])).
fof(f216,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f125,f149])).
fof(f269,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) | $spl4),
  inference(cnf_transformation,[],[f269_D])).
fof(f269_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f215,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f125,f148])).
fof(f214,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f125,f148])).
fof(f213,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f125,f146])).
fof(f212,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f125,f144])).
fof(f211,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.INJ_3'(s(fun(sK2,'type.sum.sum'(sK1,sK2)),'const.sum.INR_0'),s(fun(sK2,bool),sK4),s(fun('type.sum.sum'(sK1,sK2),bool),sK3)))) )),
  inference(resolution,[],[f125,f144])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 257
% Passive clauses: 6355
% Generated clauses: 15082
% Final active clauses: 207
% Final passive clauses: 5158
% Input formulas: 29
% Initial clauses: 82
% 
% Selected by SInE selection: 29
% SInE iterations: 1
% 
% Duplicate literals: 138
% Trivial inequalities: 511
% Fw subsumption resolutions: 1183
% Fw demodulations: 1168
% Bw demodulations: 28
% 
% Simple tautologies: 302
% Equational tautologies: 3
% Forward subsumptions: 5022
% 
% Binary resolution: 2591
% Factoring: 1
% Forward superposition: 6696
% Backward superposition: 2462
% Self superposition: 33
% Equality resolution: 84
% 
% Split clauses: 651
% Split components: 1400
% SAT solver clauses: 681
% SAT solver unit clauses: 14
% SAT solver binary clauses: 396
% 
% Sat splits: 651
% Sat splitting refutations: 21
% 
% Memory used [KB]: 8571
% Time elapsed: 0.595 s
% ------------------------------
% ------------------------------
% Success in time 12.36 s
