Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.relation
% (82953)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 108
% Passive clauses: 1632
% Generated clauses: 23504
% Final active clauses: 97
% Final passive clauses: 1356
% Input formulas: 7
% Initial clauses: 10
% 
% Duplicate literals: 1329
% Trivial inequalities: 364
% Fw subsumption resolutions: 3291
% Fw demodulations: 2983
% Bw demodulations: 9
% 
% Simple tautologies: 410
% Equational tautologies: 66
% Forward subsumptions: 10497
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 48
% Forward superposition: 8900
% Backward superposition: 5858
% Self superposition: 282
% Equality factoring: 200
% Equality resolution: 20
% 
% Split clauses: 1131
% Split components: 2292
% SAT solver clauses: 1873
% SAT solver unit clauses: 26
% SAT solver binary clauses: 222
% 
% Sat splits: 1131
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1857
% TWLsolver calls for satisfiability: 132
% 
% Memory used [KB]: 10490
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f3317,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3231,f31,f39,f831,f832,f32,f47,f65,f67,f28,f45,f33,f49,f34,f50,f68,f70,f29,f43,f71,f73,f30,f41,f129,f150,f35,f51,f37,f53])).
fof(f53,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X3),s(X1,X5))),s(X1,X2)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2),s(fun(X1,X0),X6)))) | s(X0,happ(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X4),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))) ) | $spl18),
  inference(cnf_transformation,[],[f53_D])).
fof(f53_D,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X3),s(X1,X5))),s(X1,X2)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2),s(fun(X1,X0),X6)))) | s(X0,happ(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X4),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f37,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X4),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X3),s(X1,X5))),s(X1,X2)))) | ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2),s(fun(X1,X0),X6))))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3,X4] : (! [X6] : ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2),s(fun(X1,X0),X6)))) | ! [X5] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X3),s(X1,X5))),s(X1,X2)))) | s(X0,happ(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X4),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(rectify,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2,X3,X4] : (! [X5] : ~p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2),s(fun(X1,X0),X5)))) | ! [X6] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X3),s(X1,X6))),s(X1,X2)))) | s(X0,happ(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X4),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(fun(X1,fun(X1,bool)),X3),s(X1,X6))))),s(X1,X6)))))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1,X2,X3,X4] : (? [X5] : p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2),s(fun(X1,X0),X5)))) => ! [X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X3),s(X1,X6))),s(X1,X2)))) => s(X0,happ(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X4),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(fun(X1,fun(X1,bool)),X3),s(X1,X6))))),s(X1,X6)))))),
  inference(flattening,[],[f16])).
fof(f16,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (? [X5] : p(s(bool,'const.relation.approx_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2),s(fun(X1,X0),X5)))) => ! [X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X3),s(X1,X6))),s(X1,X2)))) => s(X0,happ(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X4),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X4),s(X1,X2))),s(fun(X1,fun(X1,bool)),X3),s(X1,X6))))),s(X1,X6)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X10,X8,X9] : (? [X12] : p(s(bool,'const.relation.approx_4'(s(fun(X7,fun(X7,bool)),X8),s(fun(fun(X7,X6),fun(X7,X6)),X9),s(X7,X10),s(fun(X7,X6),X12)))) => ! [X11] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X11))),s(X7,X10)))) => s(X6,happ(s(fun(X7,X6),'const.relation.the_fun_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(fun(X7,X6),fun(X7,X6)),X9),s(X7,X10))),s(X7,X11))) = s(X6,happ(s(fun(X7,X6),happ(s(fun(fun(X7,X6),fun(X7,X6)),X9),s(fun(X7,X6),'const.relation.RESTRICT_3'(s(fun(X7,X6),'const.relation.the_fun_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(fun(X7,X6),fun(X7,X6)),X9),s(X7,X10))),s(fun(X7,fun(X7,bool)),X8),s(X7,X11))))),s(X7,X11)))))),
  file('thm.relation.relation_|49|37_prove_ltb','thm.relation.relation_|49|3|49|_prove_ltb')).
fof(f51,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4),s(fun(X0,X1),sK6(X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2))))) ) | $spl14),
  inference(cnf_transformation,[],[f51_D])).
fof(f51_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4),s(fun(X0,X1),sK6(X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f35,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4),s(fun(X0,X1),sK6(X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2))))) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) | ! [X4] : p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4),s(fun(X0,X1),sK6(X4,X3,X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) | ! [X4] : ? [X5] : p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4),s(fun(X0,X1),X5)))))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2))))) | ! [X4] : ? [X5] : p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4),s(fun(X0,X1),X5)))))),
  inference(ennf_transformation,[],[f13])).
fof(f13,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2))))) => ! [X4] : ? [X5] : p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4),s(fun(X0,X1),X5)))))),
  inference(flattening,[],[f12])).
fof(f12,plain,(
  ! [X0,X1] : ! [X2,X3] : ((p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2))))) => ! [X4] : ? [X5] : p(s(bool,'const.relation.approx_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(fun(X0,X1),fun(X0,X1)),X3),s(X0,X4),s(fun(X0,X1),X5)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7,X6] : ! [X8,X9] : ((p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),X8)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X7,fun(X7,bool)),X8))))) => ! [X10] : ? [X13] : p(s(bool,'const.relation.approx_4'(s(fun(X7,fun(X7,bool)),X8),s(fun(fun(X7,X6),fun(X7,X6)),X9),s(X7,X10),s(fun(X7,X6),X13)))))),
  file('thm.relation.relation_|49|37_prove_ltb','thm.relation.relation_|49|36_prove_ltb')).
fof(f150,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X4,happ(s(fun(X0,X4),'const.relation.the_fun_3'(s(fun(X0,fun(X0,bool)),X1),s(fun(fun(X0,X4),fun(X0,X4)),X5),s(X0,X3))),s(X0,X2))) = s(X4,happ(s(fun(X0,X4),happ(s(fun(fun(X0,X4),fun(X0,X4)),X5),s(fun(X0,X4),'const.relation.RESTRICT_3'(s(fun(X0,X4),'const.relation.the_fun_3'(s(fun(X0,fun(X0,bool)),X1),s(fun(fun(X0,X4),fun(X0,X4)),X5),s(X0,X3))),s(fun(X0,fun(X0,bool)),X1),s(X0,X2))))),s(X0,X2))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl54),
  inference(cnf_transformation,[],[f150_D])).
fof(f150_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X4,happ(s(fun(X0,X4),'const.relation.the_fun_3'(s(fun(X0,fun(X0,bool)),X1),s(fun(fun(X0,X4),fun(X0,X4)),X5),s(X0,X3))),s(X0,X2))) = s(X4,happ(s(fun(X0,X4),happ(s(fun(fun(X0,X4),fun(X0,X4)),X5),s(fun(X0,X4),'const.relation.RESTRICT_3'(s(fun(X0,X4),'const.relation.the_fun_3'(s(fun(X0,fun(X0,bool)),X1),s(fun(fun(X0,X4),fun(X0,X4)),X5),s(X0,X3))),s(fun(X0,fun(X0,bool)),X1),s(X0,X2))))),s(X0,X2))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X1))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f129,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(X4,happ(s(fun(X0,X4),'const.relation.the_fun_3'(s(fun(X0,fun(X0,bool)),X1),s(fun(fun(X0,X4),fun(X0,X4)),X5),s(X0,X3))),s(X0,X2))) = s(X4,happ(s(fun(X0,X4),happ(s(fun(fun(X0,X4),fun(X0,X4)),X5),s(fun(X0,X4),'const.relation.RESTRICT_3'(s(fun(X0,X4),'const.relation.the_fun_3'(s(fun(X0,fun(X0,bool)),X1),s(fun(fun(X0,X4),fun(X0,X4)),X5),s(X0,X3))),s(fun(X0,fun(X0,bool)),X1),s(X0,X2))))),s(X0,X2))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | ($spl14 | $spl18)),
  inference(resolution,[],[f53,f51])).
fof(f41,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK4)))) | $spl2),
  inference(cnf_transformation,[],[f41_D])).
fof(f41_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK4)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f30,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK4))))),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) & p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK4)))) & s(sK0,happ(s(fun(sK1,sK0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(sK1,sK4))),s(sK1,sK5))) != s(sK0,happ(s(fun(sK1,sK0),happ(s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(fun(sK1,sK0),'const.relation.RESTRICT_3'(s(fun(sK1,sK0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f19])).
fof(f19,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X4)))) & s(X0,happ(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))))),s(X1,X5))))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X1,fun(X1,bool)),X2))))) & (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X4)))) & s(X0,happ(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))))),s(X1,X5)))))),
  inference(ennf_transformation,[],[f10])).
fof(f10,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X1,fun(X1,bool)),X2))))) => (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X4)))) => s(X0,happ(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))))),s(X1,X5)))))),
  inference(flattening,[],[f9])).
fof(f9,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : ((p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X1,fun(X1,bool)),X2))))) => (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X4)))) => s(X0,happ(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X3),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),'const.relation.the_fun_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(fun(X1,X0),fun(X1,X0)),X3),s(X1,X4))),s(fun(X1,fun(X1,bool)),X2),s(X1,X5))))),s(X1,X5)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : ((p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),X8)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X7,fun(X7,bool)),X8))))) => (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X11))),s(X7,X10)))) => s(X6,happ(s(fun(X7,X6),'const.relation.the_fun_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(fun(X7,X6),fun(X7,X6)),X9),s(X7,X10))),s(X7,X11))) = s(X6,happ(s(fun(X7,X6),happ(s(fun(fun(X7,X6),fun(X7,X6)),X9),s(fun(X7,X6),'const.relation.RESTRICT_3'(s(fun(X7,X6),'const.relation.the_fun_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(fun(X7,X6),fun(X7,X6)),X9),s(X7,X10))),s(fun(X7,fun(X7,bool)),X8),s(X7,X11))))),s(X7,X11)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : ((p(s(bool,'const.relation.WF_1'(s(fun(X7,fun(X7,bool)),X8)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X7,fun(X7,bool)),X8))))) => (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X11))),s(X7,X10)))) => s(X6,happ(s(fun(X7,X6),'const.relation.the_fun_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(fun(X7,X6),fun(X7,X6)),X9),s(X7,X10))),s(X7,X11))) = s(X6,happ(s(fun(X7,X6),happ(s(fun(fun(X7,X6),fun(X7,X6)),X9),s(fun(X7,X6),'const.relation.RESTRICT_3'(s(fun(X7,X6),'const.relation.the_fun_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(fun(X7,X6),fun(X7,X6)),X9),s(X7,X10))),s(fun(X7,fun(X7,bool)),X8),s(X7,X11))))),s(X7,X11)))))),
  file('thm.relation.relation_|49|37_prove_ltb','thm.relation.relation_|49|37_prove_ltb')).
fof(f73,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK4))) | $spl28),
  inference(cnf_transformation,[],[f73_D])).
fof(f73_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK4))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f71,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK4))) | ($spl2 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f60,f49])).
fof(f60,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK4))) | ($spl2 | $spl12)),
  inference(superposition,[],[f41,f50])).
fof(f43,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) | $spl4),
  inference(cnf_transformation,[],[f43_D])).
fof(f43_D,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f29,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK1,fun(sK1,bool)),sK2))))),
  inference(cnf_transformation,[],[f24])).
fof(f70,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | $spl26),
  inference(cnf_transformation,[],[f70_D])).
fof(f70_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK1,fun(sK1,bool)),sK2))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f68,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | ($spl4 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f59,f49])).
fof(f59,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | ($spl4 | $spl12)),
  inference(superposition,[],[f43,f50])).
fof(f50,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f50_D])).
fof(f50_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f34,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.relation_|49|37_prove_ltb',aHL_BOOL_CASES)).
fof(f49,plain,(
  ~p(s(bool,'F_0')) | $spl11),
  inference(cnf_transformation,[],[f49_D])).
fof(f49_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f33,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f11])).
fof(f11,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.relation_|49|37_prove_ltb',aHL_FALSITY)).
fof(f45,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) | $spl6),
  inference(cnf_transformation,[],[f45_D])).
fof(f45_D,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f28,plain,(
  p(s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2))))),
  inference(cnf_transformation,[],[f24])).
fof(f67,plain,(
  s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | $spl24),
  inference(cnf_transformation,[],[f67_D])).
fof(f67_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f65,plain,(
  s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | ($spl6 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f58,f49])).
fof(f58,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | ($spl6 | $spl12)),
  inference(superposition,[],[f45,f50])).
fof(f47,plain,(
  p(s(bool,'T_0')) | $spl8),
  inference(cnf_transformation,[],[f47_D])).
fof(f47_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f32,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.relation_|49|37_prove_ltb',aHL_TRUTH)).
fof(f832,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(sK1,sK5))) = s(X0,happ(s(fun(sK1,X0),happ(s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(fun(sK1,X0),'const.relation.RESTRICT_3'(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5)))) ) | $spl192),
  inference(cnf_transformation,[],[f832_D])).
fof(f832_D,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(sK1,sK5))) = s(X0,happ(s(fun(sK1,X0),happ(s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(fun(sK1,X0),'const.relation.RESTRICT_3'(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5)))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f831,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(sK1,sK5))) = s(X0,happ(s(fun(sK1,X0),happ(s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(fun(sK1,X0),'const.relation.RESTRICT_3'(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5)))) ) | ($spl8 | $spl24 | $spl26 | $spl28 | $spl54)),
  inference(subsumption_resolution,[],[f830,f47])).
fof(f830,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(X0,happ(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(sK1,sK5))) = s(X0,happ(s(fun(sK1,X0),happ(s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(fun(sK1,X0),'const.relation.RESTRICT_3'(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5)))) ) | ($spl8 | $spl24 | $spl26 | $spl28 | $spl54)),
  inference(forward_demodulation,[],[f829,f67])).
fof(f829,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(sK1,sK5))) = s(X0,happ(s(fun(sK1,X0),happ(s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(fun(sK1,X0),'const.relation.RESTRICT_3'(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5))) | ~p(s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2))))) ) | ($spl8 | $spl26 | $spl28 | $spl54)),
  inference(subsumption_resolution,[],[f828,f47])).
fof(f828,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(X0,happ(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(sK1,sK5))) = s(X0,happ(s(fun(sK1,X0),happ(s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(fun(sK1,X0),'const.relation.RESTRICT_3'(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5))) | ~p(s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2))))) ) | ($spl8 | $spl26 | $spl28 | $spl54)),
  inference(forward_demodulation,[],[f827,f70])).
fof(f827,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(sK1,sK5))) = s(X0,happ(s(fun(sK1,X0),happ(s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(fun(sK1,X0),'const.relation.RESTRICT_3'(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2))))) ) | ($spl8 | $spl28 | $spl54)),
  inference(subsumption_resolution,[],[f800,f47])).
fof(f800,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(X0,happ(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(sK1,sK5))) = s(X0,happ(s(fun(sK1,X0),happ(s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(fun(sK1,X0),'const.relation.RESTRICT_3'(s(fun(sK1,X0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,X0),fun(sK1,X0)),X1),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) | ~p(s(bool,'const.relation.WF_1'(s(fun(sK1,fun(sK1,bool)),sK2))))) ) | ($spl28 | $spl54)),
  inference(superposition,[],[f150,f73])).
fof(f39,plain,(
  s(sK0,happ(s(fun(sK1,sK0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(sK1,sK4))),s(sK1,sK5))) != s(sK0,happ(s(fun(sK1,sK0),happ(s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(fun(sK1,sK0),'const.relation.RESTRICT_3'(s(fun(sK1,sK0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5))) | $spl1),
  inference(cnf_transformation,[],[f39_D])).
fof(f39_D,plain,(
  s(sK0,happ(s(fun(sK1,sK0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(sK1,sK4))),s(sK1,sK5))) != s(sK0,happ(s(fun(sK1,sK0),happ(s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(fun(sK1,sK0),'const.relation.RESTRICT_3'(s(fun(sK1,sK0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f31,plain,(
  s(sK0,happ(s(fun(sK1,sK0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(sK1,sK4))),s(sK1,sK5))) != s(sK0,happ(s(fun(sK1,sK0),happ(s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(fun(sK1,sK0),'const.relation.RESTRICT_3'(s(fun(sK1,sK0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(sK1,sK4))),s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))))),s(sK1,sK5)))),
  inference(cnf_transformation,[],[f24])).
fof(f3231,plain,(
  $false | ($spl1 | $spl192)),
  inference(trivial_inequality_removal,[],[f3230])).
fof(f3230,plain,(
  s(sK0,happ(s(fun(sK1,sK0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(sK1,sK4))),s(sK1,sK5))) != s(sK0,happ(s(fun(sK1,sK0),'const.relation.the_fun_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(fun(sK1,sK0),fun(sK1,sK0)),sK3),s(sK1,sK4))),s(sK1,sK5))) | ($spl1 | $spl192)),
  inference(superposition,[],[f39,f832])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 56
% Passive clauses: 318
% Generated clauses: 3204
% Final active clauses: 52
% Final passive clauses: 236
% Input formulas: 7
% Initial clauses: 10
% 
% Duplicate literals: 274
% Trivial inequalities: 14
% Fw subsumption resolutions: 167
% Fw demodulations: 332
% Bw demodulations: 28
% Condensations: 215
% 
% Simple tautologies: 53
% Equational tautologies: 22
% Forward subsumptions: 1294
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 64
% Forward superposition: 1116
% Backward superposition: 592
% Self superposition: 38
% Equality factoring: 30
% Equality resolution: 1
% 
% Split clauses: 29
% Split components: 59
% SAT solver clauses: 476
% SAT solver unit clauses: 16
% SAT solver binary clauses: 10
% SAT solver learnt clauses: 1
% 
% Sat splits: 29
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 475
% TWLsolver calls for satisfiability: 85
% 
% Memory used [KB]: 6396
% Time elapsed: 0.213 s
% ------------------------------
% ------------------------------
% Success in time 0.827 s
