Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.numpair
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numpair
% SZS output start Proof for thm.numpair
fof(f3359,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f3317])).
fof(f3317,plain,(
  s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK2),s('type.num.num','const.numpair.napp_2'(s('type.num.num',sK0),s('type.num.num',sK1))))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK2),s('type.num.num','const.numpair.napp_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))),
  inference(superposition,[],[f63,f3058])).
fof(f3058,plain,(
  ( ! [X68,X66,X67] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X68),s('type.num.num','const.numpair.napp_2'(s('type.num.num',X67),s('type.num.num',X66))))) = s('type.num.num','const.numpair.napp_2'(s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X68),s('type.num.num',X67))),s('type.num.num',X66)))) )),
  inference(forward_demodulation,[],[f3057,f88])).
fof(f88,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X16,X17] : s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X16),s('type.num.num',X17))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',X16))),s('type.num.num',X17)))),
  file('thm.numpair.napp_thm|split|1','ARITY_0_2_const.numpair.ncons')).
fof(f3057,plain,(
  ( ! [X68,X66,X67] : (s('type.num.num','const.numpair.napp_2'(s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X68),s('type.num.num',X67))),s('type.num.num',X66))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',X68))),s('type.num.num','const.numpair.napp_2'(s('type.num.num',X67),s('type.num.num',X66)))))) )),
  inference(forward_demodulation,[],[f3056,f83])).
fof(f83,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X18] : ! [X19] : s(X18,X19) = s(X18,happ(s(fun(X18,X18),'_i_0_0'),s(X18,X19)))),
  file('thm.numpair.napp_thm|split|1','_i0_ax')).
fof(f3056,plain,(
  ( ! [X68,X66,X67] : (s('type.num.num','const.numpair.napp_2'(s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X68),s('type.num.num',X67))),s('type.num.num',X66))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',X68))))),s('type.num.num','const.numpair.napp_2'(s('type.num.num',X67),s('type.num.num',X66)))))) )),
  inference(forward_demodulation,[],[f3055,f89])).
fof(f89,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.napp_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num',X1),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : s('type.num.num','const.numpair.napp_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num',X1),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'))),s('type.num.num',X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X14,X15] : s('type.num.num','const.numpair.napp_2'(s('type.num.num',X14),s('type.num.num',X15))) = s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num',X15),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'))),s('type.num.num',X14)))),
  file('thm.numpair.napp_thm|split|1','thm.numpair.napp_def')).
fof(f3055,plain,(
  ( ! [X68,X66,X67] : (s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',X68))))),s('type.num.num','const.numpair.napp_2'(s('type.num.num',X67),s('type.num.num',X66))))) = s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num',X66),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'))),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X68),s('type.num.num',X67)))))) )),
  inference(forward_demodulation,[],[f3054,f91])).
fof(f91,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))) = s(X1,happ(s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X3),s(X2,X4))),s(X0,X5)))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))) = s(X1,happ(s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X3),s(X2,X4))),s(X0,X5)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))) = s(X1,happ(s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X3),s(X2,X4))),s(X0,X5)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X18,X20,X21] : ! [X22,X19,X23] : s(X20,happ(s(fun(X21,X20),X22),s(X21,X19))) = s(X20,happ(s(fun(X18,X20),'_dst_x0x0_2'(s(fun(X21,X20),X22),s(X21,X19))),s(X18,X23)))),
  file('thm.numpair.napp_thm|split|1','_dst_x0x0_ax')).
fof(f3054,plain,(
  ( ! [X68,X66,X67] : (s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num',X66),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'))),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X68),s('type.num.num',X67))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'_dst_x0x0_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',X68))))),s('type.num.num',X67))),s('type.num.num','const.numpair.napp_2'(s('type.num.num',X67),s('type.num.num',X66)))))) )),
  inference(forward_demodulation,[],[f2958,f95])).
fof(f95,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x00x10_2'(s(fun(X2,X1),X4),s(fun(X3,X2),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x00x10_2'(s(fun(X2,X1),X4),s(fun(X3,X2),X5))),s(X3,X6)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x00x10_2'(s(fun(X2,X1),X4),s(fun(X3,X2),X5))),s(X3,X6)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X21,X20,X24,X18] : ! [X22,X19,X23] : s(fun(X21,X20),'_dst_x0x0_2'(s(fun(X24,X20),X22),s(X24,happ(s(fun(X18,X24),X19),s(X18,X23))))) = s(fun(X21,X20),happ(s(fun(X18,fun(X21,X20)),'_dst_x00x10_2'(s(fun(X24,X20),X22),s(fun(X18,X24),X19))),s(X18,X23)))),
  file('thm.numpair.napp_thm|split|1','_dst_x00x10_ax')).
fof(f2958,plain,(
  ( ! [X68,X66,X67] : (s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num',X66),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'))),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X68),s('type.num.num',X67))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),happ(s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'))),s('type.num.num',X68))),s('type.num.num',X67))),s('type.num.num','const.numpair.napp_2'(s('type.num.num',X67),s('type.num.num',X66)))))) )),
  inference(superposition,[],[f87,f89])).
fof(f87,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X1))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),happ(s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X3))),s('type.num.num',X1))),s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X1))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),happ(s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X3))),s('type.num.num',X1))),s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X1)))))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X1))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),happ(s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X3))),s('type.num.num',X1))),s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X1)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : ! [X6,X12,X8,X13] : s(X9,'const.numpair.nlistrec_3'(s(X9,X12),s(fun('type.num.num',fun('type.num.num',fun(X9,X9))),X13),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X8),s('type.num.num',X6))))) = s(X9,happ(s(fun(X9,X9),happ(s(fun('type.num.num',fun(X9,X9)),happ(s(fun('type.num.num',fun('type.num.num',fun(X9,X9))),X13),s('type.num.num',X8))),s('type.num.num',X6))),s(X9,'const.numpair.nlistrec_3'(s(X9,X12),s(fun('type.num.num',fun('type.num.num',fun(X9,X9))),X13),s('type.num.num',X6)))))),
  file('thm.numpair.napp_thm|split|1','thm.numpair.nlistrec_thm|split|1')).
fof(f63,plain,(
  s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK2),s('type.num.num','const.numpair.napp_2'(s('type.num.num',sK0),s('type.num.num',sK1))))) != s('type.num.num','const.numpair.napp_2'(s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK2),s('type.num.num',sK0))),s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK2),s('type.num.num','const.numpair.napp_2'(s('type.num.num',sK0),s('type.num.num',sK1))))) != s('type.num.num','const.numpair.napp_2'(s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK2),s('type.num.num',sK0))),s('type.num.num',sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f46])).
fof(f46,plain,(
  ? [X0,X1,X2] : s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X2),s('type.num.num','const.numpair.napp_2'(s('type.num.num',X0),s('type.num.num',X1))))) != s('type.num.num','const.numpair.napp_2'(s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2] : s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X2),s('type.num.num','const.numpair.napp_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.numpair.napp_2'(s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X8),s('type.num.num','const.numpair.napp_2'(s('type.num.num',X6),s('type.num.num',X7))))) = s('type.num.num','const.numpair.napp_2'(s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X8),s('type.num.num',X6))),s('type.num.num',X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X8),s('type.num.num','const.numpair.napp_2'(s('type.num.num',X6),s('type.num.num',X7))))) = s('type.num.num','const.numpair.napp_2'(s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X8),s('type.num.num',X6))),s('type.num.num',X7)))),
  file('thm.numpair.napp_thm|split|1','thm.numpair.napp_thm|split|1')).
% SZS output end Proof for thm.numpair
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 55
% Passive clauses: 392
% Generated clauses: 3180
% Final active clauses: 52
% Final passive clauses: 307
% Input formulas: 20
% Initial clauses: 33
% 
% Duplicate literals: 250
% Trivial inequalities: 62
% Fw subsumption resolutions: 239
% Fw demodulations: 300
% Bw demodulations: 2
% 
% Simple tautologies: 90
% Equational tautologies: 23
% Forward subsumptions: 1674
% 
% Binary resolution: 5
% Forward superposition: 1324
% Backward superposition: 797
% Self superposition: 92
% Equality factoring: 47
% Equality resolution: 17
% 
% Split clauses: 80
% Split components: 168
% SAT solver clauses: 135
% SAT solver unit clauses: 11
% SAT solver binary clauses: 43
% 
% Sat splits: 80
% 
% TWLsolver clauses: 135
% TWLsolver calls for satisfiability: 25
% 
% Memory used [KB]: 2942
% Time elapsed: 0.098 s
% ------------------------------
% ------------------------------
% Success in time 0.1 s
