Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integerRing
% (45758)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 162
% Passive clauses: 4139
% Generated clauses: 21030
% Final active clauses: 154
% Final passive clauses: 3843
% Input formulas: 123
% Initial clauses: 128
% 
% Duplicate literals: 553
% Trivial inequalities: 367
% Fw subsumption resolutions: 963
% Fw demodulations: 1305
% Bw demodulations: 4
% 
% Simple tautologies: 328
% Equational tautologies: 22
% Forward subsumptions: 12896
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 31
% Forward superposition: 9751
% Backward superposition: 7560
% Self superposition: 213
% Equality factoring: 54
% Equality resolution: 30
% 
% Split clauses: 234
% Split components: 468
% SAT solver clauses: 463
% SAT solver unit clauses: 5
% SAT solver binary clauses: 100
% 
% Sat splits: 234
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 463
% TWLsolver calls for satisfiability: 59
% 
% Memory used [KB]: 21236
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integerRing
% (45878)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 148
% Passive clauses: 1579
% Generated clauses: 19331
% Final active clauses: 147
% Final passive clauses: 1376
% Input formulas: 123
% Initial clauses: 128
% 
% Duplicate literals: 582
% Trivial inequalities: 313
% Fw subsumption resolutions: 630
% Fw demodulations: 1116
% Bw demodulations: 55
% Condensations: 717
% 
% Simple tautologies: 189
% Equational tautologies: 24
% Forward subsumptions: 11542
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 24
% Forward superposition: 7482
% Backward superposition: 6387
% Self superposition: 202
% Equality factoring: 56
% Equality resolution: 18
% 
% Split clauses: 266
% Split components: 532
% SAT solver clauses: 2718
% SAT solver unit clauses: 245
% SAT solver binary clauses: 44
% 
% Sat splits: 266
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2718
% TWLsolver calls for satisfiability: 283
% 
% Memory used [KB]: 26737
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integerRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integerRing
% SZS output start Proof for thm.integerRing
fof(f108567,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f108519])).
fof(f108519,plain,(
  s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pvar_1'(s('type.quote.index',sK0))))) != s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pvar_1'(s('type.quote.index',sK0)))))),
  inference(superposition,[],[f7854,f25262])).
fof(f25262,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),X0))) = s('type.canonical.canonical_sum'('type.integer.int'),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'('type.integer.int'),'const.ring.recordtype|46|ring_5'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_add_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_mul_0'),s(fun('type.integer.int','type.integer.int'),'const.integer.int_neg_0'))),s('type.ringNorm.polynom'('type.integer.int'),X0)))) )),
  inference(forward_demodulation,[],[f25229,f343])).
fof(f343,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),X0))) = s('type.canonical.canonical_sum'('type.integer.int'),happ(s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0'),s('type.ringNorm.polynom'('type.integer.int'),X0)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),X0))) = s('type.canonical.canonical_sum'('type.integer.int'),happ(s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0'),s('type.ringNorm.polynom'('type.integer.int'),X0)))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X47] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),X47))) = s('type.canonical.canonical_sum'('type.integer.int'),happ(s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0'),s('type.ringNorm.polynom'('type.integer.int'),X47)))),
  file('thm.integerRing.int_ring_thms|split|51','ARITY_0_1_const.integerRing.int_polynom_normalize')).
fof(f25229,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.integer.int'),happ(s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0'),s('type.ringNorm.polynom'('type.integer.int'),X0))) = s('type.canonical.canonical_sum'('type.integer.int'),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'('type.integer.int'),'const.ring.recordtype|46|ring_5'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_add_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_mul_0'),s(fun('type.integer.int','type.integer.int'),'const.integer.int_neg_0'))),s('type.ringNorm.polynom'('type.integer.int'),X0)))) )),
  inference(superposition,[],[f360,f325])).
fof(f325,plain,(
  s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0') = s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'('type.integer.int'),'const.ring.recordtype|46|ring_5'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_add_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_mul_0'),s(fun('type.integer.int','type.integer.int'),'const.integer.int_neg_0')))))),
  inference(cnf_transformation,[],[f11])).
fof(f11,axiom,(
  s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.integerRing.int_polynom_normalize_0') = s(fun('type.ringNorm.polynom'('type.integer.int'),'type.canonical.canonical_sum'('type.integer.int')),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'('type.integer.int'),'const.ring.recordtype|46|ring_5'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_add_0'),s(fun('type.integer.int',fun('type.integer.int','type.integer.int')),'const.integer.int_mul_0'),s(fun('type.integer.int','type.integer.int'),'const.integer.int_neg_0')))))),
  file('thm.integerRing.int_ring_thms|split|51','thm.integerRing.int_polynom_normalize_def')).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X7] : ! [X47,X48] : s('type.canonical.canonical_sum'(X7),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X7),X47),s('type.ringNorm.polynom'(X7),X48))) = s('type.canonical.canonical_sum'(X7),happ(s(fun('type.ringNorm.polynom'(X7),'type.canonical.canonical_sum'(X7)),'const.ringNorm.polynom_normalize_1'(s('type.ring.ring'(X7),X47))),s('type.ringNorm.polynom'(X7),X48)))),
  file('thm.integerRing.int_ring_thms|split|51','ARITY_1_2_const.ringNorm.polynom_normalize')).
fof(f7854,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pvar_1'(s('type.quote.index',sK0))))) != s('type.canonical.canonical_sum'('type.integer.int'),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'('type.integer.int'),X0),s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pvar_1'(s('type.quote.index',sK0)))))) )),
  inference(superposition,[],[f321,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X2),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X7] : ! [X18,X6] : s('type.canonical.canonical_sum'(X7),'const.ringNorm.polynom_normalize_2'(s('type.ring.ring'(X7),X18),s('type.ringNorm.polynom'(X7),'const.ringNorm.Pvar_1'(s('type.quote.index',X6))))) = s('type.canonical.canonical_sum'(X7),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0')))),
  file('thm.integerRing.int_ring_thms|split|51','thm.ringNorm.polynom_normalize_def|split|0')).
fof(f321,plain,(
  s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pvar_1'(s('type.quote.index',sK0))))) != s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',sK0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pvar_1'(s('type.quote.index',sK0))))) != s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',sK0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f312])).
fof(f312,plain,(
  ? [X0] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pvar_1'(s('type.quote.index',X0))))) != s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ~! [X0] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pvar_1'(s('type.quote.index',X0))))) = s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X0),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pvar_1'(s('type.quote.index',X6))))) = s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.canonical.canonical_sum'('type.integer.int'),'const.integerRing.int_polynom_normalize_1'(s('type.ringNorm.polynom'('type.integer.int'),'const.ringNorm.Pvar_1'(s('type.quote.index',X6))))) = s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X6),s('type.list.list'('type.quote.index'),'const.list.NIL_0'))),s('type.canonical.canonical_sum'('type.integer.int'),'const.canonical.Nil_monom_0')))),
  file('thm.integerRing.int_ring_thms|split|51','thm.integerRing.int_ring_thms|split|51')).
% SZS output end Proof for thm.integerRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 437
% Passive clauses: 6540
% Generated clauses: 107245
% Final active clauses: 298
% Final passive clauses: 2337
% Input formulas: 123
% Initial clauses: 128
% Discarded non-redundant clauses: 13654
% 
% Duplicate literals: 1941
% Trivial inequalities: 663
% Fw subsumption resolutions: 1063
% Fw demodulations: 17347
% Bw demodulations: 346
% Condensations: 1966
% Global subsumptions: 139
% 
% Simple tautologies: 374
% Equational tautologies: 390
% Forward subsumptions: 76174
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 51
% 
% Unit resulting resolution: 12352
% Forward superposition: 37880
% Backward superposition: 32132
% Self superposition: 455
% Equality factoring: 112
% Equality resolution: 33
% 
% Split clauses: 169
% Split components: 343
% SAT solver clauses: 7203
% SAT solver unit clauses: 827
% SAT solver binary clauses: 1036
% 
% Sat splits: 169
% Sat splitting refutations: 184
% 
% Memory used [KB]: 143153
% Time elapsed: 5.015 s
% ------------------------------
% ------------------------------
% Success in time 6.311 s
