Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.numpair
% (95574)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 73
% Passive clauses: 1962
% Generated clauses: 12972
% Final active clauses: 70
% Final passive clauses: 1796
% Input formulas: 21
% Initial clauses: 34
% 
% Duplicate literals: 872
% Trivial inequalities: 160
% Fw subsumption resolutions: 1602
% Fw demodulations: 592
% Bw demodulations: 1
% 
% Simple tautologies: 256
% Equational tautologies: 102
% Forward subsumptions: 5640
% 
% Binary resolution: 46
% Forward superposition: 5246
% Backward superposition: 3653
% Self superposition: 369
% Equality factoring: 283
% Equality resolution: 18
% 
% Split clauses: 657
% Split components: 1323
% SAT solver clauses: 1800
% SAT solver unit clauses: 270
% SAT solver binary clauses: 517
% 
% Sat splits: 657
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1796
% TWLsolver calls for satisfiability: 87
% 
% Memory used [KB]: 6908
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numpair
% (95589)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 46
% Passive clauses: 278
% Generated clauses: 6092
% Final active clauses: 42
% Final passive clauses: 230
% Input formulas: 21
% Initial clauses: 34
% 
% Duplicate literals: 886
% Trivial inequalities: 57
% Fw subsumption resolutions: 204
% Fw demodulations: 117
% Bw demodulations: 2
% Condensations: 1536
% Global subsumptions: 4
% 
% Simple tautologies: 115
% Equational tautologies: 62
% Forward subsumptions: 2142
% 
% Binary resolution: 16
% Forward superposition: 1547
% Backward superposition: 1087
% Self superposition: 209
% Equality factoring: 176
% Equality resolution: 9
% 
% Split clauses: 64
% Split components: 129
% SAT solver clauses: 707
% SAT solver unit clauses: 47
% SAT solver binary clauses: 74
% SAT solver learnt clauses: 8
% 
% Sat splits: 64
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 691
% TWLsolver calls for satisfiability: 104
% 
% Memory used [KB]: 6396
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numpair
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numpair
% SZS output start Proof for thm.numpair
fof(f70680,plain,(
  $false),
  inference(equality_resolution,[],[f70227])).
fof(f70227,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),sK2),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))))),s('type.num.num',X0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),sK2),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))))),s('type.num.num',X1)))) )),
  inference(backward_demodulation,[],[f70109,f324])).
fof(f324,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),sK2),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))))),s('type.num.num',X0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num','const.numpair.ncons_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK2),s('type.num.num',sK1))),s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),sK2),s('type.num.num',sK0))))),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f65,f95])).
fof(f95,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2))) | s('type.num.num',X1) = s('type.num.num',X3)) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : ((s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2))) | (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X2))) & (s('type.num.num',X1) != s('type.num.num',X3) | s('type.num.num',X0) != s('type.num.num',X2) | s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2)))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3] : ((s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2))) | (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X2))) & ((s('type.num.num',X1) != s('type.num.num',X3) | s('type.num.num',X0) != s('type.num.num',X2)) | s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2)))))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2))) <=> (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X2)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X11,X10,X6,X7] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X7),s('type.num.num',X6))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X10),s('type.num.num',X11))) <=> (s('type.num.num',X7) = s('type.num.num',X10) & s('type.num.num',X6) = s('type.num.num',X11)))),
  file('thm.numpair.nmap_thm|split|1','thm.numpair.ncons_|49||49|')).
fof(f65,plain,(
  s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),sK2),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK2),s('type.num.num',sK1))),s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),sK2),s('type.num.num',sK0)))))),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),sK2),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),sK2),s('type.num.num',sK1))),s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),sK2),s('type.num.num',sK0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f48])).
fof(f48,plain,(
  ? [X0,X1,X2] : s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num',X1))),s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X2),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2] : s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num',X1))),s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X2),s('type.num.num',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X8),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X8),s('type.num.num',X7))),s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X8),s('type.num.num',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X8),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X8),s('type.num.num',X7))),s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X8),s('type.num.num',X6)))))),
  file('thm.numpair.nmap_thm|split|1','thm.numpair.nmap_thm|split|1')).
fof(f70109,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X3),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X4),s('type.num.num',X5))))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X3),s('type.num.num',X4))),s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X3),s('type.num.num',X5)))))) )),
  inference(superposition,[],[f2140,f92])).
fof(f92,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13,X14] : s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X13),s('type.num.num',X14))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',X13))),s('type.num.num',X14)))),
  file('thm.numpair.nmap_thm|split|1','ARITY_0_2_const.numpair.ncons')).
fof(f2140,plain,(
  ( ! [X227,X225,X226] : (s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X225),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X227),s('type.num.num',X226))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X225),s('type.num.num',X227))))),s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X225),s('type.num.num',X226)))))) )),
  inference(forward_demodulation,[],[f2139,f888])).
fof(f888,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))) = s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X0))),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f794,f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X13,X14] : s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X13),s('type.num.num',X14))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X13))),s('type.num.num',X14)))),
  file('thm.numpair.nmap_thm|split|1','ARITY_1_2_const.numpair.nmap')).
fof(f794,plain,(
  ( ! [X0,X1] : (s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X0))),s('type.num.num',X1))) = s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X0))),s('type.num.num',X1)))) )),
  inference(superposition,[],[f89,f69])).
fof(f69,plain,(
  ( ! [X0] : (s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X0))) = s(fun('type.num.num','type.num.num'),'const.numpair.nlistrec_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X0)))))) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0] : s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X0))) = s(fun('type.num.num','type.num.num'),'const.numpair.nlistrec_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X0)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8] : s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X8))) = s(fun('type.num.num','type.num.num'),'const.numpair.nlistrec_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X8)))))),
  file('thm.numpair.nmap_thm|split|1','thm.numpair.nmap_def')).
fof(f89,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.numpair.nlistrec_3'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),'const.numpair.nlistrec_2'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.numpair.nlistrec_3'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),'const.numpair.nlistrec_2'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.numpair.nlistrec_3'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),'const.numpair.nlistrec_2'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : ! [X13,X14,X15] : s(X9,'const.numpair.nlistrec_3'(s(X9,X13),s(fun('type.num.num',fun('type.num.num',fun(X9,X9))),X14),s('type.num.num',X15))) = s(X9,happ(s(fun('type.num.num',X9),'const.numpair.nlistrec_2'(s(X9,X13),s(fun('type.num.num',fun('type.num.num',fun(X9,X9))),X14))),s('type.num.num',X15)))),
  file('thm.numpair.nmap_thm|split|1','ARITY_2_3_const.numpair.nlistrec')).
fof(f2139,plain,(
  ( ! [X227,X225,X226] : (s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X225))),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X227),s('type.num.num',X226))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X225),s('type.num.num',X227))))),s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X225),s('type.num.num',X226)))))) )),
  inference(forward_demodulation,[],[f2068,f1027])).
fof(f1027,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s(X33,happ(s(fun(X34,X33),X35),s(X34,happ(s(fun(X36,X34),X37),s(X36,X38))))) = s(X33,happ(s(fun(X32,X33),happ(s(fun(X36,fun(X32,X33)),'_dst_x00x10_2'(s(fun(X34,X33),X35),s(fun(X36,X34),X37))),s(X36,X38))),s(X32,X39)))) )),
  inference(superposition,[],[f94,f98])).
fof(f98,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x00x10_2'(s(fun(X2,X1),X4),s(fun(X3,X2),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x00x10_2'(s(fun(X2,X1),X4),s(fun(X3,X2),X5))),s(X3,X6)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X4),s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x00x10_2'(s(fun(X2,X1),X4),s(fun(X3,X2),X5))),s(X3,X6)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X18,X17,X22,X16] : ! [X19,X20,X21] : s(fun(X18,X17),'_dst_x0x0_2'(s(fun(X22,X17),X19),s(X22,happ(s(fun(X16,X22),X20),s(X16,X21))))) = s(fun(X18,X17),happ(s(fun(X16,fun(X18,X17)),'_dst_x00x10_2'(s(fun(X22,X17),X19),s(fun(X16,X22),X20))),s(X16,X21)))),
  file('thm.numpair.nmap_thm|split|1','_dst_x00x10_ax')).
fof(f94,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))) = s(X1,happ(s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X3),s(X2,X4))),s(X0,X5)))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))) = s(X1,happ(s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X3),s(X2,X4))),s(X0,X5)))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))) = s(X1,happ(s(fun(X0,X1),'_dst_x0x0_2'(s(fun(X2,X1),X3),s(X2,X4))),s(X0,X5)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X16,X17,X18] : ! [X19,X20,X21] : s(X17,happ(s(fun(X18,X17),X19),s(X18,X20))) = s(X17,happ(s(fun(X16,X17),'_dst_x0x0_2'(s(fun(X18,X17),X19),s(X18,X20))),s(X16,X21)))),
  file('thm.numpair.nmap_thm|split|1','_dst_x0x0_ax')).
fof(f2068,plain,(
  ( ! [X227,X225,X226] : (s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X225))),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X227),s('type.num.num',X226))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),happ(s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X225))),s('type.num.num',X227))),s('type.num.num',X226))),s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X225),s('type.num.num',X226)))))) )),
  inference(superposition,[],[f90,f888])).
fof(f90,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X1))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),happ(s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X3))),s('type.num.num',X1))),s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X1))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),happ(s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X3))),s('type.num.num',X1))),s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X1)))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X1))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),happ(s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X3))),s('type.num.num',X1))),s(X0,'const.numpair.nlistrec_3'(s(X0,X2),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X4),s('type.num.num',X1)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : ! [X6,X12,X7,X8] : s(X9,'const.numpair.nlistrec_3'(s(X9,X12),s(fun('type.num.num',fun('type.num.num',fun(X9,X9))),X8),s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s(X9,happ(s(fun(X9,X9),happ(s(fun('type.num.num',fun(X9,X9)),happ(s(fun('type.num.num',fun('type.num.num',fun(X9,X9))),X8),s('type.num.num',X7))),s('type.num.num',X6))),s(X9,'const.numpair.nlistrec_3'(s(X9,X12),s(fun('type.num.num',fun('type.num.num',fun(X9,X9))),X8),s('type.num.num',X6)))))),
  file('thm.numpair.nmap_thm|split|1','thm.numpair.nlistrec_thm|split|1')).
% SZS output end Proof for thm.numpair
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 239
% Passive clauses: 5264
% Generated clauses: 71394
% Final active clauses: 91
% Final passive clauses: 912
% Input formulas: 21
% Initial clauses: 34
% Discarded non-redundant clauses: 3084
% 
% Duplicate literals: 5573
% Trivial inequalities: 362
% Fw subsumption resolutions: 2821
% Fw demodulations: 7742
% Bw demodulations: 2267
% Condensations: 8379
% Global subsumptions: 9
% 
% Simple tautologies: 487
% Equational tautologies: 610
% Forward subsumptions: 38979
% Fw demodulations to eq. taut.: 640
% Bw demodulations to eq. taut.: 76
% 
% Unit resulting resolution: 2897
% Forward superposition: 20429
% Backward superposition: 17417
% Self superposition: 620
% Equality factoring: 295
% Equality resolution: 15
% 
% Split clauses: 233
% Split components: 471
% SAT solver clauses: 5768
% SAT solver unit clauses: 183
% SAT solver binary clauses: 795
% 
% Sat splits: 233
% Sat splitting refutations: 27
% 
% Memory used [KB]: 56161
% Time elapsed: 6.599 s
% ------------------------------
% ------------------------------
% Success in time 7.92 s
