Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (92092)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 132
% Passive clauses: 2915
% Generated clauses: 19634
% Final active clauses: 112
% Final passive clauses: 2575
% Input formulas: 27
% Initial clauses: 52
% 
% Duplicate literals: 619
% Trivial inequalities: 254
% Fw subsumption resolutions: 2476
% Fw demodulations: 3363
% Bw demodulations: 18
% 
% Simple tautologies: 317
% Equational tautologies: 39
% Forward subsumptions: 9292
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 491
% Factoring: 5
% Forward superposition: 7027
% Backward superposition: 6156
% Self superposition: 184
% Equality factoring: 78
% Equality resolution: 24
% 
% Split clauses: 248
% Split components: 508
% SAT solver clauses: 562
% SAT solver unit clauses: 16
% SAT solver binary clauses: 181
% 
% Sat splits: 248
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 560
% TWLsolver calls for satisfiability: 113
% 
% Memory used [KB]: 12025
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (92106)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 106
% Passive clauses: 850
% Generated clauses: 14096
% Final active clauses: 92
% Final passive clauses: 571
% Input formulas: 27
% Initial clauses: 52
% 
% Duplicate literals: 631
% Trivial inequalities: 123
% Fw subsumption resolutions: 1099
% Fw demodulations: 1872
% Bw demodulations: 169
% Condensations: 997
% Global subsumptions: 9
% 
% Simple tautologies: 203
% Equational tautologies: 36
% Forward subsumptions: 6722
% Bw demodulations to eq. taut.: 13
% 
% Binary resolution: 348
% Factoring: 4
% Forward superposition: 4355
% Backward superposition: 3808
% Self superposition: 117
% Equality factoring: 63
% Equality resolution: 16
% 
% Split clauses: 36
% Split components: 83
% SAT solver clauses: 1629
% SAT solver unit clauses: 70
% SAT solver binary clauses: 43
% SAT solver learnt clauses: 18
% 
% Sat splits: 36
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1593
% TWLsolver calls for satisfiability: 281
% 
% Memory used [KB]: 11769
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (92131)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 447
% Passive clauses: 5935
% Generated clauses: 138010
% Final active clauses: 142
% Final passive clauses: 919
% Input formulas: 27
% Initial clauses: 52
% Discarded non-redundant clauses: 200
% 
% Duplicate literals: 5925
% Trivial inequalities: 661
% Fw subsumption resolutions: 8311
% Fw demodulations: 21180
% Bw demodulations: 3679
% Condensations: 13239
% Global subsumptions: 2544
% 
% Simple tautologies: 1012
% Equational tautologies: 3066
% Forward subsumptions: 71987
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 138
% 
% Unit resulting resolution: 15561
% Factoring: 10
% Forward superposition: 30903
% Backward superposition: 23277
% Self superposition: 436
% Equality factoring: 211
% Equality resolution: 55
% 
% Split clauses: 330
% Split components: 667
% SAT solver clauses: 11882
% SAT solver unit clauses: 295
% SAT solver binary clauses: 1637
% 
% Sat splits: 330
% Sat splitting refutations: 2741
% 
% Memory used [KB]: 84433
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (92388)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2005
% Passive clauses: 23704
% Generated clauses: 36647
% Final active clauses: 2005
% Final passive clauses: 19834
% Input formulas: 27
% Initial clauses: 50
% 
% Selected by SInE selection: 25
% SInE iterations: 1
% 
% Duplicate literals: 1153
% Trivial inequalities: 62
% Fw subsumption resolutions: 1351
% Fw demodulations: 892
% 
% Simple tautologies: 955
% Forward subsumptions: 10430
% 
% Binary resolution: 516
% Forward superposition: 31568
% Equality resolution: 1104
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 4
% SAT solver binary clauses: 3
% 
% Sat splits: 3
% 
% Memory used [KB]: 16630
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (92600)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 91
% Passive clauses: 1413
% Generated clauses: 17490
% Final active clauses: 77
% Final passive clauses: 929
% Input formulas: 27
% Initial clauses: 52
% 
% Duplicate literals: 594
% Trivial inequalities: 331
% Fw subsumption resolutions: 2004
% Fw demodulations: 1534
% Bw demodulations: 483
% Global subsumptions: 29
% 
% Simple tautologies: 468
% Equational tautologies: 105
% Forward subsumptions: 10442
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 1233
% Forward superposition: 6136
% Backward superposition: 5334
% Self superposition: 128
% Equality factoring: 46
% Equality resolution: 15
% 
% Split clauses: 9
% Split components: 18
% SAT solver clauses: 1617
% SAT solver unit clauses: 182
% SAT solver binary clauses: 82
% 
% Sat splits: 9
% Sat splitting refutations: 163
% 
% Memory used [KB]: 12537
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (92623)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 574
% Passive clauses: 4279
% Generated clauses: 9864
% Final active clauses: 571
% Final passive clauses: 3714
% Input formulas: 27
% Initial clauses: 52
% 
% Duplicate literals: 396
% Trivial inequalities: 229
% Fw subsumption resolutions: 196
% Bw subsumption resolutions: 2
% Fw demodulations: 81
% Bw demodulations: 4
% Condensations: 290
% 
% Simple tautologies: 467
% Forward subsumptions: 3875
% Backward subsumptions: 20
% 
% Binary resolution: 1904
% Factoring: 8
% Forward superposition: 6373
% Backward superposition: 334
% Equality resolution: 64
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 4283
% SAT solver unit clauses: 9
% SAT solver binary clauses: 138
% 
% Sat splits: 2
% 
% Memory used [KB]: 18933
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f1782,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f95,f149,f153,f155,f152,f157,f151,f159,f150,f161,f165,f167,f164,f169,f163,f171,f162,f173,f177,f179,f176,f181,f175,f183,f174,f185,f190,f193,f196,f198,f187,f200,f203,f205,f210,f213,f216,f218,f207,f220,f223,f225,f226,f230,f228,f231,f235,f233,f236,f240,f238,f241,f245,f243,f250,f249,f248,f247,f246,f254,f252,f256,f257,f259,f277,f276,f275,f274,f273,f281,f279,f283,f290,f289,f288,f287,f286,f294,f292,f296,f303,f302,f301,f300,f299,f307,f305,f309,f314,f318,f316,f321,f325,f323,f328,f330,f335,f339,f337,f340,f344,f342,f349,f348,f347,f346,f345,f353,f351,f355,f356,f358,f360,f364,f362,f365,f369,f367,f370,f374,f372,f392,f354,f473,f472,f471,f470,f469,f477,f475,f479,f480,f482,f501,f478,f582,f581,f580,f579,f578,f586,f584,f588,f589,f591,f610,f587,f691,f690,f689,f688,f687,f695,f693,f697,f698,f700,f719,f696,f800,f799,f798,f797,f796,f804,f802,f806,f807,f809,f828,f805,f920,f919,f918,f917,f916,f924,f922,f926,f927,f929,f948,f925,f154,f1029,f1028,f1027,f1026,f1025,f1033,f1031,f1035,f1037,f329,f1043,f1042,f1041,f1048,f1040,f1039,f1046,f1049,f1047,f1050,f1052,f1153,f1152,f1151,f1150,f1149,f1157,f1155,f1159,f1160,f1162,f1192,f1158,f158,f1273,f1272,f1271,f1270,f1269,f1277,f1275,f1279,f1280,f1282,f1301,f1278,f156,f1382,f1381,f1380,f1379,f1378,f1386,f1384,f1388,f1390,f361,f1396,f1395,f1394,f1401,f1393,f1392,f1399,f1402,f1403,f1405,f332,f1434,f1433,f1432,f1431,f1430,f1438,f1436,f1440,f1441,f1443,f1462,f1439,f1543,f1542,f1541,f1540,f1539,f1547,f1545,f1549,f1550,f1552,f1571,f1548,f160,f1651,f1688,f1652,f1692,f1653,f1694,f1706,f1711,f1747,f1712,f1690,f1749,f1774,f1761])).
fof(f1761,plain,(
  $false | $spl186),
  inference(unit_resulting_resolution,[],[f132,f1688])).
fof(f132,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X11] : ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','thm.pred_set.NOT_IN_EMPTY')).
fof(f1774,plain,(
  $false | $spl186),
  inference(resolution,[],[f1688,f132])).
fof(f1749,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl196),
  inference(cnf_transformation,[],[f1749_D])).
fof(f1749_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f1690,plain,(
  ~p(s(bool,'T_0')) | $spl189),
  inference(cnf_transformation,[],[f1690_D])).
fof(f1690_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f1712,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl66),
  inference(superposition,[],[f138,f283])).
fof(f138,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4)))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4))))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4)))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) | p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4))))))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4)))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4)))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X7,X15,X11,X16] : (p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X6,bool),X15),s(X6,X11),s(X6,X16)))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X15)))) & p(s(bool,'const.bool.IN_2'(s(X6,X16),s(fun(X6,bool),X15)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X16))))))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','thm.pred_set.REL_RESTRICT_DEF')).
fof(f1747,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl194),
  inference(cnf_transformation,[],[f1747_D])).
fof(f1747_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f1711,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl66),
  inference(superposition,[],[f137,f283])).
fof(f137,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f1706,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))))),
  introduced(theory_axiom,[])).
fof(f1694,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl192),
  inference(cnf_transformation,[],[f1694_D])).
fof(f1694_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f1653,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl60),
  inference(superposition,[],[f138,f256])).
fof(f1692,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl190),
  inference(cnf_transformation,[],[f1692_D])).
fof(f1692_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f1652,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl60),
  inference(superposition,[],[f137,f256])).
fof(f1688,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl186),
  inference(cnf_transformation,[],[f1688_D])).
fof(f1688_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f1651,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl60),
  inference(superposition,[],[f136,f256])).
fof(f136,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f160,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) | $spl8),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1548,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl183),
  inference(cnf_transformation,[],[f1548_D])).
fof(f1548_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f1571,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) | $spl110),
  inference(superposition,[],[f144,f372])).
fof(f144,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.REL_RESTRICT_EMPTY',aHL_EXT)).
fof(f1552,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl185),
  inference(cnf_transformation,[],[f1552_D])).
fof(f1552_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f1550,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | ($spl39 | $spl112)),
  inference(backward_demodulation,[],[f373,f220])).
fof(f373,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl112),
  inference(cnf_transformation,[],[f373_D])).
fof(f373_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f1549,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl182),
  inference(cnf_transformation,[],[f1549_D])).
fof(f1549_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f1545,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl113),
  inference(trivial_inequality_removal,[],[f1544])).
fof(f1544,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl113),
  inference(superposition,[],[f374,f103])).
fof(f103,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.REL_RESTRICT_EMPTY',aHL_BOOL_CASES)).
fof(f1547,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) | $spl180),
  inference(cnf_transformation,[],[f1547_D])).
fof(f1547_D,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f1539,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) | $spl113),
  inference(unit_resulting_resolution,[],[f374,f112])).
fof(f112,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','thm.bool.EQ_CLAUSES|split|2')).
fof(f1540,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) | $spl113),
  inference(unit_resulting_resolution,[],[f374,f110])).
fof(f110,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','thm.bool.EQ_CLAUSES|split|3')).
fof(f1541,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl113),
  inference(unit_resulting_resolution,[],[f374,f103])).
fof(f1542,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) | $spl113),
  inference(unit_resulting_resolution,[],[f102,f374,f140])).
fof(f140,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','thm.bool.IMP_ANTISYM_AX')).
fof(f102,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.REL_RESTRICT_EMPTY',aHL_FALSITY)).
fof(f1543,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) | $spl113),
  inference(unit_resulting_resolution,[],[f102,f374,f140])).
fof(f1439,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl177),
  inference(cnf_transformation,[],[f1439_D])).
fof(f1439_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f1462,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) | $spl106),
  inference(superposition,[],[f144,f367])).
fof(f1443,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl179),
  inference(cnf_transformation,[],[f1443_D])).
fof(f1443_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f1441,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | ($spl35 | $spl108)),
  inference(backward_demodulation,[],[f368,f213])).
fof(f368,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl108),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1440,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl176),
  inference(cnf_transformation,[],[f1440_D])).
fof(f1440_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f1436,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl109),
  inference(trivial_inequality_removal,[],[f1435])).
fof(f1435,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl109),
  inference(superposition,[],[f369,f103])).
fof(f1438,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) | $spl174),
  inference(cnf_transformation,[],[f1438_D])).
fof(f1438_D,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f1430,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) | $spl109),
  inference(unit_resulting_resolution,[],[f369,f112])).
fof(f1431,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) | $spl109),
  inference(unit_resulting_resolution,[],[f369,f110])).
fof(f1432,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl109),
  inference(unit_resulting_resolution,[],[f369,f103])).
fof(f1433,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) | $spl109),
  inference(unit_resulting_resolution,[],[f102,f369,f140])).
fof(f1434,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0))))) | $spl109),
  inference(unit_resulting_resolution,[],[f102,f369,f140])).
fof(f332,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl87),
  inference(cnf_transformation,[],[f332_D])).
fof(f332_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f1405,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | ($spl21 | $spl150)),
  inference(backward_demodulation,[],[f1048,f181])).
fof(f1403,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | p(s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | ~p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))))),
  introduced(theory_axiom,[])).
fof(f1402,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | ~p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))))),
  introduced(theory_axiom,[])).
fof(f1399,plain,(
  ~p(s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl171),
  inference(cnf_transformation,[],[f1399_D])).
fof(f1399_D,plain,(
  ~p(s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f1392,plain,(
  ~p(s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl103),
  inference(unit_resulting_resolution,[],[f361,f108])).
fof(f108,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','thm.bool.EQ_CLAUSES|split|0')).
fof(f1393,plain,(
  ~p(s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl103),
  inference(unit_resulting_resolution,[],[f361,f106])).
fof(f106,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','thm.bool.EQ_CLAUSES|split|1')).
fof(f1401,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl172),
  inference(cnf_transformation,[],[f1401_D])).
fof(f1401_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f1394,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl103),
  inference(unit_resulting_resolution,[],[f361,f103])).
fof(f1395,plain,(
  ~p(s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl103),
  inference(unit_resulting_resolution,[],[f97,f361,f143])).
fof(f143,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f67])).
fof(f97,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.REL_RESTRICT_EMPTY',aHL_TRUTH)).
fof(f1396,plain,(
  ~p(s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl103),
  inference(unit_resulting_resolution,[],[f97,f361,f143])).
fof(f361,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl103),
  inference(cnf_transformation,[],[f361_D])).
fof(f361_D,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f1390,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | ($spl31 | $spl168)),
  inference(backward_demodulation,[],[f1388,f200])).
fof(f1388,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl168),
  inference(cnf_transformation,[],[f1388_D])).
fof(f1388_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f1384,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl105),
  inference(trivial_inequality_removal,[],[f1383])).
fof(f1383,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl105),
  inference(superposition,[],[f364,f103])).
fof(f1386,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl166),
  inference(cnf_transformation,[],[f1386_D])).
fof(f1386_D,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f1378,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl105),
  inference(unit_resulting_resolution,[],[f364,f112])).
fof(f1379,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl105),
  inference(unit_resulting_resolution,[],[f364,f110])).
fof(f1380,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl105),
  inference(unit_resulting_resolution,[],[f364,f103])).
fof(f1381,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl105),
  inference(unit_resulting_resolution,[],[f102,f364,f140])).
fof(f1382,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl105),
  inference(unit_resulting_resolution,[],[f102,f364,f140])).
fof(f156,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) | $spl4),
  inference(cnf_transformation,[],[f156_D])).
fof(f156_D,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1278,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl163),
  inference(cnf_transformation,[],[f1278_D])).
fof(f1278_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f1301,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) = s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) | $spl92),
  inference(superposition,[],[f144,f342])).
fof(f1282,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl165),
  inference(cnf_transformation,[],[f1282_D])).
fof(f1282_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl165),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl165])])).
fof(f1280,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | ($spl27 | $spl94)),
  inference(backward_demodulation,[],[f343,f193])).
fof(f343,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl94),
  inference(cnf_transformation,[],[f343_D])).
fof(f343_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f1279,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl162),
  inference(cnf_transformation,[],[f1279_D])).
fof(f1279_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f1275,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl95),
  inference(trivial_inequality_removal,[],[f1274])).
fof(f1274,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl95),
  inference(superposition,[],[f344,f103])).
fof(f1277,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl160),
  inference(cnf_transformation,[],[f1277_D])).
fof(f1277_D,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f1269,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl95),
  inference(unit_resulting_resolution,[],[f344,f112])).
fof(f1270,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl95),
  inference(unit_resulting_resolution,[],[f344,f110])).
fof(f1271,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl95),
  inference(unit_resulting_resolution,[],[f344,f103])).
fof(f1272,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl95),
  inference(unit_resulting_resolution,[],[f102,f344,f140])).
fof(f1273,plain,(
  p(s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0))))) | $spl95),
  inference(unit_resulting_resolution,[],[f102,f344,f140])).
fof(f158,plain,(
  s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) = s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) | $spl6),
  inference(cnf_transformation,[],[f158_D])).
fof(f158_D,plain,(
  s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) = s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1158,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl157),
  inference(cnf_transformation,[],[f1158_D])).
fof(f1158_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f1192,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) = s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) | $spl88),
  inference(superposition,[],[f146,f337])).
fof(f146,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK3(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK3(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK3(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK3(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X12,X6] : ! [X13,X14] : (s(fun(X6,X12),X13) = s(fun(X6,X12),X14) <=> ! [X11] : s(X12,happ(s(fun(X6,X12),X13),s(X6,X11))) = s(X12,happ(s(fun(X6,X12),X14),s(X6,X11))))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','thm.bool.FUN_EQ_THM')).
fof(f1162,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl159),
  inference(cnf_transformation,[],[f1162_D])).
fof(f1162_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f1160,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | ($spl25 | $spl90)),
  inference(backward_demodulation,[],[f338,f185])).
fof(f338,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl90),
  inference(cnf_transformation,[],[f338_D])).
fof(f338_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f1159,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl156),
  inference(cnf_transformation,[],[f1159_D])).
fof(f1159_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f1155,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl91),
  inference(trivial_inequality_removal,[],[f1154])).
fof(f1154,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl91),
  inference(superposition,[],[f339,f103])).
fof(f1157,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl154),
  inference(cnf_transformation,[],[f1157_D])).
fof(f1157_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f1149,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl91),
  inference(unit_resulting_resolution,[],[f339,f112])).
fof(f1150,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl91),
  inference(unit_resulting_resolution,[],[f339,f110])).
fof(f1151,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl91),
  inference(unit_resulting_resolution,[],[f339,f103])).
fof(f1152,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl91),
  inference(unit_resulting_resolution,[],[f102,f339,f140])).
fof(f1153,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl91),
  inference(unit_resulting_resolution,[],[f102,f339,f140])).
fof(f1052,plain,(
  p(s(bool,'F_0')) | $spl152),
  inference(cnf_transformation,[],[f1052_D])).
fof(f1052_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f1050,plain,(
  p(s(bool,'F_0')) | ($spl86 | $spl144)),
  inference(backward_demodulation,[],[f331,f1033])).
fof(f331,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl86),
  inference(cnf_transformation,[],[f331_D])).
fof(f331_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1047,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl151),
  inference(cnf_transformation,[],[f1047_D])).
fof(f1047_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f1049,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | ($spl21 | $spl86)),
  inference(backward_demodulation,[],[f331,f181])).
fof(f1046,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl149),
  inference(cnf_transformation,[],[f1046_D])).
fof(f1046_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f1039,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl85),
  inference(unit_resulting_resolution,[],[f329,f108])).
fof(f1040,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl85),
  inference(unit_resulting_resolution,[],[f329,f106])).
fof(f1048,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl150),
  inference(cnf_transformation,[],[f1048_D])).
fof(f1048_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f1041,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl85),
  inference(unit_resulting_resolution,[],[f329,f103])).
fof(f1042,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl85),
  inference(unit_resulting_resolution,[],[f97,f329,f143])).
fof(f1043,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl85),
  inference(unit_resulting_resolution,[],[f97,f329,f143])).
fof(f329,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl85),
  inference(cnf_transformation,[],[f329_D])).
fof(f329_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f1037,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | ($spl21 | $spl146)),
  inference(backward_demodulation,[],[f1035,f181])).
fof(f1035,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl146),
  inference(cnf_transformation,[],[f1035_D])).
fof(f1035_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f1031,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl87),
  inference(trivial_inequality_removal,[],[f1030])).
fof(f1030,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl87),
  inference(superposition,[],[f332,f103])).
fof(f1033,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl144),
  inference(cnf_transformation,[],[f1033_D])).
fof(f1033_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f1025,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl87),
  inference(unit_resulting_resolution,[],[f332,f112])).
fof(f1026,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl87),
  inference(unit_resulting_resolution,[],[f332,f110])).
fof(f1027,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl87),
  inference(unit_resulting_resolution,[],[f332,f103])).
fof(f1028,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl87),
  inference(unit_resulting_resolution,[],[f102,f332,f140])).
fof(f1029,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool))))) | $spl87),
  inference(unit_resulting_resolution,[],[f102,f332,f140])).
fof(f154,plain,(
  s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) = s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) | $spl2),
  inference(cnf_transformation,[],[f154_D])).
fof(f154_D,plain,(
  s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) = s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f925,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl141),
  inference(cnf_transformation,[],[f925_D])).
fof(f925_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f948,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) = s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) | $spl80),
  inference(superposition,[],[f146,f323])).
fof(f929,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl143),
  inference(cnf_transformation,[],[f929_D])).
fof(f929_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f927,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | ($spl17 | $spl82)),
  inference(backward_demodulation,[],[f324,f173])).
fof(f324,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl82),
  inference(cnf_transformation,[],[f324_D])).
fof(f324_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f926,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl140),
  inference(cnf_transformation,[],[f926_D])).
fof(f926_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f922,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl83),
  inference(trivial_inequality_removal,[],[f921])).
fof(f921,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl83),
  inference(superposition,[],[f325,f103])).
fof(f924,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) | $spl138),
  inference(cnf_transformation,[],[f924_D])).
fof(f924_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f916,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) | $spl83),
  inference(unit_resulting_resolution,[],[f325,f112])).
fof(f917,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) | $spl83),
  inference(unit_resulting_resolution,[],[f325,f110])).
fof(f918,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl83),
  inference(unit_resulting_resolution,[],[f325,f103])).
fof(f919,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) | $spl83),
  inference(unit_resulting_resolution,[],[f102,f325,f140])).
fof(f920,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) | $spl83),
  inference(unit_resulting_resolution,[],[f102,f325,f140])).
fof(f805,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl135),
  inference(cnf_transformation,[],[f805_D])).
fof(f805_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f828,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) = s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) | $spl76),
  inference(superposition,[],[f146,f316])).
fof(f809,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl137),
  inference(cnf_transformation,[],[f809_D])).
fof(f809_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f807,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | ($spl13 | $spl78)),
  inference(backward_demodulation,[],[f317,f169])).
fof(f317,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl78),
  inference(cnf_transformation,[],[f317_D])).
fof(f317_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f806,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl134),
  inference(cnf_transformation,[],[f806_D])).
fof(f806_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f802,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl79),
  inference(trivial_inequality_removal,[],[f801])).
fof(f801,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl79),
  inference(superposition,[],[f318,f103])).
fof(f804,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) | $spl132),
  inference(cnf_transformation,[],[f804_D])).
fof(f804_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f796,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) | $spl79),
  inference(unit_resulting_resolution,[],[f318,f112])).
fof(f797,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) | $spl79),
  inference(unit_resulting_resolution,[],[f318,f110])).
fof(f798,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl79),
  inference(unit_resulting_resolution,[],[f318,f103])).
fof(f799,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) | $spl79),
  inference(unit_resulting_resolution,[],[f102,f318,f140])).
fof(f800,plain,(
  p(s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool))))) | $spl79),
  inference(unit_resulting_resolution,[],[f102,f318,f140])).
fof(f696,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl129),
  inference(cnf_transformation,[],[f696_D])).
fof(f696_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f719,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) = s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) | $spl54),
  inference(superposition,[],[f144,f243])).
fof(f700,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl131),
  inference(cnf_transformation,[],[f700_D])).
fof(f700_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f698,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | ($spl23 | $spl56)),
  inference(backward_demodulation,[],[f244,f183])).
fof(f244,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl56),
  inference(cnf_transformation,[],[f244_D])).
fof(f244_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f697,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl128),
  inference(cnf_transformation,[],[f697_D])).
fof(f697_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f693,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl57),
  inference(trivial_inequality_removal,[],[f692])).
fof(f692,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl57),
  inference(superposition,[],[f245,f103])).
fof(f695,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) | $spl126),
  inference(cnf_transformation,[],[f695_D])).
fof(f695_D,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f687,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) | $spl57),
  inference(unit_resulting_resolution,[],[f245,f112])).
fof(f688,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) | $spl57),
  inference(unit_resulting_resolution,[],[f245,f110])).
fof(f689,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl57),
  inference(unit_resulting_resolution,[],[f245,f103])).
fof(f690,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) | $spl57),
  inference(unit_resulting_resolution,[],[f102,f245,f140])).
fof(f691,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) | $spl57),
  inference(unit_resulting_resolution,[],[f102,f245,f140])).
fof(f587,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl123),
  inference(cnf_transformation,[],[f587_D])).
fof(f587_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f610,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) = s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) | $spl50),
  inference(superposition,[],[f144,f238])).
fof(f591,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl125),
  inference(cnf_transformation,[],[f591_D])).
fof(f591_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f589,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | ($spl19 | $spl52)),
  inference(backward_demodulation,[],[f239,f179])).
fof(f239,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl52),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f588,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl122),
  inference(cnf_transformation,[],[f588_D])).
fof(f588_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f584,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl53),
  inference(trivial_inequality_removal,[],[f583])).
fof(f583,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl53),
  inference(superposition,[],[f240,f103])).
fof(f586,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) | $spl120),
  inference(cnf_transformation,[],[f586_D])).
fof(f586_D,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f578,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) | $spl53),
  inference(unit_resulting_resolution,[],[f240,f112])).
fof(f579,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) | $spl53),
  inference(unit_resulting_resolution,[],[f240,f110])).
fof(f580,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl53),
  inference(unit_resulting_resolution,[],[f240,f103])).
fof(f581,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) | $spl53),
  inference(unit_resulting_resolution,[],[f102,f240,f140])).
fof(f582,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0))))) | $spl53),
  inference(unit_resulting_resolution,[],[f102,f240,f140])).
fof(f478,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl117),
  inference(cnf_transformation,[],[f478_D])).
fof(f478_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f501,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) = s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) | $spl46),
  inference(superposition,[],[f144,f233])).
fof(f482,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl119),
  inference(cnf_transformation,[],[f482_D])).
fof(f482_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f480,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | ($spl15 | $spl48)),
  inference(backward_demodulation,[],[f234,f171])).
fof(f234,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl48),
  inference(cnf_transformation,[],[f234_D])).
fof(f234_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f479,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl116),
  inference(cnf_transformation,[],[f479_D])).
fof(f479_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f475,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl49),
  inference(trivial_inequality_removal,[],[f474])).
fof(f474,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl49),
  inference(superposition,[],[f235,f103])).
fof(f477,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) | $spl114),
  inference(cnf_transformation,[],[f477_D])).
fof(f477_D,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f469,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) | $spl49),
  inference(unit_resulting_resolution,[],[f235,f112])).
fof(f470,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) | $spl49),
  inference(unit_resulting_resolution,[],[f235,f110])).
fof(f471,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl49),
  inference(unit_resulting_resolution,[],[f235,f103])).
fof(f472,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) | $spl49),
  inference(unit_resulting_resolution,[],[f102,f235,f140])).
fof(f473,plain,(
  p(s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) | $spl49),
  inference(unit_resulting_resolution,[],[f102,f235,f140])).
fof(f354,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl99),
  inference(cnf_transformation,[],[f354_D])).
fof(f354_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f392,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) = s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) | $spl42),
  inference(superposition,[],[f144,f228])).
fof(f372,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl110),
  inference(cnf_transformation,[],[f372_D])).
fof(f372_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f374,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl113),
  inference(cnf_transformation,[],[f374_D])).
fof(f374_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f370,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl39),
  inference(superposition,[],[f220,f103])).
fof(f367,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl106),
  inference(cnf_transformation,[],[f367_D])).
fof(f367_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f369,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl109),
  inference(cnf_transformation,[],[f369_D])).
fof(f369_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f365,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl35),
  inference(superposition,[],[f213,f103])).
fof(f362,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl102),
  inference(cnf_transformation,[],[f362_D])).
fof(f362_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f364,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl105),
  inference(cnf_transformation,[],[f364_D])).
fof(f364_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f360,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl31),
  inference(superposition,[],[f200,f103])).
fof(f358,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl101),
  inference(cnf_transformation,[],[f358_D])).
fof(f358_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f356,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | ($spl11 | $spl44)),
  inference(backward_demodulation,[],[f229,f167])).
fof(f229,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl44),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f355,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl98),
  inference(cnf_transformation,[],[f355_D])).
fof(f355_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f351,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl45),
  inference(trivial_inequality_removal,[],[f350])).
fof(f350,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl45),
  inference(superposition,[],[f230,f103])).
fof(f353,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) | $spl96),
  inference(cnf_transformation,[],[f353_D])).
fof(f353_D,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f345,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) | $spl45),
  inference(unit_resulting_resolution,[],[f230,f112])).
fof(f346,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) | $spl45),
  inference(unit_resulting_resolution,[],[f230,f110])).
fof(f347,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl45),
  inference(unit_resulting_resolution,[],[f230,f103])).
fof(f348,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) | $spl45),
  inference(unit_resulting_resolution,[],[f102,f230,f140])).
fof(f349,plain,(
  p(s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0))))) | $spl45),
  inference(unit_resulting_resolution,[],[f102,f230,f140])).
fof(f342,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl92),
  inference(cnf_transformation,[],[f342_D])).
fof(f342_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f344,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl95),
  inference(cnf_transformation,[],[f344_D])).
fof(f344_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f340,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl27),
  inference(superposition,[],[f193,f103])).
fof(f337,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl88),
  inference(cnf_transformation,[],[f337_D])).
fof(f337_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f339,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl91),
  inference(cnf_transformation,[],[f339_D])).
fof(f339_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f335,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl25),
  inference(superposition,[],[f185,f103])).
fof(f330,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl84),
  inference(cnf_transformation,[],[f330_D])).
fof(f330_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f328,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl21),
  inference(superposition,[],[f181,f103])).
fof(f323,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl80),
  inference(cnf_transformation,[],[f323_D])).
fof(f323_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f325,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl83),
  inference(cnf_transformation,[],[f325_D])).
fof(f325_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f321,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl17),
  inference(superposition,[],[f173,f103])).
fof(f316,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl76),
  inference(cnf_transformation,[],[f316_D])).
fof(f316_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f318,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl79),
  inference(cnf_transformation,[],[f318_D])).
fof(f318_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f314,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl13),
  inference(superposition,[],[f169,f103])).
fof(f309,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl74),
  inference(cnf_transformation,[],[f309_D])).
fof(f309_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f305,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl41),
  inference(trivial_inequality_removal,[],[f304])).
fof(f304,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl41),
  inference(superposition,[],[f225,f103])).
fof(f307,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) | $spl72),
  inference(cnf_transformation,[],[f307_D])).
fof(f307_D,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f299,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) | $spl41),
  inference(unit_resulting_resolution,[],[f225,f112])).
fof(f300,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) | $spl41),
  inference(unit_resulting_resolution,[],[f225,f110])).
fof(f301,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl41),
  inference(unit_resulting_resolution,[],[f225,f103])).
fof(f302,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) | $spl41),
  inference(unit_resulting_resolution,[],[f102,f225,f140])).
fof(f303,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) | $spl41),
  inference(unit_resulting_resolution,[],[f102,f225,f140])).
fof(f296,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl70),
  inference(cnf_transformation,[],[f296_D])).
fof(f296_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f292,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl37),
  inference(trivial_inequality_removal,[],[f291])).
fof(f291,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl37),
  inference(superposition,[],[f218,f103])).
fof(f294,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) | $spl68),
  inference(cnf_transformation,[],[f294_D])).
fof(f294_D,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f286,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) | $spl37),
  inference(unit_resulting_resolution,[],[f218,f112])).
fof(f287,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) | $spl37),
  inference(unit_resulting_resolution,[],[f218,f110])).
fof(f288,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl37),
  inference(unit_resulting_resolution,[],[f218,f103])).
fof(f289,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) | $spl37),
  inference(unit_resulting_resolution,[],[f102,f218,f140])).
fof(f290,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool))))) | $spl37),
  inference(unit_resulting_resolution,[],[f102,f218,f140])).
fof(f283,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl66),
  inference(cnf_transformation,[],[f283_D])).
fof(f283_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f279,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl33),
  inference(trivial_inequality_removal,[],[f278])).
fof(f278,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl33),
  inference(superposition,[],[f205,f103])).
fof(f281,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl64),
  inference(cnf_transformation,[],[f281_D])).
fof(f281_D,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f273,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl33),
  inference(unit_resulting_resolution,[],[f205,f112])).
fof(f274,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl33),
  inference(unit_resulting_resolution,[],[f205,f110])).
fof(f275,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl33),
  inference(unit_resulting_resolution,[],[f205,f103])).
fof(f276,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl33),
  inference(unit_resulting_resolution,[],[f102,f205,f140])).
fof(f277,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl33),
  inference(unit_resulting_resolution,[],[f102,f205,f140])).
fof(f259,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl63),
  inference(cnf_transformation,[],[f259_D])).
fof(f259_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f257,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl29 | $spl60)),
  inference(backward_demodulation,[],[f256,f198])).
fof(f256,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl60),
  inference(cnf_transformation,[],[f256_D])).
fof(f256_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f252,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl29),
  inference(trivial_inequality_removal,[],[f251])).
fof(f251,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl29),
  inference(superposition,[],[f198,f103])).
fof(f254,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl58),
  inference(cnf_transformation,[],[f254_D])).
fof(f254_D,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f246,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl29),
  inference(unit_resulting_resolution,[],[f198,f112])).
fof(f247,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl29),
  inference(unit_resulting_resolution,[],[f198,f110])).
fof(f248,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl29),
  inference(unit_resulting_resolution,[],[f198,f103])).
fof(f249,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl29),
  inference(unit_resulting_resolution,[],[f102,f198,f140])).
fof(f250,plain,(
  p(s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool))))) | $spl29),
  inference(unit_resulting_resolution,[],[f102,f198,f140])).
fof(f243,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl54),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f245,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl57),
  inference(cnf_transformation,[],[f245_D])).
fof(f245_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f241,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl23),
  inference(superposition,[],[f183,f103])).
fof(f238,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl50),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f240,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl53),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f236,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl19),
  inference(superposition,[],[f179,f103])).
fof(f233,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl46),
  inference(cnf_transformation,[],[f233_D])).
fof(f233_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f235,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl49),
  inference(cnf_transformation,[],[f235_D])).
fof(f235_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f231,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl15),
  inference(superposition,[],[f171,f103])).
fof(f228,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl42),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f230,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl45),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f226,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl11),
  inference(superposition,[],[f167,f103])).
fof(f225,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl41),
  inference(cnf_transformation,[],[f225_D])).
fof(f225_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f223,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl9),
  inference(forward_demodulation,[],[f222,f133])).
fof(f133,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X11,X16] : s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X6,X11),s(X6,X16)))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','thm.relation.EMPTY_REL_DEF')).
fof(f222,plain,(
  s(bool,'const.relation.EMPTY_REL_2'(s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl9),
  inference(forward_demodulation,[],[f221,f134])).
fof(f134,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1))),s(X0,X2)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1))),s(X0,X2)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X17,X18] : s(bool,'const.relation.EMPTY_REL_2'(s(X6,X17),s(X6,X18))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0'),s(X6,X17))),s(X6,X18)))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','ARITY_0_2_const.relation.EMPTY_REL')).
fof(f221,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl9),
  inference(forward_demodulation,[],[f206,f135])).
fof(f135,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))),s(X0,X3))),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))),s(X0,X3))),s(X0,X4)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2),s(X0,X3),s(X0,X4))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))),s(X0,X3))),s(X0,X4)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X17,X18,X19,X20] : s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(X6,fun(X6,bool)),X17),s(fun(X6,bool),X18),s(X6,X19),s(X6,X20))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X6,fun(X6,bool)),X17),s(fun(X6,bool),X18))),s(X6,X19))),s(X6,X20)))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','ARITY_2_4_const.pred_set.REL_RESTRICT')).
fof(f206,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl9),
  inference(unit_resulting_resolution,[],[f161,f146])).
fof(f220,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl39),
  inference(cnf_transformation,[],[f220_D])).
fof(f220_D,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f207,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl9),
  inference(unit_resulting_resolution,[],[f161,f144])).
fof(f218,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl37),
  inference(cnf_transformation,[],[f218_D])).
fof(f218_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f216,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl9),
  inference(forward_demodulation,[],[f215,f133])).
fof(f215,plain,(
  s(bool,'const.relation.EMPTY_REL_2'(s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl9),
  inference(forward_demodulation,[],[f214,f134])).
fof(f214,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl9),
  inference(forward_demodulation,[],[f209,f135])).
fof(f209,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl9),
  inference(unit_resulting_resolution,[],[f161,f146])).
fof(f213,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl35),
  inference(cnf_transformation,[],[f213_D])).
fof(f213_D,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f210,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),bool,sK0)))) | $spl9),
  inference(unit_resulting_resolution,[],[f161,f144])).
fof(f205,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl33),
  inference(cnf_transformation,[],[f205_D])).
fof(f205_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f203,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl5),
  inference(forward_demodulation,[],[f202,f133])).
fof(f202,plain,(
  s(bool,'const.relation.EMPTY_REL_2'(s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl5),
  inference(forward_demodulation,[],[f201,f134])).
fof(f201,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl5),
  inference(forward_demodulation,[],[f186,f135])).
fof(f186,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl5),
  inference(unit_resulting_resolution,[],[f157,f146])).
fof(f200,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl31),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f187,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl5),
  inference(unit_resulting_resolution,[],[f157,f144])).
fof(f198,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl29),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f196,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl5),
  inference(forward_demodulation,[],[f195,f133])).
fof(f195,plain,(
  s(bool,'const.relation.EMPTY_REL_2'(s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl5),
  inference(forward_demodulation,[],[f194,f134])).
fof(f194,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,'const.pred_set.REL_RESTRICT_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl5),
  inference(forward_demodulation,[],[f189,f135])).
fof(f189,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),sK0,bool)))) | $spl5),
  inference(unit_resulting_resolution,[],[f157,f146])).
fof(f193,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl27),
  inference(cnf_transformation,[],[f193_D])).
fof(f193_D,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f190,plain,(
  s(bool,f(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) != s(bool,g(s(sK0,sK2(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool)))),bool,sK0)))) | $spl5),
  inference(unit_resulting_resolution,[],[f157,f144])).
fof(f185,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl25),
  inference(cnf_transformation,[],[f185_D])).
fof(f185_D,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f174,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl7),
  inference(unit_resulting_resolution,[],[f159,f146])).
fof(f183,plain,(
  s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl23),
  inference(cnf_transformation,[],[f183_D])).
fof(f183_D,plain,(
  s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f175,plain,(
  s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl7),
  inference(unit_resulting_resolution,[],[f159,f144])).
fof(f181,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl21),
  inference(cnf_transformation,[],[f181_D])).
fof(f181_D,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f176,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),sK0,bool)))) | $spl7),
  inference(unit_resulting_resolution,[],[f159,f146])).
fof(f179,plain,(
  s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl19),
  inference(cnf_transformation,[],[f179_D])).
fof(f179_D,plain,(
  s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f177,plain,(
  s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0))),bool,sK0)))) | $spl7),
  inference(unit_resulting_resolution,[],[f159,f144])).
fof(f173,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl17),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f162,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl3),
  inference(unit_resulting_resolution,[],[f155,f146])).
fof(f171,plain,(
  s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl15),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f163,plain,(
  s(bool,f(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f155,f144])).
fof(f169,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl13),
  inference(cnf_transformation,[],[f169_D])).
fof(f169_D,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f164,plain,(
  s(bool,happ(s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))),s(sK0,sK3(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),sK0,bool)))) | $spl3),
  inference(unit_resulting_resolution,[],[f155,f146])).
fof(f167,plain,(
  s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f167_D])).
fof(f167_D,plain,(
  s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f165,plain,(
  s(bool,f(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK2(f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0))),bool,sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f155,f144])).
fof(f161,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) != s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) | $spl9),
  inference(cnf_transformation,[],[f161_D])).
fof(f161_D,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) != s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f150,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) != s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) | $spl1),
  inference(unit_resulting_resolution,[],[f149,f146])).
fof(f159,plain,(
  s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) != s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) | $spl7),
  inference(cnf_transformation,[],[f159_D])).
fof(f159_D,plain,(
  s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) != s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f151,plain,(
  s(fun(sK0,bool),f(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) != s(fun(sK0,bool),g(s(sK0,sK2('const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),'const.relation.EMPTY_REL_0',fun(sK0,bool),sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f149,f144])).
fof(f157,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) != s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) | $spl5),
  inference(cnf_transformation,[],[f157_D])).
fof(f157_D,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) != s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f152,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) != s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK3('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0,fun(sK0,bool))))) | $spl1),
  inference(unit_resulting_resolution,[],[f149,f146])).
fof(f155,plain,(
  s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) != s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) | $spl3),
  inference(cnf_transformation,[],[f155_D])).
fof(f155_D,plain,(
  s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) != s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f153,plain,(
  s(fun(sK0,bool),f(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) != s(fun(sK0,bool),g(s(sK0,sK2('const.relation.EMPTY_REL_0','const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),fun(sK0,bool),sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f149,f144])).
fof(f149,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') != s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl1),
  inference(cnf_transformation,[],[f149_D])).
fof(f149_D,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') != s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f95,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') != s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') != s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f64])).
fof(f64,plain,(
  ? [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') != s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1] : s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') = s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') = s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0') = s(fun(X6,fun(X6,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0') = s(fun(X6,fun(X6,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  file('thm.pred_set.REL_RESTRICT_EMPTY','thm.pred_set.REL_RESTRICT_EMPTY')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 99
% Passive clauses: 87
% Generated clauses: 1544
% Final active clauses: 68
% Final passive clauses: 17
% Input formulas: 27
% Initial clauses: 52
% 
% Duplicate literals: 1
% Trivial inequalities: 64
% Fw demodulations: 373
% Bw demodulations: 21
% Global subsumptions: 1
% 
% Simple tautologies: 45
% Equational tautologies: 4
% Forward subsumptions: 803
% 
% Binary resolution: 6
% Unit resulting resolution: 213
% Forward superposition: 53
% Backward superposition: 702
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 447
% SAT solver unit clauses: 4
% SAT solver binary clauses: 350
% 
% Sat splits: 27
% Sat splitting refutations: 2
% 
% Memory used [KB]: 2430
% Time elapsed: 0.052 s
% ------------------------------
% ------------------------------
% Success in time 18.069 s
