Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (61935)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 156
% Passive clauses: 3322
% Generated clauses: 23861
% Final active clauses: 132
% Final passive clauses: 2838
% Input formulas: 29
% Initial clauses: 63
% 
% Duplicate literals: 895
% Trivial inequalities: 236
% Fw subsumption resolutions: 3504
% Fw demodulations: 3774
% Bw demodulations: 22
% 
% Simple tautologies: 415
% Equational tautologies: 45
% Forward subsumptions: 10436
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 976
% Factoring: 4
% Forward superposition: 7765
% Backward superposition: 6151
% Self superposition: 168
% Equality factoring: 100
% Equality resolution: 29
% 
% Split clauses: 586
% Split components: 1196
% SAT solver clauses: 1420
% SAT solver unit clauses: 25
% SAT solver binary clauses: 253
% 
% Sat splits: 586
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1415
% TWLsolver calls for satisfiability: 213
% 
% Memory used [KB]: 11257
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (61962)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 119
% Passive clauses: 1571
% Generated clauses: 15144
% Final active clauses: 116
% Final passive clauses: 1437
% Input formulas: 29
% Initial clauses: 63
% 
% Duplicate literals: 665
% Trivial inequalities: 100
% Fw subsumption resolutions: 1461
% Fw demodulations: 1467
% Bw demodulations: 13
% Condensations: 1021
% Global subsumptions: 5
% 
% Simple tautologies: 189
% Equational tautologies: 37
% Forward subsumptions: 6305
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 1007
% Factoring: 3
% Forward superposition: 4419
% Backward superposition: 3142
% Self superposition: 112
% Equality factoring: 73
% Equality resolution: 13
% 
% Split clauses: 180
% Split components: 370
% SAT solver clauses: 2345
% SAT solver unit clauses: 54
% SAT solver binary clauses: 37
% SAT solver learnt clauses: 10
% 
% Sat splits: 180
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2325
% TWLsolver calls for satisfiability: 281
% 
% Memory used [KB]: 13432
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f50086,plain,(
  $false),
  inference(subsumption_resolution,[],[f49861,f44003])).
fof(f44003,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNIV_0','const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))),
  inference(unit_resulting_resolution,[],[f4615,f4116])).
fof(f4116,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) )),
  inference(subsumption_resolution,[],[f4103,f2312])).
fof(f2312,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f111,f112,f114,f116,f119,f118,f141,f144,f160,f157,f173,f175,f183,f187,f191,f193,f220,f251,f303,f304,f318,f347,f395,f432,f435,f646,f647,f462,f1004,f1037,f1032,f1033,f1034,f1035,f1036,f1008,f1006,f776,f774,f460,f457,f393,f390,f433,f436,f779,f1011,f148,f1030,f793,f649,f279,f156,f155,f154,f153,f152,f151,f150,f149,f147,f394,f461,f305,f1024,f1013,f781,f778,f307,f306,f273,f272,f229,f228,f186,f123,f122,f2311])).
fof(f2311,plain,(
  ( ! [X39] : (~p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f1431,f318])).
fof(f1431,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'F_0')) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(X37,X38),s(fun(X37,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f193,f186])).
fof(f122,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.COMPL_CLAUSES|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f123,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f81])).
fof(f186,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f144,f123])).
fof(f228,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),X1)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f149,f116])).
fof(f229,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f149,f116])).
fof(f272,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f150,f116])).
fof(f273,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),X3)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f150,f116])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,X2,bool))) )),
  inference(superposition,[],[f173,f116])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK3(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f173,f116])).
fof(f778,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X17),s(fun(bool,bool),X18)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),X17)))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f151,f116])).
fof(f781,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,X27),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X28),s(fun(bool,bool),X29)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f151,f116])).
fof(f1013,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,X32),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X33),s(fun(bool,bool),X34)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X34)))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f152,f116])).
fof(f1024,plain,(
  ( ! [X4,X0,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X4)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f1005,f186])).
fof(f1005,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0'))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X4)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f152,f318])).
fof(f305,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0','const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.UNIV_0','const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),bool,sK0))))),
  inference(unit_resulting_resolution,[],[f111,f173])).
fof(f461,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,'T_0') != s(bool,f(s(X65,sK3(X66,X67,bool,X65)))) | s(fun(X65,bool),X66) = s(fun(X65,bool),X67) | s(bool,X68) = s(bool,g(s(X65,sK3(X66,X67,bool,X65)))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f173,f191])).
fof(f394,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X72) != s(bool,f(s(X69,sK3(X70,X71,bool,X69)))) | s(fun(X69,bool),X70) = s(fun(X69,bool),X71) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,g(s(X69,sK3(X70,X71,bool,X69))))) )),
  inference(superposition,[],[f173,f191])).
fof(f147,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X7,X10] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X10) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))),
  file('thm.pred_set.COMPL_CLAUSES|split|1','thm.pred_set.EXTENSION')).
fof(f149,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) & (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X11,X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))))),
  file('thm.pred_set.COMPL_CLAUSES|split|1','thm.pred_set.IN_COMPL')).
fof(f150,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f151,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X7,X10,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X10)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))))),
  file('thm.pred_set.COMPL_CLAUSES|split|1','thm.pred_set.IN_INTER')).
fof(f152,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f153,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f154,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X7,X10,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X10)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))))),
  file('thm.pred_set.COMPL_CLAUSES|split|1','thm.pred_set.IN_UNION')).
fof(f155,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f156,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f279,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) )),
  inference(subsumption_resolution,[],[f276,f114])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) )),
  inference(superposition,[],[f150,f116])).
fof(f649,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'const.bool.IN_2'(s(X38,X39),s(fun(X38,bool),X40)))) | s(bool,X41) = s(bool,'const.bool.IN_2'(s(X38,X39),s(fun(X38,bool),'const.pred_set.COMPL_1'(s(fun(X38,bool),X40))))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f453,f112])).
fof(f453,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X38,X39),s(fun(X38,bool),X40)))) | s(bool,X41) = s(bool,'const.bool.IN_2'(s(X38,X39),s(fun(X38,bool),'const.pred_set.COMPL_1'(s(fun(X38,bool),X40))))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f149,f191])).
fof(f793,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f785,f112])).
fof(f785,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f151,f191])).
fof(f1030,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X9)))) | s(bool,X10) = s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f1019,f112])).
fof(f1019,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X9)))) | s(bool,X10) = s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f152,f191])).
fof(f148,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f98])).
fof(f1011,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,X24),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X25),s(fun(bool,bool),X26)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X26)))) | s(bool,X24) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f152,f191])).
fof(f779,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X20),s(fun(bool,bool),X21)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X20)))) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f151,f191])).
fof(f436,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'const.bool.IN_2'(s(bool,X37),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),X38)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X38)))) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f149,f191])).
fof(f433,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),X31)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X31)))) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f150,f191])).
fof(f390,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(bool,X56),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),X57)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),X57)))) | s(bool,'T_0') = s(bool,X56) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f150,f191])).
fof(f393,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,X63),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),X64)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X62),s(fun(bool,bool),X64)))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f149,f191])).
fof(f457,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),X53)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X52),s(fun(bool,bool),X53)))) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f150,f191])).
fof(f460,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),X60)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X59),s(fun(bool,bool),X60)))) | s(bool,X59) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f149,f191])).
fof(f774,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f151,f191])).
fof(f776,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X10),s(fun(bool,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X10)))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f151,f191])).
fof(f1006,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X7)))) | s(bool,X5) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f152,f191])).
fof(f1008,plain,(
  ( ! [X14,X15,X13,X16] : (~p(s(bool,'const.bool.IN_2'(s(bool,X14),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X15),s(fun(bool,bool),X16)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X13),s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f152,f191])).
fof(f1036,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0))))),
  inference(superposition,[],[f304,f116])).
fof(f1035,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0)))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0))))) )),
  inference(superposition,[],[f304,f191])).
fof(f1034,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f304,f191])).
fof(f1033,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0)))) | s(bool,X3) = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f304,f191])).
fof(f1032,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0)))) | s(bool,X2) = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f304,f191])).
fof(f1037,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0))))),
  inference(forward_demodulation,[],[f1031,f186])).
fof(f1031,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) != s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0))))) )),
  inference(superposition,[],[f304,f318])).
fof(f1004,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) = s(bool,'const.bool.IN_2'(s(X35,X36),s(fun(X35,bool),'const.pred_set.UNIV_0'))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f978,f112])).
fof(f978,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) = s(bool,'const.bool.IN_2'(s(X35,X36),s(fun(X35,bool),'const.pred_set.UNIV_0'))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f646,f183])).
fof(f462,plain,(
  ( ! [X70,X72,X71,X69] : (s(X71,f(s(bool,'T_0'))) != s(X71,g(s(bool,'T_0'))) | s(fun(bool,X71),X69) = s(fun(bool,X71),X70) | s(bool,X72) = s(bool,sK3(X69,X70,X71,bool)) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f173,f191])).
fof(f647,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f445,f112])).
fof(f445,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f193,f191])).
fof(f646,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f444,f112])).
fof(f444,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f160,f191])).
fof(f435,plain,(
  ( ! [X35,X36] : (s(bool,X35) = s(bool,'const.bool.IN_2'(s(bool,X35),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f183,f191])).
fof(f432,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f187,f191])).
fof(f395,plain,(
  ( ! [X76,X74,X75,X73] : (s(X75,f(s(bool,X76))) != s(X75,g(s(bool,X76))) | s(fun(bool,X75),X73) = s(fun(bool,X75),X74) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,sK3(X73,X74,X75,bool))) )),
  inference(superposition,[],[f173,f191])).
fof(f347,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f191,f191])).
fof(f318,plain,(
  ( ! [X8,X7,X9] : (s(bool,X9) = s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f316,f114])).
fof(f316,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,'F_0')) | s(bool,X9) = s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f220,f186])).
fof(f304,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.UNIV_0',bool,sK0))))),
  inference(unit_resulting_resolution,[],[f111,f173])).
fof(f303,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) != s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f300,f112])).
fof(f300,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f193,f183])).
fof(f251,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f186,f116])).
fof(f220,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f217,f114])).
fof(f217,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f157,f116])).
fof(f193,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f122,f116])).
fof(f191,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f116,f116])).
fof(f187,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f144,f118])).
fof(f183,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) )),
  inference(unit_resulting_resolution,[],[f141,f119])).
fof(f175,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f114,f118])).
fof(f173,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.COMPL_CLAUSES|split|1',aHL_EXT)).
fof(f157,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.pred_set.COMPL_CLAUSES|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f160,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f74])).
fof(f144,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X11] : ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.COMPL_CLAUSES|split|1','thm.pred_set.NOT_IN_EMPTY')).
fof(f141,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X11] : p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.UNIV_0'))))),
  file('thm.pred_set.COMPL_CLAUSES|split|1','thm.pred_set.IN_UNIV')).
fof(f118,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.COMPL_CLAUSES|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f119,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f79])).
fof(f116,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.COMPL_CLAUSES|split|1',aHL_BOOL_CASES)).
fof(f114,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.COMPL_CLAUSES|split|1',aHL_FALSITY)).
fof(f112,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.COMPL_CLAUSES|split|1',aHL_TRUTH)).
fof(f111,plain,(
  s(fun(sK0,bool),'const.pred_set.UNIV_0') != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  s(fun(sK0,bool),'const.pred_set.UNIV_0') != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f71])).
fof(f71,plain,(
  ? [X0,X1] : s(fun(X0,bool),'const.pred_set.UNIV_0') != s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ~! [X0,X1] : s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,bool),'const.pred_set.UNIV_0') = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7))),s(fun(X6,bool),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,bool),'const.pred_set.UNIV_0') = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7))),s(fun(X6,bool),X7)))),
  file('thm.pred_set.COMPL_CLAUSES|split|1','thm.pred_set.COMPL_CLAUSES|split|1')).
fof(f4103,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) )),
  inference(superposition,[],[f155,f116])).
fof(f4615,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNIV_0','const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(forward_demodulation,[],[f4581,f183])).
fof(f4581,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNIV_0','const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNIV_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNIV_0','const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(unit_resulting_resolution,[],[f111,f148])).
fof(f49861,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNIV_0','const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))),
  inference(unit_resulting_resolution,[],[f44281,f150])).
fof(f44281,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.UNIV_0','const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))))),
  inference(unit_resulting_resolution,[],[f4615,f4375])).
fof(f4375,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) )),
  inference(subsumption_resolution,[],[f4359,f2312])).
fof(f4359,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) )),
  inference(superposition,[],[f156,f116])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 296
% Passive clauses: 2343
% Generated clauses: 51255
% Final active clauses: 116
% Final passive clauses: 372
% Input formulas: 29
% Initial clauses: 63
% 
% Duplicate literals: 2383
% Trivial inequalities: 424
% Fw subsumption resolutions: 4841
% Fw demodulations: 7009
% Bw demodulations: 1020
% Condensations: 5618
% Global subsumptions: 157
% 
% Simple tautologies: 574
% Equational tautologies: 712
% Forward subsumptions: 26755
% Fw demodulations to eq. taut.: 21
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 4636
% Factoring: 6
% Forward superposition: 12334
% Backward superposition: 10177
% Self superposition: 487
% Equality factoring: 310
% Equality resolution: 59
% 
% Split clauses: 114
% Split components: 246
% SAT solver clauses: 3109
% SAT solver unit clauses: 110
% SAT solver binary clauses: 341
% 
% Sat splits: 114
% Sat splitting refutations: 379
% 
% Memory used [KB]: 17014
% Time elapsed: 2.108 s
% ------------------------------
% ------------------------------
% Success in time 3.428 s
