Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.relation
% (61749)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 103
% Passive clauses: 2448
% Generated clauses: 13454
% Final active clauses: 92
% Final passive clauses: 2242
% Input formulas: 19
% Initial clauses: 46
% 
% Duplicate literals: 617
% Trivial inequalities: 260
% Fw subsumption resolutions: 1527
% Fw demodulations: 1465
% Bw demodulations: 10
% 
% Simple tautologies: 299
% Equational tautologies: 45
% Forward subsumptions: 4705
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 106
% Forward superposition: 3825
% Backward superposition: 5041
% Self superposition: 216
% Equality factoring: 145
% Equality resolution: 29
% 
% Split clauses: 758
% Split components: 1591
% SAT solver clauses: 1148
% SAT solver unit clauses: 17
% SAT solver binary clauses: 124
% 
% Sat splits: 758
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1148
% TWLsolver calls for satisfiability: 109
% 
% Memory used [KB]: 12792
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (61772)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 98
% Passive clauses: 962
% Generated clauses: 10044
% Final active clauses: 71
% Final passive clauses: 401
% Input formulas: 19
% Initial clauses: 46
% 
% Duplicate literals: 508
% Trivial inequalities: 138
% Fw subsumption resolutions: 830
% Fw demodulations: 639
% Bw demodulations: 466
% Condensations: 1138
% Global subsumptions: 4
% 
% Simple tautologies: 181
% Equational tautologies: 18
% Forward subsumptions: 4428
% Bw demodulations to eq. taut.: 17
% 
% Binary resolution: 44
% Forward superposition: 2903
% Backward superposition: 2089
% Self superposition: 121
% Equality factoring: 47
% Equality resolution: 35
% 
% Split clauses: 74
% Split components: 157
% SAT solver clauses: 1238
% SAT solver unit clauses: 50
% SAT solver binary clauses: 50
% SAT solver learnt clauses: 8
% 
% Sat splits: 74
% Sat splitting refutations: 17
% 
% TWLsolver clauses: 1222
% TWLsolver calls for satisfiability: 206
% 
% Memory used [KB]: 10874
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (61788)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 370
% Passive clauses: 11086
% Generated clauses: 88878
% Final active clauses: 61
% Final passive clauses: 759
% Input formulas: 19
% Initial clauses: 46
% Discarded non-redundant clauses: 815
% 
% Duplicate literals: 4848
% Trivial inequalities: 577
% Fw subsumption resolutions: 4341
% Fw demodulations: 7769
% Bw demodulations: 8234
% Condensations: 10406
% Global subsumptions: 95
% 
% Simple tautologies: 645
% Equational tautologies: 1555
% Forward subsumptions: 47237
% Fw demodulations to eq. taut.: 8
% Bw demodulations to eq. taut.: 249
% 
% Unit resulting resolution: 4091
% Forward superposition: 23295
% Backward superposition: 12256
% Self superposition: 502
% Equality factoring: 218
% Equality resolution: 35
% 
% Split clauses: 203
% Split components: 423
% SAT solver clauses: 11672
% SAT solver unit clauses: 563
% SAT solver binary clauses: 1719
% 
% Sat splits: 203
% Sat splitting refutations: 83
% 
% Memory used [KB]: 31086
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f23728,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f128,f129,f130,f131,f132,f134,f146,f135,f139,f144,f153,f154,f155,f156,f157,f158,f160,f165,f177,f198,f194,f208,f236,f224,f239,f241,f261,f262,f291,f288,f287,f286,f300,f301,f302,f303,f314,f383,f386,f388,f389,f393,f395,f398,f384,f400,f402,f406,f407,f411,f413,f417,f385,f433,f435,f442,f443,f444,f445,f446,f448,f147,f449,f453,f495,f502,f503,f524,f526,f527,f529,f531,f536,f560,f717,f719,f720,f721,f724,f728,f777,f779,f780,f781,f784,f788,f807,f793,f810,f808,f854,f856,f853,f857,f1021,f1104,f1106,f1107,f1109,f1111,f1116,f1157,f1158,f1161,f1163,f1166,f1167,f1168,f1170,f1171,f1172,f1175,f1177,f1178,f1179,f1180,f1182,f1186,f1187,f1188,f1190,f1200,f1201,f1202,f1204,f1221,f1222,f1225,f1227,f1230,f1231,f1232,f1234,f1235,f1236,f1239,f1241,f1242,f1243,f1244,f1246,f1250,f1251,f1252,f1254,f1264,f1265,f1266,f1268,f1330,f1332,f1333,f1334,f1337,f1341,f1358,f1346,f1361,f1359,f1556,f1558,f1593,f1594,f1591,f1595,f1824,f1826,f1827,f1828,f1831,f1835,f1854,f1840,f1857,f1855,f1862,f1864,f1865,f1866,f1869,f1873,f1962,f1963,f1966,f1968,f1971,f1972,f1973,f1975,f1976,f1977,f1980,f1982,f1983,f1984,f1985,f1987,f1991,f1992,f1993,f1995,f2005,f2006,f2007,f2009,f2047,f2048,f2051,f2053,f2056,f2057,f2058,f2060,f2061,f2062,f2065,f2067,f2068,f2069,f2070,f2072,f2076,f2077,f2078,f2080,f2090,f2091,f2092,f2094,f2189,f2191,f2192,f2193,f2196,f2200,f2221,f2205,f2224,f2222,f2277,f2279,f2280,f2282,f2284,f2289,f2313,f2315,f2316,f2318,f2320,f2325,f2383,f2384,f2387,f2389,f2392,f2393,f2394,f2396,f2397,f2398,f2401,f2403,f2404,f2405,f2406,f2408,f2412,f2413,f2414,f2416,f2426,f2427,f2428,f2430,f2468,f2469,f2472,f2474,f2477,f2478,f2479,f2481,f2482,f2483,f2486,f2488,f2489,f2490,f2491,f2493,f2497,f2498,f2499,f2501,f2511,f2512,f2513,f2515,f2673,f2675,f2676,f2678,f2680,f2685,f2745,f2747,f2748,f2749,f2752,f2756,f2773,f2775,f2776,f2777,f2780,f2784,f2924,f2925,f2928,f2930,f2933,f2934,f2935,f2937,f2938,f2939,f2942,f2944,f2945,f2946,f2947,f2949,f2953,f2954,f2955,f2957,f2967,f2968,f2969,f2971,f3029,f3030,f3033,f3035,f3038,f3039,f3040,f3042,f3043,f3044,f3047,f3049,f3050,f3051,f3052,f3054,f3058,f3059,f3060,f3062,f3072,f3073,f3074,f3076,f3190,f3192,f3193,f3195,f3197,f3202,f3493,f3497,f3500,f3502,f3506,f3508,f3512,f3514,f3517,f3518,f3524,f3530,f3533,f3538,f3543,f3547,f3564,f3568,f3612,f178,f3613,f3614,f3600,f3617,f3967,f3971,f3974,f3976,f3980,f3982,f3986,f3988,f3991,f3992,f3995,f4001,f4004,f4009,f4014,f4018,f4035,f4038,f3611,f3608,f3609,f3607,f3615,f4062,f4065,f4280,f4281,f4298,f4302,f4305,f4375,f101,f4377,f4378,f4380,f4381,f4383,f4385,f4389,f496,f4405,f4407,f4408,f4409,f4412,f4415,f470,f4436,f240,f4435,f4446,f3610,f4449,f4451,f4452,f4453,f4456,f4459,f745,f4467,f4469,f4470,f4472,f4474,f4478,f1285,f4515,f4516,f4519,f4521,f4524,f4525,f4526,f4528,f4529,f4530,f4533,f4535,f4536,f4537,f4538,f4540,f4544,f4545,f4546,f4548,f4556,f4557,f4558,f4560,f4588,f4589,f4592,f4594,f4597,f4598,f4599,f4601,f4602,f4603,f4606,f4608,f4609,f4610,f4611,f4613,f4617,f4618,f4619,f4621,f4629,f4630,f4631,f1289,f4633,f1290,f4692,f4694,f4695,f4696,f4699,f4702,f805,f4841,f4842,f4843,f4844,f4845,f4847,f4848,f4849,f4852,f4855,f1356,f5060,f5062,f5063,f5064,f5067,f5070,f1852,f5322,f5323,f5326,f5328,f5331,f5332,f5333,f5335,f5336,f5337,f5340,f5342,f5343,f5344,f5345,f5347,f5351,f5352,f5353,f5355,f5363,f5364,f5365,f5367,f5409,f5410,f5413,f5415,f5418,f5419,f5420,f5422,f5423,f5424,f5427,f5429,f5430,f5431,f5432,f5434,f5438,f5439,f5440,f5442,f5450,f5451,f5452,f2133,f5454,f2134,f5516,f5518,f5519,f5520,f5523,f5526,f1898,f5606,f5607,f5610,f5612,f5615,f5616,f5617,f5619,f5620,f5621,f5624,f5626,f5627,f5628,f5629,f5631,f5635,f5636,f5637,f5639,f5647,f5648,f5649,f5651,f5690,f5691,f5694,f5696,f5699,f5700,f5701,f5703,f5704,f5705,f5708,f5710,f5711,f5712,f5713,f5715,f5719,f5720,f5721,f5723,f5731,f5732,f5733,f5735,f2554,f5794,f5796,f5797,f5798,f5801,f5804,f2219,f5854,f5856,f5857,f5859,f5861,f5865,f5987,f5989,f5990,f5992,f5994,f5998,f6057,f6059,f6060,f6062,f6064,f6068,f2339,f6085,f6087,f6088,f6089,f6092,f6095,f6111,f6113,f6114,f6115,f6118,f6121,f6260,f6261,f6264,f6266,f6269,f6270,f6271,f6273,f6274,f6275,f6278,f6280,f6281,f6282,f6283,f6285,f6289,f6290,f6291,f6293,f6301,f6302,f6303,f6305,f6365,f6366,f6369,f6371,f6374,f6375,f6376,f6378,f6379,f6380,f6383,f6385,f6386,f6387,f6388,f6390,f6394,f6395,f6396,f6398,f6406,f6407,f6408,f3129,f6410,f3130,f6755,f6759,f6762,f6764,f6768,f6770,f2769,f6774,f3513,f6775,f6777,f6778,f3515,f6783,f6786,f6789,f6791,f6794,f6796,f2308,f6810,f3565,f6812,f3569,f7157,f7161,f7164,f7166,f7170,f7172,f2826,f7176,f3987,f7177,f7179,f7180,f3989,f7183,f7186,f7189,f7191,f7194,f7196,f2733,f7210,f4036,f7212,f4039,f7214,f7216,f7217,f7219,f7221,f7225,f3216,f7349,f7358,f7359,f7361,f7362,f7532,f7534,f7535,f7537,f7539,f7543,f2553,f7614,f7616,f7617,f7619,f7621,f7625,f7648,f7653,f7655,f7656,f7658,f7660,f7664,f7703,f7714,f7716,f7717,f7719,f7721,f7725,f7787,f7798,f7800,f7801,f7803,f7805,f7809,f7849,f7857,f7859,f7860,f7862,f7864,f7868,f7909,f7921,f7923,f7924,f7926,f7928,f7932,f7946,f7949,f7951,f7952,f7954,f7956,f7960,f8019,f8046,f8048,f8049,f8051,f8053,f8057,f8116,f8143,f8145,f8146,f8148,f8150,f8154,f8194,f8216,f8218,f8219,f8221,f8223,f8227,f8268,f8286,f8288,f8289,f8291,f8293,f8297,f8356,f8366,f8368,f8369,f8371,f8373,f8377,f8391,f8403,f8405,f8406,f8408,f8410,f8414,f8473,f8484,f8486,f8487,f8489,f8491,f8495,f8535,f8557,f8559,f8560,f8562,f8564,f8568,f8609,f8908,f8909,f8913,f8914,f8952,f8954,f8955,f8957,f8959,f8963,f9024,f9053,f9055,f9056,f9058,f9060,f9064,f9107,f9123,f9125,f9126,f9128,f9130,f9134,f9195,f9509,f9510,f9512,f9514,f9517,f9518,f9515,f9565,f9567,f9568,f9570,f9572,f9576,f9618,f9628,f9630,f9631,f9633,f9635,f9639,f9682,f9697,f9699,f9700,f9702,f9704,f9708,f9771,f9783,f9785,f9786,f9788,f9790,f9794,f9857,f9869,f9871,f9872,f9873,f9876,f9879,f9933,f9935,f9936,f10022,f10024,f10025,f10027,f10029,f10033,f10057,f10065,f10067,f10068,f10069,f10072,f10075,f10093,f10096,f10098,f10099,f10101,f10103,f10107,f10150,f10161,f10163,f10164,f10166,f10168,f10172,f10239,f11031,f11033,f11034,f11036,f11038,f11042,f11085,f11111,f11113,f11114,f11116,f11118,f11122,f11166,f11297,f11299,f11300,f11302,f11304,f11308,f11328,f11332,f11334,f11335,f11337,f11339,f11343,f11386,f11397,f11399,f11400,f11402,f11404,f11408,f11452,f11468,f11470,f11471,f11473,f11475,f11479,f11522,f11533,f11535,f11536,f11538,f11540,f11544,f11588,f11604,f11606,f11607,f11609,f11611,f11615,f11679,f11711,f11713,f11714,f11716,f11718,f11722,f11786,f11818,f11820,f11821,f11823,f11825,f11829,f11872,f11898,f11900,f11901,f11903,f11905,f11909,f11953,f11974,f11976,f11977,f11979,f11981,f11985,f12049,f12062,f12064,f12065,f12067,f12069,f12073,f12137,f12231,f12232,f12235,f12237,f12240,f12241,f12242,f12244,f12245,f12246,f12249,f12251,f12252,f12253,f12254,f12256,f12260,f12261,f12262,f12264,f12272,f12273,f12274,f12276,f12418,f12419,f12422,f12424,f12427,f12428,f12429,f12431,f12432,f12433,f12436,f12438,f12439,f12440,f12441,f12443,f12447,f12448,f12449,f12451,f12459,f12616,f12460,f12615,f12461,f12612,f12463,f12613,f12769,f12770,f12773,f12775,f12778,f12779,f12780,f12782,f12783,f12784,f12787,f12789,f12790,f12791,f12792,f12794,f12798,f12799,f12800,f12802,f12810,f12811,f12812,f12814,f12959,f12960,f12963,f12965,f12968,f12969,f12970,f12972,f12973,f12974,f12977,f12979,f12980,f12981,f12982,f12984,f12988,f12989,f12990,f12992,f13000,f13159,f13001,f13158,f13002,f13155,f13004,f13156,f13322,f13323,f13326,f13328,f13331,f13332,f13333,f13335,f13336,f13337,f13340,f13342,f13343,f13344,f13345,f13347,f13351,f13352,f13353,f13355,f13363,f13364,f13365,f13367,f13519,f13520,f13523,f13525,f13528,f13529,f13530,f13532,f13533,f13534,f13537,f13539,f13540,f13541,f13542,f13544,f13548,f13549,f13550,f13552,f13560,f13721,f13561,f13720,f13562,f13717,f13564,f13718,f13892,f13893,f13896,f13898,f13901,f13902,f13903,f13905,f13906,f13907,f13910,f13912,f13913,f13914,f13915,f13917,f13921,f13922,f13923,f13925,f13933,f13934,f13935,f13937,f14092,f14093,f14096,f14098,f14101,f14102,f14103,f14105,f14106,f14107,f14110,f14112,f14113,f14114,f14115,f14117,f14121,f14122,f14123,f14125,f14133,f14296,f14134,f14295,f14135,f14292,f14137,f14293,f14748,f237,f14749,f14796,f14797,f14803,f14804,f14897,f14899,f14900,f14901,f14904,f14907,f14986,f15441,f15442,f15444,f7350,f9513,f15445,f15448,f15449,f15446,f15576,f15578,f15579,f15580,f15583,f15586,f15637,f15677,f15679,f15680,f15681,f15684,f15687,f15739,f15789,f15791,f15792,f15793,f15796,f15799,f15887,f15962,f15964,f15965,f15966,f15969,f15972,f16057,f16127,f16129,f16130,f16131,f16134,f16137,f16188,f16252,f16254,f16255,f16256,f16259,f16262,f16314,f16376,f16378,f16379,f16380,f16383,f16386,f16471,f16517,f16519,f16520,f16521,f16524,f16527,f16614,f16739,f16741,f16742,f16743,f16746,f16749,f16835,f16983,f16985,f16986,f16987,f16990,f16993,f17078,f17329,f17331,f17332,f17333,f17336,f17339,f17389,f17568,f17570,f17571,f17572,f17575,f17578,f17629,f17745,f17747,f17748,f17749,f17752,f17755,f17840,f17947,f17949,f17950,f17951,f17954,f17957,f18043,f18172,f18174,f18175,f18176,f18179,f18182,f18232,f18360,f18362,f18363,f18364,f18367,f18370,f18421,f18533,f18535,f18536,f18538,f18540,f18544,f18633,f18768,f18770,f18771,f18773,f18775,f18779,f18868,f19155,f19157,f19158,f19160,f19162,f19166,f19255,f19277,f19279,f19280,f19281,f19284,f19287,f19297,f19377,f19379,f19380,f19382,f19384,f19388,f19477,f19614,f19616,f19617,f19618,f19621,f19624,f19634,f19765,f19767,f19768,f19770,f19772,f19776,f19832,f19874,f19876,f19877,f19878,f19881,f19884,f19894,f19911,f19913,f19914,f19916,f19918,f19922,f19979,f20015,f20017,f20018,f20019,f20022,f20025,f20035,f20052,f20054,f20055,f20057,f20059,f20063,f20119,f20146,f20148,f20149,f20151,f20153,f20157,f20214,f20301,f20303,f20304,f20305,f20308,f20311,f20322,f20324,f20328,f20326,f20673,f20681,f21169,f21172,f21175,f21176,f21675,f20674,f21682,f20682,f21842,f21844,f21845,f145,f21846,f149,f21849,f148,f21852,f21863,f22215,f22216,f22225,f3566,f22226,f21865,f22502,f22500,f22836,f21170,f22838,f22840,f22841,f21173,f23342,f23343,f23346,f855,f23349,f471,f23350,f23347,f3519,f23661,f150,f23669,f23675,f23670,f23677,f23678,f23679,f23662,f23680,f23720,f23725])).
fof(f23725,plain,(
  $false | ($spl32 | $spl70)),
  inference(trivial_inequality_removal,[],[f23693])).
fof(f23693,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl32 | $spl70)),
  inference(superposition,[],[f23684,f180])).
fof(f180,plain,(
  s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))),
  inference(subsumption_resolution,[],[f173,f126])).
fof(f126,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f98])).
fof(f98,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X7] : ! [X11] : (s(X7,X11) = s(X7,X11) <=> p(s(bool,'T_0')))),
  file('thm.relation.symmetric_inv_image','thm.bool.REFL_CLAUSE')).
fof(f173,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))),
  inference(superposition,[],[f80,f122])).
fof(f122,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.symmetric_inv_image',aHL_BOOL_CASES)).
fof(f80,plain,(
  ~p(s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))))),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  p(s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) & ~p(s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f45])).
fof(f45,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),X2)))) & ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_image_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,X1),X3)))))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),X2)))) => p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_image_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,X1),X3)))))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),X2)))) => p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),'const.relation.inv_image_2'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,X1),X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X7,fun(X7,bool)),X8)))) => p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_image_2'(s(fun(X7,fun(X7,bool)),X8),s(fun(X6,X7),X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X7,fun(X7,bool)),X8)))) => p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_image_2'(s(fun(X7,fun(X7,bool)),X8),s(fun(X6,X7),X9)))))))),
  file('thm.relation.symmetric_inv_image','thm.relation.symmetric_inv_image')).
fof(f23684,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1)) ) | ($spl32 | $spl70)),
  inference(subsumption_resolution,[],[f23683,f2341])).
fof(f2341,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3))) = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2)))) )),
  inference(subsumption_resolution,[],[f2335,f126])).
fof(f2335,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3))) = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2)))) )),
  inference(superposition,[],[f1728,f151])).
fof(f151,plain,(
  s(bool,'T_0') = s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))),
  inference(subsumption_resolution,[],[f143,f115])).
fof(f115,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.symmetric_inv_image',aHL_FALSITY)).
fof(f143,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))),
  inference(superposition,[],[f79,f122])).
fof(f79,plain,(
  p(s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))))),
  inference(cnf_transformation,[],[f56])).
fof(f1728,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) | s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X10),s(X8,X11))) = s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X11),s(X8,X10)))) )),
  inference(forward_demodulation,[],[f1644,f123])).
fof(f123,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.inv_image_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,X0),X3),s(X1,X4),s(X1,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(bool,'const.relation.inv_image_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,X0),X3),s(X1,X4),s(X1,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'const.relation.inv_image_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,X0),X3),s(X1,X4),s(X1,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6,X7] : ! [X8,X9,X11,X12] : s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X8),s(fun(X7,X6),X9),s(X7,X11),s(X7,X12))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X8),s(X6,happ(s(fun(X7,X6),X9),s(X7,X11))))),s(X6,happ(s(fun(X7,X6),X9),s(X7,X12)))))),
  file('thm.relation.symmetric_inv_image','thm.relation.inv_image_thm')).
fof(f1644,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X10),s(X8,X11))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X11))))),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10))))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7))))) )),
  inference(superposition,[],[f123,f118])).
fof(f118,plain,(
  ( ! [X4,X0,X5,X1] : (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) & (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK4(X1,X0)))),s(X0,sK5(X1,X0)))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK5(X1,X0)))),s(X0,sK4(X1,X0)))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) & (? [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & (? [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X7] : ! [X8] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X7,fun(X7,bool)),X8)))) <=> ! [X11,X12] : s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X11))),s(X7,X12))) = s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X12))),s(X7,X11))))),
  file('thm.relation.symmetric_inv_image','thm.relation.symmetric_def')).
fof(f23683,plain,(
  ( ! [X1] : (s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') != s(bool,X1)) ) | ($spl32 | $spl70)),
  inference(forward_demodulation,[],[f23682,f81])).
fof(f81,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.inv_image_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,X0),X3),s(X1,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),'const.relation.inv_image_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,X0),X3))),s(X1,X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(bool,'const.relation.inv_image_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,X0),X3),s(X1,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),'const.relation.inv_image_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,X0),X3))),s(X1,X4))),s(X1,X5)))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'const.relation.inv_image_4'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,X0),X3),s(X1,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),'const.relation.inv_image_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X1,X0),X3))),s(X1,X4))),s(X1,X5)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X7] : ! [X18,X19,X20,X21] : s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X18),s(fun(X7,X6),X19),s(X7,X20),s(X7,X21))) = s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),'const.relation.inv_image_2'(s(fun(X6,fun(X6,bool)),X18),s(fun(X7,X6),X19))),s(X7,X20))),s(X7,X21)))),
  file('thm.relation.symmetric_inv_image','ARITY_2_4_const.relation.inv_image')).
fof(f23682,plain,(
  ( ! [X1] : (s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') != s(bool,X1)) ) | ($spl32 | $spl70)),
  inference(forward_demodulation,[],[f23681,f81])).
fof(f23681,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0))))) ) | ($spl32 | $spl70)),
  inference(subsumption_resolution,[],[f23657,f496])).
fof(f23657,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | p(s(bool,X1))) ) | $spl70),
  inference(superposition,[],[f1898,f180])).
fof(f23720,plain,(
  $false | ($spl32 | $spl70)),
  inference(equality_resolution,[],[f23684])).
fof(f23680,plain,(
  ( ! [X134,X132,X135,X133] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X132,sK1),X133),s(X132,X135),s(X132,X134)))) ) | $spl368),
  inference(cnf_transformation,[],[f23680_D])).
fof(f23680_D,plain,(
  ( ! [X134,X132,X135,X133] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X132,sK1),X133),s(X132,X135),s(X132,X134)))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f23662,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5))) | s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK4(X5,X4)))),s(X4,sK5(X5,X4)))) != s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK5(X5,X4)))),s(X4,sK4(X5,X4))))) ) | $spl362),
  inference(cnf_transformation,[],[f23662_D])).
fof(f23662_D,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5))) | s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK4(X5,X4)))),s(X4,sK5(X5,X4)))) != s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK5(X5,X4)))),s(X4,sK4(X5,X4))))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f23679,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X136,fun(X136,bool)),X137))) | s(bool,happ(s(fun(X136,bool),happ(s(fun(X136,fun(X136,bool)),X137),s(X136,sK4(X137,X136)))),s(X136,sK5(X137,X136)))) != s(bool,happ(s(fun(X136,bool),happ(s(fun(X136,fun(X136,bool)),X137),s(X136,sK5(X137,X136)))),s(X136,sK4(X137,X136)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X132,sK1),X133),s(X132,X135),s(X132,X134)))) ) | ($spl34 | $spl70)),
  inference(subsumption_resolution,[],[f23646,f2438])).
fof(f2438,plain,(
  ( ! [X337,X339,X336,X338] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X336,sK1),X337),s(X336,X339),s(X336,X338)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X336,sK1),X337),s(X336,X338),s(X336,X339)))) ) | $spl34),
  inference(superposition,[],[f503,f2341])).
fof(f23646,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X136,fun(X136,bool)),X137))) | s(bool,happ(s(fun(X136,bool),happ(s(fun(X136,fun(X136,bool)),X137),s(X136,sK4(X137,X136)))),s(X136,sK5(X137,X136)))) != s(bool,happ(s(fun(X136,bool),happ(s(fun(X136,fun(X136,bool)),X137),s(X136,sK5(X137,X136)))),s(X136,sK4(X137,X136)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X132,sK1),X133),s(X132,X134),s(X132,X135)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X132,sK1),X133),s(X132,X135),s(X132,X134)))) ) | $spl70),
  inference(superposition,[],[f1898,f2343])).
fof(f2343,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X8,sK1),X9),s(X8,X11),s(X8,X10))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X8,sK1),X9),s(X8,X10),s(X8,X11)))) )),
  inference(superposition,[],[f2341,f122])).
fof(f23678,plain,(
  ( ! [X88,X89] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X89))),s(sK1,X88)))) ) | $spl366),
  inference(cnf_transformation,[],[f23678_D])).
fof(f23678_D,plain,(
  ( ! [X88,X89] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X89))),s(sK1,X88)))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f23677,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X90,fun(X90,bool)),X91))) | s(bool,happ(s(fun(X90,bool),happ(s(fun(X90,fun(X90,bool)),X91),s(X90,sK4(X91,X90)))),s(X90,sK5(X91,X90)))) != s(bool,happ(s(fun(X90,bool),happ(s(fun(X90,fun(X90,bool)),X91),s(X90,sK5(X91,X90)))),s(X90,sK4(X91,X90)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X89))),s(sK1,X88)))) ) | ($spl34 | $spl70)),
  inference(subsumption_resolution,[],[f23638,f2017])).
fof(f2017,plain,(
  ( ! [X235,X234] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X235))),s(sK1,X234)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X234))),s(sK1,X235)))) ) | $spl34),
  inference(superposition,[],[f503,f127])).
fof(f127,plain,(
  ( ! [X0,X1] : (s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1))) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0)))) )),
  inference(resolution,[],[f79,f118])).
fof(f23638,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X90,fun(X90,bool)),X91))) | s(bool,happ(s(fun(X90,bool),happ(s(fun(X90,fun(X90,bool)),X91),s(X90,sK4(X91,X90)))),s(X90,sK5(X91,X90)))) != s(bool,happ(s(fun(X90,bool),happ(s(fun(X90,fun(X90,bool)),X91),s(X90,sK5(X91,X90)))),s(X90,sK4(X91,X90)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X88))),s(sK1,X89)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X89))),s(sK1,X88)))) ) | $spl70),
  inference(superposition,[],[f1898,f1951])).
fof(f1951,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X13))),s(sK1,X12))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X12))),s(sK1,X13)))) )),
  inference(superposition,[],[f127,f122])).
fof(f23670,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(X48,fun(X48,bool)),X49),s(fun(X50,X48),X51),s(X50,X52),s(X50,X53))))) ) | $spl364),
  inference(cnf_transformation,[],[f23670_D])).
fof(f23670_D,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(X48,fun(X48,bool)),X49),s(fun(X50,X48),X51),s(X50,X52),s(X50,X53))))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f23675,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (p(s(bool,'const.relation.inv_image_4'(s(fun(X73,fun(X73,bool)),X74),s(fun(X72,X73),X75),s(X72,X76),s(X72,X77)))) | s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X78,fun(X78,bool)),X79))) | s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,sK4(X79,X78)))),s(X78,sK5(X79,X78)))) != s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,sK5(X79,X78)))),s(X78,sK4(X79,X78))))) ) | ($spl24 | $spl70)),
  inference(subsumption_resolution,[],[f23674,f385])).
fof(f23674,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (p(s(bool,'const.relation.inv_image_4'(s(fun(X73,fun(X73,bool)),X74),s(fun(X72,X73),X75),s(X72,X76),s(X72,X77)))) | s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X78,fun(X78,bool)),X79))) | s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,sK4(X79,X78)))),s(X78,sK5(X79,X78)))) != s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,sK5(X79,X78)))),s(X78,sK4(X79,X78)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X73,fun(X73,bool)),X74),s(fun(X72,X73),X75),s(X72,X76),s(X72,X77)))) ) | $spl70),
  inference(forward_demodulation,[],[f23636,f81])).
fof(f23636,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X78,fun(X78,bool)),X79))) | s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,sK4(X79,X78)))),s(X78,sK5(X79,X78)))) != s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,sK5(X79,X78)))),s(X78,sK4(X79,X78)))) | p(s(bool,happ(s(fun(X72,bool),happ(s(fun(X72,fun(X72,bool)),'const.relation.inv_image_2'(s(fun(X73,fun(X73,bool)),X74),s(fun(X72,X73),X75))),s(X72,X76))),s(X72,X77)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X73,fun(X73,bool)),X74),s(fun(X72,X73),X75),s(X72,X76),s(X72,X77)))) ) | $spl70),
  inference(superposition,[],[f1898,f1410])).
fof(f1410,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,happ(s(fun(X18,bool),happ(s(fun(X18,fun(X18,bool)),'const.relation.inv_image_2'(s(fun(X19,fun(X19,bool)),X20),s(fun(X18,X19),X21))),s(X18,X22))),s(X18,X23))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X19,fun(X19,bool)),X20),s(fun(X18,X19),X21),s(X18,X22),s(X18,X23)))) )),
  inference(superposition,[],[f81,f122])).
fof(f23669,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(X48,fun(X48,bool)),X49),s(fun(X50,X48),X51),s(X50,X52),s(X50,X53)))) | s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X54,fun(X54,bool)),X55))) | s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,sK4(X55,X54)))),s(X54,sK5(X55,X54)))) != s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,sK5(X55,X54)))),s(X54,sK4(X55,X54))))) ) | ($spl24 | $spl70)),
  inference(subsumption_resolution,[],[f23668,f385])).
fof(f23668,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(X48,fun(X48,bool)),X49),s(fun(X50,X48),X51),s(X50,X52),s(X50,X53)))) | s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X54,fun(X54,bool)),X55))) | s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,sK4(X55,X54)))),s(X54,sK5(X55,X54)))) != s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,sK5(X55,X54)))),s(X54,sK4(X55,X54)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X48,fun(X48,bool)),X49),s(fun(X50,X48),X51),s(X50,X52),s(X50,X53)))) ) | $spl70),
  inference(forward_demodulation,[],[f23633,f123])).
fof(f23633,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X54,fun(X54,bool)),X55))) | s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,sK4(X55,X54)))),s(X54,sK5(X55,X54)))) != s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,sK5(X55,X54)))),s(X54,sK4(X55,X54)))) | p(s(bool,happ(s(fun(X48,bool),happ(s(fun(X48,fun(X48,bool)),X49),s(X48,happ(s(fun(X50,X48),X51),s(X50,X52))))),s(X48,happ(s(fun(X50,X48),X51),s(X50,X53)))))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X48,fun(X48,bool)),X49),s(fun(X50,X48),X51),s(X50,X52),s(X50,X53)))) ) | $spl70),
  inference(superposition,[],[f1898,f1647])).
fof(f1647,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),happ(s(fun(X26,fun(X26,bool)),X27),s(X26,happ(s(fun(X28,X26),X29),s(X28,X30))))),s(X26,happ(s(fun(X28,X26),X29),s(X28,X31))))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X26,fun(X26,bool)),X27),s(fun(X28,X26),X29),s(X28,X30),s(X28,X31)))) )),
  inference(superposition,[],[f123,f122])).
fof(f150,plain,(
  ( ! [X14] : (p(s(bool,X14))) ) | $spl12),
  inference(cnf_transformation,[],[f150_D])).
fof(f150_D,plain,(
  ( ! [X14] : (p(s(bool,X14))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f23661,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5))) | s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK4(X5,X4)))),s(X4,sK5(X5,X4)))) != s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK5(X5,X4)))),s(X4,sK4(X5,X4)))) | p(s(bool,X3))) ) | ($spl24 | $spl70)),
  inference(subsumption_resolution,[],[f23625,f385])).
fof(f23625,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5))) | s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK4(X5,X4)))),s(X4,sK5(X5,X4)))) != s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK5(X5,X4)))),s(X4,sK4(X5,X4)))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X3)) ) | $spl70),
  inference(superposition,[],[f1898,f122])).
fof(f3519,plain,(
  ( ! [X350,X351,X348] : (s(bool,X348) != s(bool,X350) | s(bool,'F_0') = s(bool,X351) | s(bool,X350) != s(bool,X351)) ) | $spl104),
  inference(cnf_transformation,[],[f3519_D])).
fof(f3519_D,plain,(
  ( ! [X350,X351,X348] : (s(bool,X348) != s(bool,X350) | s(bool,'F_0') = s(bool,X351) | s(bool,X350) != s(bool,X351)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f23347,plain,(
  ( ! [X140,X139] : (s(bool,X139) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'T_0') = s(bool,X140) | s(bool,X139) != s(bool,X140)) ) | $spl360),
  inference(cnf_transformation,[],[f23347_D])).
fof(f23347_D,plain,(
  ( ! [X140,X139] : (s(bool,X139) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'T_0') = s(bool,X140) | s(bool,X139) != s(bool,X140)) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f23350,plain,(
  ( ! [X182,X180,X178,X181,X179] : (s(bool,X179) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X178) != s(bool,X180) | s(bool,X180) != s(bool,X181) | s(bool,'F_0') = s(bool,X181) | s(bool,X179) != s(bool,X182) | s(bool,'T_0') = s(bool,X182)) ) | ($spl27 | $spl34 | $spl38 | $spl66 | $spl338)),
  inference(subsumption_resolution,[],[f23099,f789])).
fof(f789,plain,(
  ( ! [X64,X62,X63] : (s(bool,X62) != s(bool,X64) | s(bool,X62) != s(bool,X63) | s(bool,'F_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X63)) ) | ($spl34 | $spl38)),
  inference(resolution,[],[f745,f503])).
fof(f23099,plain,(
  ( ! [X182,X180,X178,X181,X179] : (s(bool,'T_0') != s(bool,X178) | s(bool,X179) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X178) != s(bool,X180) | s(bool,X180) != s(bool,X181) | s(bool,'F_0') = s(bool,X181) | s(bool,X179) != s(bool,X182) | s(bool,'T_0') = s(bool,X182)) ) | ($spl27 | $spl34 | $spl66 | $spl338)),
  inference(superposition,[],[f2201,f21119])).
fof(f21119,plain,(
  s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | ($spl27 | $spl338)),
  inference(subsumption_resolution,[],[f20840,f435])).
fof(f20840,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | $spl338),
  inference(superposition,[],[f20327,f2343])).
fof(f20327,plain,(
  s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | $spl338),
  inference(cnf_transformation,[],[f20327_D])).
fof(f20327_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f2201,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (s(bool,X100) != s(bool,X101) | s(bool,X98) != s(bool,X100) | s(bool,X101) != s(bool,X102) | s(bool,X102) != s(bool,X103) | s(bool,'F_0') = s(bool,X103) | s(bool,X98) != s(bool,X99) | s(bool,'T_0') = s(bool,X99)) ) | ($spl34 | $spl66)),
  inference(resolution,[],[f1852,f503])).
fof(f471,plain,(
  ( ! [X4,X7] : (s(bool,'F_0') = s(bool,X4) | s(bool,X4) != s(bool,X7)) ) | $spl30),
  inference(cnf_transformation,[],[f471_D])).
fof(f471_D,plain,(
  ( ! [X4,X7] : (s(bool,'F_0') = s(bool,X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f23349,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,X156) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X155) != s(bool,X157) | s(bool,'F_0') = s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,'T_0') = s(bool,X158)) ) | ($spl6 | $spl27 | $spl34 | $spl56 | $spl338)),
  inference(subsumption_resolution,[],[f23092,f589])).
fof(f589,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f576,f126])).
fof(f576,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f497,f122])).
fof(f497,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f478,f101])).
fof(f478,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,'F_0')) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6)) ) | $spl6),
  inference(superposition,[],[f147,f122])).
fof(f23092,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,'T_0') != s(bool,X155) | s(bool,X156) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X155) != s(bool,X157) | s(bool,'F_0') = s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,'T_0') = s(bool,X158)) ) | ($spl27 | $spl34 | $spl56 | $spl338)),
  inference(superposition,[],[f1836,f21119])).
fof(f1836,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X88) != s(bool,X89) | s(bool,X86) != s(bool,X88) | s(bool,X89) != s(bool,X90) | s(bool,'F_0') = s(bool,X90) | s(bool,X86) != s(bool,X87) | s(bool,'T_0') = s(bool,X87)) ) | ($spl34 | $spl56)),
  inference(resolution,[],[f1356,f503])).
fof(f855,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,X12)) ) | $spl44),
  inference(cnf_transformation,[],[f855_D])).
fof(f855_D,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,X12)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f23346,plain,(
  ( ! [X140,X138,X139] : (s(bool,X139) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X138) | s(bool,X139) != s(bool,X140) | s(bool,'T_0') = s(bool,X140)) ) | ($spl27 | $spl34 | $spl40 | $spl338)),
  inference(subsumption_resolution,[],[f23085,f122])).
fof(f23085,plain,(
  ( ! [X140,X138,X139] : (s(bool,'T_0') != s(bool,X138) | s(bool,X139) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X138) | s(bool,X139) != s(bool,X140) | s(bool,'T_0') = s(bool,X140)) ) | ($spl27 | $spl34 | $spl40 | $spl338)),
  inference(superposition,[],[f1342,f21119])).
fof(f1342,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,X76) != s(bool,X77) | s(bool,X74) != s(bool,X76) | s(bool,'F_0') = s(bool,X77) | s(bool,X74) != s(bool,X75) | s(bool,'T_0') = s(bool,X75)) ) | ($spl34 | $spl40)),
  inference(resolution,[],[f805,f503])).
fof(f23343,plain,(
  ( ! [X125] : (s(bool,X125) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'T_0') = s(bool,X125)) ) | $spl358),
  inference(cnf_transformation,[],[f23343_D])).
fof(f23343_D,plain,(
  ( ! [X125] : (s(bool,X125) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'T_0') = s(bool,X125)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f23342,plain,(
  ( ! [X125,X124] : (s(bool,X125) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X125)) ) | ($spl27 | $spl34 | $spl38 | $spl338)),
  inference(subsumption_resolution,[],[f23077,f122])).
fof(f23077,plain,(
  ( ! [X125,X124] : (s(bool,'T_0') != s(bool,X124) | s(bool,X125) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X125)) ) | ($spl27 | $spl34 | $spl38 | $spl338)),
  inference(superposition,[],[f789,f21119])).
fof(f21173,plain,(
  ( ! [X140,X139] : (s(bool,X139) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'T_0') = s(bool,X140) | s(bool,X139) != s(bool,X140)) ) | $spl346),
  inference(cnf_transformation,[],[f21173_D])).
fof(f21173_D,plain,(
  ( ! [X140,X139] : (s(bool,X139) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'T_0') = s(bool,X140) | s(bool,X139) != s(bool,X140)) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f22841,plain,(
  ( ! [X182,X180,X178,X181,X179] : (s(bool,X179) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X178) != s(bool,X180) | s(bool,X180) != s(bool,X181) | s(bool,'F_0') = s(bool,X181) | s(bool,X179) != s(bool,X182) | s(bool,'T_0') = s(bool,X182)) ) | ($spl34 | $spl38 | $spl66 | $spl338)),
  inference(subsumption_resolution,[],[f22600,f789])).
fof(f22600,plain,(
  ( ! [X182,X180,X178,X181,X179] : (s(bool,'T_0') != s(bool,X178) | s(bool,X179) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X178) != s(bool,X180) | s(bool,X180) != s(bool,X181) | s(bool,'F_0') = s(bool,X181) | s(bool,X179) != s(bool,X182) | s(bool,'T_0') = s(bool,X182)) ) | ($spl34 | $spl66 | $spl338)),
  inference(superposition,[],[f2201,f20327])).
fof(f22840,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,X156) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X155) != s(bool,X157) | s(bool,'F_0') = s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,'T_0') = s(bool,X158)) ) | ($spl6 | $spl34 | $spl56 | $spl338)),
  inference(subsumption_resolution,[],[f22593,f589])).
fof(f22593,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,'T_0') != s(bool,X155) | s(bool,X156) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X155) != s(bool,X157) | s(bool,'F_0') = s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,'T_0') = s(bool,X158)) ) | ($spl34 | $spl56 | $spl338)),
  inference(superposition,[],[f1836,f20327])).
fof(f22838,plain,(
  ( ! [X140,X138,X139] : (s(bool,X139) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X138) | s(bool,X139) != s(bool,X140) | s(bool,'T_0') = s(bool,X140)) ) | ($spl34 | $spl40 | $spl338)),
  inference(subsumption_resolution,[],[f22586,f122])).
fof(f22586,plain,(
  ( ! [X140,X138,X139] : (s(bool,'T_0') != s(bool,X138) | s(bool,X139) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X138) | s(bool,X139) != s(bool,X140) | s(bool,'T_0') = s(bool,X140)) ) | ($spl34 | $spl40 | $spl338)),
  inference(superposition,[],[f1342,f20327])).
fof(f21170,plain,(
  ( ! [X125] : (s(bool,X125) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'T_0') = s(bool,X125)) ) | $spl344),
  inference(cnf_transformation,[],[f21170_D])).
fof(f21170_D,plain,(
  ( ! [X125] : (s(bool,X125) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'T_0') = s(bool,X125)) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f22836,plain,(
  ( ! [X125,X124] : (s(bool,X125) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X125)) ) | ($spl34 | $spl38 | $spl338)),
  inference(subsumption_resolution,[],[f22578,f122])).
fof(f22578,plain,(
  ( ! [X125,X124] : (s(bool,'T_0') != s(bool,X124) | s(bool,X125) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X125)) ) | ($spl34 | $spl38 | $spl338)),
  inference(superposition,[],[f789,f20327])).
fof(f22500,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | $spl354),
  inference(cnf_transformation,[],[f22500_D])).
fof(f22500_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f22502,plain,(
  s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | $spl357),
  inference(cnf_transformation,[],[f22502_D])).
fof(f22502_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) <=> ~$spl357),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl357])])).
fof(f21865,plain,(
  s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0))))),
  inference(forward_demodulation,[],[f21864,f81])).
fof(f21864,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0))))),
  inference(subsumption_resolution,[],[f21859,f115])).
fof(f21859,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0))))),
  inference(superposition,[],[f1897,f180])).
fof(f1897,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X10,fun(X10,bool)),X11)))) | s(bool,'F_0') != s(bool,happ(s(fun(X10,bool),happ(s(fun(X10,fun(X10,bool)),X11),s(X10,sK4(X11,X10)))),s(X10,sK5(X11,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X10,bool),happ(s(fun(X10,fun(X10,bool)),X11),s(X10,sK5(X11,X10)))),s(X10,sK4(X11,X10))))) )),
  inference(superposition,[],[f119,f122])).
fof(f119,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK4(X1,X0)))),s(X0,sK5(X1,X0)))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK5(X1,X0)))),s(X0,sK4(X1,X0))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f22226,plain,(
  ( ! [X185,X186] : (s(bool,X185) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X186) | s(bool,X185) != s(bool,X186)) ) | $spl352),
  inference(cnf_transformation,[],[f22226_D])).
fof(f22226_D,plain,(
  ( ! [X185,X186] : (s(bool,X185) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X186) | s(bool,X185) != s(bool,X186)) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f3566,plain,(
  ( ! [X333,X331,X330] : (s(bool,X331) != s(bool,X333) | s(bool,'T_0') = s(bool,X333) | s(bool,X330) != s(bool,X331)) ) | $spl108),
  inference(cnf_transformation,[],[f3566_D])).
fof(f3566_D,plain,(
  ( ! [X333,X331,X330] : (s(bool,X331) != s(bool,X333) | s(bool,'T_0') = s(bool,X333) | s(bool,X330) != s(bool,X331)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f22225,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,X183) != s(bool,X184) | s(bool,X185) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X185) != s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,X184) != s(bool,X187) | s(bool,'T_0') = s(bool,X187)) ) | ($spl34 | $spl38 | $spl66 | $spl339)),
  inference(subsumption_resolution,[],[f21970,f789])).
fof(f21970,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,'F_0') != s(bool,X183) | s(bool,X183) != s(bool,X184) | s(bool,X185) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X185) != s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,X184) != s(bool,X187) | s(bool,'T_0') = s(bool,X187)) ) | ($spl34 | $spl66 | $spl339)),
  inference(superposition,[],[f2201,f20339])).
fof(f20339,plain,(
  s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | $spl339),
  inference(trivial_inequality_removal,[],[f20335])).
fof(f20335,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | $spl339),
  inference(superposition,[],[f20328,f2343])).
fof(f22216,plain,(
  ( ! [X161] : (s(bool,X161) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X161)) ) | $spl350),
  inference(cnf_transformation,[],[f22216_D])).
fof(f22216_D,plain,(
  ( ! [X161] : (s(bool,X161) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X161)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f22215,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,X159) != s(bool,X160) | s(bool,X161) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X161) | s(bool,X160) != s(bool,X162) | s(bool,'T_0') = s(bool,X162)) ) | ($spl34 | $spl38 | $spl56 | $spl339)),
  inference(subsumption_resolution,[],[f21963,f789])).
fof(f21963,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,'F_0') != s(bool,X159) | s(bool,X159) != s(bool,X160) | s(bool,X161) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X161) | s(bool,X160) != s(bool,X162) | s(bool,'T_0') = s(bool,X162)) ) | ($spl34 | $spl56 | $spl339)),
  inference(superposition,[],[f1836,f20339])).
fof(f21863,plain,(
  ( ! [X30,X33,X31] : (s(bool,'F_0') != s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK4(X31,X30)))),s(X30,sK5(X31,X30)))) | s(bool,X33) != s(bool,'const.relation.symmetric_1'(s(fun(X30,fun(X30,bool)),X31))) | p(s(bool,X33)) | s(bool,'T_0') = s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK5(X31,X30)))),s(X30,sK4(X31,X30))))) ) | $spl348),
  inference(cnf_transformation,[],[f21863_D])).
fof(f21863_D,plain,(
  ( ! [X30,X33,X31] : (s(bool,'F_0') != s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK4(X31,X30)))),s(X30,sK5(X31,X30)))) | s(bool,X33) != s(bool,'const.relation.symmetric_1'(s(fun(X30,fun(X30,bool)),X31))) | p(s(bool,X33)) | s(bool,'T_0') = s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK5(X31,X30)))),s(X30,sK4(X31,X30))))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f21852,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'F_0') != s(bool,happ(s(fun(X74,bool),happ(s(fun(X74,fun(X74,bool)),X75),s(X74,sK4(X75,X74)))),s(X74,sK5(X75,X74)))) | s(bool,'T_0') = s(bool,happ(s(fun(X74,bool),happ(s(fun(X74,fun(X74,bool)),X75),s(X74,sK5(X75,X74)))),s(X74,sK4(X75,X74)))) | p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.relation.symmetric_1'(s(fun(X74,fun(X74,bool)),X75)))) )),
  inference(resolution,[],[f1897,f125])).
fof(f125,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f86])).
fof(f86,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11,X16,X12,X17] : ((s(bool,X11) = s(bool,X16) & (p(s(bool,X16)) => s(bool,X12) = s(bool,X17))) => ((p(s(bool,X11)) => p(s(bool,X12))) <=> (p(s(bool,X16)) => p(s(bool,X17)))))),
  file('thm.relation.symmetric_inv_image','thm.bool.IMP_CONG')).
fof(f148,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | ~p(s(bool,X9))) ) | $spl8),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | ~p(s(bool,X9))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f21849,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,'F_0') != s(bool,happ(s(fun(X65,bool),happ(s(fun(X65,fun(X65,bool)),X66),s(X65,sK4(X66,X65)))),s(X65,sK5(X66,X65)))) | s(bool,'T_0') = s(bool,happ(s(fun(X65,bool),happ(s(fun(X65,fun(X65,bool)),X66),s(X65,sK5(X66,X65)))),s(X65,sK4(X66,X65)))) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,'const.relation.symmetric_1'(s(fun(X65,fun(X65,bool)),X66))) | s(bool,X68) != s(bool,X69)) )),
  inference(resolution,[],[f1897,f89])).
fof(f89,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f149,plain,(
  ( ! [X15] : (~p(s(bool,X15))) ) | $spl10),
  inference(cnf_transformation,[],[f149_D])).
fof(f149_D,plain,(
  ( ! [X15] : (~p(s(bool,X15))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f21846,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'F_0') != s(bool,happ(s(fun(X50,bool),happ(s(fun(X50,fun(X50,bool)),X51),s(X50,sK4(X51,X50)))),s(X50,sK5(X51,X50)))) | s(bool,'T_0') = s(bool,happ(s(fun(X50,bool),happ(s(fun(X50,fun(X50,bool)),X51),s(X50,sK5(X51,X50)))),s(X50,sK4(X51,X50)))) | p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X54) != s(bool,'const.relation.symmetric_1'(s(fun(X50,fun(X50,bool)),X51)))) )),
  inference(resolution,[],[f1897,f88])).
fof(f88,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f145,plain,(
  ( ! [X4,X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,X4)) ) | $spl2),
  inference(cnf_transformation,[],[f145_D])).
fof(f145_D,plain,(
  ( ! [X4,X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,X4)) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f21845,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'F_0') != s(bool,happ(s(fun(X45,bool),happ(s(fun(X45,fun(X45,bool)),X46),s(X45,sK4(X46,X45)))),s(X45,sK5(X46,X45)))) | s(bool,'T_0') = s(bool,happ(s(fun(X45,bool),happ(s(fun(X45,fun(X45,bool)),X46),s(X45,sK5(X46,X45)))),s(X45,sK4(X46,X45)))) | p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.relation.symmetric_1'(s(fun(X45,fun(X45,bool)),X46)))) )),
  inference(resolution,[],[f1897,f87])).
fof(f87,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f21844,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,'F_0') != s(bool,happ(s(fun(X40,bool),happ(s(fun(X40,fun(X40,bool)),X41),s(X40,sK4(X41,X40)))),s(X40,sK5(X41,X40)))) | s(bool,'T_0') = s(bool,happ(s(fun(X40,bool),happ(s(fun(X40,fun(X40,bool)),X41),s(X40,sK5(X41,X40)))),s(X40,sK4(X41,X40)))) | p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.relation.symmetric_1'(s(fun(X40,fun(X40,bool)),X41))) | s(bool,X43) != s(bool,X44)) )),
  inference(resolution,[],[f1897,f85])).
fof(f85,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f21842,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK4(X31,X30)))),s(X30,sK5(X31,X30)))) | s(bool,'T_0') = s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK5(X31,X30)))),s(X30,sK4(X31,X30)))) | p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,'const.relation.symmetric_1'(s(fun(X30,fun(X30,bool)),X31)))) )),
  inference(resolution,[],[f1897,f83])).
fof(f83,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f20682,plain,(
  ( ! [X185,X186] : (s(bool,X185) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X186) | s(bool,X185) != s(bool,X186)) ) | $spl342),
  inference(cnf_transformation,[],[f20682_D])).
fof(f20682_D,plain,(
  ( ! [X185,X186] : (s(bool,X185) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X186) | s(bool,X185) != s(bool,X186)) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f21682,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,X183) != s(bool,X184) | s(bool,X185) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X185) != s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,X184) != s(bool,X187) | s(bool,'T_0') = s(bool,X187)) ) | ($spl34 | $spl38 | $spl66 | $spl339)),
  inference(subsumption_resolution,[],[f21441,f789])).
fof(f21441,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,'F_0') != s(bool,X183) | s(bool,X183) != s(bool,X184) | s(bool,X185) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X185) != s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,X184) != s(bool,X187) | s(bool,'T_0') = s(bool,X187)) ) | ($spl34 | $spl66 | $spl339)),
  inference(superposition,[],[f2201,f20338])).
fof(f20338,plain,(
  s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | $spl339),
  inference(trivial_inequality_removal,[],[f20336])).
fof(f20336,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | $spl339),
  inference(superposition,[],[f20328,f122])).
fof(f20674,plain,(
  ( ! [X161] : (s(bool,X161) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X161)) ) | $spl340),
  inference(cnf_transformation,[],[f20674_D])).
fof(f20674_D,plain,(
  ( ! [X161] : (s(bool,X161) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X161)) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f21675,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,X159) != s(bool,X160) | s(bool,X161) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X161) | s(bool,X160) != s(bool,X162) | s(bool,'T_0') = s(bool,X162)) ) | ($spl34 | $spl38 | $spl56 | $spl339)),
  inference(subsumption_resolution,[],[f21434,f789])).
fof(f21434,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,'F_0') != s(bool,X159) | s(bool,X159) != s(bool,X160) | s(bool,X161) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X161) | s(bool,X160) != s(bool,X162) | s(bool,'T_0') = s(bool,X162)) ) | ($spl34 | $spl56 | $spl339)),
  inference(superposition,[],[f1836,f20338])).
fof(f21176,plain,(
  ( ! [X182,X180,X178,X181,X179] : (s(bool,X179) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X178) != s(bool,X180) | s(bool,X180) != s(bool,X181) | s(bool,'F_0') = s(bool,X181) | s(bool,X179) != s(bool,X182) | s(bool,'T_0') = s(bool,X182)) ) | ($spl34 | $spl38 | $spl66 | $spl338)),
  inference(subsumption_resolution,[],[f20933,f789])).
fof(f20933,plain,(
  ( ! [X182,X180,X178,X181,X179] : (s(bool,'T_0') != s(bool,X178) | s(bool,X179) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X178) != s(bool,X180) | s(bool,X180) != s(bool,X181) | s(bool,'F_0') = s(bool,X181) | s(bool,X179) != s(bool,X182) | s(bool,'T_0') = s(bool,X182)) ) | ($spl34 | $spl66 | $spl338)),
  inference(superposition,[],[f2201,f20327])).
fof(f21175,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,X156) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X155) != s(bool,X157) | s(bool,'F_0') = s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,'T_0') = s(bool,X158)) ) | ($spl6 | $spl34 | $spl56 | $spl338)),
  inference(subsumption_resolution,[],[f20926,f589])).
fof(f20926,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,'T_0') != s(bool,X155) | s(bool,X156) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X155) != s(bool,X157) | s(bool,'F_0') = s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,'T_0') = s(bool,X158)) ) | ($spl34 | $spl56 | $spl338)),
  inference(superposition,[],[f1836,f20327])).
fof(f21172,plain,(
  ( ! [X140,X138,X139] : (s(bool,X139) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X138) | s(bool,X139) != s(bool,X140) | s(bool,'T_0') = s(bool,X140)) ) | ($spl34 | $spl40 | $spl338)),
  inference(subsumption_resolution,[],[f20919,f122])).
fof(f20919,plain,(
  ( ! [X140,X138,X139] : (s(bool,'T_0') != s(bool,X138) | s(bool,X139) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X138) | s(bool,X139) != s(bool,X140) | s(bool,'T_0') = s(bool,X140)) ) | ($spl34 | $spl40 | $spl338)),
  inference(superposition,[],[f1342,f20327])).
fof(f21169,plain,(
  ( ! [X125,X124] : (s(bool,X125) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X125)) ) | ($spl34 | $spl38 | $spl338)),
  inference(subsumption_resolution,[],[f20911,f122])).
fof(f20911,plain,(
  ( ! [X125,X124] : (s(bool,'T_0') != s(bool,X124) | s(bool,X125) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X125)) ) | ($spl34 | $spl38 | $spl338)),
  inference(superposition,[],[f789,f20327])).
fof(f20681,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,X183) != s(bool,X184) | s(bool,X185) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X185) != s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,X184) != s(bool,X187) | s(bool,'T_0') = s(bool,X187)) ) | ($spl34 | $spl38 | $spl66 | $spl339)),
  inference(subsumption_resolution,[],[f20439,f789])).
fof(f20439,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,'F_0') != s(bool,X183) | s(bool,X183) != s(bool,X184) | s(bool,X185) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,X185) != s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,X184) != s(bool,X187) | s(bool,'T_0') = s(bool,X187)) ) | ($spl34 | $spl66 | $spl339)),
  inference(superposition,[],[f2201,f20338])).
fof(f20673,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,X159) != s(bool,X160) | s(bool,X161) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X161) | s(bool,X160) != s(bool,X162) | s(bool,'T_0') = s(bool,X162)) ) | ($spl34 | $spl38 | $spl56 | $spl339)),
  inference(subsumption_resolution,[],[f20432,f789])).
fof(f20432,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,'F_0') != s(bool,X159) | s(bool,X159) != s(bool,X160) | s(bool,X161) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,X161) | s(bool,X160) != s(bool,X162) | s(bool,'T_0') = s(bool,X162)) ) | ($spl34 | $spl56 | $spl339)),
  inference(superposition,[],[f1836,f20338])).
fof(f20326,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | $spl336),
  inference(cnf_transformation,[],[f20326_D])).
fof(f20326_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f20328,plain,(
  s(bool,'T_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | $spl339),
  inference(cnf_transformation,[],[f20328_D])).
fof(f20328_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) <=> ~$spl339),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl339])])).
fof(f20324,plain,(
  s(bool,'T_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0))))),
  inference(forward_demodulation,[],[f20323,f81])).
fof(f20323,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0))))),
  inference(subsumption_resolution,[],[f20318,f115])).
fof(f20318,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))),s(sK0,sK5('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0)))),s(sK0,sK4('const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)),sK0))))),
  inference(superposition,[],[f1896,f180])).
fof(f1896,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X8,fun(X8,bool)),X9)))) | s(bool,'T_0') != s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X9),s(X8,sK4(X9,X8)))),s(X8,sK5(X9,X8)))) | s(bool,'F_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X9),s(X8,sK5(X9,X8)))),s(X8,sK4(X9,X8))))) )),
  inference(superposition,[],[f119,f122])).
fof(f20322,plain,(
  ( ! [X30,X33,X31] : (s(bool,'T_0') != s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK4(X31,X30)))),s(X30,sK5(X31,X30)))) | s(bool,X33) != s(bool,'const.relation.symmetric_1'(s(fun(X30,fun(X30,bool)),X31))) | p(s(bool,X33)) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK5(X31,X30)))),s(X30,sK4(X31,X30))))) ) | $spl334),
  inference(cnf_transformation,[],[f20322_D])).
fof(f20322_D,plain,(
  ( ! [X30,X33,X31] : (s(bool,'T_0') != s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK4(X31,X30)))),s(X30,sK5(X31,X30)))) | s(bool,X33) != s(bool,'const.relation.symmetric_1'(s(fun(X30,fun(X30,bool)),X31))) | p(s(bool,X33)) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK5(X31,X30)))),s(X30,sK4(X31,X30))))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f20311,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') != s(bool,happ(s(fun(X74,bool),happ(s(fun(X74,fun(X74,bool)),X75),s(X74,sK4(X75,X74)))),s(X74,sK5(X75,X74)))) | s(bool,'F_0') = s(bool,happ(s(fun(X74,bool),happ(s(fun(X74,fun(X74,bool)),X75),s(X74,sK5(X75,X74)))),s(X74,sK4(X75,X74)))) | p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.relation.symmetric_1'(s(fun(X74,fun(X74,bool)),X75)))) )),
  inference(resolution,[],[f1896,f125])).
fof(f20308,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,'T_0') != s(bool,happ(s(fun(X65,bool),happ(s(fun(X65,fun(X65,bool)),X66),s(X65,sK4(X66,X65)))),s(X65,sK5(X66,X65)))) | s(bool,'F_0') = s(bool,happ(s(fun(X65,bool),happ(s(fun(X65,fun(X65,bool)),X66),s(X65,sK5(X66,X65)))),s(X65,sK4(X66,X65)))) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,'const.relation.symmetric_1'(s(fun(X65,fun(X65,bool)),X66))) | s(bool,X68) != s(bool,X69)) )),
  inference(resolution,[],[f1896,f89])).
fof(f20305,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'T_0') != s(bool,happ(s(fun(X50,bool),happ(s(fun(X50,fun(X50,bool)),X51),s(X50,sK4(X51,X50)))),s(X50,sK5(X51,X50)))) | s(bool,'F_0') = s(bool,happ(s(fun(X50,bool),happ(s(fun(X50,fun(X50,bool)),X51),s(X50,sK5(X51,X50)))),s(X50,sK4(X51,X50)))) | p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X54) != s(bool,'const.relation.symmetric_1'(s(fun(X50,fun(X50,bool)),X51)))) )),
  inference(resolution,[],[f1896,f88])).
fof(f20304,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') != s(bool,happ(s(fun(X45,bool),happ(s(fun(X45,fun(X45,bool)),X46),s(X45,sK4(X46,X45)))),s(X45,sK5(X46,X45)))) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),happ(s(fun(X45,fun(X45,bool)),X46),s(X45,sK5(X46,X45)))),s(X45,sK4(X46,X45)))) | p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.relation.symmetric_1'(s(fun(X45,fun(X45,bool)),X46)))) )),
  inference(resolution,[],[f1896,f87])).
fof(f20303,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,'T_0') != s(bool,happ(s(fun(X40,bool),happ(s(fun(X40,fun(X40,bool)),X41),s(X40,sK4(X41,X40)))),s(X40,sK5(X41,X40)))) | s(bool,'F_0') = s(bool,happ(s(fun(X40,bool),happ(s(fun(X40,fun(X40,bool)),X41),s(X40,sK5(X41,X40)))),s(X40,sK4(X41,X40)))) | p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.relation.symmetric_1'(s(fun(X40,fun(X40,bool)),X41))) | s(bool,X43) != s(bool,X44)) )),
  inference(resolution,[],[f1896,f85])).
fof(f20301,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') != s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK4(X31,X30)))),s(X30,sK5(X31,X30)))) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,sK5(X31,X30)))),s(X30,sK4(X31,X30)))) | p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,'const.relation.symmetric_1'(s(fun(X30,fun(X30,bool)),X31)))) )),
  inference(resolution,[],[f1896,f83])).
fof(f20214,plain,(
  ( ! [X47,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,'F_0'),s(bool,X49))) | ~p(s(bool,X51)) | s(bool,'T_0') = s(bool,X48)) ) | $spl332),
  inference(cnf_transformation,[],[f20214_D])).
fof(f20214_D,plain,(
  ( ! [X47,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,'F_0'),s(bool,X49))) | ~p(s(bool,X51)) | s(bool,'T_0') = s(bool,X48)) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f20157,plain,(
  ( ! [X107,X105,X108,X106,X104] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,X105),s(bool,X106))) | s(bool,'T_0') = s(bool,X105) | p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X108) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,'F_0'),s(bool,X106)))) ) | $spl34),
  inference(resolution,[],[f11949,f125])).
fof(f11949,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,'F_0'),s(bool,X7)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X8),s(bool,X7))) | s(bool,'T_0') = s(bool,X8)) ) | $spl34),
  inference(superposition,[],[f2714,f2341])).
fof(f2714,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X3),s(bool,X5))) | s(bool,'T_0') = s(bool,X3)) ) | $spl34),
  inference(superposition,[],[f2438,f122])).
fof(f20153,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X89),s(bool,X90),s(bool,X91))) | s(bool,'T_0') = s(bool,X90) | ~p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X93) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X89),s(bool,'F_0'),s(bool,X91))) | s(bool,X92) != s(bool,X94)) ) | $spl34),
  inference(resolution,[],[f11949,f89])).
fof(f20151,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X77),s(bool,X78),s(bool,X79))) | s(bool,'T_0') = s(bool,X78) | p(s(bool,X80)) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X77),s(bool,'F_0'),s(bool,X79)))) ) | $spl34),
  inference(resolution,[],[f11949,f88])).
fof(f20149,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X65),s(bool,X66),s(bool,X67))) | s(bool,'T_0') = s(bool,X66) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X65),s(bool,'F_0'),s(bool,X67)))) ) | $spl34),
  inference(resolution,[],[f11949,f87])).
fof(f20148,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X59),s(bool,X60),s(bool,X61))) | s(bool,'T_0') = s(bool,X60) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X59),s(bool,'F_0'),s(bool,X61))) | s(bool,X62) != s(bool,X64)) ) | $spl34),
  inference(resolution,[],[f11949,f85])).
fof(f20146,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,'T_0') = s(bool,X48) | p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,X52) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,'F_0'),s(bool,X49)))) ) | $spl34),
  inference(resolution,[],[f11949,f83])).
fof(f20119,plain,(
  ( ! [X47,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,'T_0'),s(bool,X49))) | ~p(s(bool,X51)) | s(bool,'F_0') = s(bool,X48)) ) | $spl330),
  inference(cnf_transformation,[],[f20119_D])).
fof(f20119_D,plain,(
  ( ! [X47,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,'T_0'),s(bool,X49))) | ~p(s(bool,X51)) | s(bool,'F_0') = s(bool,X48)) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f20063,plain,(
  ( ! [X107,X105,X108,X106,X104] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,X105),s(bool,X106))) | s(bool,'F_0') = s(bool,X105) | p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X108) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,'T_0'),s(bool,X106)))) ) | $spl34),
  inference(resolution,[],[f11868,f125])).
fof(f11868,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,'T_0'),s(bool,X7)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X8),s(bool,X7))) | s(bool,'F_0') = s(bool,X8)) ) | $spl34),
  inference(superposition,[],[f2713,f2341])).
fof(f2713,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X2),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X0),s(bool,X2))) | s(bool,'F_0') = s(bool,X0)) ) | $spl34),
  inference(superposition,[],[f2438,f122])).
fof(f20059,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X89),s(bool,X90),s(bool,X91))) | s(bool,'F_0') = s(bool,X90) | ~p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X93) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X89),s(bool,'T_0'),s(bool,X91))) | s(bool,X92) != s(bool,X94)) ) | $spl34),
  inference(resolution,[],[f11868,f89])).
fof(f20057,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X77),s(bool,X78),s(bool,X79))) | s(bool,'F_0') = s(bool,X78) | p(s(bool,X80)) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X77),s(bool,'T_0'),s(bool,X79)))) ) | $spl34),
  inference(resolution,[],[f11868,f88])).
fof(f20055,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X65),s(bool,X66),s(bool,X67))) | s(bool,'F_0') = s(bool,X66) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X65),s(bool,'T_0'),s(bool,X67)))) ) | $spl34),
  inference(resolution,[],[f11868,f87])).
fof(f20054,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X59),s(bool,X60),s(bool,X61))) | s(bool,'F_0') = s(bool,X60) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X59),s(bool,'T_0'),s(bool,X61))) | s(bool,X62) != s(bool,X64)) ) | $spl34),
  inference(resolution,[],[f11868,f85])).
fof(f20052,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,'F_0') = s(bool,X48) | p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,X52) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,'T_0'),s(bool,X49)))) ) | $spl34),
  inference(resolution,[],[f11868,f83])).
fof(f20035,plain,(
  ( ! [X26,X24] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'F_0'))),s(bool,sK5(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK5(X24,bool)))),s(bool,'F_0'))) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24))) | s(bool,'T_0') = s(bool,sK4(X24,bool))) ) | $spl328),
  inference(cnf_transformation,[],[f20035_D])).
fof(f20035_D,plain,(
  ( ! [X26,X24] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'F_0'))),s(bool,sK5(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK5(X24,bool)))),s(bool,'F_0'))) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24))) | s(bool,'T_0') = s(bool,sK4(X24,bool))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f20025,plain,(
  ( ! [X59,X60,X58] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X58),s(bool,'F_0'))),s(bool,sK5(X58,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X58),s(bool,sK5(X58,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK4(X58,bool)) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X60) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X58)))) )),
  inference(resolution,[],[f1890,f125])).
fof(f1890,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X1)))) | s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X1),s(bool,'F_0'))),s(bool,sK5(X1,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X1),s(bool,sK5(X1,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK4(X1,bool))) )),
  inference(superposition,[],[f119,f122])).
fof(f20022,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X52),s(bool,'F_0'))),s(bool,sK5(X52,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X52),s(bool,sK5(X52,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK4(X52,bool)) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X52))) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f1890,f89])).
fof(f20019,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X40),s(bool,'F_0'))),s(bool,sK5(X40,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X40),s(bool,sK5(X40,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK4(X40,bool)) | p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X40)))) )),
  inference(resolution,[],[f1890,f88])).
fof(f20018,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X36),s(bool,'F_0'))),s(bool,sK5(X36,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X36),s(bool,sK5(X36,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK4(X36,bool)) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X36)))) )),
  inference(resolution,[],[f1890,f87])).
fof(f20017,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X32),s(bool,'F_0'))),s(bool,sK5(X32,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X32),s(bool,sK5(X32,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK4(X32,bool)) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X32))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f1890,f85])).
fof(f20015,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'F_0'))),s(bool,sK5(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK5(X24,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK4(X24,bool)) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24)))) )),
  inference(resolution,[],[f1890,f83])).
fof(f19979,plain,(
  ( ! [X47,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,'F_0'))) | ~p(s(bool,X51)) | s(bool,'T_0') = s(bool,X49)) ) | $spl326),
  inference(cnf_transformation,[],[f19979_D])).
fof(f19979_D,plain,(
  ( ! [X47,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,'F_0'))) | ~p(s(bool,X51)) | s(bool,'T_0') = s(bool,X49)) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f19922,plain,(
  ( ! [X107,X105,X108,X106,X104] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,X105),s(bool,X106))) | s(bool,'T_0') = s(bool,X106) | p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X108) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,X105),s(bool,'F_0')))) ) | $spl34),
  inference(resolution,[],[f11584,f125])).
fof(f11584,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X7),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X7),s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl34),
  inference(superposition,[],[f2698,f2341])).
fof(f2698,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,'F_0'),s(bool,X5)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl34),
  inference(superposition,[],[f2438,f122])).
fof(f19918,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X89),s(bool,X90),s(bool,X91))) | s(bool,'T_0') = s(bool,X91) | ~p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X93) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X89),s(bool,X90),s(bool,'F_0'))) | s(bool,X92) != s(bool,X94)) ) | $spl34),
  inference(resolution,[],[f11584,f89])).
fof(f19916,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X77),s(bool,X78),s(bool,X79))) | s(bool,'T_0') = s(bool,X79) | p(s(bool,X80)) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X77),s(bool,X78),s(bool,'F_0')))) ) | $spl34),
  inference(resolution,[],[f11584,f88])).
fof(f19914,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X65),s(bool,X66),s(bool,X67))) | s(bool,'T_0') = s(bool,X67) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X65),s(bool,X66),s(bool,'F_0')))) ) | $spl34),
  inference(resolution,[],[f11584,f87])).
fof(f19913,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X59),s(bool,X60),s(bool,X61))) | s(bool,'T_0') = s(bool,X61) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X59),s(bool,X60),s(bool,'F_0'))) | s(bool,X62) != s(bool,X64)) ) | $spl34),
  inference(resolution,[],[f11584,f85])).
fof(f19911,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,'T_0') = s(bool,X49) | p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,X52) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,'F_0')))) ) | $spl34),
  inference(resolution,[],[f11584,f83])).
fof(f19894,plain,(
  ( ! [X26,X24] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'T_0'))),s(bool,sK5(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK5(X24,bool)))),s(bool,'T_0'))) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24))) | s(bool,'F_0') = s(bool,sK4(X24,bool))) ) | $spl324),
  inference(cnf_transformation,[],[f19894_D])).
fof(f19894_D,plain,(
  ( ! [X26,X24] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'T_0'))),s(bool,sK5(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK5(X24,bool)))),s(bool,'T_0'))) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24))) | s(bool,'F_0') = s(bool,sK4(X24,bool))) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f19884,plain,(
  ( ! [X59,X60,X58] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X58),s(bool,'T_0'))),s(bool,sK5(X58,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X58),s(bool,sK5(X58,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK4(X58,bool)) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X60) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X58)))) )),
  inference(resolution,[],[f1889,f125])).
fof(f1889,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X0)))) | s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X0),s(bool,'T_0'))),s(bool,sK5(X0,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X0),s(bool,sK5(X0,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK4(X0,bool))) )),
  inference(superposition,[],[f119,f122])).
fof(f19881,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X52),s(bool,'T_0'))),s(bool,sK5(X52,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X52),s(bool,sK5(X52,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK4(X52,bool)) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X52))) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f1889,f89])).
fof(f19878,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X40),s(bool,'T_0'))),s(bool,sK5(X40,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X40),s(bool,sK5(X40,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK4(X40,bool)) | p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X40)))) )),
  inference(resolution,[],[f1889,f88])).
fof(f19877,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X36),s(bool,'T_0'))),s(bool,sK5(X36,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X36),s(bool,sK5(X36,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK4(X36,bool)) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X36)))) )),
  inference(resolution,[],[f1889,f87])).
fof(f19876,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X32),s(bool,'T_0'))),s(bool,sK5(X32,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X32),s(bool,sK5(X32,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK4(X32,bool)) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X32))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f1889,f85])).
fof(f19874,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'T_0'))),s(bool,sK5(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK5(X24,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK4(X24,bool)) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24)))) )),
  inference(resolution,[],[f1889,f83])).
fof(f19832,plain,(
  ( ! [X47,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,'T_0'))) | ~p(s(bool,X51)) | s(bool,'F_0') = s(bool,X49)) ) | $spl322),
  inference(cnf_transformation,[],[f19832_D])).
fof(f19832_D,plain,(
  ( ! [X47,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,'T_0'))) | ~p(s(bool,X51)) | s(bool,'F_0') = s(bool,X49)) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f19776,plain,(
  ( ! [X107,X105,X108,X106,X104] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,X105),s(bool,X106))) | s(bool,'F_0') = s(bool,X106) | p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X108) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,X105),s(bool,'T_0')))) ) | $spl34),
  inference(resolution,[],[f11518,f125])).
fof(f11518,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X7),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X7),s(bool,X8))) | s(bool,'F_0') = s(bool,X8)) ) | $spl34),
  inference(superposition,[],[f2697,f2341])).
fof(f2697,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,'T_0'),s(bool,X2)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X2),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) ) | $spl34),
  inference(superposition,[],[f2438,f122])).
fof(f19772,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X89),s(bool,X90),s(bool,X91))) | s(bool,'F_0') = s(bool,X91) | ~p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X93) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X89),s(bool,X90),s(bool,'T_0'))) | s(bool,X92) != s(bool,X94)) ) | $spl34),
  inference(resolution,[],[f11518,f89])).
fof(f19770,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X77),s(bool,X78),s(bool,X79))) | s(bool,'F_0') = s(bool,X79) | p(s(bool,X80)) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X77),s(bool,X78),s(bool,'T_0')))) ) | $spl34),
  inference(resolution,[],[f11518,f88])).
fof(f19768,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X65),s(bool,X66),s(bool,X67))) | s(bool,'F_0') = s(bool,X67) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X65),s(bool,X66),s(bool,'T_0')))) ) | $spl34),
  inference(resolution,[],[f11518,f87])).
fof(f19767,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X59),s(bool,X60),s(bool,X61))) | s(bool,'F_0') = s(bool,X61) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X59),s(bool,X60),s(bool,'T_0'))) | s(bool,X62) != s(bool,X64)) ) | $spl34),
  inference(resolution,[],[f11518,f85])).
fof(f19765,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,X49))) | s(bool,'F_0') = s(bool,X49) | p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,X52) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X47),s(bool,X48),s(bool,'T_0')))) ) | $spl34),
  inference(resolution,[],[f11518,f83])).
fof(f19634,plain,(
  ( ! [X26,X24] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'F_0'))),s(bool,sK4(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK4(X24,bool)))),s(bool,'F_0'))) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24))) | s(bool,'T_0') = s(bool,sK5(X24,bool))) ) | $spl320),
  inference(cnf_transformation,[],[f19634_D])).
fof(f19634_D,plain,(
  ( ! [X26,X24] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'F_0'))),s(bool,sK4(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK4(X24,bool)))),s(bool,'F_0'))) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24))) | s(bool,'T_0') = s(bool,sK5(X24,bool))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f19624,plain,(
  ( ! [X59,X60,X58] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X58),s(bool,'F_0'))),s(bool,sK4(X58,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X58),s(bool,sK4(X58,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X58,bool)) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X60) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X58)))) )),
  inference(resolution,[],[f1884,f125])).
fof(f1884,plain,(
  ( ! [X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X1)))) | s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X1),s(bool,'F_0'))),s(bool,sK4(X1,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X1),s(bool,sK4(X1,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X1,bool))) )),
  inference(superposition,[],[f119,f122])).
fof(f19621,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X52),s(bool,'F_0'))),s(bool,sK4(X52,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X52),s(bool,sK4(X52,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X52,bool)) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X52))) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f1884,f89])).
fof(f19618,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X40),s(bool,'F_0'))),s(bool,sK4(X40,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X40),s(bool,sK4(X40,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X40,bool)) | p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X40)))) )),
  inference(resolution,[],[f1884,f88])).
fof(f19617,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X36),s(bool,'F_0'))),s(bool,sK4(X36,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X36),s(bool,sK4(X36,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X36,bool)) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X36)))) )),
  inference(resolution,[],[f1884,f87])).
fof(f19616,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X32),s(bool,'F_0'))),s(bool,sK4(X32,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X32),s(bool,sK4(X32,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X32,bool)) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X32))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f1884,f85])).
fof(f19614,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'F_0'))),s(bool,sK4(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK4(X24,bool)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X24,bool)) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24)))) )),
  inference(resolution,[],[f1884,f83])).
fof(f19477,plain,(
  ( ! [X52,X55,X53,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X51),s(bool,X53))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,'F_0'),s(bool,X53))) | ~p(s(bool,X55))) ) | $spl318),
  inference(cnf_transformation,[],[f19477_D])).
fof(f19477_D,plain,(
  ( ! [X52,X55,X53,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X51),s(bool,X53))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,'F_0'),s(bool,X53))) | ~p(s(bool,X55))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f19388,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,'T_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X109),s(bool,'F_0'),s(bool,X110))) | p(s(bool,X111)) | ~p(s(bool,X112)) | s(bool,X112) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X109),s(bool,X108),s(bool,X110)))) ) | $spl32),
  inference(resolution,[],[f11279,f125])).
fof(f11279,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X2),s(bool,X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,'F_0'),s(bool,X1)))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f11272])).
fof(f11272,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X2),s(bool,X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,'F_0'),s(bool,X1)))) ) | $spl32),
  inference(superposition,[],[f2645,f2343])).
fof(f2645,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,'F_0'))) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X3),s(bool,X5)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl32),
  inference(superposition,[],[f2435,f122])).
fof(f2435,plain,(
  ( ! [X325,X323,X324,X322] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X322,sK1),X323),s(X322,X325),s(X322,X324))) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X322,sK1),X323),s(X322,X324),s(X322,X325))))) ) | $spl32),
  inference(superposition,[],[f496,f2341])).
fof(f19384,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,'F_0'),s(bool,X95))) | ~p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X97) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X93),s(bool,X95))) | s(bool,X96) != s(bool,X98)) ) | $spl32),
  inference(resolution,[],[f11279,f89])).
fof(f19382,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X82),s(bool,'F_0'),s(bool,X83))) | p(s(bool,X84)) | ~p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X82),s(bool,X81),s(bool,X83)))) ) | $spl32),
  inference(resolution,[],[f11279,f88])).
fof(f19380,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,'T_0') = s(bool,X69) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X70),s(bool,'F_0'),s(bool,X71))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X72) != s(bool,X74) | s(bool,X73) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X70),s(bool,X69),s(bool,X71)))) ) | $spl32),
  inference(resolution,[],[f11279,f87])).
fof(f19379,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X64),s(bool,'F_0'),s(bool,X65))) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X67) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X64),s(bool,X63),s(bool,X65))) | s(bool,X66) != s(bool,X68)) ) | $spl32),
  inference(resolution,[],[f11279,f85])).
fof(f19377,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,'F_0'),s(bool,X53))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,X56) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X51),s(bool,X53)))) ) | $spl32),
  inference(resolution,[],[f11279,f83])).
fof(f19297,plain,(
  ( ! [X26,X24] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'T_0'))),s(bool,sK4(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK4(X24,bool)))),s(bool,'T_0'))) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24))) | s(bool,'F_0') = s(bool,sK5(X24,bool))) ) | $spl316),
  inference(cnf_transformation,[],[f19297_D])).
fof(f19297_D,plain,(
  ( ! [X26,X24] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'T_0'))),s(bool,sK4(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK4(X24,bool)))),s(bool,'T_0'))) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24))) | s(bool,'F_0') = s(bool,sK5(X24,bool))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f19287,plain,(
  ( ! [X59,X60,X58] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X58),s(bool,'T_0'))),s(bool,sK4(X58,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X58),s(bool,sK4(X58,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK5(X58,bool)) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X60) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X58)))) )),
  inference(resolution,[],[f1883,f125])).
fof(f1883,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X0)))) | s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X0),s(bool,'T_0'))),s(bool,sK4(X0,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X0),s(bool,sK4(X0,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK5(X0,bool))) )),
  inference(superposition,[],[f119,f122])).
fof(f19284,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X52),s(bool,'T_0'))),s(bool,sK4(X52,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X52),s(bool,sK4(X52,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK5(X52,bool)) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X52))) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f1883,f89])).
fof(f19281,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X40),s(bool,'T_0'))),s(bool,sK4(X40,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X40),s(bool,sK4(X40,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK5(X40,bool)) | p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X40)))) )),
  inference(resolution,[],[f1883,f88])).
fof(f19280,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X36),s(bool,'T_0'))),s(bool,sK4(X36,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X36),s(bool,sK4(X36,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK5(X36,bool)) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X36)))) )),
  inference(resolution,[],[f1883,f87])).
fof(f19279,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X32),s(bool,'T_0'))),s(bool,sK4(X32,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X32),s(bool,sK4(X32,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK5(X32,bool)) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X32))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f1883,f85])).
fof(f19277,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,'T_0'))),s(bool,sK4(X24,bool)))) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X24),s(bool,sK4(X24,bool)))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK5(X24,bool)) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(bool,fun(bool,bool)),X24)))) )),
  inference(resolution,[],[f1883,f83])).
fof(f19255,plain,(
  ( ! [X52,X55,X53,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X51),s(bool,X53))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,'T_0'),s(bool,X53))) | ~p(s(bool,X55))) ) | $spl314),
  inference(cnf_transformation,[],[f19255_D])).
fof(f19255_D,plain,(
  ( ! [X52,X55,X53,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X51),s(bool,X53))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,'T_0'),s(bool,X53))) | ~p(s(bool,X55))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f19166,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,'F_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X109),s(bool,'T_0'),s(bool,X110))) | p(s(bool,X111)) | ~p(s(bool,X112)) | s(bool,X112) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X109),s(bool,X108),s(bool,X110)))) ) | $spl32),
  inference(resolution,[],[f11221,f125])).
fof(f11221,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X2),s(bool,X1)))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,'T_0'),s(bool,X1)))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f11214])).
fof(f11214,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X2),s(bool,X1)))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,'T_0'),s(bool,X1)))) ) | $spl32),
  inference(superposition,[],[f2644,f2343])).
fof(f2644,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X2),s(bool,'T_0'))) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X0),s(bool,X2)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl32),
  inference(superposition,[],[f2435,f122])).
fof(f19162,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,'F_0') = s(bool,X93) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,'T_0'),s(bool,X95))) | ~p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X97) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X93),s(bool,X95))) | s(bool,X96) != s(bool,X98)) ) | $spl32),
  inference(resolution,[],[f11221,f89])).
fof(f19160,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X82),s(bool,'T_0'),s(bool,X83))) | p(s(bool,X84)) | ~p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X82),s(bool,X81),s(bool,X83)))) ) | $spl32),
  inference(resolution,[],[f11221,f88])).
fof(f19158,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X70),s(bool,'T_0'),s(bool,X71))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X72) != s(bool,X74) | s(bool,X73) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X70),s(bool,X69),s(bool,X71)))) ) | $spl32),
  inference(resolution,[],[f11221,f87])).
fof(f19157,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X64),s(bool,'T_0'),s(bool,X65))) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X67) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X64),s(bool,X63),s(bool,X65))) | s(bool,X66) != s(bool,X68)) ) | $spl32),
  inference(resolution,[],[f11221,f85])).
fof(f19155,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,'T_0'),s(bool,X53))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,X56) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X51),s(bool,X53)))) ) | $spl32),
  inference(resolution,[],[f11221,f83])).
fof(f18868,plain,(
  ( ! [X52,X55,X53,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,'F_0'))) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,X51))) | ~p(s(bool,X55))) ) | $spl312),
  inference(cnf_transformation,[],[f18868_D])).
fof(f18868_D,plain,(
  ( ! [X52,X55,X53,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,'F_0'))) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,X51))) | ~p(s(bool,X55))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f18779,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,'T_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X109),s(bool,X110),s(bool,'F_0'))) | p(s(bool,X111)) | ~p(s(bool,X112)) | s(bool,X112) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X109),s(bool,X110),s(bool,X108)))) ) | $spl32),
  inference(resolution,[],[f11018,f125])).
fof(f11018,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X1),s(bool,X2)))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X1),s(bool,'F_0')))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f11011])).
fof(f11011,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X1),s(bool,X2)))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X1),s(bool,'F_0')))) ) | $spl32),
  inference(superposition,[],[f2629,f2343])).
fof(f2629,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,'F_0'),s(bool,X5))) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,X3)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl32),
  inference(superposition,[],[f2435,f122])).
fof(f18775,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X95),s(bool,'F_0'))) | ~p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X97) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X95),s(bool,X93))) | s(bool,X96) != s(bool,X98)) ) | $spl32),
  inference(resolution,[],[f11018,f89])).
fof(f18773,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X82),s(bool,X83),s(bool,'F_0'))) | p(s(bool,X84)) | ~p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X82),s(bool,X83),s(bool,X81)))) ) | $spl32),
  inference(resolution,[],[f11018,f88])).
fof(f18771,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,'T_0') = s(bool,X69) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X70),s(bool,X71),s(bool,'F_0'))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X72) != s(bool,X74) | s(bool,X73) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X70),s(bool,X71),s(bool,X69)))) ) | $spl32),
  inference(resolution,[],[f11018,f87])).
fof(f18770,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X64),s(bool,X65),s(bool,'F_0'))) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X67) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X64),s(bool,X65),s(bool,X63))) | s(bool,X66) != s(bool,X68)) ) | $spl32),
  inference(resolution,[],[f11018,f85])).
fof(f18768,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,'F_0'))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,X56) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,X51)))) ) | $spl32),
  inference(resolution,[],[f11018,f83])).
fof(f18633,plain,(
  ( ! [X52,X55,X53,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,'T_0'))) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,X51))) | ~p(s(bool,X55))) ) | $spl310),
  inference(cnf_transformation,[],[f18633_D])).
fof(f18633_D,plain,(
  ( ! [X52,X55,X53,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,'T_0'))) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,X51))) | ~p(s(bool,X55))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f18544,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,'F_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X109),s(bool,X110),s(bool,'T_0'))) | p(s(bool,X111)) | ~p(s(bool,X112)) | s(bool,X112) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X109),s(bool,X110),s(bool,X108)))) ) | $spl32),
  inference(resolution,[],[f10975,f125])).
fof(f10975,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X1),s(bool,'T_0')))) ) | $spl32),
  inference(trivial_inequality_removal,[],[f10968])).
fof(f10968,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X0),s(bool,X1),s(bool,'T_0')))) ) | $spl32),
  inference(superposition,[],[f2628,f2343])).
fof(f2628,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,'T_0'),s(bool,X2))) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X2),s(bool,X0)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl32),
  inference(superposition,[],[f2435,f122])).
fof(f18540,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,'F_0') = s(bool,X93) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X95),s(bool,'T_0'))) | ~p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X97) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X95),s(bool,X93))) | s(bool,X96) != s(bool,X98)) ) | $spl32),
  inference(resolution,[],[f10975,f89])).
fof(f18538,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X82),s(bool,X83),s(bool,'T_0'))) | p(s(bool,X84)) | ~p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X82),s(bool,X83),s(bool,X81)))) ) | $spl32),
  inference(resolution,[],[f10975,f88])).
fof(f18536,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X70),s(bool,X71),s(bool,'T_0'))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X72) != s(bool,X74) | s(bool,X73) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X70),s(bool,X71),s(bool,X69)))) ) | $spl32),
  inference(resolution,[],[f10975,f87])).
fof(f18535,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X64),s(bool,X65),s(bool,'T_0'))) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X67) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X64),s(bool,X65),s(bool,X63))) | s(bool,X66) != s(bool,X68)) ) | $spl32),
  inference(resolution,[],[f10975,f85])).
fof(f18533,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,'T_0'))) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,X56) | s(bool,X55) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X52),s(bool,X53),s(bool,X51)))) ) | $spl32),
  inference(resolution,[],[f10975,f83])).
fof(f18421,plain,(
  ( ! [X94,X92,X93,X96] : (s(bool,'T_0') = s(bool,X92) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,'F_0'),s(bool,X94))) | p(s(bool,X96)) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X92),s(bool,X94)))) ) | $spl308),
  inference(cnf_transformation,[],[f18421_D])).
fof(f18421_D,plain,(
  ( ! [X94,X92,X93,X96] : (s(bool,'T_0') = s(bool,X92) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,'F_0'),s(bool,X94))) | p(s(bool,X96)) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X92),s(bool,X94)))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f18370,plain,(
  ( ! [X146,X149,X147,X150,X148] : (s(bool,'T_0') = s(bool,X146) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X147),s(bool,X146),s(bool,X148))) | p(s(bool,X149)) | p(s(bool,X150)) | s(bool,X150) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X147),s(bool,'F_0'),s(bool,X148)))) ) | $spl24),
  inference(resolution,[],[f10929,f125])).
fof(f10929,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f10923])).
fof(f10923,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl24),
  inference(superposition,[],[f2598,f2343])).
fof(f2598,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X11),s(bool,X9))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,'F_0'),s(bool,X11)))) | s(bool,'T_0') = s(bool,X9)) ) | $spl24),
  inference(superposition,[],[f2433,f122])).
fof(f2433,plain,(
  ( ! [X315,X316,X314,X317] : (s(bool,'T_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X314,sK1),X315),s(X314,X317),s(X314,X316))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X314,sK1),X315),s(X314,X316),s(X314,X317))))) ) | $spl24),
  inference(superposition,[],[f385,f2341])).
fof(f18367,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(bool,'T_0') = s(bool,X134) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X135),s(bool,X134),s(bool,X136))) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X135),s(bool,'F_0'),s(bool,X136))) | s(bool,X138) != s(bool,X139)) ) | $spl24),
  inference(resolution,[],[f10929,f89])).
fof(f18364,plain,(
  ( ! [X121,X118,X116,X120,X119,X117] : (s(bool,'T_0') = s(bool,X116) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X117),s(bool,X116),s(bool,X118))) | p(s(bool,X119)) | ~p(s(bool,X120)) | p(s(bool,X121)) | s(bool,X121) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X117),s(bool,'F_0'),s(bool,X118)))) ) | $spl24),
  inference(resolution,[],[f10929,f88])).
fof(f18363,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (s(bool,'T_0') = s(bool,X110) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X111),s(bool,X110),s(bool,X112))) | p(s(bool,X113)) | p(s(bool,X114)) | s(bool,X113) != s(bool,X115) | s(bool,X114) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X111),s(bool,'F_0'),s(bool,X112)))) ) | $spl24),
  inference(resolution,[],[f10929,f87])).
fof(f18362,plain,(
  ( ! [X109,X107,X105,X108,X106,X104] : (s(bool,'T_0') = s(bool,X104) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X105),s(bool,X104),s(bool,X106))) | p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X105),s(bool,'F_0'),s(bool,X106))) | s(bool,X108) != s(bool,X109)) ) | $spl24),
  inference(resolution,[],[f10929,f85])).
fof(f18360,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,'T_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X92),s(bool,X94))) | p(s(bool,X95)) | p(s(bool,X96)) | s(bool,X95) != s(bool,X97) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,'F_0'),s(bool,X94)))) ) | $spl24),
  inference(resolution,[],[f10929,f83])).
fof(f18232,plain,(
  ( ! [X94,X92,X93,X96] : (s(bool,'F_0') = s(bool,X92) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,'T_0'),s(bool,X94))) | p(s(bool,X96)) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X92),s(bool,X94)))) ) | $spl306),
  inference(cnf_transformation,[],[f18232_D])).
fof(f18232_D,plain,(
  ( ! [X94,X92,X93,X96] : (s(bool,'F_0') = s(bool,X92) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,'T_0'),s(bool,X94))) | p(s(bool,X96)) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X92),s(bool,X94)))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f18182,plain,(
  ( ! [X146,X149,X147,X150,X148] : (s(bool,'F_0') = s(bool,X146) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X147),s(bool,X146),s(bool,X148))) | p(s(bool,X149)) | p(s(bool,X150)) | s(bool,X150) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X147),s(bool,'T_0'),s(bool,X148)))) ) | $spl24),
  inference(resolution,[],[f10872,f125])).
fof(f10872,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl24),
  inference(subsumption_resolution,[],[f10858,f122])).
fof(f10858,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl24),
  inference(superposition,[],[f2597,f2343])).
fof(f2597,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X8),s(bool,X6))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,'T_0'),s(bool,X8)))) | s(bool,'F_0') = s(bool,X6)) ) | $spl24),
  inference(superposition,[],[f2433,f122])).
fof(f18179,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(bool,'F_0') = s(bool,X134) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X135),s(bool,X134),s(bool,X136))) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X135),s(bool,'T_0'),s(bool,X136))) | s(bool,X138) != s(bool,X139)) ) | $spl24),
  inference(resolution,[],[f10872,f89])).
fof(f18176,plain,(
  ( ! [X121,X118,X116,X120,X119,X117] : (s(bool,'F_0') = s(bool,X116) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X117),s(bool,X116),s(bool,X118))) | p(s(bool,X119)) | ~p(s(bool,X120)) | p(s(bool,X121)) | s(bool,X121) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X117),s(bool,'T_0'),s(bool,X118)))) ) | $spl24),
  inference(resolution,[],[f10872,f88])).
fof(f18175,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (s(bool,'F_0') = s(bool,X110) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X111),s(bool,X110),s(bool,X112))) | p(s(bool,X113)) | p(s(bool,X114)) | s(bool,X113) != s(bool,X115) | s(bool,X114) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X111),s(bool,'T_0'),s(bool,X112)))) ) | $spl24),
  inference(resolution,[],[f10872,f87])).
fof(f18174,plain,(
  ( ! [X109,X107,X105,X108,X106,X104] : (s(bool,'F_0') = s(bool,X104) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X105),s(bool,X104),s(bool,X106))) | p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X105),s(bool,'T_0'),s(bool,X106))) | s(bool,X108) != s(bool,X109)) ) | $spl24),
  inference(resolution,[],[f10872,f85])).
fof(f18172,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X92),s(bool,X94))) | p(s(bool,X95)) | p(s(bool,X96)) | s(bool,X95) != s(bool,X97) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,'T_0'),s(bool,X94)))) ) | $spl24),
  inference(resolution,[],[f10872,f83])).
fof(f18043,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'T_0') = s(bool,X100) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X100),s(bool,X102))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,'F_0'),s(bool,X102))) | p(s(bool,X104))) ) | $spl304),
  inference(cnf_transformation,[],[f18043_D])).
fof(f18043_D,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'T_0') = s(bool,X100) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X100),s(bool,X102))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,'F_0'),s(bool,X102))) | p(s(bool,X104))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f17957,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,'T_0') = s(bool,X154) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X155),s(bool,'F_0'),s(bool,X156))) | p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X155),s(bool,X154),s(bool,X156)))) ) | $spl24),
  inference(resolution,[],[f10792,f125])).
fof(f10792,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f10786])).
fof(f10786,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) ) | $spl24),
  inference(superposition,[],[f2596,f2343])).
fof(f2596,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,'F_0'))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X3),s(bool,X5)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl24),
  inference(superposition,[],[f2433,f122])).
fof(f17954,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (s(bool,'T_0') = s(bool,X142) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X143),s(bool,'F_0'),s(bool,X144))) | p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X143),s(bool,X142),s(bool,X144))) | s(bool,X146) != s(bool,X147)) ) | $spl24),
  inference(resolution,[],[f10792,f89])).
fof(f17951,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (s(bool,'T_0') = s(bool,X124) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X125),s(bool,'F_0'),s(bool,X126))) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X129) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X125),s(bool,X124),s(bool,X126)))) ) | $spl24),
  inference(resolution,[],[f10792,f88])).
fof(f17950,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X119),s(bool,'F_0'),s(bool,X120))) | p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X122) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X119),s(bool,X118),s(bool,X120)))) ) | $spl24),
  inference(resolution,[],[f10792,f87])).
fof(f17949,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (s(bool,'T_0') = s(bool,X112) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X113),s(bool,'F_0'),s(bool,X114))) | p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X115) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X113),s(bool,X112),s(bool,X114))) | s(bool,X116) != s(bool,X117)) ) | $spl24),
  inference(resolution,[],[f10792,f85])).
fof(f17947,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (s(bool,'T_0') = s(bool,X100) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,'F_0'),s(bool,X102))) | p(s(bool,X103)) | p(s(bool,X104)) | s(bool,X103) != s(bool,X105) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X100),s(bool,X102)))) ) | $spl24),
  inference(resolution,[],[f10792,f83])).
fof(f17840,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,X100) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X100),s(bool,X102))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,'T_0'),s(bool,X102))) | p(s(bool,X104))) ) | $spl302),
  inference(cnf_transformation,[],[f17840_D])).
fof(f17840_D,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,X100) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X100),s(bool,X102))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,'T_0'),s(bool,X102))) | p(s(bool,X104))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f17755,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,'F_0') = s(bool,X154) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X155),s(bool,'T_0'),s(bool,X156))) | p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X155),s(bool,X154),s(bool,X156)))) ) | $spl24),
  inference(resolution,[],[f10734,f125])).
fof(f10734,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f10728])).
fof(f10728,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) ) | $spl24),
  inference(superposition,[],[f2595,f2343])).
fof(f2595,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X2),s(bool,'T_0'))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X0),s(bool,X2)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f2433,f122])).
fof(f17752,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (s(bool,'F_0') = s(bool,X142) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X143),s(bool,'T_0'),s(bool,X144))) | p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X143),s(bool,X142),s(bool,X144))) | s(bool,X146) != s(bool,X147)) ) | $spl24),
  inference(resolution,[],[f10734,f89])).
fof(f17749,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (s(bool,'F_0') = s(bool,X124) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X125),s(bool,'T_0'),s(bool,X126))) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X129) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X125),s(bool,X124),s(bool,X126)))) ) | $spl24),
  inference(resolution,[],[f10734,f88])).
fof(f17748,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (s(bool,'F_0') = s(bool,X118) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X119),s(bool,'T_0'),s(bool,X120))) | p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X122) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X119),s(bool,X118),s(bool,X120)))) ) | $spl24),
  inference(resolution,[],[f10734,f87])).
fof(f17747,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (s(bool,'F_0') = s(bool,X112) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X113),s(bool,'T_0'),s(bool,X114))) | p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X115) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X113),s(bool,X112),s(bool,X114))) | s(bool,X116) != s(bool,X117)) ) | $spl24),
  inference(resolution,[],[f10734,f85])).
fof(f17745,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (s(bool,'F_0') = s(bool,X100) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,'T_0'),s(bool,X102))) | p(s(bool,X103)) | p(s(bool,X104)) | s(bool,X103) != s(bool,X105) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X100),s(bool,X102)))) ) | $spl24),
  inference(resolution,[],[f10734,f83])).
fof(f17629,plain,(
  ( ! [X94,X97,X95,X93] : (s(bool,'T_0') = s(bool,X93) | s(bool,X97) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X95),s(bool,'F_0'))) | p(s(bool,X97)) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X95),s(bool,X93)))) ) | $spl300),
  inference(cnf_transformation,[],[f17629_D])).
fof(f17629_D,plain,(
  ( ! [X94,X97,X95,X93] : (s(bool,'T_0') = s(bool,X93) | s(bool,X97) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X95),s(bool,'F_0'))) | p(s(bool,X97)) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X95),s(bool,X93)))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f17578,plain,(
  ( ! [X151,X149,X147,X150,X148] : (s(bool,'T_0') = s(bool,X147) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X148),s(bool,X149),s(bool,X147))) | p(s(bool,X150)) | p(s(bool,X151)) | s(bool,X151) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X148),s(bool,X149),s(bool,'F_0')))) ) | $spl24),
  inference(resolution,[],[f10636,f125])).
fof(f10636,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f10630])).
fof(f10630,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl24),
  inference(superposition,[],[f2582,f2343])).
fof(f2582,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X9),s(bool,X11))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X11),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X9)) ) | $spl24),
  inference(superposition,[],[f2433,f122])).
fof(f17575,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (s(bool,'T_0') = s(bool,X135) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X136),s(bool,X137),s(bool,X135))) | p(s(bool,X138)) | ~p(s(bool,X139)) | s(bool,X138) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X136),s(bool,X137),s(bool,'F_0'))) | s(bool,X139) != s(bool,X140)) ) | $spl24),
  inference(resolution,[],[f10636,f89])).
fof(f17572,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (s(bool,'T_0') = s(bool,X117) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X118),s(bool,X119),s(bool,X117))) | p(s(bool,X120)) | ~p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X122) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X118),s(bool,X119),s(bool,'F_0')))) ) | $spl24),
  inference(resolution,[],[f10636,f88])).
fof(f17571,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (s(bool,'T_0') = s(bool,X111) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X112),s(bool,X113),s(bool,X111))) | p(s(bool,X114)) | p(s(bool,X115)) | s(bool,X114) != s(bool,X116) | s(bool,X115) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X112),s(bool,X113),s(bool,'F_0')))) ) | $spl24),
  inference(resolution,[],[f10636,f87])).
fof(f17570,plain,(
  ( ! [X109,X107,X105,X110,X108,X106] : (s(bool,'T_0') = s(bool,X105) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X106),s(bool,X107),s(bool,X105))) | p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X108) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X106),s(bool,X107),s(bool,'F_0'))) | s(bool,X109) != s(bool,X110)) ) | $spl24),
  inference(resolution,[],[f10636,f85])).
fof(f17568,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,'T_0') = s(bool,X93) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X95),s(bool,X93))) | p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X96) != s(bool,X98) | s(bool,X97) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X94),s(bool,X95),s(bool,'F_0')))) ) | $spl24),
  inference(resolution,[],[f10636,f83])).
fof(f17389,plain,(
  ( ! [X94,X92,X90,X91] : (s(bool,'F_0') = s(bool,X90) | s(bool,X94) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X91),s(bool,X92),s(bool,'T_0'))) | p(s(bool,X94)) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X91),s(bool,X92),s(bool,X90)))) ) | $spl298),
  inference(cnf_transformation,[],[f17389_D])).
fof(f17389_D,plain,(
  ( ! [X94,X92,X90,X91] : (s(bool,'F_0') = s(bool,X90) | s(bool,X94) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X91),s(bool,X92),s(bool,'T_0'))) | p(s(bool,X94)) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X91),s(bool,X92),s(bool,X90)))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f17339,plain,(
  ( ! [X146,X144,X147,X145,X148] : (s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X145),s(bool,X146),s(bool,X144))) | p(s(bool,X147)) | p(s(bool,X148)) | s(bool,X148) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X145),s(bool,X146),s(bool,'T_0')))) ) | $spl24),
  inference(resolution,[],[f10579,f125])).
fof(f10579,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl24),
  inference(subsumption_resolution,[],[f10546,f122])).
fof(f10546,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl24),
  inference(superposition,[],[f2581,f2343])).
fof(f2581,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X6),s(bool,X8))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X8),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X6)) ) | $spl24),
  inference(superposition,[],[f2433,f122])).
fof(f17336,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'F_0') = s(bool,X132) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X133),s(bool,X134),s(bool,X132))) | p(s(bool,X135)) | ~p(s(bool,X136)) | s(bool,X135) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X133),s(bool,X134),s(bool,'T_0'))) | s(bool,X136) != s(bool,X137)) ) | $spl24),
  inference(resolution,[],[f10579,f89])).
fof(f17333,plain,(
  ( ! [X118,X116,X114,X119,X117,X115] : (s(bool,'F_0') = s(bool,X114) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X115),s(bool,X116),s(bool,X114))) | p(s(bool,X117)) | ~p(s(bool,X118)) | p(s(bool,X119)) | s(bool,X119) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X115),s(bool,X116),s(bool,'T_0')))) ) | $spl24),
  inference(resolution,[],[f10579,f88])).
fof(f17332,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,'F_0') = s(bool,X108) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X109),s(bool,X110),s(bool,X108))) | p(s(bool,X111)) | p(s(bool,X112)) | s(bool,X111) != s(bool,X113) | s(bool,X112) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X109),s(bool,X110),s(bool,'T_0')))) ) | $spl24),
  inference(resolution,[],[f10579,f87])).
fof(f17331,plain,(
  ( ! [X103,X107,X105,X102,X106,X104] : (s(bool,'F_0') = s(bool,X102) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X103),s(bool,X104),s(bool,X102))) | p(s(bool,X105)) | ~p(s(bool,X106)) | s(bool,X105) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X103),s(bool,X104),s(bool,'T_0'))) | s(bool,X106) != s(bool,X107)) ) | $spl24),
  inference(resolution,[],[f10579,f85])).
fof(f17329,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'F_0') = s(bool,X90) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X91),s(bool,X92),s(bool,X90))) | p(s(bool,X93)) | p(s(bool,X94)) | s(bool,X93) != s(bool,X95) | s(bool,X94) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X91),s(bool,X92),s(bool,'T_0')))) ) | $spl24),
  inference(resolution,[],[f10579,f83])).
fof(f17078,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'T_0') = s(bool,X100) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,'F_0'))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,X100))) | p(s(bool,X104))) ) | $spl296),
  inference(cnf_transformation,[],[f17078_D])).
fof(f17078_D,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'T_0') = s(bool,X100) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,'F_0'))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,X100))) | p(s(bool,X104))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f16993,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,'T_0') = s(bool,X154) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X155),s(bool,X156),s(bool,'F_0'))) | p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X155),s(bool,X156),s(bool,X154)))) ) | $spl24),
  inference(resolution,[],[f10485,f125])).
fof(f10485,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0')))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f10479])).
fof(f10479,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0')))) ) | $spl24),
  inference(superposition,[],[f2580,f2343])).
fof(f2580,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,'F_0'),s(bool,X5))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,X3)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl24),
  inference(superposition,[],[f2433,f122])).
fof(f16990,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (s(bool,'T_0') = s(bool,X142) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X143),s(bool,X144),s(bool,'F_0'))) | p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X143),s(bool,X144),s(bool,X142))) | s(bool,X146) != s(bool,X147)) ) | $spl24),
  inference(resolution,[],[f10485,f89])).
fof(f16987,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (s(bool,'T_0') = s(bool,X124) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X125),s(bool,X126),s(bool,'F_0'))) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X129) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X125),s(bool,X126),s(bool,X124)))) ) | $spl24),
  inference(resolution,[],[f10485,f88])).
fof(f16986,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X119),s(bool,X120),s(bool,'F_0'))) | p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X122) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X119),s(bool,X120),s(bool,X118)))) ) | $spl24),
  inference(resolution,[],[f10485,f87])).
fof(f16985,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (s(bool,'T_0') = s(bool,X112) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X113),s(bool,X114),s(bool,'F_0'))) | p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X115) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X113),s(bool,X114),s(bool,X112))) | s(bool,X116) != s(bool,X117)) ) | $spl24),
  inference(resolution,[],[f10485,f85])).
fof(f16983,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (s(bool,'T_0') = s(bool,X100) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,'F_0'))) | p(s(bool,X103)) | p(s(bool,X104)) | s(bool,X103) != s(bool,X105) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,X100)))) ) | $spl24),
  inference(resolution,[],[f10485,f83])).
fof(f16835,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,X100) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,'T_0'))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,X100))) | p(s(bool,X104))) ) | $spl294),
  inference(cnf_transformation,[],[f16835_D])).
fof(f16835_D,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,X100) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,'T_0'))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,X100))) | p(s(bool,X104))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f16749,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,'F_0') = s(bool,X154) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X155),s(bool,X156),s(bool,'T_0'))) | p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X155),s(bool,X156),s(bool,X154)))) ) | $spl24),
  inference(resolution,[],[f10442,f125])).
fof(f10442,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0')))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f10436])).
fof(f10436,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0')))) ) | $spl24),
  inference(superposition,[],[f2579,f2343])).
fof(f2579,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,'T_0'),s(bool,X2))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X2),s(bool,X0)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f2433,f122])).
fof(f16746,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (s(bool,'F_0') = s(bool,X142) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X143),s(bool,X144),s(bool,'T_0'))) | p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X143),s(bool,X144),s(bool,X142))) | s(bool,X146) != s(bool,X147)) ) | $spl24),
  inference(resolution,[],[f10442,f89])).
fof(f16743,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (s(bool,'F_0') = s(bool,X124) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X125),s(bool,X126),s(bool,'T_0'))) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X129) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X125),s(bool,X126),s(bool,X124)))) ) | $spl24),
  inference(resolution,[],[f10442,f88])).
fof(f16742,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (s(bool,'F_0') = s(bool,X118) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X119),s(bool,X120),s(bool,'T_0'))) | p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X122) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X119),s(bool,X120),s(bool,X118)))) ) | $spl24),
  inference(resolution,[],[f10442,f87])).
fof(f16741,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (s(bool,'F_0') = s(bool,X112) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X113),s(bool,X114),s(bool,'T_0'))) | p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X115) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X113),s(bool,X114),s(bool,X112))) | s(bool,X116) != s(bool,X117)) ) | $spl24),
  inference(resolution,[],[f10442,f85])).
fof(f16739,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (s(bool,'F_0') = s(bool,X100) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,'T_0'))) | p(s(bool,X103)) | p(s(bool,X104)) | s(bool,X103) != s(bool,X105) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X101),s(bool,X102),s(bool,X100)))) ) | $spl24),
  inference(resolution,[],[f10442,f83])).
fof(f16614,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,'F_0'),s(bool,X101))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,X102))) | p(s(bool,X104)) | s(bool,'T_0') = s(bool,X102)) ) | $spl292),
  inference(cnf_transformation,[],[f16614_D])).
fof(f16614_D,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,'F_0'),s(bool,X101))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,X102))) | p(s(bool,X104)) | s(bool,'T_0') = s(bool,X102)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f16527,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X154),s(bool,'F_0'),s(bool,X155))) | s(bool,'T_0') = s(bool,X156) | p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X154),s(bool,X155),s(bool,X156)))) ) | $spl24),
  inference(resolution,[],[f2809,f125])).
fof(f2809,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X11),s(bool,X9)))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,'F_0'),s(bool,X11))) | s(bool,'T_0') = s(bool,X9)) ) | $spl24),
  inference(superposition,[],[f2434,f122])).
fof(f2434,plain,(
  ( ! [X321,X318,X320,X319] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X318,sK1),X319),s(X318,X321),s(X318,X320)))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X318,sK1),X319),s(X318,X320),s(X318,X321)))) ) | $spl24),
  inference(superposition,[],[f432,f2341])).
fof(f432,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(trivial_inequality_removal,[],[f420])).
fof(f420,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f385,f122])).
fof(f16524,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X142),s(bool,'F_0'),s(bool,X143))) | s(bool,'T_0') = s(bool,X144) | p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X142),s(bool,X143),s(bool,X144))) | s(bool,X146) != s(bool,X147)) ) | $spl24),
  inference(resolution,[],[f2809,f89])).
fof(f16521,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X124),s(bool,'F_0'),s(bool,X125))) | s(bool,'T_0') = s(bool,X126) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X129) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X124),s(bool,X125),s(bool,X126)))) ) | $spl24),
  inference(resolution,[],[f2809,f88])).
fof(f16520,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X118),s(bool,'F_0'),s(bool,X119))) | s(bool,'T_0') = s(bool,X120) | p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X122) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X118),s(bool,X119),s(bool,X120)))) ) | $spl24),
  inference(resolution,[],[f2809,f87])).
fof(f16519,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X112),s(bool,'F_0'),s(bool,X113))) | s(bool,'T_0') = s(bool,X114) | p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X115) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X112),s(bool,X113),s(bool,X114))) | s(bool,X116) != s(bool,X117)) ) | $spl24),
  inference(resolution,[],[f2809,f85])).
fof(f16517,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,'F_0'),s(bool,X101))) | s(bool,'T_0') = s(bool,X102) | p(s(bool,X103)) | p(s(bool,X104)) | s(bool,X103) != s(bool,X105) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,X102)))) ) | $spl24),
  inference(resolution,[],[f2809,f83])).
fof(f16471,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,'T_0'),s(bool,X101))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,X102))) | p(s(bool,X104)) | s(bool,'F_0') = s(bool,X102)) ) | $spl290),
  inference(cnf_transformation,[],[f16471_D])).
fof(f16471_D,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,'T_0'),s(bool,X101))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,X102))) | p(s(bool,X104)) | s(bool,'F_0') = s(bool,X102)) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f16386,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X154),s(bool,'T_0'),s(bool,X155))) | s(bool,'F_0') = s(bool,X156) | p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X154),s(bool,X155),s(bool,X156)))) ) | $spl24),
  inference(resolution,[],[f2808,f125])).
fof(f2808,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X8),s(bool,X6)))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,'T_0'),s(bool,X8))) | s(bool,'F_0') = s(bool,X6)) ) | $spl24),
  inference(superposition,[],[f2434,f122])).
fof(f16383,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X142),s(bool,'T_0'),s(bool,X143))) | s(bool,'F_0') = s(bool,X144) | p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X142),s(bool,X143),s(bool,X144))) | s(bool,X146) != s(bool,X147)) ) | $spl24),
  inference(resolution,[],[f2808,f89])).
fof(f16380,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X124),s(bool,'T_0'),s(bool,X125))) | s(bool,'F_0') = s(bool,X126) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X129) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X124),s(bool,X125),s(bool,X126)))) ) | $spl24),
  inference(resolution,[],[f2808,f88])).
fof(f16379,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X118),s(bool,'T_0'),s(bool,X119))) | s(bool,'F_0') = s(bool,X120) | p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X122) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X118),s(bool,X119),s(bool,X120)))) ) | $spl24),
  inference(resolution,[],[f2808,f87])).
fof(f16378,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X112),s(bool,'T_0'),s(bool,X113))) | s(bool,'F_0') = s(bool,X114) | p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X115) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X112),s(bool,X113),s(bool,X114))) | s(bool,X116) != s(bool,X117)) ) | $spl24),
  inference(resolution,[],[f2808,f85])).
fof(f16376,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,'T_0'),s(bool,X101))) | s(bool,'F_0') = s(bool,X102) | p(s(bool,X103)) | p(s(bool,X104)) | s(bool,X103) != s(bool,X105) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,X102)))) ) | $spl24),
  inference(resolution,[],[f2808,f83])).
fof(f16314,plain,(
  ( ! [X94,X97,X95,X93] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X94),s(bool,X95))) | s(bool,X97) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X95),s(bool,'F_0'))) | p(s(bool,X97)) | s(bool,'T_0') = s(bool,X94)) ) | $spl288),
  inference(cnf_transformation,[],[f16314_D])).
fof(f16314_D,plain,(
  ( ! [X94,X97,X95,X93] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X94),s(bool,X95))) | s(bool,X97) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X95),s(bool,'F_0'))) | p(s(bool,X97)) | s(bool,'T_0') = s(bool,X94)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f16262,plain,(
  ( ! [X151,X149,X147,X150,X148] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X147),s(bool,X148),s(bool,X149))) | s(bool,'T_0') = s(bool,X148) | p(s(bool,X150)) | p(s(bool,X151)) | s(bool,X151) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X147),s(bool,X149),s(bool,'F_0')))) ) | $spl24),
  inference(resolution,[],[f2807,f125])).
fof(f2807,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X3),s(bool,X5))) | s(bool,'T_0') = s(bool,X3)) ) | $spl24),
  inference(superposition,[],[f2434,f122])).
fof(f16259,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X135),s(bool,X136),s(bool,X137))) | s(bool,'T_0') = s(bool,X136) | p(s(bool,X138)) | ~p(s(bool,X139)) | s(bool,X138) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X135),s(bool,X137),s(bool,'F_0'))) | s(bool,X139) != s(bool,X140)) ) | $spl24),
  inference(resolution,[],[f2807,f89])).
fof(f16256,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X117),s(bool,X118),s(bool,X119))) | s(bool,'T_0') = s(bool,X118) | p(s(bool,X120)) | ~p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X122) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X117),s(bool,X119),s(bool,'F_0')))) ) | $spl24),
  inference(resolution,[],[f2807,f88])).
fof(f16255,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X111),s(bool,X112),s(bool,X113))) | s(bool,'T_0') = s(bool,X112) | p(s(bool,X114)) | p(s(bool,X115)) | s(bool,X114) != s(bool,X116) | s(bool,X115) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X111),s(bool,X113),s(bool,'F_0')))) ) | $spl24),
  inference(resolution,[],[f2807,f87])).
fof(f16254,plain,(
  ( ! [X109,X107,X105,X110,X108,X106] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X105),s(bool,X106),s(bool,X107))) | s(bool,'T_0') = s(bool,X106) | p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X108) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X105),s(bool,X107),s(bool,'F_0'))) | s(bool,X109) != s(bool,X110)) ) | $spl24),
  inference(resolution,[],[f2807,f85])).
fof(f16252,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X94),s(bool,X95))) | s(bool,'T_0') = s(bool,X94) | p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X96) != s(bool,X98) | s(bool,X97) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X93),s(bool,X95),s(bool,'F_0')))) ) | $spl24),
  inference(resolution,[],[f2807,f83])).
fof(f16188,plain,(
  ( ! [X94,X92,X90,X91] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X90),s(bool,X91),s(bool,X92))) | s(bool,X94) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X90),s(bool,X92),s(bool,'T_0'))) | p(s(bool,X94)) | s(bool,'F_0') = s(bool,X91)) ) | $spl286),
  inference(cnf_transformation,[],[f16188_D])).
fof(f16188_D,plain,(
  ( ! [X94,X92,X90,X91] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X90),s(bool,X91),s(bool,X92))) | s(bool,X94) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X90),s(bool,X92),s(bool,'T_0'))) | p(s(bool,X94)) | s(bool,'F_0') = s(bool,X91)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f16137,plain,(
  ( ! [X146,X144,X147,X145,X148] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X144),s(bool,X145),s(bool,X146))) | s(bool,'F_0') = s(bool,X145) | p(s(bool,X147)) | p(s(bool,X148)) | s(bool,X148) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X144),s(bool,X146),s(bool,'T_0')))) ) | $spl24),
  inference(resolution,[],[f2806,f125])).
fof(f2806,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X2),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X0),s(bool,X2))) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f2434,f122])).
fof(f16134,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X132),s(bool,X133),s(bool,X134))) | s(bool,'F_0') = s(bool,X133) | p(s(bool,X135)) | ~p(s(bool,X136)) | s(bool,X135) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X132),s(bool,X134),s(bool,'T_0'))) | s(bool,X136) != s(bool,X137)) ) | $spl24),
  inference(resolution,[],[f2806,f89])).
fof(f16131,plain,(
  ( ! [X118,X116,X114,X119,X117,X115] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X114),s(bool,X115),s(bool,X116))) | s(bool,'F_0') = s(bool,X115) | p(s(bool,X117)) | ~p(s(bool,X118)) | p(s(bool,X119)) | s(bool,X119) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X114),s(bool,X116),s(bool,'T_0')))) ) | $spl24),
  inference(resolution,[],[f2806,f88])).
fof(f16130,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X108),s(bool,X109),s(bool,X110))) | s(bool,'F_0') = s(bool,X109) | p(s(bool,X111)) | p(s(bool,X112)) | s(bool,X111) != s(bool,X113) | s(bool,X112) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X108),s(bool,X110),s(bool,'T_0')))) ) | $spl24),
  inference(resolution,[],[f2806,f87])).
fof(f16129,plain,(
  ( ! [X103,X107,X105,X102,X106,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X102),s(bool,X103),s(bool,X104))) | s(bool,'F_0') = s(bool,X103) | p(s(bool,X105)) | ~p(s(bool,X106)) | s(bool,X105) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X102),s(bool,X104),s(bool,'T_0'))) | s(bool,X106) != s(bool,X107)) ) | $spl24),
  inference(resolution,[],[f2806,f85])).
fof(f16127,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X90),s(bool,X91),s(bool,X92))) | s(bool,'F_0') = s(bool,X91) | p(s(bool,X93)) | p(s(bool,X94)) | s(bool,X93) != s(bool,X95) | s(bool,X94) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X90),s(bool,X92),s(bool,'T_0')))) ) | $spl24),
  inference(resolution,[],[f2806,f83])).
fof(f16057,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,'F_0'))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X102),s(bool,X101))) | p(s(bool,X104)) | s(bool,'T_0') = s(bool,X102)) ) | $spl284),
  inference(cnf_transformation,[],[f16057_D])).
fof(f16057_D,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,'F_0'))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X102),s(bool,X101))) | p(s(bool,X104)) | s(bool,'T_0') = s(bool,X102)) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f15972,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X154),s(bool,X155),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X156) | p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X154),s(bool,X156),s(bool,X155)))) ) | $spl24),
  inference(resolution,[],[f2793,f125])).
fof(f2793,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X9),s(bool,X11)))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X11),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl24),
  inference(superposition,[],[f2434,f122])).
fof(f15969,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X142),s(bool,X143),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X144) | p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X142),s(bool,X144),s(bool,X143))) | s(bool,X146) != s(bool,X147)) ) | $spl24),
  inference(resolution,[],[f2793,f89])).
fof(f15966,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X124),s(bool,X125),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X126) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X129) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X124),s(bool,X126),s(bool,X125)))) ) | $spl24),
  inference(resolution,[],[f2793,f88])).
fof(f15965,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X118),s(bool,X119),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X120) | p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X122) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X118),s(bool,X120),s(bool,X119)))) ) | $spl24),
  inference(resolution,[],[f2793,f87])).
fof(f15964,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X112),s(bool,X113),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X114) | p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X115) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X112),s(bool,X114),s(bool,X113))) | s(bool,X116) != s(bool,X117)) ) | $spl24),
  inference(resolution,[],[f2793,f85])).
fof(f15962,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X102) | p(s(bool,X103)) | p(s(bool,X104)) | s(bool,X103) != s(bool,X105) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X102),s(bool,X101)))) ) | $spl24),
  inference(resolution,[],[f2793,f83])).
fof(f15887,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,'T_0'))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X102),s(bool,X101))) | p(s(bool,X104)) | s(bool,'F_0') = s(bool,X102)) ) | $spl282),
  inference(cnf_transformation,[],[f15887_D])).
fof(f15887_D,plain,(
  ( ! [X101,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,'T_0'))) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X102),s(bool,X101))) | p(s(bool,X104)) | s(bool,'F_0') = s(bool,X102)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f15799,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X154),s(bool,X155),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X156) | p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X154),s(bool,X156),s(bool,X155)))) ) | $spl24),
  inference(resolution,[],[f2792,f125])).
fof(f2792,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X6),s(bool,X8)))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X8),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X6)) ) | $spl24),
  inference(superposition,[],[f2434,f122])).
fof(f15796,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X142),s(bool,X143),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X144) | p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X142),s(bool,X144),s(bool,X143))) | s(bool,X146) != s(bool,X147)) ) | $spl24),
  inference(resolution,[],[f2792,f89])).
fof(f15793,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X124),s(bool,X125),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X126) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X129) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X124),s(bool,X126),s(bool,X125)))) ) | $spl24),
  inference(resolution,[],[f2792,f88])).
fof(f15792,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X118),s(bool,X119),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X120) | p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X122) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X118),s(bool,X120),s(bool,X119)))) ) | $spl24),
  inference(resolution,[],[f2792,f87])).
fof(f15791,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X112),s(bool,X113),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X114) | p(s(bool,X115)) | ~p(s(bool,X116)) | s(bool,X115) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X112),s(bool,X114),s(bool,X113))) | s(bool,X116) != s(bool,X117)) ) | $spl24),
  inference(resolution,[],[f2792,f85])).
fof(f15789,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X101),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X102) | p(s(bool,X103)) | p(s(bool,X104)) | s(bool,X103) != s(bool,X105) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X100),s(bool,X102),s(bool,X101)))) ) | $spl24),
  inference(resolution,[],[f2792,f83])).
fof(f15739,plain,(
  ( ! [X94,X92,X93,X96] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,X93),s(bool,X94))) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,'F_0'),s(bool,X93))) | p(s(bool,X96)) | s(bool,'T_0') = s(bool,X94)) ) | $spl280),
  inference(cnf_transformation,[],[f15739_D])).
fof(f15739_D,plain,(
  ( ! [X94,X92,X93,X96] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,X93),s(bool,X94))) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,'F_0'),s(bool,X93))) | p(s(bool,X96)) | s(bool,'T_0') = s(bool,X94)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f15687,plain,(
  ( ! [X146,X149,X147,X150,X148] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X146),s(bool,X147),s(bool,X148))) | s(bool,'T_0') = s(bool,X148) | p(s(bool,X149)) | p(s(bool,X150)) | s(bool,X150) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X146),s(bool,'F_0'),s(bool,X147)))) ) | $spl24),
  inference(resolution,[],[f2791,f125])).
fof(f2791,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,'F_0'),s(bool,X5)))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl24),
  inference(superposition,[],[f2434,f122])).
fof(f15684,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X134),s(bool,X135),s(bool,X136))) | s(bool,'T_0') = s(bool,X136) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X134),s(bool,'F_0'),s(bool,X135))) | s(bool,X138) != s(bool,X139)) ) | $spl24),
  inference(resolution,[],[f2791,f89])).
fof(f15681,plain,(
  ( ! [X121,X118,X116,X120,X119,X117] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X116),s(bool,X117),s(bool,X118))) | s(bool,'T_0') = s(bool,X118) | p(s(bool,X119)) | ~p(s(bool,X120)) | p(s(bool,X121)) | s(bool,X121) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X116),s(bool,'F_0'),s(bool,X117)))) ) | $spl24),
  inference(resolution,[],[f2791,f88])).
fof(f15680,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X110),s(bool,X111),s(bool,X112))) | s(bool,'T_0') = s(bool,X112) | p(s(bool,X113)) | p(s(bool,X114)) | s(bool,X113) != s(bool,X115) | s(bool,X114) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X110),s(bool,'F_0'),s(bool,X111)))) ) | $spl24),
  inference(resolution,[],[f2791,f87])).
fof(f15679,plain,(
  ( ! [X109,X107,X105,X108,X106,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,X105),s(bool,X106))) | s(bool,'T_0') = s(bool,X106) | p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,'F_0'),s(bool,X105))) | s(bool,X108) != s(bool,X109)) ) | $spl24),
  inference(resolution,[],[f2791,f85])).
fof(f15677,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,X93),s(bool,X94))) | s(bool,'T_0') = s(bool,X94) | p(s(bool,X95)) | p(s(bool,X96)) | s(bool,X95) != s(bool,X97) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,'F_0'),s(bool,X93)))) ) | $spl24),
  inference(resolution,[],[f2791,f83])).
fof(f15637,plain,(
  ( ! [X94,X92,X93,X96] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,X93),s(bool,X94))) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,'T_0'),s(bool,X93))) | p(s(bool,X96)) | s(bool,'F_0') = s(bool,X94)) ) | $spl278),
  inference(cnf_transformation,[],[f15637_D])).
fof(f15637_D,plain,(
  ( ! [X94,X92,X93,X96] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,X93),s(bool,X94))) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,'T_0'),s(bool,X93))) | p(s(bool,X96)) | s(bool,'F_0') = s(bool,X94)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f15586,plain,(
  ( ! [X146,X149,X147,X150,X148] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X146),s(bool,X147),s(bool,X148))) | s(bool,'F_0') = s(bool,X148) | p(s(bool,X149)) | p(s(bool,X150)) | s(bool,X150) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X146),s(bool,'T_0'),s(bool,X147)))) ) | $spl24),
  inference(resolution,[],[f2790,f125])).
fof(f2790,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,'T_0'),s(bool,X2)))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X2),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f2434,f122])).
fof(f15583,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X134),s(bool,X135),s(bool,X136))) | s(bool,'F_0') = s(bool,X136) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X134),s(bool,'T_0'),s(bool,X135))) | s(bool,X138) != s(bool,X139)) ) | $spl24),
  inference(resolution,[],[f2790,f89])).
fof(f15580,plain,(
  ( ! [X121,X118,X116,X120,X119,X117] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X116),s(bool,X117),s(bool,X118))) | s(bool,'F_0') = s(bool,X118) | p(s(bool,X119)) | ~p(s(bool,X120)) | p(s(bool,X121)) | s(bool,X121) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X116),s(bool,'T_0'),s(bool,X117)))) ) | $spl24),
  inference(resolution,[],[f2790,f88])).
fof(f15579,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X110),s(bool,X111),s(bool,X112))) | s(bool,'F_0') = s(bool,X112) | p(s(bool,X113)) | p(s(bool,X114)) | s(bool,X113) != s(bool,X115) | s(bool,X114) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X110),s(bool,'T_0'),s(bool,X111)))) ) | $spl24),
  inference(resolution,[],[f2790,f87])).
fof(f15578,plain,(
  ( ! [X109,X107,X105,X108,X106,X104] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,X105),s(bool,X106))) | s(bool,'F_0') = s(bool,X106) | p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X104),s(bool,'T_0'),s(bool,X105))) | s(bool,X108) != s(bool,X109)) ) | $spl24),
  inference(resolution,[],[f2790,f85])).
fof(f15576,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,X93),s(bool,X94))) | s(bool,'F_0') = s(bool,X94) | p(s(bool,X95)) | p(s(bool,X96)) | s(bool,X95) != s(bool,X97) | s(bool,X96) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X92),s(bool,'T_0'),s(bool,X93)))) ) | $spl24),
  inference(resolution,[],[f2790,f83])).
fof(f15446,plain,(
  ( ! [X705,X707,X708,X710,X704,X706,X711,X703] : (s(bool,X710) != s(bool,happ(s(fun(X703,bool),happ(s(fun(X703,fun(X703,bool)),X704),s(X703,happ(s(fun(X705,X703),X706),s(X705,X707))))),s(X703,happ(s(fun(X705,X703),X706),s(X705,X708))))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X703,fun(X703,bool)),X704),s(fun(X705,X703),X706),s(X705,X707),s(X705,X708))) | s(bool,'T_0') = s(bool,X711) | s(bool,X710) != s(bool,X711)) ) | $spl276),
  inference(cnf_transformation,[],[f15446_D])).
fof(f15446_D,plain,(
  ( ! [X705,X707,X708,X710,X704,X706,X711,X703] : (s(bool,X710) != s(bool,happ(s(fun(X703,bool),happ(s(fun(X703,fun(X703,bool)),X704),s(X703,happ(s(fun(X705,X703),X706),s(X705,X707))))),s(X703,happ(s(fun(X705,X703),X706),s(X705,X708))))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X703,fun(X703,bool)),X704),s(fun(X705,X703),X706),s(X705,X707),s(X705,X708))) | s(bool,'T_0') = s(bool,X711) | s(bool,X710) != s(bool,X711)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f15449,plain,(
  ( ! [X825,X819,X823,X821,X826,X824,X828,X818,X822,X820,X827] : (s(bool,X825) != s(bool,happ(s(fun(X818,bool),happ(s(fun(X818,fun(X818,bool)),X819),s(X818,happ(s(fun(X820,X818),X821),s(X820,X822))))),s(X818,happ(s(fun(X820,X818),X821),s(X820,X823))))) | s(bool,X824) != s(bool,X826) | s(bool,X826) != s(bool,X827) | s(bool,'F_0') = s(bool,X827) | s(bool,X825) != s(bool,X828) | s(bool,'T_0') = s(bool,X828) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X818,fun(X818,bool)),X819),s(fun(X820,X818),X821),s(X820,X822),s(X820,X823)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f15234,f789])).
fof(f15234,plain,(
  ( ! [X825,X819,X823,X821,X826,X824,X828,X818,X822,X820,X827] : (s(bool,'T_0') != s(bool,X824) | s(bool,X825) != s(bool,happ(s(fun(X818,bool),happ(s(fun(X818,fun(X818,bool)),X819),s(X818,happ(s(fun(X820,X818),X821),s(X820,X822))))),s(X818,happ(s(fun(X820,X818),X821),s(X820,X823))))) | s(bool,X824) != s(bool,X826) | s(bool,X826) != s(bool,X827) | s(bool,'F_0') = s(bool,X827) | s(bool,X825) != s(bool,X828) | s(bool,'T_0') = s(bool,X828) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X818,fun(X818,bool)),X819),s(fun(X820,X818),X821),s(X820,X822),s(X820,X823)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f1648])).
fof(f1648,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X32,bool),happ(s(fun(X32,fun(X32,bool)),X33),s(X32,happ(s(fun(X34,X32),X35),s(X34,X36))))),s(X32,happ(s(fun(X34,X32),X35),s(X34,X37))))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X32,fun(X32,bool)),X33),s(fun(X34,X32),X35),s(X34,X36),s(X34,X37)))) )),
  inference(superposition,[],[f123,f122])).
fof(f15448,plain,(
  ( ! [X765,X767,X763,X771,X769,X764,X766,X762,X770,X768] : (s(bool,X769) != s(bool,happ(s(fun(X762,bool),happ(s(fun(X762,fun(X762,bool)),X763),s(X762,happ(s(fun(X764,X762),X765),s(X764,X766))))),s(X762,happ(s(fun(X764,X762),X765),s(X764,X767))))) | s(bool,X768) != s(bool,X770) | s(bool,'F_0') = s(bool,X770) | s(bool,X769) != s(bool,X771) | s(bool,'T_0') = s(bool,X771) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X762,fun(X762,bool)),X763),s(fun(X764,X762),X765),s(X764,X766),s(X764,X767)))) ) | ($spl6 | $spl34 | $spl56)),
  inference(subsumption_resolution,[],[f15228,f589])).
fof(f15228,plain,(
  ( ! [X765,X767,X763,X771,X769,X764,X766,X762,X770,X768] : (s(bool,'T_0') != s(bool,X768) | s(bool,X769) != s(bool,happ(s(fun(X762,bool),happ(s(fun(X762,fun(X762,bool)),X763),s(X762,happ(s(fun(X764,X762),X765),s(X764,X766))))),s(X762,happ(s(fun(X764,X762),X765),s(X764,X767))))) | s(bool,X768) != s(bool,X770) | s(bool,'F_0') = s(bool,X770) | s(bool,X769) != s(bool,X771) | s(bool,'T_0') = s(bool,X771) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X762,fun(X762,bool)),X763),s(fun(X764,X762),X765),s(X764,X766),s(X764,X767)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f1648])).
fof(f15445,plain,(
  ( ! [X705,X707,X708,X710,X704,X706,X709,X711,X703] : (s(bool,X710) != s(bool,happ(s(fun(X703,bool),happ(s(fun(X703,fun(X703,bool)),X704),s(X703,happ(s(fun(X705,X703),X706),s(X705,X707))))),s(X703,happ(s(fun(X705,X703),X706),s(X705,X708))))) | s(bool,'F_0') = s(bool,X709) | s(bool,X710) != s(bool,X711) | s(bool,'T_0') = s(bool,X711) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X703,fun(X703,bool)),X704),s(fun(X705,X703),X706),s(X705,X707),s(X705,X708)))) ) | ($spl34 | $spl40)),
  inference(subsumption_resolution,[],[f15221,f122])).
fof(f15221,plain,(
  ( ! [X705,X707,X708,X710,X704,X706,X709,X711,X703] : (s(bool,'T_0') != s(bool,X709) | s(bool,X710) != s(bool,happ(s(fun(X703,bool),happ(s(fun(X703,fun(X703,bool)),X704),s(X703,happ(s(fun(X705,X703),X706),s(X705,X707))))),s(X703,happ(s(fun(X705,X703),X706),s(X705,X708))))) | s(bool,'F_0') = s(bool,X709) | s(bool,X710) != s(bool,X711) | s(bool,'T_0') = s(bool,X711) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X703,fun(X703,bool)),X704),s(fun(X705,X703),X706),s(X705,X707),s(X705,X708)))) ) | ($spl34 | $spl40)),
  inference(superposition,[],[f1342,f1648])).
fof(f9513,plain,(
  ( ! [X584,X580,X582,X581,X583,X579] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X580,fun(X580,bool)),X581),s(fun(X579,X580),X582),s(X579,X583),s(X579,X584)))) ) | $spl184),
  inference(cnf_transformation,[],[f9513_D])).
fof(f9513_D,plain,(
  ( ! [X584,X580,X582,X581,X583,X579] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X580,fun(X580,bool)),X581),s(fun(X579,X580),X582),s(X579,X583),s(X579,X584)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f7350,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),X2))) | s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X4))),s(X1,X3)))) ) | $spl136),
  inference(cnf_transformation,[],[f7350_D])).
fof(f7350_D,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),X2))) | s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X4))),s(X1,X3)))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f15444,plain,(
  ( ! [X697,X699,X693,X695,X700,X702,X696,X698,X694,X701] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X699,fun(X699,bool)),X700))) | s(bool,happ(s(fun(X699,bool),happ(s(fun(X699,fun(X699,bool)),X700),s(X699,X701))),s(X699,X702))) = s(bool,happ(s(fun(X699,bool),happ(s(fun(X699,fun(X699,bool)),X700),s(X699,X702))),s(X699,X701))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X693,fun(X693,bool)),X694),s(fun(X695,X693),X696),s(X695,X697),s(X695,X698)))) ) | ($spl24 | $spl50)),
  inference(subsumption_resolution,[],[f15443,f432])).
fof(f15443,plain,(
  ( ! [X697,X699,X693,X695,X700,X702,X696,X698,X694,X701] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(X693,fun(X693,bool)),X694),s(fun(X695,X693),X696),s(X695,X697),s(X695,X698)))) | s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X699,fun(X699,bool)),X700))) | s(bool,happ(s(fun(X699,bool),happ(s(fun(X699,fun(X699,bool)),X700),s(X699,X701))),s(X699,X702))) = s(bool,happ(s(fun(X699,bool),happ(s(fun(X699,fun(X699,bool)),X700),s(X699,X702))),s(X699,X701))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X693,fun(X693,bool)),X694),s(fun(X695,X693),X696),s(X695,X697),s(X695,X698)))) ) | $spl50),
  inference(forward_demodulation,[],[f15220,f123])).
fof(f15220,plain,(
  ( ! [X697,X699,X693,X695,X700,X702,X696,X698,X694,X701] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X699,fun(X699,bool)),X700))) | s(bool,happ(s(fun(X699,bool),happ(s(fun(X699,fun(X699,bool)),X700),s(X699,X701))),s(X699,X702))) = s(bool,happ(s(fun(X699,bool),happ(s(fun(X699,fun(X699,bool)),X700),s(X699,X702))),s(X699,X701))) | ~p(s(bool,happ(s(fun(X693,bool),happ(s(fun(X693,fun(X693,bool)),X694),s(X693,happ(s(fun(X695,X693),X696),s(X695,X697))))),s(X693,happ(s(fun(X695,X693),X696),s(X695,X698)))))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X693,fun(X693,bool)),X694),s(fun(X695,X693),X696),s(X695,X697),s(X695,X698)))) ) | $spl50),
  inference(superposition,[],[f1285,f1648])).
fof(f15442,plain,(
  ( ! [X648,X644,X646,X642,X645,X641,X643] : (s(bool,X648) != s(bool,happ(s(fun(X641,bool),happ(s(fun(X641,fun(X641,bool)),X642),s(X641,happ(s(fun(X643,X641),X644),s(X643,X645))))),s(X641,happ(s(fun(X643,X641),X644),s(X643,X646))))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X641,fun(X641,bool)),X642),s(fun(X643,X641),X644),s(X643,X645),s(X643,X646))) | s(bool,'T_0') = s(bool,X648)) ) | $spl274),
  inference(cnf_transformation,[],[f15442_D])).
fof(f15442_D,plain,(
  ( ! [X648,X644,X646,X642,X645,X641,X643] : (s(bool,X648) != s(bool,happ(s(fun(X641,bool),happ(s(fun(X641,fun(X641,bool)),X642),s(X641,happ(s(fun(X643,X641),X644),s(X643,X645))))),s(X641,happ(s(fun(X643,X641),X644),s(X643,X646))))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X641,fun(X641,bool)),X642),s(fun(X643,X641),X644),s(X643,X645),s(X643,X646))) | s(bool,'T_0') = s(bool,X648)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f15441,plain,(
  ( ! [X648,X644,X646,X642,X645,X647,X641,X643] : (s(bool,X648) != s(bool,happ(s(fun(X641,bool),happ(s(fun(X641,fun(X641,bool)),X642),s(X641,happ(s(fun(X643,X641),X644),s(X643,X645))))),s(X641,happ(s(fun(X643,X641),X644),s(X643,X646))))) | s(bool,'F_0') = s(bool,X647) | s(bool,'T_0') = s(bool,X648) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X641,fun(X641,bool)),X642),s(fun(X643,X641),X644),s(X643,X645),s(X643,X646)))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f15213,f122])).
fof(f15213,plain,(
  ( ! [X648,X644,X646,X642,X645,X647,X641,X643] : (s(bool,'T_0') != s(bool,X647) | s(bool,X648) != s(bool,happ(s(fun(X641,bool),happ(s(fun(X641,fun(X641,bool)),X642),s(X641,happ(s(fun(X643,X641),X644),s(X643,X645))))),s(X641,happ(s(fun(X643,X641),X644),s(X643,X646))))) | s(bool,'F_0') = s(bool,X647) | s(bool,'T_0') = s(bool,X648) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X641,fun(X641,bool)),X642),s(fun(X643,X641),X644),s(X643,X645),s(X643,X646)))) ) | ($spl34 | $spl38)),
  inference(superposition,[],[f789,f1648])).
fof(f14986,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X3),s(X1,X4))) | s(bool,X0) != s(bool,X6) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X4),s(X1,X3)))) ) | $spl272),
  inference(cnf_transformation,[],[f14986_D])).
fof(f14986_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X3),s(X1,X4))) | s(bool,X0) != s(bool,X6) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X4),s(X1,X3)))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f14907,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (s(bool,X74) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X75,sK1),X76),s(X75,X77),s(X75,X78))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X75,sK1),X76),s(X75,X78),s(X75,X77))) | p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X74) != s(bool,X80)) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f7727,f125])).
fof(f7727,plain,(
  ( ! [X101,X99,X102,X100,X98] : (p(s(bool,X98)) | s(bool,X98) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X99,sK1),X100),s(X99,X101),s(X99,X102))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X99,sK1),X100),s(X99,X102),s(X99,X101)))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f2432,f432])).
fof(f2432,plain,(
  ( ! [X313,X310,X312,X309,X311] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X309,sK1),X310),s(X309,X311),s(X309,X312)))) | p(s(bool,X313)) | s(bool,X313) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X309,sK1),X310),s(X309,X312),s(X309,X311)))) ) | $spl6),
  inference(superposition,[],[f147,f2341])).
fof(f14904,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X57,sK1),X58),s(X57,X59),s(X57,X60))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X57,sK1),X58),s(X57,X60),s(X57,X59))) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X56) != s(bool,X61) | s(bool,X62) != s(bool,X63)) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f7727,f89])).
fof(f14901,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s(bool,X32) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X33,sK1),X34),s(X33,X35),s(X33,X36))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X33,sK1),X34),s(X33,X36),s(X33,X35))) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X32) != s(bool,X39)) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f7727,f88])).
fof(f14900,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X25,sK1),X26),s(X25,X27),s(X25,X28))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X25,sK1),X26),s(X25,X28),s(X25,X27))) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X24) != s(bool,X30)) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f7727,f87])).
fof(f14899,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X17,sK1),X18),s(X17,X19),s(X17,X20))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X17,sK1),X18),s(X17,X20),s(X17,X19))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X16) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f7727,f85])).
fof(f14897,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X3),s(X1,X4))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X4),s(X1,X3))) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X0) != s(bool,X6)) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f7727,f83])).
fof(f14804,plain,(
  ( ! [X825,X831,X823,X826,X824,X828,X832,X827] : (s(bool,X831) != s(bool,happ(s(fun(X823,bool),happ(s(fun(X823,fun(X823,bool)),X824),s(X823,happ(s(fun(X825,X823),X826),s(X825,X827))))),s(X823,happ(s(fun(X825,X823),X826),s(X825,X828))))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X823,fun(X823,bool)),X824),s(fun(X825,X823),X826),s(X825,X827),s(X825,X828))) | s(bool,'F_0') = s(bool,X832) | s(bool,X831) != s(bool,X832)) ) | $spl270),
  inference(cnf_transformation,[],[f14804_D])).
fof(f14804_D,plain,(
  ( ! [X825,X831,X823,X826,X824,X828,X832,X827] : (s(bool,X831) != s(bool,happ(s(fun(X823,bool),happ(s(fun(X823,fun(X823,bool)),X824),s(X823,happ(s(fun(X825,X823),X826),s(X825,X827))))),s(X823,happ(s(fun(X825,X823),X826),s(X825,X828))))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X823,fun(X823,bool)),X824),s(fun(X825,X823),X826),s(X825,X827),s(X825,X828))) | s(bool,'F_0') = s(bool,X832) | s(bool,X831) != s(bool,X832)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f14803,plain,(
  ( ! [X825,X831,X829,X823,X833,X826,X824,X830,X828,X832,X827] : (s(bool,X829) != s(bool,X830) | s(bool,X831) != s(bool,happ(s(fun(X823,bool),happ(s(fun(X823,fun(X823,bool)),X824),s(X823,happ(s(fun(X825,X823),X826),s(X825,X827))))),s(X823,happ(s(fun(X825,X823),X826),s(X825,X828))))) | s(bool,X831) != s(bool,X832) | s(bool,'F_0') = s(bool,X832) | s(bool,X830) != s(bool,X833) | s(bool,'T_0') = s(bool,X833) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X823,fun(X823,bool)),X824),s(fun(X825,X823),X826),s(X825,X827),s(X825,X828)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f14592,f789])).
fof(f14592,plain,(
  ( ! [X825,X831,X829,X823,X833,X826,X824,X830,X828,X832,X827] : (s(bool,'F_0') != s(bool,X829) | s(bool,X829) != s(bool,X830) | s(bool,X831) != s(bool,happ(s(fun(X823,bool),happ(s(fun(X823,fun(X823,bool)),X824),s(X823,happ(s(fun(X825,X823),X826),s(X825,X827))))),s(X823,happ(s(fun(X825,X823),X826),s(X825,X828))))) | s(bool,X831) != s(bool,X832) | s(bool,'F_0') = s(bool,X832) | s(bool,X830) != s(bool,X833) | s(bool,'T_0') = s(bool,X833) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X823,fun(X823,bool)),X824),s(fun(X825,X823),X826),s(X825,X827),s(X825,X828)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f1647])).
fof(f14797,plain,(
  ( ! [X767,X771,X769,X766,X770,X768,X774] : (s(bool,X774) != s(bool,happ(s(fun(X766,bool),happ(s(fun(X766,fun(X766,bool)),X767),s(X766,happ(s(fun(X768,X766),X769),s(X768,X770))))),s(X766,happ(s(fun(X768,X766),X769),s(X768,X771))))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X766,fun(X766,bool)),X767),s(fun(X768,X766),X769),s(X768,X770),s(X768,X771))) | s(bool,'F_0') = s(bool,X774)) ) | $spl268),
  inference(cnf_transformation,[],[f14797_D])).
fof(f14797_D,plain,(
  ( ! [X767,X771,X769,X766,X770,X768,X774] : (s(bool,X774) != s(bool,happ(s(fun(X766,bool),happ(s(fun(X766,fun(X766,bool)),X767),s(X766,happ(s(fun(X768,X766),X769),s(X768,X770))))),s(X766,happ(s(fun(X768,X766),X769),s(X768,X771))))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X766,fun(X766,bool)),X767),s(fun(X768,X766),X769),s(X768,X770),s(X768,X771))) | s(bool,'F_0') = s(bool,X774)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f14796,plain,(
  ( ! [X767,X771,X769,X775,X773,X766,X770,X768,X774,X772] : (s(bool,X772) != s(bool,X773) | s(bool,X774) != s(bool,happ(s(fun(X766,bool),happ(s(fun(X766,fun(X766,bool)),X767),s(X766,happ(s(fun(X768,X766),X769),s(X768,X770))))),s(X766,happ(s(fun(X768,X766),X769),s(X768,X771))))) | s(bool,'F_0') = s(bool,X774) | s(bool,X773) != s(bool,X775) | s(bool,'T_0') = s(bool,X775) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X766,fun(X766,bool)),X767),s(fun(X768,X766),X769),s(X768,X770),s(X768,X771)))) ) | ($spl34 | $spl38 | $spl56)),
  inference(subsumption_resolution,[],[f14586,f789])).
fof(f14586,plain,(
  ( ! [X767,X771,X769,X775,X773,X766,X770,X768,X774,X772] : (s(bool,'F_0') != s(bool,X772) | s(bool,X772) != s(bool,X773) | s(bool,X774) != s(bool,happ(s(fun(X766,bool),happ(s(fun(X766,fun(X766,bool)),X767),s(X766,happ(s(fun(X768,X766),X769),s(X768,X770))))),s(X766,happ(s(fun(X768,X766),X769),s(X768,X771))))) | s(bool,'F_0') = s(bool,X774) | s(bool,X773) != s(bool,X775) | s(bool,'T_0') = s(bool,X775) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X766,fun(X766,bool)),X767),s(fun(X768,X766),X769),s(X768,X770),s(X768,X771)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f1647])).
fof(f14749,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,happ(s(fun(X123,sK1),X124),s(X123,X126))))),s(sK1,happ(s(fun(X123,sK1),X124),s(X123,X125))))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X123,sK1),X124),s(X123,X125),s(X123,X126)))) ) | $spl266),
  inference(cnf_transformation,[],[f14749_D])).
fof(f14749_D,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,happ(s(fun(X123,sK1),X124),s(X123,X126))))),s(sK1,happ(s(fun(X123,sK1),X124),s(X123,X125))))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X123,sK1),X124),s(X123,X125),s(X123,X126)))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f237,plain,(
  ( ! [X53] : (s(bool,'F_0') != s(bool,X53)) ) | $spl18),
  inference(cnf_transformation,[],[f237_D])).
fof(f237_D,plain,(
  ( ! [X53] : (s(bool,'F_0') != s(bool,X53)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f14748,plain,(
  ( ! [X127,X125,X123,X126,X124] : (s(bool,'F_0') != s(bool,X127) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,happ(s(fun(X123,sK1),X124),s(X123,X126))))),s(sK1,happ(s(fun(X123,sK1),X124),s(X123,X125))))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X123,sK1),X124),s(X123,X125),s(X123,X126)))) ) | ($spl6 | $spl32)),
  inference(subsumption_resolution,[],[f14507,f496])).
fof(f14507,plain,(
  ( ! [X127,X125,X123,X126,X124] : (s(bool,'F_0') != s(bool,X127) | p(s(bool,X127)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,happ(s(fun(X123,sK1),X124),s(X123,X126))))),s(sK1,happ(s(fun(X123,sK1),X124),s(X123,X125))))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X123,sK1),X124),s(X123,X125),s(X123,X126)))) ) | $spl6),
  inference(superposition,[],[f7652,f1647])).
fof(f7652,plain,(
  ( ! [X30,X31,X29] : (s(bool,X31) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X30))),s(sK1,X29))) | p(s(bool,X31)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,X30)))) ) | $spl6),
  inference(subsumption_resolution,[],[f7646,f126])).
fof(f7646,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X30))),s(sK1,X29))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,X30)))) ) | $spl6),
  inference(superposition,[],[f2011,f122])).
fof(f2011,plain,(
  ( ! [X220,X219,X221] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X219))),s(sK1,X220)))) | p(s(bool,X221)) | s(bool,X221) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X220))),s(sK1,X219)))) ) | $spl6),
  inference(superposition,[],[f147,f127])).
fof(f14293,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s(bool,X28) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X23),s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X27))) | p(s(bool,X28)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X23),s(fun(X24,bool),X25),s(X24,X26),s(X24,X27))))) ) | $spl260),
  inference(cnf_transformation,[],[f14293_D])).
fof(f14293_D,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s(bool,X28) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X23),s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X27))) | p(s(bool,X28)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X23),s(fun(X24,bool),X25),s(X24,X26),s(X24,X27))))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f14137,plain,(
  ( ! [X397,X399,X401,X403,X398,X400,X402] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X397),s(fun(X398,bool),X399),s(X398,X400),s(X398,X401)))) | p(s(bool,X402)) | p(s(bool,X403)) | s(bool,X403) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X397),s(bool,happ(s(fun(X398,bool),X399),s(X398,X400))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X398,bool),X399),s(X398,X401)))) )),
  inference(superposition,[],[f125,f1641])).
fof(f1641,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X36),s(fun(X33,bool),X34),s(X33,X37),s(X33,X35))) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X36),s(bool,happ(s(fun(X33,bool),X34),s(X33,X37))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),X34),s(X33,X35)))) )),
  inference(superposition,[],[f123,f122])).
fof(f14292,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,X20) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X15),s(bool,happ(s(fun(X16,bool),X17),s(X16,X18))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,X19))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X15),s(fun(X16,bool),X17),s(X16,X18),s(X16,X19)))) | ~p(s(bool,X20))) ) | $spl258),
  inference(cnf_transformation,[],[f14292_D])).
fof(f14292_D,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,X20) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X15),s(bool,happ(s(fun(X16,bool),X17),s(X16,X18))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,X19))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X15),s(fun(X16,bool),X17),s(X16,X18),s(X16,X19)))) | ~p(s(bool,X20))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f14135,plain,(
  ( ! [X383,X389,X385,X387,X388,X384,X386] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X383),s(fun(X384,bool),X385),s(X384,X386),s(X384,X387)))) | p(s(bool,X388)) | ~p(s(bool,X389)) | s(bool,X389) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X383),s(bool,happ(s(fun(X384,bool),X385),s(X384,X386))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X384,bool),X385),s(X384,X387)))) )),
  inference(superposition,[],[f125,f1641])).
fof(f14295,plain,(
  ( ! [X47,X50,X48,X53,X51,X49] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X47),s(bool,happ(s(fun(X48,bool),X49),s(X48,X50))))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,X51))) | s(bool,X53) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X47),s(fun(X48,bool),X49),s(X48,X50),s(X48,X51))) | ~p(s(bool,X53))) ) | $spl262),
  inference(cnf_transformation,[],[f14295_D])).
fof(f14295_D,plain,(
  ( ! [X47,X50,X48,X53,X51,X49] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X47),s(bool,happ(s(fun(X48,bool),X49),s(X48,X50))))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,X51))) | s(bool,X53) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X47),s(fun(X48,bool),X49),s(X48,X50),s(X48,X51))) | ~p(s(bool,X53))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f14134,plain,(
  ( ! [X377,X379,X380,X382,X376,X378,X381] : (s(bool,X381) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X376),s(fun(X377,bool),X378),s(X377,X379),s(X377,X380))) | p(s(bool,X382)) | ~p(s(bool,X381)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X376),s(bool,happ(s(fun(X377,bool),X378),s(X377,X379))))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(X377,bool),X378),s(X377,X380)))) )),
  inference(superposition,[],[f125,f1641])).
fof(f14296,plain,(
  ( ! [X66,X64,X69,X67,X65,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X63),s(bool,happ(s(fun(X64,bool),X65),s(X64,X66))))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(X64,bool),X65),s(X64,X67))) | s(bool,X69) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X63),s(fun(X64,bool),X65),s(X64,X66),s(X64,X67))) | p(s(bool,X69))) ) | $spl264),
  inference(cnf_transformation,[],[f14296_D])).
fof(f14296_D,plain,(
  ( ! [X66,X64,X69,X67,X65,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X63),s(bool,happ(s(fun(X64,bool),X65),s(X64,X66))))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(X64,bool),X65),s(X64,X67))) | s(bool,X69) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X63),s(fun(X64,bool),X65),s(X64,X66),s(X64,X67))) | p(s(bool,X69))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f14133,plain,(
  ( ! [X372,X374,X370,X373,X375,X369,X371] : (s(bool,X374) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X369),s(fun(X370,bool),X371),s(X370,X372),s(X370,X373))) | p(s(bool,X375)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X369),s(bool,happ(s(fun(X370,bool),X371),s(X370,X372))))),s(bool,'F_0')))) | p(s(bool,X374)) | s(bool,'T_0') = s(bool,happ(s(fun(X370,bool),X371),s(X370,X373)))) )),
  inference(superposition,[],[f125,f1641])).
fof(f14125,plain,(
  ( ! [X327,X333,X329,X331,X326,X332,X328,X330] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X326),s(fun(X327,bool),X328),s(X327,X329),s(X327,X330)))) | p(s(bool,X331)) | ~p(s(bool,X332)) | s(bool,X331) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X326),s(bool,happ(s(fun(X327,bool),X328),s(X327,X329))))),s(bool,'F_0'))) | s(bool,X332) != s(bool,X333) | s(bool,'T_0') = s(bool,happ(s(fun(X327,bool),X328),s(X327,X330)))) )),
  inference(superposition,[],[f89,f1641])).
fof(f14123,plain,(
  ( ! [X313,X315,X310,X316,X312,X314,X311,X317] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X310),s(fun(X311,bool),X312),s(X311,X313),s(X311,X314)))) | ~p(s(bool,X315)) | ~p(s(bool,X316)) | s(bool,X316) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X310),s(bool,happ(s(fun(X311,bool),X312),s(X311,X313))))),s(bool,'F_0'))) | s(bool,X315) != s(bool,X317) | s(bool,'T_0') = s(bool,happ(s(fun(X311,bool),X312),s(X311,X314)))) )),
  inference(superposition,[],[f89,f1641])).
fof(f14122,plain,(
  ( ! [X302,X308,X304,X306,X303,X309,X305,X307] : (s(bool,X307) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X302),s(fun(X303,bool),X304),s(X303,X305),s(X303,X306))) | p(s(bool,X307)) | ~p(s(bool,X308)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X302),s(bool,happ(s(fun(X303,bool),X304),s(X303,X305))))),s(bool,'F_0')))) | s(bool,X308) != s(bool,X309) | s(bool,'T_0') = s(bool,happ(s(fun(X303,bool),X304),s(X303,X306)))) )),
  inference(superposition,[],[f89,f1641])).
fof(f14121,plain,(
  ( ! [X294,X300,X296,X298,X295,X301,X297,X299] : (s(bool,X299) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X294),s(fun(X295,bool),X296),s(X295,X297),s(X295,X298))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X294),s(bool,happ(s(fun(X295,bool),X296),s(X295,X297))))),s(bool,'F_0')))) | ~p(s(bool,X300)) | ~p(s(bool,X299)) | s(bool,X300) != s(bool,X301) | s(bool,'T_0') = s(bool,happ(s(fun(X295,bool),X296),s(X295,X298)))) )),
  inference(superposition,[],[f89,f1641])).
fof(f14117,plain,(
  ( ! [X263,X269,X265,X267,X262,X268,X264,X266] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X262),s(fun(X263,bool),X264),s(X263,X265),s(X263,X266)))) | p(s(bool,X267)) | ~p(s(bool,X268)) | p(s(bool,X269)) | s(bool,X269) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X262),s(bool,happ(s(fun(X263,bool),X264),s(X263,X265))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X263,bool),X264),s(X263,X266)))) )),
  inference(superposition,[],[f88,f1641])).
fof(f14115,plain,(
  ( ! [X253,X251,X249,X246,X252,X250,X248,X247] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X246),s(fun(X247,bool),X248),s(X247,X249),s(X247,X250)))) | p(s(bool,X251)) | ~p(s(bool,X252)) | ~p(s(bool,X253)) | s(bool,X252) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X246),s(bool,happ(s(fun(X247,bool),X248),s(X247,X249))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X247,bool),X248),s(X247,X250)))) )),
  inference(superposition,[],[f88,f1641])).
fof(f14114,plain,(
  ( ! [X243,X241,X239,X244,X242,X240,X238,X245] : (s(bool,X243) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X238),s(fun(X239,bool),X240),s(X239,X241),s(X239,X242))) | p(s(bool,X244)) | ~p(s(bool,X243)) | ~p(s(bool,X245)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X238),s(bool,happ(s(fun(X239,bool),X240),s(X239,X241))))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(X239,bool),X240),s(X239,X242)))) )),
  inference(superposition,[],[f88,f1641])).
fof(f14113,plain,(
  ( ! [X231,X237,X235,X233,X230,X236,X234,X232] : (s(bool,X235) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X230),s(fun(X231,bool),X232),s(X231,X233),s(X231,X234))) | p(s(bool,X236)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X230),s(bool,happ(s(fun(X231,bool),X232),s(X231,X233))))),s(bool,'F_0')))) | ~p(s(bool,X237)) | p(s(bool,X235)) | s(bool,'T_0') = s(bool,happ(s(fun(X231,bool),X232),s(X231,X234)))) )),
  inference(superposition,[],[f88,f1641])).
fof(f14112,plain,(
  ( ! [X222,X229,X227,X225,X223,X228,X226,X224] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X222),s(fun(X223,bool),X224),s(X223,X225),s(X223,X226)))) | p(s(bool,X227)) | p(s(bool,X228)) | s(bool,X227) != s(bool,X229) | s(bool,X228) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X222),s(bool,happ(s(fun(X223,bool),X224),s(X223,X225))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X223,bool),X224),s(X223,X226)))) )),
  inference(superposition,[],[f87,f1641])).
fof(f14110,plain,(
  ( ! [X212,X210,X208,X206,X213,X211,X209,X207] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X206),s(fun(X207,bool),X208),s(X207,X209),s(X207,X210)))) | p(s(bool,X211)) | ~p(s(bool,X212)) | s(bool,X211) != s(bool,X213) | s(bool,X212) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X206),s(bool,happ(s(fun(X207,bool),X208),s(X207,X209))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X207,bool),X208),s(X207,X210)))) )),
  inference(superposition,[],[f87,f1641])).
fof(f14107,plain,(
  ( ! [X189,X187,X185,X182,X188,X186,X184,X183] : (s(bool,X187) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X182),s(fun(X183,bool),X184),s(X183,X185),s(X183,X186))) | p(s(bool,X188)) | ~p(s(bool,X187)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X182),s(bool,happ(s(fun(X183,bool),X184),s(X183,X185))))),s(bool,'F_0')))) | s(bool,X188) != s(bool,X189) | s(bool,'T_0') = s(bool,happ(s(fun(X183,bool),X184),s(X183,X186)))) )),
  inference(superposition,[],[f87,f1641])).
fof(f14106,plain,(
  ( ! [X177,X175,X180,X178,X176,X174,X181,X179] : (s(bool,X179) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X174),s(fun(X175,bool),X176),s(X175,X177),s(X175,X178))) | p(s(bool,X180)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X174),s(bool,happ(s(fun(X175,bool),X176),s(X175,X177))))),s(bool,'F_0')))) | p(s(bool,X179)) | s(bool,X180) != s(bool,X181) | s(bool,'T_0') = s(bool,happ(s(fun(X175,bool),X176),s(X175,X178)))) )),
  inference(superposition,[],[f87,f1641])).
fof(f14105,plain,(
  ( ! [X167,X173,X171,X169,X166,X172,X170,X168] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X166),s(fun(X167,bool),X168),s(X167,X169),s(X167,X170)))) | p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X171) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X166),s(bool,happ(s(fun(X167,bool),X168),s(X167,X169))))),s(bool,'F_0'))) | s(bool,X172) != s(bool,X173) | s(bool,'T_0') = s(bool,happ(s(fun(X167,bool),X168),s(X167,X170)))) )),
  inference(superposition,[],[f85,f1641])).
fof(f14103,plain,(
  ( ! [X156,X154,X152,X151,X157,X155,X153,X150] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X150),s(fun(X151,bool),X152),s(X151,X153),s(X151,X154)))) | ~p(s(bool,X155)) | ~p(s(bool,X156)) | s(bool,X156) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X150),s(bool,happ(s(fun(X151,bool),X152),s(X151,X153))))),s(bool,'F_0'))) | s(bool,X155) != s(bool,X157) | s(bool,'T_0') = s(bool,happ(s(fun(X151,bool),X152),s(X151,X154)))) )),
  inference(superposition,[],[f85,f1641])).
fof(f14102,plain,(
  ( ! [X146,X144,X142,X149,X147,X145,X143,X148] : (s(bool,X147) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X142),s(fun(X143,bool),X144),s(X143,X145),s(X143,X146))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X142),s(bool,happ(s(fun(X143,bool),X144),s(X143,X145))))),s(bool,'F_0')))) | ~p(s(bool,X148)) | ~p(s(bool,X147)) | s(bool,X148) != s(bool,X149) | s(bool,'T_0') = s(bool,happ(s(fun(X143,bool),X144),s(X143,X146)))) )),
  inference(superposition,[],[f85,f1641])).
fof(f14101,plain,(
  ( ! [X134,X140,X138,X136,X135,X141,X139,X137] : (s(bool,X139) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X134),s(fun(X135,bool),X136),s(X135,X137),s(X135,X138))) | p(s(bool,X139)) | ~p(s(bool,X140)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X134),s(bool,happ(s(fun(X135,bool),X136),s(X135,X137))))),s(bool,'F_0')))) | s(bool,X140) != s(bool,X141) | s(bool,'T_0') = s(bool,happ(s(fun(X135,bool),X136),s(X135,X138)))) )),
  inference(superposition,[],[f85,f1641])).
fof(f14098,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X110),s(fun(X111,bool),X112),s(X111,X113),s(X111,X114)))) | p(s(bool,X115)) | p(s(bool,X116)) | s(bool,X115) != s(bool,X117) | s(bool,X116) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X110),s(bool,happ(s(fun(X111,bool),X112),s(X111,X113))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X111,bool),X112),s(X111,X114)))) )),
  inference(superposition,[],[f83,f1641])).
fof(f14096,plain,(
  ( ! [X94,X101,X99,X97,X95,X100,X98,X96] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X94),s(fun(X95,bool),X96),s(X95,X97),s(X95,X98)))) | p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X99) != s(bool,X101) | s(bool,X100) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X94),s(bool,happ(s(fun(X95,bool),X96),s(X95,X97))))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X95,bool),X96),s(X95,X98)))) )),
  inference(superposition,[],[f83,f1641])).
fof(f14093,plain,(
  ( ! [X70,X76,X74,X72,X71,X77,X75,X73] : (s(bool,X75) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X70),s(fun(X71,bool),X72),s(X71,X73),s(X71,X74))) | p(s(bool,X76)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X70),s(bool,happ(s(fun(X71,bool),X72),s(X71,X73))))),s(bool,'F_0')))) | p(s(bool,X75)) | s(bool,X76) != s(bool,X77) | s(bool,'T_0') = s(bool,happ(s(fun(X71,bool),X72),s(X71,X74)))) )),
  inference(superposition,[],[f83,f1641])).
fof(f14092,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (s(bool,X67) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X62),s(fun(X63,bool),X64),s(X63,X65),s(X63,X66))) | p(s(bool,X68)) | ~p(s(bool,X67)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X62),s(bool,happ(s(fun(X63,bool),X64),s(X63,X65))))),s(bool,'F_0')))) | s(bool,X68) != s(bool,X69) | s(bool,'T_0') = s(bool,happ(s(fun(X63,bool),X64),s(X63,X66)))) )),
  inference(superposition,[],[f83,f1641])).
fof(f13937,plain,(
  ( ! [X350,X356,X352,X354,X351,X353,X355] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X350),s(bool,happ(s(fun(X351,bool),X352),s(X351,X353))))),s(bool,'F_0')))) | p(s(bool,X355)) | p(s(bool,X356)) | s(bool,X356) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X350),s(fun(X351,bool),X352),s(X351,X353),s(X351,X354))) | s(bool,'T_0') = s(bool,happ(s(fun(X351,bool),X352),s(X351,X354)))) )),
  inference(superposition,[],[f125,f1641])).
fof(f13935,plain,(
  ( ! [X341,X337,X339,X340,X342,X336,X338] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X336),s(bool,happ(s(fun(X337,bool),X338),s(X337,X339))))),s(bool,'F_0')))) | p(s(bool,X341)) | ~p(s(bool,X342)) | s(bool,X342) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X336),s(fun(X337,bool),X338),s(X337,X339),s(X337,X340))) | s(bool,'T_0') = s(bool,happ(s(fun(X337,bool),X338),s(X337,X340)))) )),
  inference(superposition,[],[f125,f1641])).
fof(f13934,plain,(
  ( ! [X333,X335,X329,X331,X332,X334,X330] : (s(bool,X334) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X329),s(bool,happ(s(fun(X330,bool),X331),s(X330,X332))))),s(bool,'F_0'))) | p(s(bool,X335)) | ~p(s(bool,X334)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X329),s(fun(X330,bool),X331),s(X330,X332),s(X330,X333)))) | s(bool,'T_0') = s(bool,happ(s(fun(X330,bool),X331),s(X330,X333)))) )),
  inference(superposition,[],[f125,f1641])).
fof(f13933,plain,(
  ( ! [X325,X327,X323,X324,X326,X322,X328] : (s(bool,X327) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X322),s(bool,happ(s(fun(X323,bool),X324),s(X323,X325))))),s(bool,'F_0'))) | p(s(bool,X328)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X322),s(fun(X323,bool),X324),s(X323,X325),s(X323,X326)))) | p(s(bool,X327)) | s(bool,'T_0') = s(bool,happ(s(fun(X323,bool),X324),s(X323,X326)))) )),
  inference(superposition,[],[f125,f1641])).
fof(f13925,plain,(
  ( ! [X280,X282,X279,X285,X281,X283,X284,X286] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X279),s(bool,happ(s(fun(X280,bool),X281),s(X280,X282))))),s(bool,'F_0')))) | p(s(bool,X284)) | ~p(s(bool,X285)) | s(bool,X284) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X279),s(fun(X280,bool),X281),s(X280,X282),s(X280,X283))) | s(bool,X285) != s(bool,X286) | s(bool,'T_0') = s(bool,happ(s(fun(X280,bool),X281),s(X280,X283)))) )),
  inference(superposition,[],[f89,f1641])).
fof(f13923,plain,(
  ( ! [X263,X269,X265,X267,X268,X270,X264,X266] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X263),s(bool,happ(s(fun(X264,bool),X265),s(X264,X266))))),s(bool,'F_0')))) | ~p(s(bool,X268)) | ~p(s(bool,X269)) | s(bool,X269) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X263),s(fun(X264,bool),X265),s(X264,X266),s(X264,X267))) | s(bool,X268) != s(bool,X270) | s(bool,'T_0') = s(bool,happ(s(fun(X264,bool),X265),s(X264,X267)))) )),
  inference(superposition,[],[f89,f1641])).
fof(f13922,plain,(
  ( ! [X255,X261,X257,X259,X260,X262,X256,X258] : (s(bool,X260) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X255),s(bool,happ(s(fun(X256,bool),X257),s(X256,X258))))),s(bool,'F_0'))) | p(s(bool,X260)) | ~p(s(bool,X261)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X255),s(fun(X256,bool),X257),s(X256,X258),s(X256,X259)))) | s(bool,X261) != s(bool,X262) | s(bool,'T_0') = s(bool,happ(s(fun(X256,bool),X257),s(X256,X259)))) )),
  inference(superposition,[],[f89,f1641])).
fof(f13921,plain,(
  ( ! [X253,X251,X249,X254,X252,X250,X248,X247] : (s(bool,X252) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X247),s(bool,happ(s(fun(X248,bool),X249),s(X248,X250))))),s(bool,'F_0'))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X247),s(fun(X248,bool),X249),s(X248,X250),s(X248,X251)))) | ~p(s(bool,X253)) | ~p(s(bool,X252)) | s(bool,X253) != s(bool,X254) | s(bool,'T_0') = s(bool,happ(s(fun(X248,bool),X249),s(X248,X251)))) )),
  inference(superposition,[],[f89,f1641])).
fof(f13917,plain,(
  ( ! [X218,X216,X222,X220,X215,X219,X217,X221] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X215),s(bool,happ(s(fun(X216,bool),X217),s(X216,X218))))),s(bool,'F_0')))) | p(s(bool,X220)) | ~p(s(bool,X221)) | p(s(bool,X222)) | s(bool,X222) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X215),s(fun(X216,bool),X217),s(X216,X218),s(X216,X219))) | s(bool,'T_0') = s(bool,happ(s(fun(X216,bool),X217),s(X216,X219)))) )),
  inference(superposition,[],[f88,f1641])).
fof(f13915,plain,(
  ( ! [X206,X204,X202,X200,X199,X205,X203,X201] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X199),s(bool,happ(s(fun(X200,bool),X201),s(X200,X202))))),s(bool,'F_0')))) | p(s(bool,X204)) | ~p(s(bool,X205)) | ~p(s(bool,X206)) | s(bool,X205) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X199),s(fun(X200,bool),X201),s(X200,X202),s(X200,X203))) | s(bool,'T_0') = s(bool,happ(s(fun(X200,bool),X201),s(X200,X203)))) )),
  inference(superposition,[],[f88,f1641])).
fof(f13914,plain,(
  ( ! [X191,X198,X196,X194,X192,X197,X195,X193] : (s(bool,X196) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X191),s(bool,happ(s(fun(X192,bool),X193),s(X192,X194))))),s(bool,'F_0'))) | p(s(bool,X197)) | ~p(s(bool,X196)) | ~p(s(bool,X198)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X191),s(fun(X192,bool),X193),s(X192,X194),s(X192,X195)))) | s(bool,'T_0') = s(bool,happ(s(fun(X192,bool),X193),s(X192,X195)))) )),
  inference(superposition,[],[f88,f1641])).
fof(f13913,plain,(
  ( ! [X189,X187,X185,X190,X188,X186,X184,X183] : (s(bool,X188) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X183),s(bool,happ(s(fun(X184,bool),X185),s(X184,X186))))),s(bool,'F_0'))) | p(s(bool,X189)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X183),s(fun(X184,bool),X185),s(X184,X186),s(X184,X187)))) | ~p(s(bool,X190)) | p(s(bool,X188)) | s(bool,'T_0') = s(bool,happ(s(fun(X184,bool),X185),s(X184,X187)))) )),
  inference(superposition,[],[f88,f1641])).
fof(f13912,plain,(
  ( ! [X177,X175,X182,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X175),s(bool,happ(s(fun(X176,bool),X177),s(X176,X178))))),s(bool,'F_0')))) | p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X182) | s(bool,X181) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X175),s(fun(X176,bool),X177),s(X176,X178),s(X176,X179))) | s(bool,'T_0') = s(bool,happ(s(fun(X176,bool),X177),s(X176,X179)))) )),
  inference(superposition,[],[f87,f1641])).
fof(f13910,plain,(
  ( ! [X165,X163,X161,X159,X166,X164,X162,X160] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X159),s(bool,happ(s(fun(X160,bool),X161),s(X160,X162))))),s(bool,'F_0')))) | p(s(bool,X164)) | ~p(s(bool,X165)) | s(bool,X164) != s(bool,X166) | s(bool,X165) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X159),s(fun(X160,bool),X161),s(X160,X162),s(X160,X163))) | s(bool,'T_0') = s(bool,happ(s(fun(X160,bool),X161),s(X160,X163)))) )),
  inference(superposition,[],[f87,f1641])).
fof(f13907,plain,(
  ( ! [X142,X140,X138,X136,X135,X141,X139,X137] : (s(bool,X140) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X135),s(bool,happ(s(fun(X136,bool),X137),s(X136,X138))))),s(bool,'F_0'))) | p(s(bool,X141)) | ~p(s(bool,X140)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X135),s(fun(X136,bool),X137),s(X136,X138),s(X136,X139)))) | s(bool,X141) != s(bool,X142) | s(bool,'T_0') = s(bool,happ(s(fun(X136,bool),X137),s(X136,X139)))) )),
  inference(superposition,[],[f87,f1641])).
fof(f13906,plain,(
  ( ! [X127,X134,X132,X130,X128,X133,X131,X129] : (s(bool,X132) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X127),s(bool,happ(s(fun(X128,bool),X129),s(X128,X130))))),s(bool,'F_0'))) | p(s(bool,X133)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X127),s(fun(X128,bool),X129),s(X128,X130),s(X128,X131)))) | p(s(bool,X132)) | s(bool,X133) != s(bool,X134) | s(bool,'T_0') = s(bool,happ(s(fun(X128,bool),X129),s(X128,X131)))) )),
  inference(superposition,[],[f87,f1641])).
fof(f13905,plain,(
  ( ! [X125,X123,X121,X126,X124,X122,X120,X119] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X119),s(bool,happ(s(fun(X120,bool),X121),s(X120,X122))))),s(bool,'F_0')))) | p(s(bool,X124)) | ~p(s(bool,X125)) | s(bool,X124) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X119),s(fun(X120,bool),X121),s(X120,X122),s(X120,X123))) | s(bool,X125) != s(bool,X126) | s(bool,'T_0') = s(bool,happ(s(fun(X120,bool),X121),s(X120,X123)))) )),
  inference(superposition,[],[f85,f1641])).
fof(f13903,plain,(
  ( ! [X103,X109,X107,X105,X110,X108,X106,X104] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X103),s(bool,happ(s(fun(X104,bool),X105),s(X104,X106))))),s(bool,'F_0')))) | ~p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X109) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X103),s(fun(X104,bool),X105),s(X104,X106),s(X104,X107))) | s(bool,X108) != s(bool,X110) | s(bool,'T_0') = s(bool,happ(s(fun(X104,bool),X105),s(X104,X107)))) )),
  inference(superposition,[],[f85,f1641])).
fof(f13902,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (s(bool,X100) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X95),s(bool,happ(s(fun(X96,bool),X97),s(X96,X98))))),s(bool,'F_0'))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X95),s(fun(X96,bool),X97),s(X96,X98),s(X96,X99)))) | ~p(s(bool,X101)) | ~p(s(bool,X100)) | s(bool,X101) != s(bool,X102) | s(bool,'T_0') = s(bool,happ(s(fun(X96,bool),X97),s(X96,X99)))) )),
  inference(superposition,[],[f85,f1641])).
fof(f13901,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (s(bool,X92) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X87),s(bool,happ(s(fun(X88,bool),X89),s(X88,X90))))),s(bool,'F_0'))) | p(s(bool,X92)) | ~p(s(bool,X93)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X87),s(fun(X88,bool),X89),s(X88,X90),s(X88,X91)))) | s(bool,X93) != s(bool,X94) | s(bool,'T_0') = s(bool,happ(s(fun(X88,bool),X89),s(X88,X91)))) )),
  inference(superposition,[],[f85,f1641])).
fof(f13898,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X63),s(bool,happ(s(fun(X64,bool),X65),s(X64,X66))))),s(bool,'F_0')))) | p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X63),s(fun(X64,bool),X65),s(X64,X66),s(X64,X67))) | s(bool,'T_0') = s(bool,happ(s(fun(X64,bool),X65),s(X64,X67)))) )),
  inference(superposition,[],[f83,f1641])).
fof(f13896,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X47),s(bool,happ(s(fun(X48,bool),X49),s(X48,X50))))),s(bool,'F_0')))) | p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,X54) | s(bool,X53) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X47),s(fun(X48,bool),X49),s(X48,X50),s(X48,X51))) | s(bool,'T_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,X51)))) )),
  inference(superposition,[],[f83,f1641])).
fof(f13893,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (s(bool,X28) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X23),s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))))),s(bool,'F_0'))) | p(s(bool,X29)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X23),s(fun(X24,bool),X25),s(X24,X26),s(X24,X27)))) | p(s(bool,X28)) | s(bool,X29) != s(bool,X30) | s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X27)))) )),
  inference(superposition,[],[f83,f1641])).
fof(f13892,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (s(bool,X20) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X15),s(bool,happ(s(fun(X16,bool),X17),s(X16,X18))))),s(bool,'F_0'))) | p(s(bool,X21)) | ~p(s(bool,X20)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X15),s(fun(X16,bool),X17),s(X16,X18),s(X16,X19)))) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,X19)))) )),
  inference(superposition,[],[f83,f1641])).
fof(f13718,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X23) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X18),s(bool,happ(s(fun(X19,bool),X20),s(X19,X21))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,X22))) | p(s(bool,X23)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X18),s(fun(X19,bool),X20),s(X19,X21),s(X19,X22))))) ) | $spl252),
  inference(cnf_transformation,[],[f13718_D])).
fof(f13718_D,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X23) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X18),s(bool,happ(s(fun(X19,bool),X20),s(X19,X21))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,X22))) | p(s(bool,X23)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X18),s(fun(X19,bool),X20),s(X19,X21),s(X19,X22))))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f13564,plain,(
  ( ! [X391,X393,X395,X390,X396,X392,X394] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X390),s(fun(X391,bool),X392),s(X391,X393),s(X391,X394)))) | p(s(bool,X395)) | p(s(bool,X396)) | s(bool,X396) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X390),s(bool,happ(s(fun(X391,bool),X392),s(X391,X393))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X391,bool),X392),s(X391,X394)))) )),
  inference(superposition,[],[f125,f1640])).
fof(f1640,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X31),s(fun(X28,bool),X29),s(X28,X32),s(X28,X30))) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X31),s(bool,happ(s(fun(X28,bool),X29),s(X28,X32))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X28,bool),X29),s(X28,X30)))) )),
  inference(superposition,[],[f123,f122])).
fof(f13717,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (s(bool,X15) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X10),s(bool,happ(s(fun(X11,bool),X12),s(X11,X13))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X11,bool),X12),s(X11,X14))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X10),s(fun(X11,bool),X12),s(X11,X13),s(X11,X14)))) | ~p(s(bool,X15))) ) | $spl250),
  inference(cnf_transformation,[],[f13717_D])).
fof(f13717_D,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (s(bool,X15) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X10),s(bool,happ(s(fun(X11,bool),X12),s(X11,X13))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X11,bool),X12),s(X11,X14))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X10),s(fun(X11,bool),X12),s(X11,X13),s(X11,X14)))) | ~p(s(bool,X15))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f13562,plain,(
  ( ! [X377,X379,X380,X382,X376,X378,X381] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X376),s(fun(X377,bool),X378),s(X377,X379),s(X377,X380)))) | p(s(bool,X381)) | ~p(s(bool,X382)) | s(bool,X382) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X376),s(bool,happ(s(fun(X377,bool),X378),s(X377,X379))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X377,bool),X378),s(X377,X380)))) )),
  inference(superposition,[],[f125,f1640])).
fof(f13720,plain,(
  ( ! [X45,X43,X48,X46,X44,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X42),s(bool,happ(s(fun(X43,bool),X44),s(X43,X45))))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,happ(s(fun(X43,bool),X44),s(X43,X46))) | s(bool,X48) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X42),s(fun(X43,bool),X44),s(X43,X45),s(X43,X46))) | ~p(s(bool,X48))) ) | $spl254),
  inference(cnf_transformation,[],[f13720_D])).
fof(f13720_D,plain,(
  ( ! [X45,X43,X48,X46,X44,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X42),s(bool,happ(s(fun(X43,bool),X44),s(X43,X45))))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,happ(s(fun(X43,bool),X44),s(X43,X46))) | s(bool,X48) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X42),s(fun(X43,bool),X44),s(X43,X45),s(X43,X46))) | ~p(s(bool,X48))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f13561,plain,(
  ( ! [X372,X374,X370,X373,X375,X369,X371] : (s(bool,X374) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X369),s(fun(X370,bool),X371),s(X370,X372),s(X370,X373))) | p(s(bool,X375)) | ~p(s(bool,X374)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X369),s(bool,happ(s(fun(X370,bool),X371),s(X370,X372))))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,happ(s(fun(X370,bool),X371),s(X370,X373)))) )),
  inference(superposition,[],[f125,f1640])).
fof(f13721,plain,(
  ( ! [X61,X59,X64,X62,X60,X58] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X58),s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,happ(s(fun(X59,bool),X60),s(X59,X62))) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X58),s(fun(X59,bool),X60),s(X59,X61),s(X59,X62))) | p(s(bool,X64))) ) | $spl256),
  inference(cnf_transformation,[],[f13721_D])).
fof(f13721_D,plain,(
  ( ! [X61,X59,X64,X62,X60,X58] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X58),s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,happ(s(fun(X59,bool),X60),s(X59,X62))) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X58),s(fun(X59,bool),X60),s(X59,X61),s(X59,X62))) | p(s(bool,X64))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f13560,plain,(
  ( ! [X364,X366,X362,X368,X365,X367,X363] : (s(bool,X367) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X362),s(fun(X363,bool),X364),s(X363,X365),s(X363,X366))) | p(s(bool,X368)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X362),s(bool,happ(s(fun(X363,bool),X364),s(X363,X365))))),s(bool,'T_0')))) | p(s(bool,X367)) | s(bool,'F_0') = s(bool,happ(s(fun(X363,bool),X364),s(X363,X366)))) )),
  inference(superposition,[],[f125,f1640])).
fof(f13552,plain,(
  ( ! [X325,X321,X323,X324,X326,X320,X322,X319] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X319),s(fun(X320,bool),X321),s(X320,X322),s(X320,X323)))) | p(s(bool,X324)) | ~p(s(bool,X325)) | s(bool,X324) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X319),s(bool,happ(s(fun(X320,bool),X321),s(X320,X322))))),s(bool,'T_0'))) | s(bool,X325) != s(bool,X326) | s(bool,'F_0') = s(bool,happ(s(fun(X320,bool),X321),s(X320,X323)))) )),
  inference(superposition,[],[f89,f1640])).
fof(f13550,plain,(
  ( ! [X308,X310,X304,X306,X303,X309,X305,X307] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X303),s(fun(X304,bool),X305),s(X304,X306),s(X304,X307)))) | ~p(s(bool,X308)) | ~p(s(bool,X309)) | s(bool,X309) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X303),s(bool,happ(s(fun(X304,bool),X305),s(X304,X306))))),s(bool,'T_0'))) | s(bool,X308) != s(bool,X310) | s(bool,'F_0') = s(bool,happ(s(fun(X304,bool),X305),s(X304,X307)))) )),
  inference(superposition,[],[f89,f1640])).
fof(f13549,plain,(
  ( ! [X300,X302,X296,X298,X295,X301,X297,X299] : (s(bool,X300) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X295),s(fun(X296,bool),X297),s(X296,X298),s(X296,X299))) | p(s(bool,X300)) | ~p(s(bool,X301)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X295),s(bool,happ(s(fun(X296,bool),X297),s(X296,X298))))),s(bool,'T_0')))) | s(bool,X301) != s(bool,X302) | s(bool,'F_0') = s(bool,happ(s(fun(X296,bool),X297),s(X296,X299)))) )),
  inference(superposition,[],[f89,f1640])).
fof(f13548,plain,(
  ( ! [X292,X294,X288,X290,X287,X293,X289,X291] : (s(bool,X292) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X287),s(fun(X288,bool),X289),s(X288,X290),s(X288,X291))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X287),s(bool,happ(s(fun(X288,bool),X289),s(X288,X290))))),s(bool,'T_0')))) | ~p(s(bool,X293)) | ~p(s(bool,X292)) | s(bool,X293) != s(bool,X294) | s(bool,'F_0') = s(bool,happ(s(fun(X288,bool),X289),s(X288,X291)))) )),
  inference(superposition,[],[f89,f1640])).
fof(f13544,plain,(
  ( ! [X255,X261,X257,X259,X260,X262,X256,X258] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X255),s(fun(X256,bool),X257),s(X256,X258),s(X256,X259)))) | p(s(bool,X260)) | ~p(s(bool,X261)) | p(s(bool,X262)) | s(bool,X262) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X255),s(bool,happ(s(fun(X256,bool),X257),s(X256,X258))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X256,bool),X257),s(X256,X259)))) )),
  inference(superposition,[],[f88,f1640])).
fof(f13542,plain,(
  ( ! [X243,X241,X239,X246,X244,X242,X240,X245] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X239),s(fun(X240,bool),X241),s(X240,X242),s(X240,X243)))) | p(s(bool,X244)) | ~p(s(bool,X245)) | ~p(s(bool,X246)) | s(bool,X245) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X239),s(bool,happ(s(fun(X240,bool),X241),s(X240,X242))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X240,bool),X241),s(X240,X243)))) )),
  inference(superposition,[],[f88,f1640])).
fof(f13541,plain,(
  ( ! [X231,X237,X235,X233,X238,X236,X234,X232] : (s(bool,X236) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X231),s(fun(X232,bool),X233),s(X232,X234),s(X232,X235))) | p(s(bool,X237)) | ~p(s(bool,X236)) | ~p(s(bool,X238)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X231),s(bool,happ(s(fun(X232,bool),X233),s(X232,X234))))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,happ(s(fun(X232,bool),X233),s(X232,X235)))) )),
  inference(superposition,[],[f88,f1640])).
fof(f13540,plain,(
  ( ! [X229,X227,X225,X223,X230,X228,X226,X224] : (s(bool,X228) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X223),s(fun(X224,bool),X225),s(X224,X226),s(X224,X227))) | p(s(bool,X229)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X223),s(bool,happ(s(fun(X224,bool),X225),s(X224,X226))))),s(bool,'T_0')))) | ~p(s(bool,X230)) | p(s(bool,X228)) | s(bool,'F_0') = s(bool,happ(s(fun(X224,bool),X225),s(X224,X227)))) )),
  inference(superposition,[],[f88,f1640])).
fof(f13539,plain,(
  ( ! [X218,X216,X222,X220,X215,X219,X217,X221] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X215),s(fun(X216,bool),X217),s(X216,X218),s(X216,X219)))) | p(s(bool,X220)) | p(s(bool,X221)) | s(bool,X220) != s(bool,X222) | s(bool,X221) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X215),s(bool,happ(s(fun(X216,bool),X217),s(X216,X218))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X216,bool),X217),s(X216,X219)))) )),
  inference(superposition,[],[f87,f1640])).
fof(f13537,plain,(
  ( ! [X206,X204,X202,X200,X199,X205,X203,X201] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X199),s(fun(X200,bool),X201),s(X200,X202),s(X200,X203)))) | p(s(bool,X204)) | ~p(s(bool,X205)) | s(bool,X204) != s(bool,X206) | s(bool,X205) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X199),s(bool,happ(s(fun(X200,bool),X201),s(X200,X202))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X200,bool),X201),s(X200,X203)))) )),
  inference(superposition,[],[f87,f1640])).
fof(f13534,plain,(
  ( ! [X177,X175,X182,X180,X178,X176,X181,X179] : (s(bool,X180) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X175),s(fun(X176,bool),X177),s(X176,X178),s(X176,X179))) | p(s(bool,X181)) | ~p(s(bool,X180)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X175),s(bool,happ(s(fun(X176,bool),X177),s(X176,X178))))),s(bool,'T_0')))) | s(bool,X181) != s(bool,X182) | s(bool,'F_0') = s(bool,happ(s(fun(X176,bool),X177),s(X176,X179)))) )),
  inference(superposition,[],[f87,f1640])).
fof(f13533,plain,(
  ( ! [X167,X173,X171,X169,X174,X172,X170,X168] : (s(bool,X172) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X167),s(fun(X168,bool),X169),s(X168,X170),s(X168,X171))) | p(s(bool,X173)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X167),s(bool,happ(s(fun(X168,bool),X169),s(X168,X170))))),s(bool,'T_0')))) | p(s(bool,X172)) | s(bool,X173) != s(bool,X174) | s(bool,'F_0') = s(bool,happ(s(fun(X168,bool),X169),s(X168,X171)))) )),
  inference(superposition,[],[f87,f1640])).
fof(f13532,plain,(
  ( ! [X165,X163,X161,X159,X166,X164,X162,X160] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X159),s(fun(X160,bool),X161),s(X160,X162),s(X160,X163)))) | p(s(bool,X164)) | ~p(s(bool,X165)) | s(bool,X164) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X159),s(bool,happ(s(fun(X160,bool),X161),s(X160,X162))))),s(bool,'T_0'))) | s(bool,X165) != s(bool,X166) | s(bool,'F_0') = s(bool,happ(s(fun(X160,bool),X161),s(X160,X163)))) )),
  inference(superposition,[],[f85,f1640])).
fof(f13530,plain,(
  ( ! [X146,X144,X149,X147,X145,X143,X150,X148] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X143),s(fun(X144,bool),X145),s(X144,X146),s(X144,X147)))) | ~p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X149) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X143),s(bool,happ(s(fun(X144,bool),X145),s(X144,X146))))),s(bool,'T_0'))) | s(bool,X148) != s(bool,X150) | s(bool,'F_0') = s(bool,happ(s(fun(X144,bool),X145),s(X144,X147)))) )),
  inference(superposition,[],[f85,f1640])).
fof(f13529,plain,(
  ( ! [X142,X140,X138,X136,X135,X141,X139,X137] : (s(bool,X140) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X135),s(fun(X136,bool),X137),s(X136,X138),s(X136,X139))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X135),s(bool,happ(s(fun(X136,bool),X137),s(X136,X138))))),s(bool,'T_0')))) | ~p(s(bool,X141)) | ~p(s(bool,X140)) | s(bool,X141) != s(bool,X142) | s(bool,'F_0') = s(bool,happ(s(fun(X136,bool),X137),s(X136,X139)))) )),
  inference(superposition,[],[f85,f1640])).
fof(f13528,plain,(
  ( ! [X127,X134,X132,X130,X128,X133,X131,X129] : (s(bool,X132) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X127),s(fun(X128,bool),X129),s(X128,X130),s(X128,X131))) | p(s(bool,X132)) | ~p(s(bool,X133)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X127),s(bool,happ(s(fun(X128,bool),X129),s(X128,X130))))),s(bool,'T_0')))) | s(bool,X133) != s(bool,X134) | s(bool,'F_0') = s(bool,happ(s(fun(X128,bool),X129),s(X128,X131)))) )),
  inference(superposition,[],[f85,f1640])).
fof(f13525,plain,(
  ( ! [X103,X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X103),s(fun(X104,bool),X105),s(X104,X106),s(X104,X107)))) | p(s(bool,X108)) | p(s(bool,X109)) | s(bool,X108) != s(bool,X110) | s(bool,X109) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X103),s(bool,happ(s(fun(X104,bool),X105),s(X104,X106))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X104,bool),X105),s(X104,X107)))) )),
  inference(superposition,[],[f83,f1640])).
fof(f13523,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X87),s(fun(X88,bool),X89),s(X88,X90),s(X88,X91)))) | p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X92) != s(bool,X94) | s(bool,X93) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X87),s(bool,happ(s(fun(X88,bool),X89),s(X88,X90))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X88,bool),X89),s(X88,X91)))) )),
  inference(superposition,[],[f83,f1640])).
fof(f13520,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s(bool,X68) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X63),s(fun(X64,bool),X65),s(X64,X66),s(X64,X67))) | p(s(bool,X69)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X63),s(bool,happ(s(fun(X64,bool),X65),s(X64,X66))))),s(bool,'T_0')))) | p(s(bool,X68)) | s(bool,X69) != s(bool,X70) | s(bool,'F_0') = s(bool,happ(s(fun(X64,bool),X65),s(X64,X67)))) )),
  inference(superposition,[],[f83,f1640])).
fof(f13519,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X55] : (s(bool,X60) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X55),s(fun(X56,bool),X57),s(X56,X58),s(X56,X59))) | p(s(bool,X61)) | ~p(s(bool,X60)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X55),s(bool,happ(s(fun(X56,bool),X57),s(X56,X58))))),s(bool,'T_0')))) | s(bool,X61) != s(bool,X62) | s(bool,'F_0') = s(bool,happ(s(fun(X56,bool),X57),s(X56,X59)))) )),
  inference(superposition,[],[f83,f1640])).
fof(f13367,plain,(
  ( ! [X350,X346,X349,X351,X345,X347,X348] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X345),s(bool,happ(s(fun(X346,bool),X347),s(X346,X348))))),s(bool,'T_0')))) | p(s(bool,X350)) | p(s(bool,X351)) | s(bool,X351) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X345),s(fun(X346,bool),X347),s(X346,X348),s(X346,X349))) | s(bool,'F_0') = s(bool,happ(s(fun(X346,bool),X347),s(X346,X349)))) )),
  inference(superposition,[],[f125,f1640])).
fof(f13365,plain,(
  ( ! [X333,X335,X331,X337,X332,X334,X336] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X331),s(bool,happ(s(fun(X332,bool),X333),s(X332,X334))))),s(bool,'T_0')))) | p(s(bool,X336)) | ~p(s(bool,X337)) | s(bool,X337) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X331),s(fun(X332,bool),X333),s(X332,X334),s(X332,X335))) | s(bool,'F_0') = s(bool,happ(s(fun(X332,bool),X333),s(X332,X335)))) )),
  inference(superposition,[],[f125,f1640])).
fof(f13364,plain,(
  ( ! [X325,X327,X329,X324,X326,X328,X330] : (s(bool,X329) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X324),s(bool,happ(s(fun(X325,bool),X326),s(X325,X327))))),s(bool,'T_0'))) | p(s(bool,X330)) | ~p(s(bool,X329)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X324),s(fun(X325,bool),X326),s(X325,X327),s(X325,X328)))) | s(bool,'F_0') = s(bool,happ(s(fun(X325,bool),X326),s(X325,X328)))) )),
  inference(superposition,[],[f125,f1640])).
fof(f13363,plain,(
  ( ! [X321,X323,X318,X320,X322,X317,X319] : (s(bool,X322) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X317),s(bool,happ(s(fun(X318,bool),X319),s(X318,X320))))),s(bool,'T_0'))) | p(s(bool,X323)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X317),s(fun(X318,bool),X319),s(X318,X320),s(X318,X321)))) | p(s(bool,X322)) | s(bool,'F_0') = s(bool,happ(s(fun(X318,bool),X319),s(X318,X321)))) )),
  inference(superposition,[],[f125,f1640])).
fof(f13355,plain,(
  ( ! [X280,X277,X279,X275,X281,X276,X278,X274] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X274),s(bool,happ(s(fun(X275,bool),X276),s(X275,X277))))),s(bool,'T_0')))) | p(s(bool,X279)) | ~p(s(bool,X280)) | s(bool,X279) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X274),s(fun(X275,bool),X276),s(X275,X277),s(X275,X278))) | s(bool,X280) != s(bool,X281) | s(bool,'F_0') = s(bool,happ(s(fun(X275,bool),X276),s(X275,X278)))) )),
  inference(superposition,[],[f89,f1640])).
fof(f13353,plain,(
  ( ! [X261,X263,X259,X265,X260,X262,X258,X264] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X258),s(bool,happ(s(fun(X259,bool),X260),s(X259,X261))))),s(bool,'T_0')))) | ~p(s(bool,X263)) | ~p(s(bool,X264)) | s(bool,X264) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X258),s(fun(X259,bool),X260),s(X259,X261),s(X259,X262))) | s(bool,X263) != s(bool,X265) | s(bool,'F_0') = s(bool,happ(s(fun(X259,bool),X260),s(X259,X262)))) )),
  inference(superposition,[],[f89,f1640])).
fof(f13352,plain,(
  ( ! [X255,X253,X251,X257,X254,X252,X250,X256] : (s(bool,X255) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X250),s(bool,happ(s(fun(X251,bool),X252),s(X251,X253))))),s(bool,'T_0'))) | p(s(bool,X255)) | ~p(s(bool,X256)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X250),s(fun(X251,bool),X252),s(X251,X253),s(X251,X254)))) | s(bool,X256) != s(bool,X257) | s(bool,'F_0') = s(bool,happ(s(fun(X251,bool),X252),s(X251,X254)))) )),
  inference(superposition,[],[f89,f1640])).
fof(f13351,plain,(
  ( ! [X243,X249,X246,X244,X242,X248,X247,X245] : (s(bool,X247) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X242),s(bool,happ(s(fun(X243,bool),X244),s(X243,X245))))),s(bool,'T_0'))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X242),s(fun(X243,bool),X244),s(X243,X245),s(X243,X246)))) | ~p(s(bool,X248)) | ~p(s(bool,X247)) | s(bool,X248) != s(bool,X249) | s(bool,'F_0') = s(bool,happ(s(fun(X243,bool),X244),s(X243,X246)))) )),
  inference(superposition,[],[f89,f1640])).
fof(f13347,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X217,X214] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X210),s(bool,happ(s(fun(X211,bool),X212),s(X211,X213))))),s(bool,'T_0')))) | p(s(bool,X215)) | ~p(s(bool,X216)) | p(s(bool,X217)) | s(bool,X217) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X210),s(fun(X211,bool),X212),s(X211,X213),s(X211,X214))) | s(bool,'F_0') = s(bool,happ(s(fun(X211,bool),X212),s(X211,X214)))) )),
  inference(superposition,[],[f88,f1640])).
fof(f13345,plain,(
  ( ! [X198,X196,X194,X200,X199,X197,X195,X201] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X194),s(bool,happ(s(fun(X195,bool),X196),s(X195,X197))))),s(bool,'T_0')))) | p(s(bool,X199)) | ~p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X200) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X194),s(fun(X195,bool),X196),s(X195,X197),s(X195,X198))) | s(bool,'F_0') = s(bool,happ(s(fun(X195,bool),X196),s(X195,X198)))) )),
  inference(superposition,[],[f88,f1640])).
fof(f13344,plain,(
  ( ! [X191,X189,X187,X192,X190,X188,X186,X193] : (s(bool,X191) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X186),s(bool,happ(s(fun(X187,bool),X188),s(X187,X189))))),s(bool,'T_0'))) | p(s(bool,X192)) | ~p(s(bool,X191)) | ~p(s(bool,X193)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X186),s(fun(X187,bool),X188),s(X187,X189),s(X187,X190)))) | s(bool,'F_0') = s(bool,happ(s(fun(X187,bool),X188),s(X187,X190)))) )),
  inference(superposition,[],[f88,f1640])).
fof(f13343,plain,(
  ( ! [X185,X182,X180,X178,X184,X183,X181,X179] : (s(bool,X183) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X178),s(bool,happ(s(fun(X179,bool),X180),s(X179,X181))))),s(bool,'T_0'))) | p(s(bool,X184)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X178),s(fun(X179,bool),X180),s(X179,X181),s(X179,X182)))) | ~p(s(bool,X185)) | p(s(bool,X183)) | s(bool,'F_0') = s(bool,happ(s(fun(X179,bool),X180),s(X179,X182)))) )),
  inference(superposition,[],[f88,f1640])).
fof(f13342,plain,(
  ( ! [X177,X175,X173,X171,X176,X174,X172,X170] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X170),s(bool,happ(s(fun(X171,bool),X172),s(X171,X173))))),s(bool,'T_0')))) | p(s(bool,X175)) | p(s(bool,X176)) | s(bool,X175) != s(bool,X177) | s(bool,X176) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X170),s(fun(X171,bool),X172),s(X171,X173),s(X171,X174))) | s(bool,'F_0') = s(bool,happ(s(fun(X171,bool),X172),s(X171,X174)))) )),
  inference(superposition,[],[f87,f1640])).
fof(f13340,plain,(
  ( ! [X158,X156,X154,X161,X159,X157,X155,X160] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X154),s(bool,happ(s(fun(X155,bool),X156),s(X155,X157))))),s(bool,'T_0')))) | p(s(bool,X159)) | ~p(s(bool,X160)) | s(bool,X159) != s(bool,X161) | s(bool,X160) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X154),s(fun(X155,bool),X156),s(X155,X157),s(X155,X158))) | s(bool,'F_0') = s(bool,happ(s(fun(X155,bool),X156),s(X155,X158)))) )),
  inference(superposition,[],[f87,f1640])).
fof(f13337,plain,(
  ( ! [X134,X132,X130,X136,X135,X133,X131,X137] : (s(bool,X135) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X130),s(bool,happ(s(fun(X131,bool),X132),s(X131,X133))))),s(bool,'T_0'))) | p(s(bool,X136)) | ~p(s(bool,X135)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X130),s(fun(X131,bool),X132),s(X131,X133),s(X131,X134)))) | s(bool,X136) != s(bool,X137) | s(bool,'F_0') = s(bool,happ(s(fun(X131,bool),X132),s(X131,X134)))) )),
  inference(superposition,[],[f87,f1640])).
fof(f13336,plain,(
  ( ! [X127,X125,X123,X128,X126,X124,X122,X129] : (s(bool,X127) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X122),s(bool,happ(s(fun(X123,bool),X124),s(X123,X125))))),s(bool,'T_0'))) | p(s(bool,X128)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X122),s(fun(X123,bool),X124),s(X123,X125),s(X123,X126)))) | p(s(bool,X127)) | s(bool,X128) != s(bool,X129) | s(bool,'F_0') = s(bool,happ(s(fun(X123,bool),X124),s(X123,X126)))) )),
  inference(superposition,[],[f87,f1640])).
fof(f13335,plain,(
  ( ! [X121,X118,X116,X114,X120,X119,X117,X115] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X114),s(bool,happ(s(fun(X115,bool),X116),s(X115,X117))))),s(bool,'T_0')))) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X114),s(fun(X115,bool),X116),s(X115,X117),s(X115,X118))) | s(bool,X120) != s(bool,X121) | s(bool,'F_0') = s(bool,happ(s(fun(X115,bool),X116),s(X115,X118)))) )),
  inference(superposition,[],[f85,f1640])).
fof(f13333,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X98,X104] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X98),s(bool,happ(s(fun(X99,bool),X100),s(X99,X101))))),s(bool,'T_0')))) | ~p(s(bool,X103)) | ~p(s(bool,X104)) | s(bool,X104) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X98),s(fun(X99,bool),X100),s(X99,X101),s(X99,X102))) | s(bool,X103) != s(bool,X105) | s(bool,'F_0') = s(bool,happ(s(fun(X99,bool),X100),s(X99,X102)))) )),
  inference(superposition,[],[f85,f1640])).
fof(f13332,plain,(
  ( ! [X94,X92,X90,X97,X95,X93,X91,X96] : (s(bool,X95) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X90),s(bool,happ(s(fun(X91,bool),X92),s(X91,X93))))),s(bool,'T_0'))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X90),s(fun(X91,bool),X92),s(X91,X93),s(X91,X94)))) | ~p(s(bool,X96)) | ~p(s(bool,X95)) | s(bool,X96) != s(bool,X97) | s(bool,'F_0') = s(bool,happ(s(fun(X91,bool),X92),s(X91,X94)))) )),
  inference(superposition,[],[f85,f1640])).
fof(f13331,plain,(
  ( ! [X88,X87,X85,X83,X89,X86,X84,X82] : (s(bool,X87) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X82),s(bool,happ(s(fun(X83,bool),X84),s(X83,X85))))),s(bool,'T_0'))) | p(s(bool,X87)) | ~p(s(bool,X88)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X82),s(fun(X83,bool),X84),s(X83,X85),s(X83,X86)))) | s(bool,X88) != s(bool,X89) | s(bool,'F_0') = s(bool,happ(s(fun(X83,bool),X84),s(X83,X86)))) )),
  inference(superposition,[],[f85,f1640])).
fof(f13328,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X65,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X58),s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))))),s(bool,'T_0')))) | p(s(bool,X63)) | p(s(bool,X64)) | s(bool,X63) != s(bool,X65) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X58),s(fun(X59,bool),X60),s(X59,X61),s(X59,X62))) | s(bool,'F_0') = s(bool,happ(s(fun(X59,bool),X60),s(X59,X62)))) )),
  inference(superposition,[],[f83,f1640])).
fof(f13326,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42,X49] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X42),s(bool,happ(s(fun(X43,bool),X44),s(X43,X45))))),s(bool,'T_0')))) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X42),s(fun(X43,bool),X44),s(X43,X45),s(X43,X46))) | s(bool,'F_0') = s(bool,happ(s(fun(X43,bool),X44),s(X43,X46)))) )),
  inference(superposition,[],[f83,f1640])).
fof(f13323,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,X23) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X18),s(bool,happ(s(fun(X19,bool),X20),s(X19,X21))))),s(bool,'T_0'))) | p(s(bool,X24)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X18),s(fun(X19,bool),X20),s(X19,X21),s(X19,X22)))) | p(s(bool,X23)) | s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,X22)))) )),
  inference(superposition,[],[f83,f1640])).
fof(f13322,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(bool,X15) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X10),s(bool,happ(s(fun(X11,bool),X12),s(X11,X13))))),s(bool,'T_0'))) | p(s(bool,X16)) | ~p(s(bool,X15)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X10),s(fun(X11,bool),X12),s(X11,X13),s(X11,X14)))) | s(bool,X16) != s(bool,X17) | s(bool,'F_0') = s(bool,happ(s(fun(X11,bool),X12),s(X11,X14)))) )),
  inference(superposition,[],[f83,f1640])).
fof(f13156,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X13),s(bool,'F_0'))),s(bool,happ(s(fun(X14,bool),X15),s(X14,X17))))) | s(bool,'T_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,X16))) | p(s(bool,X18)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X13),s(fun(X14,bool),X15),s(X14,X16),s(X14,X17))))) ) | $spl244),
  inference(cnf_transformation,[],[f13156_D])).
fof(f13156_D,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X13),s(bool,'F_0'))),s(bool,happ(s(fun(X14,bool),X15),s(X14,X17))))) | s(bool,'T_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,X16))) | p(s(bool,X18)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X13),s(fun(X14,bool),X15),s(X14,X16),s(X14,X17))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f13004,plain,(
  ( ! [X383,X389,X385,X387,X388,X384,X386] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X383),s(fun(X384,bool),X385),s(X384,X387),s(X384,X386)))) | p(s(bool,X388)) | p(s(bool,X389)) | s(bool,X389) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X383),s(bool,'F_0'))),s(bool,happ(s(fun(X384,bool),X385),s(X384,X386))))) | s(bool,'T_0') = s(bool,happ(s(fun(X384,bool),X385),s(X384,X387)))) )),
  inference(superposition,[],[f125,f1617])).
fof(f1617,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X36),s(fun(X33,bool),X34),s(X33,X35),s(X33,X37))) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X36),s(bool,'F_0'))),s(bool,happ(s(fun(X33,bool),X34),s(X33,X37))))) | s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),X34),s(X33,X35)))) )),
  inference(superposition,[],[f123,f122])).
fof(f13155,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s(bool,X10) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X5),s(bool,'F_0'))),s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X5),s(fun(X6,bool),X7),s(X6,X8),s(X6,X9)))) | ~p(s(bool,X10))) ) | $spl242),
  inference(cnf_transformation,[],[f13155_D])).
fof(f13155_D,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s(bool,X10) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X5),s(bool,'F_0'))),s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X5),s(fun(X6,bool),X7),s(X6,X8),s(X6,X9)))) | ~p(s(bool,X10))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f13002,plain,(
  ( ! [X372,X374,X370,X373,X375,X369,X371] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X369),s(fun(X370,bool),X371),s(X370,X373),s(X370,X372)))) | p(s(bool,X374)) | ~p(s(bool,X375)) | s(bool,X375) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X369),s(bool,'F_0'))),s(bool,happ(s(fun(X370,bool),X371),s(X370,X372))))) | s(bool,'T_0') = s(bool,happ(s(fun(X370,bool),X371),s(X370,X373)))) )),
  inference(superposition,[],[f125,f1617])).
fof(f13158,plain,(
  ( ! [X39,X37,X43,X41,X38,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X37),s(bool,'F_0'))),s(bool,happ(s(fun(X38,bool),X39),s(X38,X41)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X38,bool),X39),s(X38,X40))) | s(bool,X43) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X37),s(fun(X38,bool),X39),s(X38,X40),s(X38,X41))) | ~p(s(bool,X43))) ) | $spl246),
  inference(cnf_transformation,[],[f13158_D])).
fof(f13158_D,plain,(
  ( ! [X39,X37,X43,X41,X38,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X37),s(bool,'F_0'))),s(bool,happ(s(fun(X38,bool),X39),s(X38,X41)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X38,bool),X39),s(X38,X40))) | s(bool,X43) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X37),s(fun(X38,bool),X39),s(X38,X40),s(X38,X41))) | ~p(s(bool,X43))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f13001,plain,(
  ( ! [X364,X366,X362,X368,X365,X367,X363] : (s(bool,X367) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X362),s(fun(X363,bool),X364),s(X363,X366),s(X363,X365))) | p(s(bool,X368)) | ~p(s(bool,X367)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X362),s(bool,'F_0'))),s(bool,happ(s(fun(X363,bool),X364),s(X363,X365)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X363,bool),X364),s(X363,X366)))) )),
  inference(superposition,[],[f125,f1617])).
fof(f13159,plain,(
  ( ! [X59,X57,X54,X56,X55,X53] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X53),s(bool,'F_0'))),s(bool,happ(s(fun(X54,bool),X55),s(X54,X57)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X54,bool),X55),s(X54,X56))) | s(bool,X59) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X53),s(fun(X54,bool),X55),s(X54,X56),s(X54,X57))) | p(s(bool,X59))) ) | $spl248),
  inference(cnf_transformation,[],[f13159_D])).
fof(f13159_D,plain,(
  ( ! [X59,X57,X54,X56,X55,X53] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X53),s(bool,'F_0'))),s(bool,happ(s(fun(X54,bool),X55),s(X54,X57)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X54,bool),X55),s(X54,X56))) | s(bool,X59) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X53),s(fun(X54,bool),X55),s(X54,X56),s(X54,X57))) | p(s(bool,X59))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f13000,plain,(
  ( ! [X356,X358,X360,X357,X359,X355,X361] : (s(bool,X360) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X355),s(fun(X356,bool),X357),s(X356,X359),s(X356,X358))) | p(s(bool,X361)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X355),s(bool,'F_0'))),s(bool,happ(s(fun(X356,bool),X357),s(X356,X358)))))) | p(s(bool,X360)) | s(bool,'T_0') = s(bool,happ(s(fun(X356,bool),X357),s(X356,X359)))) )),
  inference(superposition,[],[f125,f1617])).
fof(f12992,plain,(
  ( ! [X313,X315,X316,X318,X312,X314,X317,X319] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X312),s(fun(X313,bool),X314),s(X313,X316),s(X313,X315)))) | p(s(bool,X317)) | ~p(s(bool,X318)) | s(bool,X317) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X312),s(bool,'F_0'))),s(bool,happ(s(fun(X313,bool),X314),s(X313,X315))))) | s(bool,X318) != s(bool,X319) | s(bool,'T_0') = s(bool,happ(s(fun(X313,bool),X314),s(X313,X316)))) )),
  inference(superposition,[],[f89,f1617])).
fof(f12990,plain,(
  ( ! [X300,X302,X296,X298,X301,X303,X297,X299] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X296),s(fun(X297,bool),X298),s(X297,X300),s(X297,X299)))) | ~p(s(bool,X301)) | ~p(s(bool,X302)) | s(bool,X302) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X296),s(bool,'F_0'))),s(bool,happ(s(fun(X297,bool),X298),s(X297,X299))))) | s(bool,X301) != s(bool,X303) | s(bool,'T_0') = s(bool,happ(s(fun(X297,bool),X298),s(X297,X300)))) )),
  inference(superposition,[],[f89,f1617])).
fof(f12989,plain,(
  ( ! [X292,X294,X288,X290,X293,X295,X289,X291] : (s(bool,X293) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X288),s(fun(X289,bool),X290),s(X289,X292),s(X289,X291))) | p(s(bool,X293)) | ~p(s(bool,X294)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X288),s(bool,'F_0'))),s(bool,happ(s(fun(X289,bool),X290),s(X289,X291)))))) | s(bool,X294) != s(bool,X295) | s(bool,'T_0') = s(bool,happ(s(fun(X289,bool),X290),s(X289,X292)))) )),
  inference(superposition,[],[f89,f1617])).
fof(f12988,plain,(
  ( ! [X280,X282,X285,X287,X281,X283,X284,X286] : (s(bool,X285) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X280),s(fun(X281,bool),X282),s(X281,X284),s(X281,X283))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X280),s(bool,'F_0'))),s(bool,happ(s(fun(X281,bool),X282),s(X281,X283)))))) | ~p(s(bool,X286)) | ~p(s(bool,X285)) | s(bool,X286) != s(bool,X287) | s(bool,'T_0') = s(bool,happ(s(fun(X281,bool),X282),s(X281,X284)))) )),
  inference(superposition,[],[f89,f1617])).
fof(f12984,plain,(
  ( ! [X255,X253,X251,X249,X254,X252,X250,X248] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X248),s(fun(X249,bool),X250),s(X249,X252),s(X249,X251)))) | p(s(bool,X253)) | ~p(s(bool,X254)) | p(s(bool,X255)) | s(bool,X255) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X248),s(bool,'F_0'))),s(bool,happ(s(fun(X249,bool),X250),s(X249,X251))))) | s(bool,'T_0') = s(bool,happ(s(fun(X249,bool),X250),s(X249,X252)))) )),
  inference(superposition,[],[f88,f1617])).
fof(f12982,plain,(
  ( ! [X239,X237,X235,X233,X238,X236,X234,X232] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X232),s(fun(X233,bool),X234),s(X233,X236),s(X233,X235)))) | p(s(bool,X237)) | ~p(s(bool,X238)) | ~p(s(bool,X239)) | s(bool,X238) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X232),s(bool,'F_0'))),s(bool,happ(s(fun(X233,bool),X234),s(X233,X235))))) | s(bool,'T_0') = s(bool,happ(s(fun(X233,bool),X234),s(X233,X236)))) )),
  inference(superposition,[],[f88,f1617])).
fof(f12981,plain,(
  ( ! [X231,X229,X227,X225,X230,X228,X226,X224] : (s(bool,X229) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X224),s(fun(X225,bool),X226),s(X225,X228),s(X225,X227))) | p(s(bool,X230)) | ~p(s(bool,X229)) | ~p(s(bool,X231)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X224),s(bool,'F_0'))),s(bool,happ(s(fun(X225,bool),X226),s(X225,X227)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X225,bool),X226),s(X225,X228)))) )),
  inference(superposition,[],[f88,f1617])).
fof(f12980,plain,(
  ( ! [X218,X216,X222,X220,X219,X217,X223,X221] : (s(bool,X221) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X216),s(fun(X217,bool),X218),s(X217,X220),s(X217,X219))) | p(s(bool,X222)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X216),s(bool,'F_0'))),s(bool,happ(s(fun(X217,bool),X218),s(X217,X219)))))) | ~p(s(bool,X223)) | p(s(bool,X221)) | s(bool,'T_0') = s(bool,happ(s(fun(X217,bool),X218),s(X217,X220)))) )),
  inference(superposition,[],[f88,f1617])).
fof(f12979,plain,(
  ( ! [X212,X210,X208,X215,X213,X211,X209,X214] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X208),s(fun(X209,bool),X210),s(X209,X212),s(X209,X211)))) | p(s(bool,X213)) | p(s(bool,X214)) | s(bool,X213) != s(bool,X215) | s(bool,X214) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X208),s(bool,'F_0'))),s(bool,happ(s(fun(X209,bool),X210),s(X209,X211))))) | s(bool,'T_0') = s(bool,happ(s(fun(X209,bool),X210),s(X209,X212)))) )),
  inference(superposition,[],[f87,f1617])).
fof(f12977,plain,(
  ( ! [X198,X196,X194,X192,X199,X197,X195,X193] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X192),s(fun(X193,bool),X194),s(X193,X196),s(X193,X195)))) | p(s(bool,X197)) | ~p(s(bool,X198)) | s(bool,X197) != s(bool,X199) | s(bool,X198) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X192),s(bool,'F_0'))),s(bool,happ(s(fun(X193,bool),X194),s(X193,X195))))) | s(bool,'T_0') = s(bool,happ(s(fun(X193,bool),X194),s(X193,X196)))) )),
  inference(superposition,[],[f87,f1617])).
fof(f12974,plain,(
  ( ! [X175,X173,X171,X169,X174,X172,X170,X168] : (s(bool,X173) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X168),s(fun(X169,bool),X170),s(X169,X172),s(X169,X171))) | p(s(bool,X174)) | ~p(s(bool,X173)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X168),s(bool,'F_0'))),s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))))) | s(bool,X174) != s(bool,X175) | s(bool,'T_0') = s(bool,happ(s(fun(X169,bool),X170),s(X169,X172)))) )),
  inference(superposition,[],[f87,f1617])).
fof(f12973,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162,X160] : (s(bool,X165) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X160),s(fun(X161,bool),X162),s(X161,X164),s(X161,X163))) | p(s(bool,X166)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X160),s(bool,'F_0'))),s(bool,happ(s(fun(X161,bool),X162),s(X161,X163)))))) | p(s(bool,X165)) | s(bool,X166) != s(bool,X167) | s(bool,'T_0') = s(bool,happ(s(fun(X161,bool),X162),s(X161,X164)))) )),
  inference(superposition,[],[f87,f1617])).
fof(f12972,plain,(
  ( ! [X158,X156,X154,X152,X159,X157,X155,X153] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X152),s(fun(X153,bool),X154),s(X153,X156),s(X153,X155)))) | p(s(bool,X157)) | ~p(s(bool,X158)) | s(bool,X157) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X152),s(bool,'F_0'))),s(bool,happ(s(fun(X153,bool),X154),s(X153,X155))))) | s(bool,X158) != s(bool,X159) | s(bool,'T_0') = s(bool,happ(s(fun(X153,bool),X154),s(X153,X156)))) )),
  inference(superposition,[],[f85,f1617])).
fof(f12970,plain,(
  ( ! [X142,X140,X138,X136,X143,X141,X139,X137] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X136),s(fun(X137,bool),X138),s(X137,X140),s(X137,X139)))) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X142) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X136),s(bool,'F_0'))),s(bool,happ(s(fun(X137,bool),X138),s(X137,X139))))) | s(bool,X141) != s(bool,X143) | s(bool,'T_0') = s(bool,happ(s(fun(X137,bool),X138),s(X137,X140)))) )),
  inference(superposition,[],[f85,f1617])).
fof(f12969,plain,(
  ( ! [X134,X132,X130,X128,X135,X133,X131,X129] : (s(bool,X133) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X128),s(fun(X129,bool),X130),s(X129,X132),s(X129,X131))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X128),s(bool,'F_0'))),s(bool,happ(s(fun(X129,bool),X130),s(X129,X131)))))) | ~p(s(bool,X134)) | ~p(s(bool,X133)) | s(bool,X134) != s(bool,X135) | s(bool,'T_0') = s(bool,happ(s(fun(X129,bool),X130),s(X129,X132)))) )),
  inference(superposition,[],[f85,f1617])).
fof(f12968,plain,(
  ( ! [X127,X125,X123,X121,X126,X124,X122,X120] : (s(bool,X125) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X120),s(fun(X121,bool),X122),s(X121,X124),s(X121,X123))) | p(s(bool,X125)) | ~p(s(bool,X126)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X120),s(bool,'F_0'))),s(bool,happ(s(fun(X121,bool),X122),s(X121,X123)))))) | s(bool,X126) != s(bool,X127) | s(bool,'T_0') = s(bool,happ(s(fun(X121,bool),X122),s(X121,X124)))) )),
  inference(superposition,[],[f85,f1617])).
fof(f12965,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X96),s(fun(X97,bool),X98),s(X97,X100),s(X97,X99)))) | p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X96),s(bool,'F_0'))),s(bool,happ(s(fun(X97,bool),X98),s(X97,X99))))) | s(bool,'T_0') = s(bool,happ(s(fun(X97,bool),X98),s(X97,X100)))) )),
  inference(superposition,[],[f83,f1617])).
fof(f12963,plain,(
  ( ! [X80,X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X80),s(fun(X81,bool),X82),s(X81,X84),s(X81,X83)))) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X87) | s(bool,X86) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X80),s(bool,'F_0'))),s(bool,happ(s(fun(X81,bool),X82),s(X81,X83))))) | s(bool,'T_0') = s(bool,happ(s(fun(X81,bool),X82),s(X81,X84)))) )),
  inference(superposition,[],[f83,f1617])).
fof(f12960,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X61) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X56),s(fun(X57,bool),X58),s(X57,X60),s(X57,X59))) | p(s(bool,X62)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X56),s(bool,'F_0'))),s(bool,happ(s(fun(X57,bool),X58),s(X57,X59)))))) | p(s(bool,X61)) | s(bool,X62) != s(bool,X63) | s(bool,'T_0') = s(bool,happ(s(fun(X57,bool),X58),s(X57,X60)))) )),
  inference(superposition,[],[f83,f1617])).
fof(f12959,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (s(bool,X53) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X48),s(fun(X49,bool),X50),s(X49,X52),s(X49,X51))) | p(s(bool,X54)) | ~p(s(bool,X53)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X48),s(bool,'F_0'))),s(bool,happ(s(fun(X49,bool),X50),s(X49,X51)))))) | s(bool,X54) != s(bool,X55) | s(bool,'T_0') = s(bool,happ(s(fun(X49,bool),X50),s(X49,X52)))) )),
  inference(superposition,[],[f83,f1617])).
fof(f12814,plain,(
  ( ! [X344,X346,X341,X343,X345,X340,X342] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X340),s(bool,'F_0'))),s(bool,happ(s(fun(X341,bool),X342),s(X341,X344)))))) | p(s(bool,X345)) | p(s(bool,X346)) | s(bool,X346) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X340),s(fun(X341,bool),X342),s(X341,X343),s(X341,X344))) | s(bool,'T_0') = s(bool,happ(s(fun(X341,bool),X342),s(X341,X343)))) )),
  inference(superposition,[],[f125,f1617])).
fof(f12812,plain,(
  ( ! [X327,X329,X331,X326,X332,X328,X330] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X326),s(bool,'F_0'))),s(bool,happ(s(fun(X327,bool),X328),s(X327,X330)))))) | p(s(bool,X331)) | ~p(s(bool,X332)) | s(bool,X332) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X326),s(fun(X327,bool),X328),s(X327,X329),s(X327,X330))) | s(bool,'T_0') = s(bool,happ(s(fun(X327,bool),X328),s(X327,X329)))) )),
  inference(superposition,[],[f125,f1617])).
fof(f12811,plain,(
  ( ! [X325,X321,X323,X324,X320,X322,X319] : (s(bool,X324) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X319),s(bool,'F_0'))),s(bool,happ(s(fun(X320,bool),X321),s(X320,X323))))) | p(s(bool,X325)) | ~p(s(bool,X324)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X319),s(fun(X320,bool),X321),s(X320,X322),s(X320,X323)))) | s(bool,'T_0') = s(bool,happ(s(fun(X320,bool),X321),s(X320,X322)))) )),
  inference(superposition,[],[f125,f1617])).
fof(f12810,plain,(
  ( ! [X313,X315,X316,X318,X312,X314,X317] : (s(bool,X317) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X312),s(bool,'F_0'))),s(bool,happ(s(fun(X313,bool),X314),s(X313,X316))))) | p(s(bool,X318)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X312),s(fun(X313,bool),X314),s(X313,X315),s(X313,X316)))) | p(s(bool,X317)) | s(bool,'T_0') = s(bool,happ(s(fun(X313,bool),X314),s(X313,X315)))) )),
  inference(superposition,[],[f125,f1617])).
fof(f12802,plain,(
  ( ! [X269,X271,X273,X275,X270,X276,X272,X274] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X269),s(bool,'F_0'))),s(bool,happ(s(fun(X270,bool),X271),s(X270,X273)))))) | p(s(bool,X274)) | ~p(s(bool,X275)) | s(bool,X274) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X269),s(fun(X270,bool),X271),s(X270,X272),s(X270,X273))) | s(bool,X275) != s(bool,X276) | s(bool,'T_0') = s(bool,happ(s(fun(X270,bool),X271),s(X270,X272)))) )),
  inference(superposition,[],[f89,f1617])).
fof(f12800,plain,(
  ( ! [X255,X253,X257,X259,X260,X254,X256,X258] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X253),s(bool,'F_0'))),s(bool,happ(s(fun(X254,bool),X255),s(X254,X257)))))) | ~p(s(bool,X258)) | ~p(s(bool,X259)) | s(bool,X259) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X253),s(fun(X254,bool),X255),s(X254,X256),s(X254,X257))) | s(bool,X258) != s(bool,X260) | s(bool,'T_0') = s(bool,happ(s(fun(X254,bool),X255),s(X254,X256)))) )),
  inference(superposition,[],[f89,f1617])).
fof(f12799,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247,X245] : (s(bool,X250) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X245),s(bool,'F_0'))),s(bool,happ(s(fun(X246,bool),X247),s(X246,X249))))) | p(s(bool,X250)) | ~p(s(bool,X251)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X245),s(fun(X246,bool),X247),s(X246,X248),s(X246,X249)))) | s(bool,X251) != s(bool,X252) | s(bool,'T_0') = s(bool,happ(s(fun(X246,bool),X247),s(X246,X248)))) )),
  inference(superposition,[],[f89,f1617])).
fof(f12798,plain,(
  ( ! [X243,X241,X239,X237,X244,X242,X240,X238] : (s(bool,X242) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X237),s(bool,'F_0'))),s(bool,happ(s(fun(X238,bool),X239),s(X238,X241))))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X237),s(fun(X238,bool),X239),s(X238,X240),s(X238,X241)))) | ~p(s(bool,X243)) | ~p(s(bool,X242)) | s(bool,X243) != s(bool,X244) | s(bool,'T_0') = s(bool,happ(s(fun(X238,bool),X239),s(X238,X240)))) )),
  inference(superposition,[],[f89,f1617])).
fof(f12794,plain,(
  ( ! [X212,X210,X208,X206,X211,X209,X207,X205] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X205),s(bool,'F_0'))),s(bool,happ(s(fun(X206,bool),X207),s(X206,X209)))))) | p(s(bool,X210)) | ~p(s(bool,X211)) | p(s(bool,X212)) | s(bool,X212) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X205),s(fun(X206,bool),X207),s(X206,X208),s(X206,X209))) | s(bool,'T_0') = s(bool,happ(s(fun(X206,bool),X207),s(X206,X208)))) )),
  inference(superposition,[],[f88,f1617])).
fof(f12792,plain,(
  ( ! [X191,X189,X196,X194,X192,X190,X195,X193] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X189),s(bool,'F_0'))),s(bool,happ(s(fun(X190,bool),X191),s(X190,X193)))))) | p(s(bool,X194)) | ~p(s(bool,X195)) | ~p(s(bool,X196)) | s(bool,X195) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X189),s(fun(X190,bool),X191),s(X190,X192),s(X190,X193))) | s(bool,'T_0') = s(bool,happ(s(fun(X190,bool),X191),s(X190,X192)))) )),
  inference(superposition,[],[f88,f1617])).
fof(f12791,plain,(
  ( ! [X187,X185,X182,X188,X186,X184,X183,X181] : (s(bool,X186) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X181),s(bool,'F_0'))),s(bool,happ(s(fun(X182,bool),X183),s(X182,X185))))) | p(s(bool,X187)) | ~p(s(bool,X186)) | ~p(s(bool,X188)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X181),s(fun(X182,bool),X183),s(X182,X184),s(X182,X185)))) | s(bool,'T_0') = s(bool,happ(s(fun(X182,bool),X183),s(X182,X184)))) )),
  inference(superposition,[],[f88,f1617])).
fof(f12790,plain,(
  ( ! [X177,X175,X173,X180,X178,X176,X174,X179] : (s(bool,X178) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X173),s(bool,'F_0'))),s(bool,happ(s(fun(X174,bool),X175),s(X174,X177))))) | p(s(bool,X179)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X173),s(fun(X174,bool),X175),s(X174,X176),s(X174,X177)))) | ~p(s(bool,X180)) | p(s(bool,X178)) | s(bool,'T_0') = s(bool,happ(s(fun(X174,bool),X175),s(X174,X176)))) )),
  inference(superposition,[],[f88,f1617])).
fof(f12789,plain,(
  ( ! [X167,X165,X171,X169,X166,X172,X170,X168] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X165),s(bool,'F_0'))),s(bool,happ(s(fun(X166,bool),X167),s(X166,X169)))))) | p(s(bool,X170)) | p(s(bool,X171)) | s(bool,X170) != s(bool,X172) | s(bool,X171) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X165),s(fun(X166,bool),X167),s(X166,X168),s(X166,X169))) | s(bool,'T_0') = s(bool,happ(s(fun(X166,bool),X167),s(X166,X168)))) )),
  inference(superposition,[],[f87,f1617])).
fof(f12787,plain,(
  ( ! [X156,X154,X152,X151,X149,X155,X153,X150] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X149),s(bool,'F_0'))),s(bool,happ(s(fun(X150,bool),X151),s(X150,X153)))))) | p(s(bool,X154)) | ~p(s(bool,X155)) | s(bool,X154) != s(bool,X156) | s(bool,X155) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X149),s(fun(X150,bool),X151),s(X150,X152),s(X150,X153))) | s(bool,'T_0') = s(bool,happ(s(fun(X150,bool),X151),s(X150,X152)))) )),
  inference(superposition,[],[f87,f1617])).
fof(f12784,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X131,X129] : (s(bool,X130) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X125),s(bool,'F_0'))),s(bool,happ(s(fun(X126,bool),X127),s(X126,X129))))) | p(s(bool,X131)) | ~p(s(bool,X130)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X125),s(fun(X126,bool),X127),s(X126,X128),s(X126,X129)))) | s(bool,X131) != s(bool,X132) | s(bool,'T_0') = s(bool,happ(s(fun(X126,bool),X127),s(X126,X128)))) )),
  inference(superposition,[],[f87,f1617])).
fof(f12783,plain,(
  ( ! [X123,X121,X118,X124,X122,X120,X119,X117] : (s(bool,X122) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X117),s(bool,'F_0'))),s(bool,happ(s(fun(X118,bool),X119),s(X118,X121))))) | p(s(bool,X123)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X117),s(fun(X118,bool),X119),s(X118,X120),s(X118,X121)))) | p(s(bool,X122)) | s(bool,X123) != s(bool,X124) | s(bool,'T_0') = s(bool,happ(s(fun(X118,bool),X119),s(X118,X120)))) )),
  inference(superposition,[],[f87,f1617])).
fof(f12782,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X109),s(bool,'F_0'))),s(bool,happ(s(fun(X110,bool),X111),s(X110,X113)))))) | p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X114) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X109),s(fun(X110,bool),X111),s(X110,X112),s(X110,X113))) | s(bool,X115) != s(bool,X116) | s(bool,'T_0') = s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) )),
  inference(superposition,[],[f85,f1617])).
fof(f12780,plain,(
  ( ! [X94,X99,X97,X95,X93,X100,X98,X96] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X93),s(bool,'F_0'))),s(bool,happ(s(fun(X94,bool),X95),s(X94,X97)))))) | ~p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X99) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X93),s(fun(X94,bool),X95),s(X94,X96),s(X94,X97))) | s(bool,X98) != s(bool,X100) | s(bool,'T_0') = s(bool,happ(s(fun(X94,bool),X95),s(X94,X96)))) )),
  inference(superposition,[],[f85,f1617])).
fof(f12779,plain,(
  ( ! [X92,X90,X88,X87,X85,X91,X89,X86] : (s(bool,X90) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X85),s(bool,'F_0'))),s(bool,happ(s(fun(X86,bool),X87),s(X86,X89))))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X85),s(fun(X86,bool),X87),s(X86,X88),s(X86,X89)))) | ~p(s(bool,X91)) | ~p(s(bool,X90)) | s(bool,X91) != s(bool,X92) | s(bool,'T_0') = s(bool,happ(s(fun(X86,bool),X87),s(X86,X88)))) )),
  inference(superposition,[],[f85,f1617])).
fof(f12778,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X84,X82] : (s(bool,X82) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X77),s(bool,'F_0'))),s(bool,happ(s(fun(X78,bool),X79),s(X78,X81))))) | p(s(bool,X82)) | ~p(s(bool,X83)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X77),s(fun(X78,bool),X79),s(X78,X80),s(X78,X81)))) | s(bool,X83) != s(bool,X84) | s(bool,'T_0') = s(bool,happ(s(fun(X78,bool),X79),s(X78,X80)))) )),
  inference(superposition,[],[f85,f1617])).
fof(f12775,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X53),s(bool,'F_0'))),s(bool,happ(s(fun(X54,bool),X55),s(X54,X57)))))) | p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X58) != s(bool,X60) | s(bool,X59) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X53),s(fun(X54,bool),X55),s(X54,X56),s(X54,X57))) | s(bool,'T_0') = s(bool,happ(s(fun(X54,bool),X55),s(X54,X56)))) )),
  inference(superposition,[],[f83,f1617])).
fof(f12773,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X37),s(bool,'F_0'))),s(bool,happ(s(fun(X38,bool),X39),s(X38,X41)))))) | p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,X44) | s(bool,X43) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X37),s(fun(X38,bool),X39),s(X38,X40),s(X38,X41))) | s(bool,'T_0') = s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) )),
  inference(superposition,[],[f83,f1617])).
fof(f12770,plain,(
  ( ! [X14,X19,X17,X15,X13,X20,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X13),s(bool,'F_0'))),s(bool,happ(s(fun(X14,bool),X15),s(X14,X17))))) | p(s(bool,X19)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X13),s(fun(X14,bool),X15),s(X14,X16),s(X14,X17)))) | p(s(bool,X18)) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,X16)))) )),
  inference(superposition,[],[f83,f1617])).
fof(f12769,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X11,X9] : (s(bool,X10) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X5),s(bool,'F_0'))),s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))) | p(s(bool,X11)) | ~p(s(bool,X10)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X5),s(fun(X6,bool),X7),s(X6,X8),s(X6,X9)))) | s(bool,X11) != s(bool,X12) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) )),
  inference(superposition,[],[f83,f1617])).
fof(f12613,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,X13) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,'T_0'))),s(bool,happ(s(fun(X9,bool),X10),s(X9,X12))))) | s(bool,'F_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))) | p(s(bool,X13)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X8),s(fun(X9,bool),X10),s(X9,X11),s(X9,X12))))) ) | $spl236),
  inference(cnf_transformation,[],[f12613_D])).
fof(f12613_D,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,X13) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,'T_0'))),s(bool,happ(s(fun(X9,bool),X10),s(X9,X12))))) | s(bool,'F_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))) | p(s(bool,X13)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X8),s(fun(X9,bool),X10),s(X9,X11),s(X9,X12))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f12463,plain,(
  ( ! [X377,X379,X380,X382,X376,X378,X381] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X376),s(fun(X377,bool),X378),s(X377,X380),s(X377,X379)))) | p(s(bool,X381)) | p(s(bool,X382)) | s(bool,X382) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X376),s(bool,'T_0'))),s(bool,happ(s(fun(X377,bool),X378),s(X377,X379))))) | s(bool,'F_0') = s(bool,happ(s(fun(X377,bool),X378),s(X377,X380)))) )),
  inference(superposition,[],[f125,f1616])).
fof(f1616,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X31),s(fun(X28,bool),X29),s(X28,X30),s(X28,X32))) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X31),s(bool,'T_0'))),s(bool,happ(s(fun(X28,bool),X29),s(X28,X32))))) | s(bool,'F_0') = s(bool,happ(s(fun(X28,bool),X29),s(X28,X30)))) )),
  inference(superposition,[],[f123,f122])).
fof(f12612,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X5) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X0),s(bool,'T_0'))),s(bool,happ(s(fun(X1,bool),X2),s(X1,X4))))) | s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X0),s(fun(X1,bool),X2),s(X1,X3),s(X1,X4)))) | ~p(s(bool,X5))) ) | $spl234),
  inference(cnf_transformation,[],[f12612_D])).
fof(f12612_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X5) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X0),s(bool,'T_0'))),s(bool,happ(s(fun(X1,bool),X2),s(X1,X4))))) | s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X0),s(fun(X1,bool),X2),s(X1,X3),s(X1,X4)))) | ~p(s(bool,X5))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f12461,plain,(
  ( ! [X364,X366,X362,X368,X365,X367,X363] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X362),s(fun(X363,bool),X364),s(X363,X366),s(X363,X365)))) | p(s(bool,X367)) | ~p(s(bool,X368)) | s(bool,X368) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X362),s(bool,'T_0'))),s(bool,happ(s(fun(X363,bool),X364),s(X363,X365))))) | s(bool,'F_0') = s(bool,happ(s(fun(X363,bool),X364),s(X363,X366)))) )),
  inference(superposition,[],[f125,f1616])).
fof(f12615,plain,(
  ( ! [X35,X33,X38,X36,X34,X32] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X32),s(bool,'T_0'))),s(bool,happ(s(fun(X33,bool),X34),s(X33,X36)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X33,bool),X34),s(X33,X35))) | s(bool,X38) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X32),s(fun(X33,bool),X34),s(X33,X35),s(X33,X36))) | ~p(s(bool,X38))) ) | $spl238),
  inference(cnf_transformation,[],[f12615_D])).
fof(f12615_D,plain,(
  ( ! [X35,X33,X38,X36,X34,X32] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X32),s(bool,'T_0'))),s(bool,happ(s(fun(X33,bool),X34),s(X33,X36)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X33,bool),X34),s(X33,X35))) | s(bool,X38) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X32),s(fun(X33,bool),X34),s(X33,X35),s(X33,X36))) | ~p(s(bool,X38))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f12460,plain,(
  ( ! [X356,X358,X360,X357,X359,X355,X361] : (s(bool,X360) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X355),s(fun(X356,bool),X357),s(X356,X359),s(X356,X358))) | p(s(bool,X361)) | ~p(s(bool,X360)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X355),s(bool,'T_0'))),s(bool,happ(s(fun(X356,bool),X357),s(X356,X358)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X356,bool),X357),s(X356,X359)))) )),
  inference(superposition,[],[f125,f1616])).
fof(f12616,plain,(
  ( ! [X54,X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X48),s(bool,'T_0'))),s(bool,happ(s(fun(X49,bool),X50),s(X49,X52)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X49,bool),X50),s(X49,X51))) | s(bool,X54) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X48),s(fun(X49,bool),X50),s(X49,X51),s(X49,X52))) | p(s(bool,X54))) ) | $spl240),
  inference(cnf_transformation,[],[f12616_D])).
fof(f12616_D,plain,(
  ( ! [X54,X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X48),s(bool,'T_0'))),s(bool,happ(s(fun(X49,bool),X50),s(X49,X52)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X49,bool),X50),s(X49,X51))) | s(bool,X54) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X48),s(fun(X49,bool),X50),s(X49,X51),s(X49,X52))) | p(s(bool,X54))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f12459,plain,(
  ( ! [X350,X352,X354,X349,X351,X353,X348] : (s(bool,X353) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X348),s(fun(X349,bool),X350),s(X349,X352),s(X349,X351))) | p(s(bool,X354)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X348),s(bool,'T_0'))),s(bool,happ(s(fun(X349,bool),X350),s(X349,X351)))))) | p(s(bool,X353)) | s(bool,'F_0') = s(bool,happ(s(fun(X349,bool),X350),s(X349,X352)))) )),
  inference(superposition,[],[f125,f1616])).
fof(f12451,plain,(
  ( ! [X308,X310,X306,X312,X309,X311,X305,X307] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X305),s(fun(X306,bool),X307),s(X306,X309),s(X306,X308)))) | p(s(bool,X310)) | ~p(s(bool,X311)) | s(bool,X310) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X305),s(bool,'T_0'))),s(bool,happ(s(fun(X306,bool),X307),s(X306,X308))))) | s(bool,X311) != s(bool,X312) | s(bool,'F_0') = s(bool,happ(s(fun(X306,bool),X307),s(X306,X309)))) )),
  inference(superposition,[],[f89,f1616])).
fof(f12449,plain,(
  ( ! [X292,X294,X290,X296,X293,X295,X289,X291] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X289),s(fun(X290,bool),X291),s(X290,X293),s(X290,X292)))) | ~p(s(bool,X294)) | ~p(s(bool,X295)) | s(bool,X295) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X289),s(bool,'T_0'))),s(bool,happ(s(fun(X290,bool),X291),s(X290,X292))))) | s(bool,X294) != s(bool,X296) | s(bool,'F_0') = s(bool,happ(s(fun(X290,bool),X291),s(X290,X293)))) )),
  inference(superposition,[],[f89,f1616])).
fof(f12448,plain,(
  ( ! [X282,X288,X285,X287,X281,X283,X284,X286] : (s(bool,X286) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X281),s(fun(X282,bool),X283),s(X282,X285),s(X282,X284))) | p(s(bool,X286)) | ~p(s(bool,X287)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X281),s(bool,'T_0'))),s(bool,happ(s(fun(X282,bool),X283),s(X282,X284)))))) | s(bool,X287) != s(bool,X288) | s(bool,'F_0') = s(bool,happ(s(fun(X282,bool),X283),s(X282,X285)))) )),
  inference(superposition,[],[f89,f1616])).
fof(f12447,plain,(
  ( ! [X280,X277,X279,X273,X275,X276,X278,X274] : (s(bool,X278) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X273),s(fun(X274,bool),X275),s(X274,X277),s(X274,X276))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X273),s(bool,'T_0'))),s(bool,happ(s(fun(X274,bool),X275),s(X274,X276)))))) | ~p(s(bool,X279)) | ~p(s(bool,X278)) | s(bool,X279) != s(bool,X280) | s(bool,'F_0') = s(bool,happ(s(fun(X274,bool),X275),s(X274,X277)))) )),
  inference(superposition,[],[f89,f1616])).
fof(f12443,plain,(
  ( ! [X243,X241,X246,X244,X242,X248,X247,X245] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X241),s(fun(X242,bool),X243),s(X242,X245),s(X242,X244)))) | p(s(bool,X246)) | ~p(s(bool,X247)) | p(s(bool,X248)) | s(bool,X248) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X241),s(bool,'T_0'))),s(bool,happ(s(fun(X242,bool),X243),s(X242,X244))))) | s(bool,'F_0') = s(bool,happ(s(fun(X242,bool),X243),s(X242,X245)))) )),
  inference(superposition,[],[f88,f1616])).
fof(f12441,plain,(
  ( ! [X231,X229,X227,X225,X230,X228,X226,X232] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X225),s(fun(X226,bool),X227),s(X226,X229),s(X226,X228)))) | p(s(bool,X230)) | ~p(s(bool,X231)) | ~p(s(bool,X232)) | s(bool,X231) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X225),s(bool,'T_0'))),s(bool,happ(s(fun(X226,bool),X227),s(X226,X228))))) | s(bool,'F_0') = s(bool,happ(s(fun(X226,bool),X227),s(X226,X229)))) )),
  inference(superposition,[],[f88,f1616])).
fof(f12440,plain,(
  ( ! [X218,X222,X220,X219,X217,X223,X221,X224] : (s(bool,X222) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X217),s(fun(X218,bool),X219),s(X218,X221),s(X218,X220))) | p(s(bool,X223)) | ~p(s(bool,X222)) | ~p(s(bool,X224)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X217),s(bool,'T_0'))),s(bool,happ(s(fun(X218,bool),X219),s(X218,X220)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X218,bool),X219),s(X218,X221)))) )),
  inference(superposition,[],[f88,f1616])).
fof(f12439,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X209,X214] : (s(bool,X214) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X209),s(fun(X210,bool),X211),s(X210,X213),s(X210,X212))) | p(s(bool,X215)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X209),s(bool,'T_0'))),s(bool,happ(s(fun(X210,bool),X211),s(X210,X212)))))) | ~p(s(bool,X216)) | p(s(bool,X214)) | s(bool,'F_0') = s(bool,happ(s(fun(X210,bool),X211),s(X210,X213)))) )),
  inference(superposition,[],[f88,f1616])).
fof(f12438,plain,(
  ( ! [X208,X206,X204,X202,X207,X205,X203,X201] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X201),s(fun(X202,bool),X203),s(X202,X205),s(X202,X204)))) | p(s(bool,X206)) | p(s(bool,X207)) | s(bool,X206) != s(bool,X208) | s(bool,X207) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X201),s(bool,'T_0'))),s(bool,happ(s(fun(X202,bool),X203),s(X202,X204))))) | s(bool,'F_0') = s(bool,happ(s(fun(X202,bool),X203),s(X202,X205)))) )),
  inference(superposition,[],[f87,f1616])).
fof(f12436,plain,(
  ( ! [X191,X189,X187,X185,X192,X190,X188,X186] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X185),s(fun(X186,bool),X187),s(X186,X189),s(X186,X188)))) | p(s(bool,X190)) | ~p(s(bool,X191)) | s(bool,X190) != s(bool,X192) | s(bool,X191) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X185),s(bool,'T_0'))),s(bool,happ(s(fun(X186,bool),X187),s(X186,X188))))) | s(bool,'F_0') = s(bool,happ(s(fun(X186,bool),X187),s(X186,X189)))) )),
  inference(superposition,[],[f87,f1616])).
fof(f12433,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162,X168] : (s(bool,X166) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X161),s(fun(X162,bool),X163),s(X162,X165),s(X162,X164))) | p(s(bool,X167)) | ~p(s(bool,X166)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X161),s(bool,'T_0'))),s(bool,happ(s(fun(X162,bool),X163),s(X162,X164)))))) | s(bool,X167) != s(bool,X168) | s(bool,'F_0') = s(bool,happ(s(fun(X162,bool),X163),s(X162,X165)))) )),
  inference(superposition,[],[f87,f1616])).
fof(f12432,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X153,X160] : (s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X153),s(fun(X154,bool),X155),s(X154,X157),s(X154,X156))) | p(s(bool,X159)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X153),s(bool,'T_0'))),s(bool,happ(s(fun(X154,bool),X155),s(X154,X156)))))) | p(s(bool,X158)) | s(bool,X159) != s(bool,X160) | s(bool,'F_0') = s(bool,happ(s(fun(X154,bool),X155),s(X154,X157)))) )),
  inference(superposition,[],[f87,f1616])).
fof(f12431,plain,(
  ( ! [X146,X152,X151,X149,X147,X145,X150,X148] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X145),s(fun(X146,bool),X147),s(X146,X149),s(X146,X148)))) | p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X145),s(bool,'T_0'))),s(bool,happ(s(fun(X146,bool),X147),s(X146,X148))))) | s(bool,X151) != s(bool,X152) | s(bool,'F_0') = s(bool,happ(s(fun(X146,bool),X147),s(X146,X149)))) )),
  inference(superposition,[],[f85,f1616])).
fof(f12429,plain,(
  ( ! [X134,X132,X130,X136,X135,X133,X131,X129] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X129),s(fun(X130,bool),X131),s(X130,X133),s(X130,X132)))) | ~p(s(bool,X134)) | ~p(s(bool,X135)) | s(bool,X135) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X129),s(bool,'T_0'))),s(bool,happ(s(fun(X130,bool),X131),s(X130,X132))))) | s(bool,X134) != s(bool,X136) | s(bool,'F_0') = s(bool,happ(s(fun(X130,bool),X131),s(X130,X133)))) )),
  inference(superposition,[],[f85,f1616])).
fof(f12428,plain,(
  ( ! [X127,X125,X123,X121,X128,X126,X124,X122] : (s(bool,X126) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X121),s(fun(X122,bool),X123),s(X122,X125),s(X122,X124))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X121),s(bool,'T_0'))),s(bool,happ(s(fun(X122,bool),X123),s(X122,X124)))))) | ~p(s(bool,X127)) | ~p(s(bool,X126)) | s(bool,X127) != s(bool,X128) | s(bool,'F_0') = s(bool,happ(s(fun(X122,bool),X123),s(X122,X125)))) )),
  inference(superposition,[],[f85,f1616])).
fof(f12427,plain,(
  ( ! [X118,X116,X114,X120,X119,X117,X115,X113] : (s(bool,X118) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X113),s(fun(X114,bool),X115),s(X114,X117),s(X114,X116))) | p(s(bool,X118)) | ~p(s(bool,X119)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X113),s(bool,'T_0'))),s(bool,happ(s(fun(X114,bool),X115),s(X114,X116)))))) | s(bool,X119) != s(bool,X120) | s(bool,'F_0') = s(bool,happ(s(fun(X114,bool),X115),s(X114,X117)))) )),
  inference(superposition,[],[f85,f1616])).
fof(f12424,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89,X96] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X89),s(fun(X90,bool),X91),s(X90,X93),s(X90,X92)))) | p(s(bool,X94)) | p(s(bool,X95)) | s(bool,X94) != s(bool,X96) | s(bool,X95) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X89),s(bool,'T_0'))),s(bool,happ(s(fun(X90,bool),X91),s(X90,X92))))) | s(bool,'F_0') = s(bool,happ(s(fun(X90,bool),X91),s(X90,X93)))) )),
  inference(superposition,[],[f83,f1616])).
fof(f12422,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75,X73] : (p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X73),s(fun(X74,bool),X75),s(X74,X77),s(X74,X76)))) | p(s(bool,X78)) | ~p(s(bool,X79)) | s(bool,X78) != s(bool,X80) | s(bool,X79) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X73),s(bool,'T_0'))),s(bool,happ(s(fun(X74,bool),X75),s(X74,X76))))) | s(bool,'F_0') = s(bool,happ(s(fun(X74,bool),X75),s(X74,X77)))) )),
  inference(superposition,[],[f83,f1616])).
fof(f12419,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (s(bool,X54) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X49),s(fun(X50,bool),X51),s(X50,X53),s(X50,X52))) | p(s(bool,X55)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X49),s(bool,'T_0'))),s(bool,happ(s(fun(X50,bool),X51),s(X50,X52)))))) | p(s(bool,X54)) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,happ(s(fun(X50,bool),X51),s(X50,X53)))) )),
  inference(superposition,[],[f83,f1616])).
fof(f12418,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42] : (s(bool,X46) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X41),s(fun(X42,bool),X43),s(X42,X45),s(X42,X44))) | p(s(bool,X47)) | ~p(s(bool,X46)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X41),s(bool,'T_0'))),s(bool,happ(s(fun(X42,bool),X43),s(X42,X44)))))) | s(bool,X47) != s(bool,X48) | s(bool,'F_0') = s(bool,happ(s(fun(X42,bool),X43),s(X42,X45)))) )),
  inference(superposition,[],[f83,f1616])).
fof(f12276,plain,(
  ( ! [X335,X341,X337,X339,X340,X336,X338] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X335),s(bool,'T_0'))),s(bool,happ(s(fun(X336,bool),X337),s(X336,X339)))))) | p(s(bool,X340)) | p(s(bool,X341)) | s(bool,X341) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X335),s(fun(X336,bool),X337),s(X336,X338),s(X336,X339))) | s(bool,'F_0') = s(bool,happ(s(fun(X336,bool),X337),s(X336,X338)))) )),
  inference(superposition,[],[f125,f1616])).
fof(f12274,plain,(
  ( ! [X325,X327,X321,X323,X324,X326,X322] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X321),s(bool,'T_0'))),s(bool,happ(s(fun(X322,bool),X323),s(X322,X325)))))) | p(s(bool,X326)) | ~p(s(bool,X327)) | s(bool,X327) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X321),s(fun(X322,bool),X323),s(X322,X324),s(X322,X325))) | s(bool,'F_0') = s(bool,happ(s(fun(X322,bool),X323),s(X322,X324)))) )),
  inference(superposition,[],[f125,f1616])).
fof(f12273,plain,(
  ( ! [X315,X316,X318,X314,X320,X317,X319] : (s(bool,X319) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X314),s(bool,'T_0'))),s(bool,happ(s(fun(X315,bool),X316),s(X315,X318))))) | p(s(bool,X320)) | ~p(s(bool,X319)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X314),s(fun(X315,bool),X316),s(X315,X317),s(X315,X318)))) | s(bool,'F_0') = s(bool,happ(s(fun(X315,bool),X316),s(X315,X317)))) )),
  inference(superposition,[],[f125,f1616])).
fof(f12272,plain,(
  ( ! [X313,X308,X310,X312,X309,X311,X307] : (s(bool,X312) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X307),s(bool,'T_0'))),s(bool,happ(s(fun(X308,bool),X309),s(X308,X311))))) | p(s(bool,X313)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X307),s(fun(X308,bool),X309),s(X308,X310),s(X308,X311)))) | p(s(bool,X312)) | s(bool,'F_0') = s(bool,happ(s(fun(X308,bool),X309),s(X308,X310)))) )),
  inference(superposition,[],[f125,f1616])).
fof(f12264,plain,(
  ( ! [X269,X271,X265,X267,X268,X270,X264,X266] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X264),s(bool,'T_0'))),s(bool,happ(s(fun(X265,bool),X266),s(X265,X268)))))) | p(s(bool,X269)) | ~p(s(bool,X270)) | s(bool,X269) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X264),s(fun(X265,bool),X266),s(X265,X267),s(X265,X268))) | s(bool,X270) != s(bool,X271) | s(bool,'F_0') = s(bool,happ(s(fun(X265,bool),X266),s(X265,X267)))) )),
  inference(superposition,[],[f89,f1616])).
fof(f12262,plain,(
  ( ! [X255,X253,X251,X249,X254,X252,X250,X248] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X248),s(bool,'T_0'))),s(bool,happ(s(fun(X249,bool),X250),s(X249,X252)))))) | ~p(s(bool,X253)) | ~p(s(bool,X254)) | s(bool,X254) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X248),s(fun(X249,bool),X250),s(X249,X251),s(X249,X252))) | s(bool,X253) != s(bool,X255) | s(bool,'F_0') = s(bool,happ(s(fun(X249,bool),X250),s(X249,X251)))) )),
  inference(superposition,[],[f89,f1616])).
fof(f12261,plain,(
  ( ! [X243,X241,X246,X244,X242,X240,X247,X245] : (s(bool,X245) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X240),s(bool,'T_0'))),s(bool,happ(s(fun(X241,bool),X242),s(X241,X244))))) | p(s(bool,X245)) | ~p(s(bool,X246)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X240),s(fun(X241,bool),X242),s(X241,X243),s(X241,X244)))) | s(bool,X246) != s(bool,X247) | s(bool,'F_0') = s(bool,happ(s(fun(X241,bool),X242),s(X241,X243)))) )),
  inference(superposition,[],[f89,f1616])).
fof(f12260,plain,(
  ( ! [X239,X237,X235,X233,X238,X236,X234,X232] : (s(bool,X237) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X232),s(bool,'T_0'))),s(bool,happ(s(fun(X233,bool),X234),s(X233,X236))))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X232),s(fun(X233,bool),X234),s(X233,X235),s(X233,X236)))) | ~p(s(bool,X238)) | ~p(s(bool,X237)) | s(bool,X238) != s(bool,X239) | s(bool,'F_0') = s(bool,happ(s(fun(X233,bool),X234),s(X233,X235)))) )),
  inference(superposition,[],[f89,f1616])).
fof(f12256,plain,(
  ( ! [X206,X204,X202,X200,X207,X205,X203,X201] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X200),s(bool,'T_0'))),s(bool,happ(s(fun(X201,bool),X202),s(X201,X204)))))) | p(s(bool,X205)) | ~p(s(bool,X206)) | p(s(bool,X207)) | s(bool,X207) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X200),s(fun(X201,bool),X202),s(X201,X203),s(X201,X204))) | s(bool,'F_0') = s(bool,happ(s(fun(X201,bool),X202),s(X201,X203)))) )),
  inference(superposition,[],[f88,f1616])).
fof(f12254,plain,(
  ( ! [X191,X189,X187,X185,X190,X188,X186,X184] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X184),s(bool,'T_0'))),s(bool,happ(s(fun(X185,bool),X186),s(X185,X188)))))) | p(s(bool,X189)) | ~p(s(bool,X190)) | ~p(s(bool,X191)) | s(bool,X190) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X184),s(fun(X185,bool),X186),s(X185,X187),s(X185,X188))) | s(bool,'F_0') = s(bool,happ(s(fun(X185,bool),X186),s(X185,X187)))) )),
  inference(superposition,[],[f88,f1616])).
fof(f12253,plain,(
  ( ! [X177,X182,X180,X178,X176,X183,X181,X179] : (s(bool,X181) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X176),s(bool,'T_0'))),s(bool,happ(s(fun(X177,bool),X178),s(X177,X180))))) | p(s(bool,X182)) | ~p(s(bool,X181)) | ~p(s(bool,X183)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X176),s(fun(X177,bool),X178),s(X177,X179),s(X177,X180)))) | s(bool,'F_0') = s(bool,happ(s(fun(X177,bool),X178),s(X177,X179)))) )),
  inference(superposition,[],[f88,f1616])).
fof(f12252,plain,(
  ( ! [X175,X173,X171,X169,X174,X172,X170,X168] : (s(bool,X173) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X168),s(bool,'T_0'))),s(bool,happ(s(fun(X169,bool),X170),s(X169,X172))))) | p(s(bool,X174)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X168),s(fun(X169,bool),X170),s(X169,X171),s(X169,X172)))) | ~p(s(bool,X175)) | p(s(bool,X173)) | s(bool,'F_0') = s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))) )),
  inference(superposition,[],[f88,f1616])).
fof(f12251,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162,X160] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X160),s(bool,'T_0'))),s(bool,happ(s(fun(X161,bool),X162),s(X161,X164)))))) | p(s(bool,X165)) | p(s(bool,X166)) | s(bool,X165) != s(bool,X167) | s(bool,X166) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X160),s(fun(X161,bool),X162),s(X161,X163),s(X161,X164))) | s(bool,'F_0') = s(bool,happ(s(fun(X161,bool),X162),s(X161,X163)))) )),
  inference(superposition,[],[f87,f1616])).
fof(f12249,plain,(
  ( ! [X146,X144,X151,X149,X147,X145,X150,X148] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X144),s(bool,'T_0'))),s(bool,happ(s(fun(X145,bool),X146),s(X145,X148)))))) | p(s(bool,X149)) | ~p(s(bool,X150)) | s(bool,X149) != s(bool,X151) | s(bool,X150) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X144),s(fun(X145,bool),X146),s(X145,X147),s(X145,X148))) | s(bool,'F_0') = s(bool,happ(s(fun(X145,bool),X146),s(X145,X147)))) )),
  inference(superposition,[],[f87,f1616])).
fof(f12246,plain,(
  ( ! [X127,X125,X123,X121,X126,X124,X122,X120] : (s(bool,X125) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X120),s(bool,'T_0'))),s(bool,happ(s(fun(X121,bool),X122),s(X121,X124))))) | p(s(bool,X126)) | ~p(s(bool,X125)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X120),s(fun(X121,bool),X122),s(X121,X123),s(X121,X124)))) | s(bool,X126) != s(bool,X127) | s(bool,'F_0') = s(bool,happ(s(fun(X121,bool),X122),s(X121,X123)))) )),
  inference(superposition,[],[f87,f1616])).
fof(f12245,plain,(
  ( ! [X118,X116,X114,X112,X119,X117,X115,X113] : (s(bool,X117) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X112),s(bool,'T_0'))),s(bool,happ(s(fun(X113,bool),X114),s(X113,X116))))) | p(s(bool,X118)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X112),s(fun(X113,bool),X114),s(X113,X115),s(X113,X116)))) | p(s(bool,X117)) | s(bool,X118) != s(bool,X119) | s(bool,'F_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,X115)))) )),
  inference(superposition,[],[f87,f1616])).
fof(f12244,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X104),s(bool,'T_0'))),s(bool,happ(s(fun(X105,bool),X106),s(X105,X108)))))) | p(s(bool,X109)) | ~p(s(bool,X110)) | s(bool,X109) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X104),s(fun(X105,bool),X106),s(X105,X107),s(X105,X108))) | s(bool,X110) != s(bool,X111) | s(bool,'F_0') = s(bool,happ(s(fun(X105,bool),X106),s(X105,X107)))) )),
  inference(superposition,[],[f85,f1616])).
fof(f12242,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X88),s(bool,'T_0'))),s(bool,happ(s(fun(X89,bool),X90),s(X89,X92)))))) | ~p(s(bool,X93)) | ~p(s(bool,X94)) | s(bool,X94) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X88),s(fun(X89,bool),X90),s(X89,X91),s(X89,X92))) | s(bool,X93) != s(bool,X95) | s(bool,'F_0') = s(bool,happ(s(fun(X89,bool),X90),s(X89,X91)))) )),
  inference(superposition,[],[f85,f1616])).
fof(f12241,plain,(
  ( ! [X80,X87,X85,X83,X81,X86,X84,X82] : (s(bool,X85) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X80),s(bool,'T_0'))),s(bool,happ(s(fun(X81,bool),X82),s(X81,X84))))) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X80),s(fun(X81,bool),X82),s(X81,X83),s(X81,X84)))) | ~p(s(bool,X86)) | ~p(s(bool,X85)) | s(bool,X86) != s(bool,X87) | s(bool,'F_0') = s(bool,happ(s(fun(X81,bool),X82),s(X81,X83)))) )),
  inference(superposition,[],[f85,f1616])).
fof(f12240,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (s(bool,X77) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X72),s(bool,'T_0'))),s(bool,happ(s(fun(X73,bool),X74),s(X73,X76))))) | p(s(bool,X77)) | ~p(s(bool,X78)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X72),s(fun(X73,bool),X74),s(X73,X75),s(X73,X76)))) | s(bool,X78) != s(bool,X79) | s(bool,'F_0') = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75)))) )),
  inference(superposition,[],[f85,f1616])).
fof(f12237,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X48),s(bool,'T_0'))),s(bool,happ(s(fun(X49,bool),X50),s(X49,X52)))))) | p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X48),s(fun(X49,bool),X50),s(X49,X51),s(X49,X52))) | s(bool,'F_0') = s(bool,happ(s(fun(X49,bool),X50),s(X49,X51)))) )),
  inference(superposition,[],[f83,f1616])).
fof(f12235,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X32),s(bool,'T_0'))),s(bool,happ(s(fun(X33,bool),X34),s(X33,X36)))))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X32),s(fun(X33,bool),X34),s(X33,X35),s(X33,X36))) | s(bool,'F_0') = s(bool,happ(s(fun(X33,bool),X34),s(X33,X35)))) )),
  inference(superposition,[],[f83,f1616])).
fof(f12232,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(bool,X13) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X8),s(bool,'T_0'))),s(bool,happ(s(fun(X9,bool),X10),s(X9,X12))))) | p(s(bool,X14)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X8),s(fun(X9,bool),X10),s(X9,X11),s(X9,X12)))) | p(s(bool,X13)) | s(bool,X14) != s(bool,X15) | s(bool,'F_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) )),
  inference(superposition,[],[f83,f1616])).
fof(f12231,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X5) != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),X0),s(bool,'T_0'))),s(bool,happ(s(fun(X1,bool),X2),s(X1,X4))))) | p(s(bool,X6)) | ~p(s(bool,X5)) | p(s(bool,'const.relation.inv_image_4'(s(fun(bool,fun(bool,bool)),X0),s(fun(X1,bool),X2),s(X1,X3),s(X1,X4)))) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) )),
  inference(superposition,[],[f83,f1616])).
fof(f12137,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) | $spl232),
  inference(cnf_transformation,[],[f12137_D])).
fof(f12137_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f12073,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'F_0'),s(bool,X61))) | s(bool,'T_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) ) | $spl34),
  inference(resolution,[],[f2716,f125])).
fof(f2716,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X11),s(bool,X9)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,'F_0'),s(bool,X11))) | s(bool,'T_0') = s(bool,X9)) ) | $spl34),
  inference(superposition,[],[f2438,f122])).
fof(f12069,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'F_0'),s(bool,X46))) | s(bool,'T_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,X48) != s(bool,X50)) ) | $spl34),
  inference(resolution,[],[f2716,f89])).
fof(f12067,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'F_0'),s(bool,X34))) | s(bool,'T_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) ) | $spl34),
  inference(resolution,[],[f2716,f88])).
fof(f12065,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'F_0'),s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) ) | $spl34),
  inference(resolution,[],[f2716,f87])).
fof(f12064,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'F_0'),s(bool,X16))) | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,X18) != s(bool,X20)) ) | $spl34),
  inference(resolution,[],[f2716,f85])).
fof(f12062,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) ) | $spl34),
  inference(resolution,[],[f2716,f83])).
fof(f12049,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) | $spl230),
  inference(cnf_transformation,[],[f12049_D])).
fof(f12049_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f11985,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'T_0'),s(bool,X61))) | s(bool,'F_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) ) | $spl34),
  inference(resolution,[],[f2715,f125])).
fof(f2715,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X8),s(bool,X6)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,'T_0'),s(bool,X8))) | s(bool,'F_0') = s(bool,X6)) ) | $spl34),
  inference(superposition,[],[f2438,f122])).
fof(f11981,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'T_0'),s(bool,X46))) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,X48) != s(bool,X50)) ) | $spl34),
  inference(resolution,[],[f2715,f89])).
fof(f11979,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'T_0'),s(bool,X34))) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) ) | $spl34),
  inference(resolution,[],[f2715,f88])).
fof(f11977,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'T_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) ) | $spl34),
  inference(resolution,[],[f2715,f87])).
fof(f11976,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'T_0'),s(bool,X16))) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,X18) != s(bool,X20)) ) | $spl34),
  inference(resolution,[],[f2715,f85])).
fof(f11974,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4))) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) ) | $spl34),
  inference(resolution,[],[f2715,f83])).
fof(f11953,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'F_0'))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X4)) ) | $spl228),
  inference(cnf_transformation,[],[f11953_D])).
fof(f11953_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'F_0'))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f11909,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62))) | s(bool,'T_0') = s(bool,X61) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X62),s(bool,'F_0')))) ) | $spl34),
  inference(resolution,[],[f2714,f125])).
fof(f11905,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,'T_0') = s(bool,X46) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X47),s(bool,'F_0'))) | s(bool,X48) != s(bool,X50)) ) | $spl34),
  inference(resolution,[],[f2714,f89])).
fof(f11903,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35))) | s(bool,'T_0') = s(bool,X34) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X35),s(bool,'F_0')))) ) | $spl34),
  inference(resolution,[],[f2714,f88])).
fof(f11901,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23))) | s(bool,'T_0') = s(bool,X22) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X23),s(bool,'F_0')))) ) | $spl34),
  inference(resolution,[],[f2714,f87])).
fof(f11900,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,'T_0') = s(bool,X16) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X17),s(bool,'F_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl34),
  inference(resolution,[],[f2714,f85])).
fof(f11898,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,'T_0') = s(bool,X4) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'F_0')))) ) | $spl34),
  inference(resolution,[],[f2714,f83])).
fof(f11872,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'T_0'))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X4)) ) | $spl226),
  inference(cnf_transformation,[],[f11872_D])).
fof(f11872_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'T_0'))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X4)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f11829,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62))) | s(bool,'F_0') = s(bool,X61) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X62),s(bool,'T_0')))) ) | $spl34),
  inference(resolution,[],[f2713,f125])).
fof(f11825,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,'F_0') = s(bool,X46) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X47),s(bool,'T_0'))) | s(bool,X48) != s(bool,X50)) ) | $spl34),
  inference(resolution,[],[f2713,f89])).
fof(f11823,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35))) | s(bool,'F_0') = s(bool,X34) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X35),s(bool,'T_0')))) ) | $spl34),
  inference(resolution,[],[f2713,f88])).
fof(f11821,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23))) | s(bool,'F_0') = s(bool,X22) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X23),s(bool,'T_0')))) ) | $spl34),
  inference(resolution,[],[f2713,f87])).
fof(f11820,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,'F_0') = s(bool,X16) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X17),s(bool,'T_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl34),
  inference(resolution,[],[f2713,f85])).
fof(f11818,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,'F_0') = s(bool,X4) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'T_0')))) ) | $spl34),
  inference(resolution,[],[f2713,f83])).
fof(f11786,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0'))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) | $spl224),
  inference(cnf_transformation,[],[f11786_D])).
fof(f11786_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0'))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f11722,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X62),s(bool,X61)))) ) | $spl34),
  inference(resolution,[],[f2700,f125])).
fof(f2700,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X9),s(bool,X11)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X11),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl34),
  inference(superposition,[],[f2438,f122])).
fof(f11718,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X47),s(bool,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl34),
  inference(resolution,[],[f2700,f89])).
fof(f11716,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X35),s(bool,X34)))) ) | $spl34),
  inference(resolution,[],[f2700,f88])).
fof(f11714,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X23),s(bool,X22)))) ) | $spl34),
  inference(resolution,[],[f2700,f87])).
fof(f11713,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X17),s(bool,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl34),
  inference(resolution,[],[f2700,f85])).
fof(f11711,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl34),
  inference(resolution,[],[f2700,f83])).
fof(f11679,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0'))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) | $spl222),
  inference(cnf_transformation,[],[f11679_D])).
fof(f11679_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0'))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f11615,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X62),s(bool,X61)))) ) | $spl34),
  inference(resolution,[],[f2699,f125])).
fof(f2699,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X6),s(bool,X8)))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X8),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X6)) ) | $spl34),
  inference(superposition,[],[f2438,f122])).
fof(f11611,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X47),s(bool,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl34),
  inference(resolution,[],[f2699,f89])).
fof(f11609,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X35),s(bool,X34)))) ) | $spl34),
  inference(resolution,[],[f2699,f88])).
fof(f11607,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X23),s(bool,X22)))) ) | $spl34),
  inference(resolution,[],[f2699,f87])).
fof(f11606,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X17),s(bool,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl34),
  inference(resolution,[],[f2699,f85])).
fof(f11604,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl34),
  inference(resolution,[],[f2699,f83])).
fof(f11588,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) | $spl220),
  inference(cnf_transformation,[],[f11588_D])).
fof(f11588_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f11544,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62))) | s(bool,'T_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'F_0'),s(bool,X61)))) ) | $spl34),
  inference(resolution,[],[f2698,f125])).
fof(f11540,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,'T_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'F_0'),s(bool,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl34),
  inference(resolution,[],[f2698,f89])).
fof(f11538,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35))) | s(bool,'T_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'F_0'),s(bool,X34)))) ) | $spl34),
  inference(resolution,[],[f2698,f88])).
fof(f11536,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23))) | s(bool,'T_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'F_0'),s(bool,X22)))) ) | $spl34),
  inference(resolution,[],[f2698,f87])).
fof(f11535,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'F_0'),s(bool,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl34),
  inference(resolution,[],[f2698,f85])).
fof(f11533,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) ) | $spl34),
  inference(resolution,[],[f2698,f83])).
fof(f11522,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) | $spl218),
  inference(cnf_transformation,[],[f11522_D])).
fof(f11522_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f11479,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62))) | s(bool,'F_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'T_0'),s(bool,X61)))) ) | $spl34),
  inference(resolution,[],[f2697,f125])).
fof(f11475,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'T_0'),s(bool,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl34),
  inference(resolution,[],[f2697,f89])).
fof(f11473,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35))) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'T_0'),s(bool,X34)))) ) | $spl34),
  inference(resolution,[],[f2697,f88])).
fof(f11471,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'T_0'),s(bool,X22)))) ) | $spl34),
  inference(resolution,[],[f2697,f87])).
fof(f11470,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'T_0'),s(bool,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl34),
  inference(resolution,[],[f2697,f85])).
fof(f11468,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) ) | $spl34),
  inference(resolution,[],[f2697,f83])).
fof(f11452,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,X3) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,X3))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,'F_0'),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X3)) ) | $spl216),
  inference(cnf_transformation,[],[f11452_D])).
fof(f11452_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,X3) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,X3))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,'F_0'),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X3)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f11408,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,X60) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X61),s(bool,X62),s(bool,X60))) | s(bool,'T_0') = s(bool,X60) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X61),s(bool,'F_0'),s(bool,X62)))) ) | $spl32),
  inference(resolution,[],[f2647,f125])).
fof(f2647,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,'F_0'),s(bool,X11)))) | s(bool,X9) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X11),s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl32),
  inference(superposition,[],[f2435,f122])).
fof(f11404,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,X45) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X46),s(bool,X47),s(bool,X45))) | s(bool,'T_0') = s(bool,X45) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X46),s(bool,'F_0'),s(bool,X47))) | s(bool,X48) != s(bool,X50)) ) | $spl32),
  inference(resolution,[],[f2647,f89])).
fof(f11402,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,X33) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X34),s(bool,X35),s(bool,X33))) | s(bool,'T_0') = s(bool,X33) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X34),s(bool,'F_0'),s(bool,X35)))) ) | $spl32),
  inference(resolution,[],[f2647,f88])).
fof(f11400,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,X21) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X22),s(bool,X23),s(bool,X21))) | s(bool,'T_0') = s(bool,X21) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X22),s(bool,'F_0'),s(bool,X23)))) ) | $spl32),
  inference(resolution,[],[f2647,f87])).
fof(f11399,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,X15) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X16),s(bool,X17),s(bool,X15))) | s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X16),s(bool,'F_0'),s(bool,X17))) | s(bool,X18) != s(bool,X20)) ) | $spl32),
  inference(resolution,[],[f2647,f85])).
fof(f11397,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,X3) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,X3))) | s(bool,'T_0') = s(bool,X3) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,'F_0'),s(bool,X5)))) ) | $spl32),
  inference(resolution,[],[f2647,f83])).
fof(f11386,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) | $spl214),
  inference(cnf_transformation,[],[f11386_D])).
fof(f11386_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f11343,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62))) | s(bool,'F_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'T_0'),s(bool,X61)))) ) | $spl32),
  inference(resolution,[],[f2646,f125])).
fof(f2646,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,'T_0'),s(bool,X8)))) | s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X8),s(bool,X6))) | s(bool,'F_0') = s(bool,X6)) ) | $spl32),
  inference(superposition,[],[f2435,f122])).
fof(f11339,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'T_0'),s(bool,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl32),
  inference(resolution,[],[f2646,f89])).
fof(f11337,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35))) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'T_0'),s(bool,X34)))) ) | $spl32),
  inference(resolution,[],[f2646,f88])).
fof(f11335,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'T_0'),s(bool,X22)))) ) | $spl32),
  inference(resolution,[],[f2646,f87])).
fof(f11334,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'T_0'),s(bool,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl32),
  inference(resolution,[],[f2646,f85])).
fof(f11332,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) ) | $spl32),
  inference(resolution,[],[f2646,f83])).
fof(f11328,plain,(
  ( ! [X6,X10,X8,X7,X13,X11,X9] : (s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X10),s(X8,X11))) = s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X11),s(X8,X10))) | s(bool,X13) != s(bool,'const.relation.symmetric_1'(s(fun(X8,fun(X8,bool)),'const.relation.inv_image_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9))))) | ~p(s(bool,X13))) ) | $spl212),
  inference(cnf_transformation,[],[f11328_D])).
fof(f11328_D,plain,(
  ( ! [X6,X10,X8,X7,X13,X11,X9] : (s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X10),s(X8,X11))) = s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X11),s(X8,X10))) | s(bool,X13) != s(bool,'const.relation.symmetric_1'(s(fun(X8,fun(X8,bool)),'const.relation.inv_image_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9))))) | ~p(s(bool,X13))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f11308,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (s(bool,'const.relation.inv_image_4'(s(fun(X96,fun(X96,bool)),X97),s(fun(X98,X96),X99),s(X98,X100),s(X98,X101))) = s(bool,'const.relation.inv_image_4'(s(fun(X96,fun(X96,bool)),X97),s(fun(X98,X96),X99),s(X98,X101),s(X98,X100))) | p(s(bool,X102)) | ~p(s(bool,X103)) | s(bool,X103) != s(bool,'const.relation.symmetric_1'(s(fun(X98,fun(X98,bool)),'const.relation.inv_image_2'(s(fun(X96,fun(X96,bool)),X97),s(fun(X98,X96),X99)))))) )),
  inference(resolution,[],[f1483,f125])).
fof(f1483,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_image_2'(s(fun(X7,fun(X7,bool)),X8),s(fun(X6,X7),X9)))))) | s(bool,'const.relation.inv_image_4'(s(fun(X7,fun(X7,bool)),X8),s(fun(X6,X7),X9),s(X6,X10),s(X6,X11))) = s(bool,'const.relation.inv_image_4'(s(fun(X7,fun(X7,bool)),X8),s(fun(X6,X7),X9),s(X6,X11),s(X6,X10)))) )),
  inference(forward_demodulation,[],[f1408,f81])).
fof(f1408,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'const.relation.inv_image_4'(s(fun(X7,fun(X7,bool)),X8),s(fun(X6,X7),X9),s(X6,X10),s(X6,X11))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.inv_image_2'(s(fun(X7,fun(X7,bool)),X8),s(fun(X6,X7),X9))),s(X6,X11))),s(X6,X10))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),'const.relation.inv_image_2'(s(fun(X7,fun(X7,bool)),X8),s(fun(X6,X7),X9))))))) )),
  inference(superposition,[],[f81,f118])).
fof(f11304,plain,(
  ( ! [X70,X76,X74,X72,X71,X69,X77,X75,X73] : (s(bool,'const.relation.inv_image_4'(s(fun(X69,fun(X69,bool)),X70),s(fun(X71,X69),X72),s(X71,X73),s(X71,X74))) = s(bool,'const.relation.inv_image_4'(s(fun(X69,fun(X69,bool)),X70),s(fun(X71,X69),X72),s(X71,X74),s(X71,X73))) | ~p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X76) != s(bool,'const.relation.symmetric_1'(s(fun(X71,fun(X71,bool)),'const.relation.inv_image_2'(s(fun(X69,fun(X69,bool)),X70),s(fun(X71,X69),X72))))) | s(bool,X75) != s(bool,X77)) )),
  inference(resolution,[],[f1483,f89])).
fof(f11302,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53,X51] : (s(bool,'const.relation.inv_image_4'(s(fun(X51,fun(X51,bool)),X52),s(fun(X53,X51),X54),s(X53,X55),s(X53,X56))) = s(bool,'const.relation.inv_image_4'(s(fun(X51,fun(X51,bool)),X52),s(fun(X53,X51),X54),s(X53,X56),s(X53,X55))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,'const.relation.symmetric_1'(s(fun(X53,fun(X53,bool)),'const.relation.inv_image_2'(s(fun(X51,fun(X51,bool)),X52),s(fun(X53,X51),X54)))))) )),
  inference(resolution,[],[f1483,f88])).
fof(f11300,plain,(
  ( ! [X39,X37,X35,X33,X41,X38,X36,X34,X40] : (s(bool,'const.relation.inv_image_4'(s(fun(X33,fun(X33,bool)),X34),s(fun(X35,X33),X36),s(X35,X37),s(X35,X38))) = s(bool,'const.relation.inv_image_4'(s(fun(X33,fun(X33,bool)),X34),s(fun(X35,X33),X36),s(X35,X38),s(X35,X37))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.relation.symmetric_1'(s(fun(X35,fun(X35,bool)),'const.relation.inv_image_2'(s(fun(X33,fun(X33,bool)),X34),s(fun(X35,X33),X36)))))) )),
  inference(resolution,[],[f1483,f87])).
fof(f11299,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25,X32] : (s(bool,'const.relation.inv_image_4'(s(fun(X24,fun(X24,bool)),X25),s(fun(X26,X24),X27),s(X26,X28),s(X26,X29))) = s(bool,'const.relation.inv_image_4'(s(fun(X24,fun(X24,bool)),X25),s(fun(X26,X24),X27),s(X26,X29),s(X26,X28))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.relation.symmetric_1'(s(fun(X26,fun(X26,bool)),'const.relation.inv_image_2'(s(fun(X24,fun(X24,bool)),X25),s(fun(X26,X24),X27))))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f1483,f85])).
fof(f11297,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X13,X11,X9] : (s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X10),s(X8,X11))) = s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X11),s(X8,X10))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.relation.symmetric_1'(s(fun(X8,fun(X8,bool)),'const.relation.inv_image_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9)))))) )),
  inference(resolution,[],[f1483,f83])).
fof(f11166,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,X3) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X3),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,'F_0'))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X3)) ) | $spl210),
  inference(cnf_transformation,[],[f11166_D])).
fof(f11166_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,X3) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X3),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,'F_0'))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X3)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f11122,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,X60) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X61),s(bool,X60),s(bool,X62))) | s(bool,'T_0') = s(bool,X60) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X61),s(bool,X62),s(bool,'F_0')))) ) | $spl32),
  inference(resolution,[],[f2631,f125])).
fof(f2631,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X11),s(bool,'F_0')))) | s(bool,X9) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X9),s(bool,X11))) | s(bool,'T_0') = s(bool,X9)) ) | $spl32),
  inference(superposition,[],[f2435,f122])).
fof(f11118,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,X45) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X46),s(bool,X45),s(bool,X47))) | s(bool,'T_0') = s(bool,X45) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X46),s(bool,X47),s(bool,'F_0'))) | s(bool,X48) != s(bool,X50)) ) | $spl32),
  inference(resolution,[],[f2631,f89])).
fof(f11116,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,X33) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X34),s(bool,X33),s(bool,X35))) | s(bool,'T_0') = s(bool,X33) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X34),s(bool,X35),s(bool,'F_0')))) ) | $spl32),
  inference(resolution,[],[f2631,f88])).
fof(f11114,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,X21) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X22),s(bool,X21),s(bool,X23))) | s(bool,'T_0') = s(bool,X21) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X22),s(bool,X23),s(bool,'F_0')))) ) | $spl32),
  inference(resolution,[],[f2631,f87])).
fof(f11113,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,X15) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X16),s(bool,X15),s(bool,X17))) | s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X16),s(bool,X17),s(bool,'F_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl32),
  inference(resolution,[],[f2631,f85])).
fof(f11111,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,X3) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X3),s(bool,X5))) | s(bool,'T_0') = s(bool,X3) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,'F_0')))) ) | $spl32),
  inference(resolution,[],[f2631,f83])).
fof(f11085,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'T_0'))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X4)) ) | $spl208),
  inference(cnf_transformation,[],[f11085_D])).
fof(f11085_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'T_0'))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X4)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f11042,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62))) | s(bool,'F_0') = s(bool,X61) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X62),s(bool,'T_0')))) ) | $spl32),
  inference(resolution,[],[f2630,f125])).
fof(f2630,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X8),s(bool,'T_0')))) | s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X6),s(bool,X8))) | s(bool,'F_0') = s(bool,X6)) ) | $spl32),
  inference(superposition,[],[f2435,f122])).
fof(f11038,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,'F_0') = s(bool,X46) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X47),s(bool,'T_0'))) | s(bool,X48) != s(bool,X50)) ) | $spl32),
  inference(resolution,[],[f2630,f89])).
fof(f11036,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35))) | s(bool,'F_0') = s(bool,X34) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X35),s(bool,'T_0')))) ) | $spl32),
  inference(resolution,[],[f2630,f88])).
fof(f11034,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23))) | s(bool,'F_0') = s(bool,X22) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X23),s(bool,'T_0')))) ) | $spl32),
  inference(resolution,[],[f2630,f87])).
fof(f11033,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,'F_0') = s(bool,X16) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X17),s(bool,'T_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl32),
  inference(resolution,[],[f2630,f85])).
fof(f11031,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,'F_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | s(bool,'F_0') = s(bool,X4) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'T_0')))) ) | $spl32),
  inference(resolution,[],[f2630,f83])).
fof(f10239,plain,(
  ( ! [X6,X8,X7,X5,X11,X9] : (s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X8),s(X6,X9))) | s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X9),s(X6,X8))) | ~p(s(bool,X11)) | s(bool,'T_0') = s(bool,X5)) ) | $spl206),
  inference(cnf_transformation,[],[f10239_D])).
fof(f10239_D,plain,(
  ( ! [X6,X8,X7,X5,X11,X9] : (s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X8),s(X6,X9))) | s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X9),s(X6,X8))) | ~p(s(bool,X11)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f10172,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,X84) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X85,sK1),X86),s(X85,X87),s(X85,X88))) | s(bool,'T_0') = s(bool,X84) | p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X90) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X85,sK1),X86),s(X85,X88),s(X85,X87)))) ) | $spl6),
  inference(resolution,[],[f2437,f125])).
fof(f2437,plain,(
  ( ! [X333,X335,X331,X332,X334] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X331,sK1),X332),s(X331,X333),s(X331,X334)))) | s(bool,X335) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X331,sK1),X332),s(X331,X334),s(X331,X333))) | s(bool,'T_0') = s(bool,X335)) ) | $spl6),
  inference(superposition,[],[f497,f2341])).
fof(f10168,plain,(
  ( ! [X61,X68,X66,X64,X62,X67,X65,X63] : (s(bool,X61) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X62,sK1),X63),s(X62,X64),s(X62,X65))) | s(bool,'T_0') = s(bool,X61) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X67) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X62,sK1),X63),s(X62,X65),s(X62,X64))) | s(bool,X66) != s(bool,X68)) ) | $spl6),
  inference(resolution,[],[f2437,f89])).
fof(f10166,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X51,X49] : (s(bool,X45) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X46,sK1),X47),s(X46,X48),s(X46,X49))) | s(bool,'T_0') = s(bool,X45) | p(s(bool,X50)) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X46,sK1),X47),s(X46,X49),s(X46,X48)))) ) | $spl6),
  inference(resolution,[],[f2437,f88])).
fof(f10164,plain,(
  ( ! [X30,X35,X33,X31,X29,X36,X34,X32] : (s(bool,X29) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X30,sK1),X31),s(X30,X32),s(X30,X33))) | s(bool,'T_0') = s(bool,X29) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X30,sK1),X31),s(X30,X33),s(X30,X32)))) ) | $spl6),
  inference(resolution,[],[f2437,f87])).
fof(f10163,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X22,sK1),X23),s(X22,X24),s(X22,X25))) | s(bool,'T_0') = s(bool,X21) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X22,sK1),X23),s(X22,X25),s(X22,X24))) | s(bool,X26) != s(bool,X28)) ) | $spl6),
  inference(resolution,[],[f2437,f85])).
fof(f10161,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X11,X9] : (s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X8),s(X6,X9))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X9),s(X6,X8)))) ) | $spl6),
  inference(resolution,[],[f2437,f83])).
fof(f10150,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X3),s(X1,X4))) | s(bool,X0) != s(bool,X6) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X4),s(X1,X3)))) ) | $spl204),
  inference(cnf_transformation,[],[f10150_D])).
fof(f10150_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X3),s(X1,X4))) | s(bool,X0) != s(bool,X6) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X4),s(X1,X3)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f10107,plain,(
  ( ! [X80,X85,X83,X81,X79,X84,X82] : (s(bool,X79) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X80,sK1),X81),s(X80,X82),s(X80,X83))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X80,sK1),X81),s(X80,X83),s(X80,X82))) | p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X79) != s(bool,X85)) ) | $spl6),
  inference(resolution,[],[f2436,f125])).
fof(f2436,plain,(
  ( ! [X327,X329,X326,X328,X330] : (~p(s(bool,X330)) | s(bool,X330) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X326,sK1),X327),s(X326,X329),s(X326,X328))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X326,sK1),X327),s(X326,X328),s(X326,X329)))) ) | $spl6),
  inference(superposition,[],[f497,f2341])).
fof(f10103,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X57,sK1),X58),s(X57,X59),s(X57,X60))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X57,sK1),X58),s(X57,X60),s(X57,X59))) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X56) != s(bool,X62) | s(bool,X61) != s(bool,X63)) ) | $spl6),
  inference(resolution,[],[f2436,f89])).
fof(f10101,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,X40) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X41,sK1),X42),s(X41,X43),s(X41,X44))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X41,sK1),X42),s(X41,X44),s(X41,X43))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X40) != s(bool,X46)) ) | $spl6),
  inference(resolution,[],[f2436,f88])).
fof(f10099,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X25,sK1),X26),s(X25,X27),s(X25,X28))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X25,sK1),X26),s(X25,X28),s(X25,X27))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X24) != s(bool,X30)) ) | $spl6),
  inference(resolution,[],[f2436,f87])).
fof(f10098,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X17,sK1),X18),s(X17,X19),s(X17,X20))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X17,sK1),X18),s(X17,X20),s(X17,X19))) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X16) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl6),
  inference(resolution,[],[f2436,f85])).
fof(f10096,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X3),s(X1,X4))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X1,sK1),X2),s(X1,X4),s(X1,X3))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X0) != s(bool,X6)) ) | $spl6),
  inference(resolution,[],[f2436,f83])).
fof(f10093,plain,(
  ( ! [X14,X15,X18,X16] : (s(bool,X14) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X15))),s(sK1,X16))) | s(bool,X18) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X16))),s(sK1,X15))) | p(s(bool,X18)) | s(bool,'F_0') = s(bool,X14)) ) | $spl202),
  inference(cnf_transformation,[],[f10093_D])).
fof(f10093_D,plain,(
  ( ! [X14,X15,X18,X16] : (s(bool,X14) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X15))),s(sK1,X16))) | s(bool,X18) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X16))),s(sK1,X15))) | p(s(bool,X18)) | s(bool,'F_0') = s(bool,X14)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f10075,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,X68) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X69))),s(sK1,X70))) | s(bool,'F_0') = s(bool,X68) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X72) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X70))),s(sK1,X69)))) ) | $spl6),
  inference(resolution,[],[f2018,f125])).
fof(f2018,plain,(
  ( ! [X237,X238,X236] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X237))),s(sK1,X236)))) | s(bool,X238) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X236))),s(sK1,X237))) | s(bool,'F_0') = s(bool,X238)) ) | $spl6),
  inference(superposition,[],[f504,f127])).
fof(f504,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f486,f126])).
fof(f486,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f147,f122])).
fof(f10072,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X57))),s(sK1,X58))) | s(bool,'F_0') = s(bool,X56) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X58))),s(sK1,X57))) | s(bool,X60) != s(bool,X61)) ) | $spl6),
  inference(resolution,[],[f2018,f89])).
fof(f10069,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (s(bool,X38) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X39))),s(sK1,X40))) | s(bool,'F_0') = s(bool,X38) | p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X43) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X40))),s(sK1,X39)))) ) | $spl6),
  inference(resolution,[],[f2018,f88])).
fof(f10068,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (s(bool,X32) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X33))),s(sK1,X34))) | s(bool,'F_0') = s(bool,X32) | p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X34))),s(sK1,X33)))) ) | $spl6),
  inference(resolution,[],[f2018,f87])).
fof(f10067,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(bool,X26) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X27))),s(sK1,X28))) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X28))),s(sK1,X27))) | s(bool,X30) != s(bool,X31)) ) | $spl6),
  inference(resolution,[],[f2018,f85])).
fof(f10065,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(bool,X14) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X15))),s(sK1,X16))) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X16))),s(sK1,X15)))) ) | $spl6),
  inference(resolution,[],[f2018,f83])).
fof(f10057,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,X5))) | s(bool,X7) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X3)) ) | $spl200),
  inference(cnf_transformation,[],[f10057_D])).
fof(f10057_D,plain,(
  ( ! [X4,X7,X5,X3] : (s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,X5))) | s(bool,X7) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X3)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f10033,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,X60) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X61))),s(sK1,X62))) | s(bool,'T_0') = s(bool,X60) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X62))),s(sK1,X61)))) ) | $spl6),
  inference(resolution,[],[f2016,f125])).
fof(f2016,plain,(
  ( ! [X231,X233,X232] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X231))),s(sK1,X232)))) | s(bool,X233) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X232))),s(sK1,X231))) | s(bool,'T_0') = s(bool,X233)) ) | $spl6),
  inference(superposition,[],[f497,f127])).
fof(f10029,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,X45) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X46))),s(sK1,X47))) | s(bool,'T_0') = s(bool,X45) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X47))),s(sK1,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f2016,f89])).
fof(f10027,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,X33) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X34))),s(sK1,X35))) | s(bool,'T_0') = s(bool,X33) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X35))),s(sK1,X34)))) ) | $spl6),
  inference(resolution,[],[f2016,f88])).
fof(f10025,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,X21) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,X23))) | s(bool,'T_0') = s(bool,X21) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X23))),s(sK1,X22)))) ) | $spl6),
  inference(resolution,[],[f2016,f87])).
fof(f10024,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,X15) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X16))),s(sK1,X17))) | s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X17))),s(sK1,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f2016,f85])).
fof(f10022,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,X5))) | s(bool,'T_0') = s(bool,X3) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,X4)))) ) | $spl6),
  inference(resolution,[],[f2016,f83])).
fof(f9936,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X23) != s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,X26) != s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X25) != s(bool,X26) | s(bool,X22) != s(bool,X23)) ) | $spl198),
  inference(cnf_transformation,[],[f9936_D])).
fof(f9936_D,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X23) != s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,X26) != s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X25) != s(bool,X26) | s(bool,X22) != s(bool,X23)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f9935,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,'F_0') != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) != s(bool,X27) | s(bool,X23) != s(bool,X24)) ) | ($spl24 | $spl76)),
  inference(subsumption_resolution,[],[f9914,f385])).
fof(f9914,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,'F_0') != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) != s(bool,X27) | s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X21)) ) | $spl76),
  inference(superposition,[],[f2219,f122])).
fof(f9933,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X8) | p(s(bool,X8)) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,'F_0') = s(bool,X6) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) | $spl196),
  inference(cnf_transformation,[],[f9933_D])).
fof(f9933_D,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X8) | p(s(bool,X8)) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,'F_0') = s(bool,X6) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f9879,plain,(
  ( ! [X94,X101,X99,X97,X95,X102,X100,X98,X96] : (s(bool,X94) != s(bool,X95) | s(bool,X94) != s(bool,X96) | s(bool,X97) != s(bool,X98) | s(bool,X98) != s(bool,X99) | s(bool,'F_0') = s(bool,X100) | s(bool,X99) != s(bool,X100) | s(bool,X96) != s(bool,X97) | p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X95) != s(bool,X102)) ) | $spl76),
  inference(resolution,[],[f2219,f125])).
fof(f9876,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X79,X77,X75,X73] : (s(bool,X70) != s(bool,X71) | s(bool,X70) != s(bool,X72) | s(bool,X73) != s(bool,X74) | s(bool,X74) != s(bool,X75) | s(bool,'F_0') = s(bool,X76) | s(bool,X75) != s(bool,X76) | s(bool,X72) != s(bool,X73) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X71) != s(bool,X77) | s(bool,X78) != s(bool,X79)) ) | $spl76),
  inference(resolution,[],[f2219,f89])).
fof(f9873,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X40,X49] : (s(bool,X40) != s(bool,X41) | s(bool,X40) != s(bool,X42) | s(bool,X43) != s(bool,X44) | s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X46) | s(bool,X45) != s(bool,X46) | s(bool,X42) != s(bool,X43) | p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X41) != s(bool,X49)) ) | $spl76),
  inference(resolution,[],[f2219,f88])).
fof(f9872,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X30) != s(bool,X32) | s(bool,X33) != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X36) | s(bool,X35) != s(bool,X36) | s(bool,X32) != s(bool,X33) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X31) != s(bool,X38)) ) | $spl76),
  inference(resolution,[],[f2219,f87])).
fof(f9871,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X20) != s(bool,X22) | s(bool,X23) != s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,X26) | s(bool,X25) != s(bool,X26) | s(bool,X22) != s(bool,X23) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X21) != s(bool,X27) | s(bool,X28) != s(bool,X29)) ) | $spl76),
  inference(resolution,[],[f2219,f85])).
fof(f9869,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X6) | s(bool,X5) != s(bool,X6) | s(bool,X2) != s(bool,X3) | p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X1) != s(bool,X8)) ) | $spl76),
  inference(resolution,[],[f2219,f83])).
fof(f9857,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) | $spl194),
  inference(cnf_transformation,[],[f9857_D])).
fof(f9857_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f9794,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'F_0'),s(bool,X61)))) | s(bool,'T_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X62),s(bool,X61)))) ) | $spl6),
  inference(resolution,[],[f8468,f125])).
fof(f8468,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X8),s(bool,X7)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,'F_0'),s(bool,X7)))) | s(bool,'T_0') = s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f7580,f2341])).
fof(f7580,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X11),s(bool,X9)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,'F_0'),s(bool,X11)))) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f7530,f122])).
fof(f7530,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X16,sK1),X17),s(X16,X18),s(X16,X19)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X16,sK1),X17),s(X16,X19),s(X16,X18))))) ) | $spl6),
  inference(forward_demodulation,[],[f7527,f123])).
fof(f7527,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X16,sK1),X17),s(X16,X18),s(X16,X19)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,happ(s(fun(X16,sK1),X17),s(X16,X19))))),s(sK1,happ(s(fun(X16,sK1),X17),s(X16,X18))))))) ) | $spl6),
  inference(superposition,[],[f7482,f123])).
fof(f7482,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1))))) ) | $spl6),
  inference(equality_resolution,[],[f2010])).
fof(f2010,plain,(
  ( ! [X218,X216,X217] : (s(bool,X218) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X217))),s(sK1,X216))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X216))),s(sK1,X217)))) | ~p(s(bool,X218))) ) | $spl6),
  inference(superposition,[],[f147,f127])).
fof(f9790,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'F_0'),s(bool,X46)))) | s(bool,'T_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X47),s(bool,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f8468,f89])).
fof(f9788,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'F_0'),s(bool,X34)))) | s(bool,'T_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X35),s(bool,X34)))) ) | $spl6),
  inference(resolution,[],[f8468,f88])).
fof(f9786,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'F_0'),s(bool,X22)))) | s(bool,'T_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X23),s(bool,X22)))) ) | $spl6),
  inference(resolution,[],[f8468,f87])).
fof(f9785,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'F_0'),s(bool,X16)))) | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X17),s(bool,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f8468,f85])).
fof(f9783,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl6),
  inference(resolution,[],[f8468,f83])).
fof(f9771,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) | $spl192),
  inference(cnf_transformation,[],[f9771_D])).
fof(f9771_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f9708,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'T_0'),s(bool,X61)))) | s(bool,'F_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X62),s(bool,X61)))) ) | $spl6),
  inference(resolution,[],[f8351,f125])).
fof(f8351,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X8),s(bool,X7)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,'T_0'),s(bool,X7)))) | s(bool,'F_0') = s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f7579,f2341])).
fof(f7579,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X8),s(bool,X6)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,'T_0'),s(bool,X8)))) | s(bool,'F_0') = s(bool,X6)) ) | $spl6),
  inference(superposition,[],[f7530,f122])).
fof(f9704,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'T_0'),s(bool,X46)))) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X47),s(bool,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f8351,f89])).
fof(f9702,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'T_0'),s(bool,X34)))) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X35),s(bool,X34)))) ) | $spl6),
  inference(resolution,[],[f8351,f88])).
fof(f9700,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'T_0'),s(bool,X22)))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X23),s(bool,X22)))) ) | $spl6),
  inference(resolution,[],[f8351,f87])).
fof(f9699,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'T_0'),s(bool,X16)))) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X17),s(bool,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f8351,f85])).
fof(f9697,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl6),
  inference(resolution,[],[f8351,f83])).
fof(f9682,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X4)) ) | $spl190),
  inference(cnf_transformation,[],[f9682_D])).
fof(f9682_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f9639,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) | s(bool,'T_0') = s(bool,X61) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'F_0'),s(bool,X62)))) ) | $spl6),
  inference(resolution,[],[f8264,f125])).
fof(f8264,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,'F_0'),s(bool,X7)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X8),s(bool,X7)))) | s(bool,'T_0') = s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f7578,f2341])).
fof(f7578,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,'F_0')))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X3),s(bool,X5)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl6),
  inference(superposition,[],[f7530,f122])).
fof(f9635,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47)))) | s(bool,'T_0') = s(bool,X46) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'F_0'),s(bool,X47))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f8264,f89])).
fof(f9633,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) | s(bool,'T_0') = s(bool,X34) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'F_0'),s(bool,X35)))) ) | $spl6),
  inference(resolution,[],[f8264,f88])).
fof(f9631,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) | s(bool,'T_0') = s(bool,X22) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'F_0'),s(bool,X23)))) ) | $spl6),
  inference(resolution,[],[f8264,f87])).
fof(f9630,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17)))) | s(bool,'T_0') = s(bool,X16) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'F_0'),s(bool,X17))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f8264,f85])).
fof(f9628,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'T_0') = s(bool,X4) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X5)))) ) | $spl6),
  inference(resolution,[],[f8264,f83])).
fof(f9618,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X4)) ) | $spl188),
  inference(cnf_transformation,[],[f9618_D])).
fof(f9618_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X4)) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f9576,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) | s(bool,'F_0') = s(bool,X61) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'T_0'),s(bool,X62)))) ) | $spl6),
  inference(resolution,[],[f8190,f125])).
fof(f8190,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,'T_0'),s(bool,X7)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X8),s(bool,X7)))) | s(bool,'F_0') = s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f7577,f2341])).
fof(f7577,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X2),s(bool,'T_0')))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X0),s(bool,X2)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f7530,f122])).
fof(f9572,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47)))) | s(bool,'F_0') = s(bool,X46) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'T_0'),s(bool,X47))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f8190,f89])).
fof(f9570,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) | s(bool,'F_0') = s(bool,X34) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'T_0'),s(bool,X35)))) ) | $spl6),
  inference(resolution,[],[f8190,f88])).
fof(f9568,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) | s(bool,'F_0') = s(bool,X22) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'T_0'),s(bool,X23)))) ) | $spl6),
  inference(resolution,[],[f8190,f87])).
fof(f9567,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17)))) | s(bool,'F_0') = s(bool,X16) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'T_0'),s(bool,X17))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f8190,f85])).
fof(f9565,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'F_0') = s(bool,X4) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X5)))) ) | $spl6),
  inference(resolution,[],[f8190,f83])).
fof(f9515,plain,(
  ( ! [X596,X592,X594,X590,X597,X593,X589,X591] : (s(bool,X596) != s(bool,happ(s(fun(X589,bool),happ(s(fun(X589,fun(X589,bool)),'const.relation.inv_image_2'(s(fun(X590,fun(X590,bool)),X591),s(fun(X589,X590),X592))),s(X589,X593))),s(X589,X594))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X590,fun(X590,bool)),X591),s(fun(X589,X590),X592),s(X589,X593),s(X589,X594))) | s(bool,'T_0') = s(bool,X597) | s(bool,X596) != s(bool,X597)) ) | $spl186),
  inference(cnf_transformation,[],[f9515_D])).
fof(f9515_D,plain,(
  ( ! [X596,X592,X594,X590,X597,X593,X589,X591] : (s(bool,X596) != s(bool,happ(s(fun(X589,bool),happ(s(fun(X589,fun(X589,bool)),'const.relation.inv_image_2'(s(fun(X590,fun(X590,bool)),X591),s(fun(X589,X590),X592))),s(X589,X593))),s(X589,X594))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X590,fun(X590,bool)),X591),s(fun(X589,X590),X592),s(X589,X593),s(X589,X594))) | s(bool,'T_0') = s(bool,X597) | s(bool,X596) != s(bool,X597)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f9518,plain,(
  ( ! [X697,X699,X705,X700,X702,X696,X698,X704,X706,X701,X703] : (s(bool,X703) != s(bool,happ(s(fun(X696,bool),happ(s(fun(X696,fun(X696,bool)),'const.relation.inv_image_2'(s(fun(X697,fun(X697,bool)),X698),s(fun(X696,X697),X699))),s(X696,X700))),s(X696,X701))) | s(bool,X702) != s(bool,X704) | s(bool,X704) != s(bool,X705) | s(bool,'F_0') = s(bool,X705) | s(bool,X703) != s(bool,X706) | s(bool,'T_0') = s(bool,X706) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X697,fun(X697,bool)),X698),s(fun(X696,X697),X699),s(X696,X700),s(X696,X701)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f9372,f789])).
fof(f9372,plain,(
  ( ! [X697,X699,X705,X700,X702,X696,X698,X704,X706,X701,X703] : (s(bool,'T_0') != s(bool,X702) | s(bool,X703) != s(bool,happ(s(fun(X696,bool),happ(s(fun(X696,fun(X696,bool)),'const.relation.inv_image_2'(s(fun(X697,fun(X697,bool)),X698),s(fun(X696,X697),X699))),s(X696,X700))),s(X696,X701))) | s(bool,X702) != s(bool,X704) | s(bool,X704) != s(bool,X705) | s(bool,'F_0') = s(bool,X705) | s(bool,X703) != s(bool,X706) | s(bool,'T_0') = s(bool,X706) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X697,fun(X697,bool)),X698),s(fun(X696,X697),X699),s(X696,X700),s(X696,X701)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f1411])).
fof(f1411,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),happ(s(fun(X24,fun(X24,bool)),'const.relation.inv_image_2'(s(fun(X25,fun(X25,bool)),X26),s(fun(X24,X25),X27))),s(X24,X28))),s(X24,X29))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X25,fun(X25,bool)),X26),s(fun(X24,X25),X27),s(X24,X28),s(X24,X29)))) )),
  inference(superposition,[],[f81,f122])).
fof(f9517,plain,(
  ( ! [X656,X652,X654,X648,X650,X657,X653,X655,X649,X651] : (s(bool,X655) != s(bool,happ(s(fun(X648,bool),happ(s(fun(X648,fun(X648,bool)),'const.relation.inv_image_2'(s(fun(X649,fun(X649,bool)),X650),s(fun(X648,X649),X651))),s(X648,X652))),s(X648,X653))) | s(bool,X654) != s(bool,X656) | s(bool,'F_0') = s(bool,X656) | s(bool,X655) != s(bool,X657) | s(bool,'T_0') = s(bool,X657) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X649,fun(X649,bool)),X650),s(fun(X648,X649),X651),s(X648,X652),s(X648,X653)))) ) | ($spl6 | $spl34 | $spl56)),
  inference(subsumption_resolution,[],[f9367,f589])).
fof(f9367,plain,(
  ( ! [X656,X652,X654,X648,X650,X657,X653,X655,X649,X651] : (s(bool,'T_0') != s(bool,X654) | s(bool,X655) != s(bool,happ(s(fun(X648,bool),happ(s(fun(X648,fun(X648,bool)),'const.relation.inv_image_2'(s(fun(X649,fun(X649,bool)),X650),s(fun(X648,X649),X651))),s(X648,X652))),s(X648,X653))) | s(bool,X654) != s(bool,X656) | s(bool,'F_0') = s(bool,X656) | s(bool,X655) != s(bool,X657) | s(bool,'T_0') = s(bool,X657) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X649,fun(X649,bool)),X650),s(fun(X648,X649),X651),s(X648,X652),s(X648,X653)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f1411])).
fof(f9514,plain,(
  ( ! [X596,X592,X594,X590,X597,X593,X595,X589,X591] : (s(bool,X596) != s(bool,happ(s(fun(X589,bool),happ(s(fun(X589,fun(X589,bool)),'const.relation.inv_image_2'(s(fun(X590,fun(X590,bool)),X591),s(fun(X589,X590),X592))),s(X589,X593))),s(X589,X594))) | s(bool,'F_0') = s(bool,X595) | s(bool,X596) != s(bool,X597) | s(bool,'T_0') = s(bool,X597) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X590,fun(X590,bool)),X591),s(fun(X589,X590),X592),s(X589,X593),s(X589,X594)))) ) | ($spl34 | $spl40)),
  inference(subsumption_resolution,[],[f9360,f122])).
fof(f9360,plain,(
  ( ! [X596,X592,X594,X590,X597,X593,X595,X589,X591] : (s(bool,'T_0') != s(bool,X595) | s(bool,X596) != s(bool,happ(s(fun(X589,bool),happ(s(fun(X589,fun(X589,bool)),'const.relation.inv_image_2'(s(fun(X590,fun(X590,bool)),X591),s(fun(X589,X590),X592))),s(X589,X593))),s(X589,X594))) | s(bool,'F_0') = s(bool,X595) | s(bool,X596) != s(bool,X597) | s(bool,'T_0') = s(bool,X597) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X590,fun(X590,bool)),X591),s(fun(X589,X590),X592),s(X589,X593),s(X589,X594)))) ) | ($spl34 | $spl40)),
  inference(superposition,[],[f1342,f1411])).
fof(f9512,plain,(
  ( ! [X588,X584,X586,X580,X582,X585,X587,X581,X583,X579] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X585,fun(X585,bool)),X586))) | s(bool,happ(s(fun(X585,bool),happ(s(fun(X585,fun(X585,bool)),X586),s(X585,X587))),s(X585,X588))) = s(bool,happ(s(fun(X585,bool),happ(s(fun(X585,fun(X585,bool)),X586),s(X585,X588))),s(X585,X587))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X580,fun(X580,bool)),X581),s(fun(X579,X580),X582),s(X579,X583),s(X579,X584)))) ) | ($spl24 | $spl50)),
  inference(subsumption_resolution,[],[f9511,f432])).
fof(f9511,plain,(
  ( ! [X588,X584,X586,X580,X582,X585,X587,X581,X583,X579] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(X580,fun(X580,bool)),X581),s(fun(X579,X580),X582),s(X579,X583),s(X579,X584)))) | s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X585,fun(X585,bool)),X586))) | s(bool,happ(s(fun(X585,bool),happ(s(fun(X585,fun(X585,bool)),X586),s(X585,X587))),s(X585,X588))) = s(bool,happ(s(fun(X585,bool),happ(s(fun(X585,fun(X585,bool)),X586),s(X585,X588))),s(X585,X587))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X580,fun(X580,bool)),X581),s(fun(X579,X580),X582),s(X579,X583),s(X579,X584)))) ) | $spl50),
  inference(forward_demodulation,[],[f9359,f81])).
fof(f9359,plain,(
  ( ! [X588,X584,X586,X580,X582,X585,X587,X581,X583,X579] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X585,fun(X585,bool)),X586))) | s(bool,happ(s(fun(X585,bool),happ(s(fun(X585,fun(X585,bool)),X586),s(X585,X587))),s(X585,X588))) = s(bool,happ(s(fun(X585,bool),happ(s(fun(X585,fun(X585,bool)),X586),s(X585,X588))),s(X585,X587))) | ~p(s(bool,happ(s(fun(X579,bool),happ(s(fun(X579,fun(X579,bool)),'const.relation.inv_image_2'(s(fun(X580,fun(X580,bool)),X581),s(fun(X579,X580),X582))),s(X579,X583))),s(X579,X584)))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X580,fun(X580,bool)),X581),s(fun(X579,X580),X582),s(X579,X583),s(X579,X584)))) ) | $spl50),
  inference(superposition,[],[f1285,f1411])).
fof(f9510,plain,(
  ( ! [X532,X534,X528,X530,X529,X531,X527] : (s(bool,X534) != s(bool,happ(s(fun(X527,bool),happ(s(fun(X527,fun(X527,bool)),'const.relation.inv_image_2'(s(fun(X528,fun(X528,bool)),X529),s(fun(X527,X528),X530))),s(X527,X531))),s(X527,X532))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X528,fun(X528,bool)),X529),s(fun(X527,X528),X530),s(X527,X531),s(X527,X532))) | s(bool,'T_0') = s(bool,X534)) ) | $spl182),
  inference(cnf_transformation,[],[f9510_D])).
fof(f9510_D,plain,(
  ( ! [X532,X534,X528,X530,X529,X531,X527] : (s(bool,X534) != s(bool,happ(s(fun(X527,bool),happ(s(fun(X527,fun(X527,bool)),'const.relation.inv_image_2'(s(fun(X528,fun(X528,bool)),X529),s(fun(X527,X528),X530))),s(X527,X531))),s(X527,X532))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X528,fun(X528,bool)),X529),s(fun(X527,X528),X530),s(X527,X531),s(X527,X532))) | s(bool,'T_0') = s(bool,X534)) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f9509,plain,(
  ( ! [X532,X534,X528,X530,X533,X529,X531,X527] : (s(bool,X534) != s(bool,happ(s(fun(X527,bool),happ(s(fun(X527,fun(X527,bool)),'const.relation.inv_image_2'(s(fun(X528,fun(X528,bool)),X529),s(fun(X527,X528),X530))),s(X527,X531))),s(X527,X532))) | s(bool,'F_0') = s(bool,X533) | s(bool,'T_0') = s(bool,X534) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X528,fun(X528,bool)),X529),s(fun(X527,X528),X530),s(X527,X531),s(X527,X532)))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f9352,f122])).
fof(f9352,plain,(
  ( ! [X532,X534,X528,X530,X533,X529,X531,X527] : (s(bool,'T_0') != s(bool,X533) | s(bool,X534) != s(bool,happ(s(fun(X527,bool),happ(s(fun(X527,fun(X527,bool)),'const.relation.inv_image_2'(s(fun(X528,fun(X528,bool)),X529),s(fun(X527,X528),X530))),s(X527,X531))),s(X527,X532))) | s(bool,'F_0') = s(bool,X533) | s(bool,'T_0') = s(bool,X534) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(X528,fun(X528,bool)),X529),s(fun(X527,X528),X530),s(X527,X531),s(X527,X532)))) ) | ($spl34 | $spl38)),
  inference(superposition,[],[f789,f1411])).
fof(f9195,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0')))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) | $spl180),
  inference(cnf_transformation,[],[f9195_D])).
fof(f9195_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0')))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f9134,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) ) | $spl6),
  inference(resolution,[],[f8111,f125])).
fof(f8111,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X8),s(bool,X7)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X8),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X7)) ) | $spl6),
  inference(superposition,[],[f7558,f2341])).
fof(f7558,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X9),s(bool,X11)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X10),s(bool,X11),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f7530,f122])).
fof(f9130,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f8111,f89])).
fof(f9128,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) ) | $spl6),
  inference(resolution,[],[f8111,f88])).
fof(f9126,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) ) | $spl6),
  inference(resolution,[],[f8111,f87])).
fof(f9125,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f8111,f85])).
fof(f9123,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) ) | $spl6),
  inference(resolution,[],[f8111,f83])).
fof(f9107,plain,(
  ( ! [X2,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X2),s(bool,'T_0'),s(bool,'F_0')))) | s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X2),s(bool,X3),s(bool,'F_0'))) | ~p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3)) ) | $spl178),
  inference(cnf_transformation,[],[f9107_D])).
fof(f9107_D,plain,(
  ( ! [X2,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X2),s(bool,'T_0'),s(bool,'F_0')))) | s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X2),s(bool,X3),s(bool,'F_0'))) | ~p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3)) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f9064,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X48),s(bool,'T_0'),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X49) | p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X48),s(bool,X49),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f8026,f125])).
fof(f8026,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X14),s(bool,X15),s(bool,'F_0')))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X14),s(bool,'T_0'),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X15)) ) | $spl6),
  inference(forward_demodulation,[],[f8025,f180])).
fof(f8025,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X14),s(bool,'T_0'),s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))))) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X14),s(bool,X15),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X15)) ) | $spl6),
  inference(forward_demodulation,[],[f7997,f2341])).
fof(f7997,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X14),s(bool,X15),s(bool,'F_0')))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X14),s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X15)) ) | $spl6),
  inference(superposition,[],[f7557,f180])).
fof(f7557,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X6),s(bool,X8)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X7),s(bool,X8),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X6)) ) | $spl6),
  inference(superposition,[],[f7530,f122])).
fof(f9060,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X37),s(bool,'T_0'),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X38) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X37),s(bool,X38),s(bool,'F_0'))) | s(bool,X39) != s(bool,X41)) ) | $spl6),
  inference(resolution,[],[f8026,f89])).
fof(f9058,plain,(
  ( ! [X30,X28,X31,X29,X27] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X27),s(bool,'T_0'),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X27),s(bool,X28),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f8026,f88])).
fof(f9056,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X17),s(bool,'T_0'),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X18) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X17),s(bool,X18),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f8026,f87])).
fof(f9055,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X12),s(bool,'T_0'),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X13) | ~p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X15) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X12),s(bool,X13),s(bool,'F_0'))) | s(bool,X14) != s(bool,X16)) ) | $spl6),
  inference(resolution,[],[f8026,f85])).
fof(f9053,plain,(
  ( ! [X6,X4,X2,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X2),s(bool,'T_0'),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X2),s(bool,X3),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f8026,f83])).
fof(f9024,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0')))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) | $spl176),
  inference(cnf_transformation,[],[f9024_D])).
fof(f9024_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0')))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f8963,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) ) | $spl6),
  inference(resolution,[],[f8014,f125])).
fof(f8014,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X8),s(bool,X7)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X8),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X7)) ) | $spl6),
  inference(superposition,[],[f7557,f2341])).
fof(f8959,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f8014,f89])).
fof(f8957,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) ) | $spl6),
  inference(resolution,[],[f8014,f88])).
fof(f8955,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) ) | $spl6),
  inference(resolution,[],[f8014,f87])).
fof(f8954,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f8014,f85])).
fof(f8952,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) ) | $spl6),
  inference(resolution,[],[f8014,f83])).
fof(f8914,plain,(
  ( ! [X705,X710,X702,X704,X706,X709,X701,X703] : (s(bool,X709) != s(bool,happ(s(fun(X701,bool),happ(s(fun(X701,fun(X701,bool)),'const.relation.inv_image_2'(s(fun(X702,fun(X702,bool)),X703),s(fun(X701,X702),X704))),s(X701,X705))),s(X701,X706))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X702,fun(X702,bool)),X703),s(fun(X701,X702),X704),s(X701,X705),s(X701,X706))) | s(bool,'F_0') = s(bool,X710) | s(bool,X709) != s(bool,X710)) ) | $spl174),
  inference(cnf_transformation,[],[f8914_D])).
fof(f8914_D,plain,(
  ( ! [X705,X710,X702,X704,X706,X709,X701,X703] : (s(bool,X709) != s(bool,happ(s(fun(X701,bool),happ(s(fun(X701,fun(X701,bool)),'const.relation.inv_image_2'(s(fun(X702,fun(X702,bool)),X703),s(fun(X701,X702),X704))),s(X701,X705))),s(X701,X706))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X702,fun(X702,bool)),X703),s(fun(X701,X702),X704),s(X701,X705),s(X701,X706))) | s(bool,'F_0') = s(bool,X710) | s(bool,X709) != s(bool,X710)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f8913,plain,(
  ( ! [X707,X705,X711,X708,X710,X702,X704,X706,X709,X701,X703] : (s(bool,X707) != s(bool,X708) | s(bool,X709) != s(bool,happ(s(fun(X701,bool),happ(s(fun(X701,fun(X701,bool)),'const.relation.inv_image_2'(s(fun(X702,fun(X702,bool)),X703),s(fun(X701,X702),X704))),s(X701,X705))),s(X701,X706))) | s(bool,X709) != s(bool,X710) | s(bool,'F_0') = s(bool,X710) | s(bool,X708) != s(bool,X711) | s(bool,'T_0') = s(bool,X711) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X702,fun(X702,bool)),X703),s(fun(X701,X702),X704),s(X701,X705),s(X701,X706)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f8773,f789])).
fof(f8773,plain,(
  ( ! [X707,X705,X711,X708,X710,X702,X704,X706,X709,X701,X703] : (s(bool,'F_0') != s(bool,X707) | s(bool,X707) != s(bool,X708) | s(bool,X709) != s(bool,happ(s(fun(X701,bool),happ(s(fun(X701,fun(X701,bool)),'const.relation.inv_image_2'(s(fun(X702,fun(X702,bool)),X703),s(fun(X701,X702),X704))),s(X701,X705))),s(X701,X706))) | s(bool,X709) != s(bool,X710) | s(bool,'F_0') = s(bool,X710) | s(bool,X708) != s(bool,X711) | s(bool,'T_0') = s(bool,X711) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X702,fun(X702,bool)),X703),s(fun(X701,X702),X704),s(X701,X705),s(X701,X706)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f1410])).
fof(f8909,plain,(
  ( ! [X660,X656,X652,X654,X657,X653,X655] : (s(bool,X660) != s(bool,happ(s(fun(X652,bool),happ(s(fun(X652,fun(X652,bool)),'const.relation.inv_image_2'(s(fun(X653,fun(X653,bool)),X654),s(fun(X652,X653),X655))),s(X652,X656))),s(X652,X657))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X653,fun(X653,bool)),X654),s(fun(X652,X653),X655),s(X652,X656),s(X652,X657))) | s(bool,'F_0') = s(bool,X660)) ) | $spl172),
  inference(cnf_transformation,[],[f8909_D])).
fof(f8909_D,plain,(
  ( ! [X660,X656,X652,X654,X657,X653,X655] : (s(bool,X660) != s(bool,happ(s(fun(X652,bool),happ(s(fun(X652,fun(X652,bool)),'const.relation.inv_image_2'(s(fun(X653,fun(X653,bool)),X654),s(fun(X652,X653),X655))),s(X652,X656))),s(X652,X657))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X653,fun(X653,bool)),X654),s(fun(X652,X653),X655),s(X652,X656),s(X652,X657))) | s(bool,'F_0') = s(bool,X660)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f8908,plain,(
  ( ! [X660,X656,X658,X652,X654,X661,X657,X659,X653,X655] : (s(bool,X658) != s(bool,X659) | s(bool,X660) != s(bool,happ(s(fun(X652,bool),happ(s(fun(X652,fun(X652,bool)),'const.relation.inv_image_2'(s(fun(X653,fun(X653,bool)),X654),s(fun(X652,X653),X655))),s(X652,X656))),s(X652,X657))) | s(bool,'F_0') = s(bool,X660) | s(bool,X659) != s(bool,X661) | s(bool,'T_0') = s(bool,X661) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X653,fun(X653,bool)),X654),s(fun(X652,X653),X655),s(X652,X656),s(X652,X657)))) ) | ($spl34 | $spl38 | $spl56)),
  inference(subsumption_resolution,[],[f8768,f789])).
fof(f8768,plain,(
  ( ! [X660,X656,X658,X652,X654,X661,X657,X659,X653,X655] : (s(bool,'F_0') != s(bool,X658) | s(bool,X658) != s(bool,X659) | s(bool,X660) != s(bool,happ(s(fun(X652,bool),happ(s(fun(X652,fun(X652,bool)),'const.relation.inv_image_2'(s(fun(X653,fun(X653,bool)),X654),s(fun(X652,X653),X655))),s(X652,X656))),s(X652,X657))) | s(bool,'F_0') = s(bool,X660) | s(bool,X659) != s(bool,X661) | s(bool,'T_0') = s(bool,X661) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(X653,fun(X653,bool)),X654),s(fun(X652,X653),X655),s(X652,X656),s(X652,X657)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f1410])).
fof(f8609,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0'))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) | $spl170),
  inference(cnf_transformation,[],[f8609_D])).
fof(f8609_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0'))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f8568,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) | s(bool,'T_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f7905,f125])).
fof(f7905,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X7),s(bool,'F_0')))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X7),s(bool,X8)))) | s(bool,'T_0') = s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f7556,f2341])).
fof(f7556,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,'F_0'),s(bool,X5)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X4),s(bool,X5),s(bool,X3)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl6),
  inference(superposition,[],[f7530,f122])).
fof(f8564,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47)))) | s(bool,'T_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,'F_0'))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f7905,f89])).
fof(f8562,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) | s(bool,'T_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f7905,f88])).
fof(f8560,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) | s(bool,'T_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f7905,f87])).
fof(f8559,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17)))) | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,'F_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f7905,f85])).
fof(f8557,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f7905,f83])).
fof(f8535,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0'))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) | $spl168),
  inference(cnf_transformation,[],[f8535_D])).
fof(f8535_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0'))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f8495,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) | s(bool,'F_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,'T_0')))) ) | $spl6),
  inference(resolution,[],[f7845,f125])).
fof(f7845,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X7),s(bool,'T_0')))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X6),s(bool,X7),s(bool,X8)))) | s(bool,'F_0') = s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f7555,f2341])).
fof(f7555,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,'T_0'),s(bool,X2)))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X1),s(bool,X2),s(bool,X0)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f7530,f122])).
fof(f8491,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47)))) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,'T_0'))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f7845,f89])).
fof(f8489,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,'T_0')))) ) | $spl6),
  inference(resolution,[],[f7845,f88])).
fof(f8487,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,'T_0')))) ) | $spl6),
  inference(resolution,[],[f7845,f87])).
fof(f8486,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17)))) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,'T_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f7845,f85])).
fof(f8484,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0')))) ) | $spl6),
  inference(resolution,[],[f7845,f83])).
fof(f8473,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) | $spl166),
  inference(cnf_transformation,[],[f8473_D])).
fof(f8473_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f8414,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'F_0'),s(bool,X61)))) | s(bool,'T_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) ) | $spl6),
  inference(resolution,[],[f7580,f125])).
fof(f8410,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'F_0'),s(bool,X46)))) | s(bool,'T_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f7580,f89])).
fof(f8408,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'F_0'),s(bool,X34)))) | s(bool,'T_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) ) | $spl6),
  inference(resolution,[],[f7580,f88])).
fof(f8406,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'F_0'),s(bool,X22)))) | s(bool,'T_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) ) | $spl6),
  inference(resolution,[],[f7580,f87])).
fof(f8405,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'F_0'),s(bool,X16)))) | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f7580,f85])).
fof(f8403,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) ) | $spl6),
  inference(resolution,[],[f7580,f83])).
fof(f8391,plain,(
  ( ! [X6,X8,X7,X5,X11,X9] : (s(bool,X5) != s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9))) | s(bool,X11) != s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7))) | ~p(s(bool,X11)) | ~p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X8)))) | p(s(bool,X5))) ) | $spl164),
  inference(cnf_transformation,[],[f8391_D])).
fof(f8391_D,plain,(
  ( ! [X6,X8,X7,X5,X11,X9] : (s(bool,X5) != s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9))) | s(bool,X11) != s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7))) | ~p(s(bool,X11)) | ~p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X8)))) | p(s(bool,X5))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f8377,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,X84) != s(bool,happ(s(fun(X85,bool),happ(s(fun(X85,fun(X85,bool)),X86),s(X85,X87))),s(X85,X88))) | p(s(bool,X84)) | ~p(s(bool,happ(s(fun(X85,bool),happ(s(fun(X85,fun(X85,bool)),X86),s(X85,X88))),s(X85,X87)))) | p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X90) != s(bool,'const.relation.symmetric_1'(s(fun(X85,fun(X85,bool)),X86)))) ) | $spl54),
  inference(resolution,[],[f1290,f125])).
fof(f8373,plain,(
  ( ! [X61,X68,X66,X64,X62,X67,X65,X63] : (s(bool,X61) != s(bool,happ(s(fun(X62,bool),happ(s(fun(X62,fun(X62,bool)),X63),s(X62,X64))),s(X62,X65))) | p(s(bool,X61)) | ~p(s(bool,happ(s(fun(X62,bool),happ(s(fun(X62,fun(X62,bool)),X63),s(X62,X65))),s(X62,X64)))) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X67) != s(bool,'const.relation.symmetric_1'(s(fun(X62,fun(X62,bool)),X63))) | s(bool,X66) != s(bool,X68)) ) | $spl54),
  inference(resolution,[],[f1290,f89])).
fof(f8371,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X51,X49] : (s(bool,X45) != s(bool,happ(s(fun(X46,bool),happ(s(fun(X46,fun(X46,bool)),X47),s(X46,X48))),s(X46,X49))) | p(s(bool,X45)) | ~p(s(bool,happ(s(fun(X46,bool),happ(s(fun(X46,fun(X46,bool)),X47),s(X46,X49))),s(X46,X48)))) | p(s(bool,X50)) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.relation.symmetric_1'(s(fun(X46,fun(X46,bool)),X47)))) ) | $spl54),
  inference(resolution,[],[f1290,f88])).
fof(f8369,plain,(
  ( ! [X30,X35,X33,X31,X29,X36,X34,X32] : (s(bool,X29) != s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,X32))),s(X30,X33))) | p(s(bool,X29)) | ~p(s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,X33))),s(X30,X32)))) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.relation.symmetric_1'(s(fun(X30,fun(X30,bool)),X31)))) ) | $spl54),
  inference(resolution,[],[f1290,f87])).
fof(f8368,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,happ(s(fun(X22,bool),happ(s(fun(X22,fun(X22,bool)),X23),s(X22,X24))),s(X22,X25))) | p(s(bool,X21)) | ~p(s(bool,happ(s(fun(X22,bool),happ(s(fun(X22,fun(X22,bool)),X23),s(X22,X25))),s(X22,X24)))) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.relation.symmetric_1'(s(fun(X22,fun(X22,bool)),X23))) | s(bool,X26) != s(bool,X28)) ) | $spl54),
  inference(resolution,[],[f1290,f85])).
fof(f8366,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X11,X9] : (s(bool,X5) != s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9))) | p(s(bool,X5)) | ~p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X8)))) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) ) | $spl54),
  inference(resolution,[],[f1290,f83])).
fof(f8356,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) | $spl162),
  inference(cnf_transformation,[],[f8356_D])).
fof(f8356_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f8297,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'T_0'),s(bool,X61)))) | s(bool,'F_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) ) | $spl6),
  inference(resolution,[],[f7579,f125])).
fof(f8293,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'T_0'),s(bool,X46)))) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f7579,f89])).
fof(f8291,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'T_0'),s(bool,X34)))) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) ) | $spl6),
  inference(resolution,[],[f7579,f88])).
fof(f8289,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'T_0'),s(bool,X22)))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) ) | $spl6),
  inference(resolution,[],[f7579,f87])).
fof(f8288,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'T_0'),s(bool,X16)))) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f7579,f85])).
fof(f8286,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) ) | $spl6),
  inference(resolution,[],[f7579,f83])).
fof(f8268,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'F_0'))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X4)) ) | $spl160),
  inference(cnf_transformation,[],[f8268_D])).
fof(f8268_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'F_0'))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f8227,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) | s(bool,'T_0') = s(bool,X61) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X62),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f7578,f125])).
fof(f8223,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47)))) | s(bool,'T_0') = s(bool,X46) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X47),s(bool,'F_0'))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f7578,f89])).
fof(f8221,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) | s(bool,'T_0') = s(bool,X34) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X35),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f7578,f88])).
fof(f8219,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) | s(bool,'T_0') = s(bool,X22) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X23),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f7578,f87])).
fof(f8218,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17)))) | s(bool,'T_0') = s(bool,X16) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X17),s(bool,'F_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f7578,f85])).
fof(f8216,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'T_0') = s(bool,X4) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'F_0')))) ) | $spl6),
  inference(resolution,[],[f7578,f83])).
fof(f8194,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'T_0'))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X4)) ) | $spl158),
  inference(cnf_transformation,[],[f8194_D])).
fof(f8194_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'T_0'))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X4)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f8154,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) | s(bool,'F_0') = s(bool,X61) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X62),s(bool,'T_0')))) ) | $spl6),
  inference(resolution,[],[f7577,f125])).
fof(f8150,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47)))) | s(bool,'F_0') = s(bool,X46) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X47),s(bool,'T_0'))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f7577,f89])).
fof(f8148,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) | s(bool,'F_0') = s(bool,X34) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X35),s(bool,'T_0')))) ) | $spl6),
  inference(resolution,[],[f7577,f88])).
fof(f8146,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) | s(bool,'F_0') = s(bool,X22) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X23),s(bool,'T_0')))) ) | $spl6),
  inference(resolution,[],[f7577,f87])).
fof(f8145,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17)))) | s(bool,'F_0') = s(bool,X16) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X17),s(bool,'T_0'))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f7577,f85])).
fof(f8143,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'F_0') = s(bool,X4) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,'T_0')))) ) | $spl6),
  inference(resolution,[],[f7577,f83])).
fof(f8116,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0')))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) | $spl156),
  inference(cnf_transformation,[],[f8116_D])).
fof(f8116_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0')))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f8057,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X62),s(bool,X61)))) ) | $spl6),
  inference(resolution,[],[f7558,f125])).
fof(f8053,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X47),s(bool,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f7558,f89])).
fof(f8051,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X35),s(bool,X34)))) ) | $spl6),
  inference(resolution,[],[f7558,f88])).
fof(f8049,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X23),s(bool,X22)))) ) | $spl6),
  inference(resolution,[],[f7558,f87])).
fof(f8048,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X17),s(bool,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f7558,f85])).
fof(f8046,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl6),
  inference(resolution,[],[f7558,f83])).
fof(f8019,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0')))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) | $spl154),
  inference(cnf_transformation,[],[f8019_D])).
fof(f8019_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0')))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f7960,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X62),s(bool,X61)))) ) | $spl6),
  inference(resolution,[],[f7557,f125])).
fof(f7956,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X47),s(bool,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f7557,f89])).
fof(f7954,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X35),s(bool,X34)))) ) | $spl6),
  inference(resolution,[],[f7557,f88])).
fof(f7952,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X23),s(bool,X22)))) ) | $spl6),
  inference(resolution,[],[f7557,f87])).
fof(f7951,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X17),s(bool,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f7557,f85])).
fof(f7949,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X5),s(bool,X4)))) ) | $spl6),
  inference(resolution,[],[f7557,f83])).
fof(f7946,plain,(
  ( ! [X6,X8,X7,X5,X11,X9] : (s(bool,X5) != s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9))) | s(bool,X11) != s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7))) | ~p(s(bool,X11)) | ~p(s(bool,X5)) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X8))))) ) | $spl152),
  inference(cnf_transformation,[],[f7946_D])).
fof(f7946_D,plain,(
  ( ! [X6,X8,X7,X5,X11,X9] : (s(bool,X5) != s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9))) | s(bool,X11) != s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7))) | ~p(s(bool,X11)) | ~p(s(bool,X5)) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X8))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f7932,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,X84) != s(bool,happ(s(fun(X85,bool),happ(s(fun(X85,fun(X85,bool)),X86),s(X85,X87))),s(X85,X88))) | p(s(bool,happ(s(fun(X85,bool),happ(s(fun(X85,fun(X85,bool)),X86),s(X85,X88))),s(X85,X87)))) | ~p(s(bool,X84)) | p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X90) != s(bool,'const.relation.symmetric_1'(s(fun(X85,fun(X85,bool)),X86)))) ) | $spl52),
  inference(resolution,[],[f1289,f125])).
fof(f7928,plain,(
  ( ! [X61,X68,X66,X64,X62,X67,X65,X63] : (s(bool,X61) != s(bool,happ(s(fun(X62,bool),happ(s(fun(X62,fun(X62,bool)),X63),s(X62,X64))),s(X62,X65))) | p(s(bool,happ(s(fun(X62,bool),happ(s(fun(X62,fun(X62,bool)),X63),s(X62,X65))),s(X62,X64)))) | ~p(s(bool,X61)) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X67) != s(bool,'const.relation.symmetric_1'(s(fun(X62,fun(X62,bool)),X63))) | s(bool,X66) != s(bool,X68)) ) | $spl52),
  inference(resolution,[],[f1289,f89])).
fof(f7926,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X51,X49] : (s(bool,X45) != s(bool,happ(s(fun(X46,bool),happ(s(fun(X46,fun(X46,bool)),X47),s(X46,X48))),s(X46,X49))) | p(s(bool,happ(s(fun(X46,bool),happ(s(fun(X46,fun(X46,bool)),X47),s(X46,X49))),s(X46,X48)))) | ~p(s(bool,X45)) | p(s(bool,X50)) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.relation.symmetric_1'(s(fun(X46,fun(X46,bool)),X47)))) ) | $spl52),
  inference(resolution,[],[f1289,f88])).
fof(f7924,plain,(
  ( ! [X30,X35,X33,X31,X29,X36,X34,X32] : (s(bool,X29) != s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,X32))),s(X30,X33))) | p(s(bool,happ(s(fun(X30,bool),happ(s(fun(X30,fun(X30,bool)),X31),s(X30,X33))),s(X30,X32)))) | ~p(s(bool,X29)) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.relation.symmetric_1'(s(fun(X30,fun(X30,bool)),X31)))) ) | $spl52),
  inference(resolution,[],[f1289,f87])).
fof(f7923,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,happ(s(fun(X22,bool),happ(s(fun(X22,fun(X22,bool)),X23),s(X22,X24))),s(X22,X25))) | p(s(bool,happ(s(fun(X22,bool),happ(s(fun(X22,fun(X22,bool)),X23),s(X22,X25))),s(X22,X24)))) | ~p(s(bool,X21)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.relation.symmetric_1'(s(fun(X22,fun(X22,bool)),X23))) | s(bool,X26) != s(bool,X28)) ) | $spl52),
  inference(resolution,[],[f1289,f85])).
fof(f7921,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X11,X9] : (s(bool,X5) != s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X8)))) | ~p(s(bool,X5)) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) ) | $spl52),
  inference(resolution,[],[f1289,f83])).
fof(f7909,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) | $spl150),
  inference(cnf_transformation,[],[f7909_D])).
fof(f7909_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f7868,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) | s(bool,'T_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'F_0'),s(bool,X61)))) ) | $spl6),
  inference(resolution,[],[f7556,f125])).
fof(f7864,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47)))) | s(bool,'T_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'F_0'),s(bool,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f7556,f89])).
fof(f7862,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) | s(bool,'T_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'F_0'),s(bool,X34)))) ) | $spl6),
  inference(resolution,[],[f7556,f88])).
fof(f7860,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) | s(bool,'T_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'F_0'),s(bool,X22)))) ) | $spl6),
  inference(resolution,[],[f7556,f87])).
fof(f7859,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17)))) | s(bool,'T_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'F_0'),s(bool,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f7556,f85])).
fof(f7857,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'F_0'),s(bool,X4)))) ) | $spl6),
  inference(resolution,[],[f7556,f83])).
fof(f7849,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) | $spl148),
  inference(cnf_transformation,[],[f7849_D])).
fof(f7849_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4))) | ~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f7809,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,X61),s(bool,X62)))) | s(bool,'F_0') = s(bool,X62) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X60),s(bool,'T_0'),s(bool,X61)))) ) | $spl6),
  inference(resolution,[],[f7555,f125])).
fof(f7805,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,X46),s(bool,X47)))) | s(bool,'F_0') = s(bool,X47) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X45),s(bool,'T_0'),s(bool,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f7555,f89])).
fof(f7803,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,X34),s(bool,X35)))) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X33),s(bool,'T_0'),s(bool,X34)))) ) | $spl6),
  inference(resolution,[],[f7555,f88])).
fof(f7801,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,X22),s(bool,X23)))) | s(bool,'F_0') = s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X21),s(bool,'T_0'),s(bool,X22)))) ) | $spl6),
  inference(resolution,[],[f7555,f87])).
fof(f7800,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,X16),s(bool,X17)))) | s(bool,'F_0') = s(bool,X17) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X15),s(bool,'T_0'),s(bool,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f7555,f85])).
fof(f7798,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,X4),s(bool,X5)))) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(bool,sK1),X3),s(bool,'T_0'),s(bool,X4)))) ) | $spl6),
  inference(resolution,[],[f7555,f83])).
fof(f7787,plain,(
  ( ! [X6,X8,X7,X5,X11,X9] : (p(s(bool,X5)) | s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X9),s(X6,X8))) | ~p(s(bool,X11)) | s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X8),s(X6,X9)))) ) | $spl146),
  inference(cnf_transformation,[],[f7787_D])).
fof(f7787_D,plain,(
  ( ! [X6,X8,X7,X5,X11,X9] : (p(s(bool,X5)) | s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X9),s(X6,X8))) | ~p(s(bool,X11)) | s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X8),s(X6,X9)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f7725,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (p(s(bool,X84)) | s(bool,X84) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X85,sK1),X86),s(X85,X87),s(X85,X88))) | p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X90) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X85,sK1),X86),s(X85,X88),s(X85,X87)))) ) | $spl6),
  inference(resolution,[],[f2432,f125])).
fof(f7721,plain,(
  ( ! [X61,X68,X66,X64,X62,X67,X65,X63] : (p(s(bool,X61)) | s(bool,X61) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X62,sK1),X63),s(X62,X64),s(X62,X65))) | ~p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X67) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X62,sK1),X63),s(X62,X65),s(X62,X64))) | s(bool,X66) != s(bool,X68)) ) | $spl6),
  inference(resolution,[],[f2432,f89])).
fof(f7719,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X51,X49] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X46,sK1),X47),s(X46,X48),s(X46,X49))) | p(s(bool,X50)) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X46,sK1),X47),s(X46,X49),s(X46,X48)))) ) | $spl6),
  inference(resolution,[],[f2432,f88])).
fof(f7717,plain,(
  ( ! [X30,X35,X33,X31,X29,X36,X34,X32] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X30,sK1),X31),s(X30,X32),s(X30,X33))) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X30,sK1),X31),s(X30,X33),s(X30,X32)))) ) | $spl6),
  inference(resolution,[],[f2432,f87])).
fof(f7716,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,X21) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X22,sK1),X23),s(X22,X24),s(X22,X25))) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X22,sK1),X23),s(X22,X25),s(X22,X24))) | s(bool,X26) != s(bool,X28)) ) | $spl6),
  inference(resolution,[],[f2432,f85])).
fof(f7714,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X11,X9] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X8),s(X6,X9))) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X6,sK1),X7),s(X6,X9),s(X6,X8)))) ) | $spl6),
  inference(resolution,[],[f2432,f83])).
fof(f7703,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3)))) | s(bool,X4) != s(bool,X6) | ~p(s(bool,X6)) | s(bool,X4) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2)))) ) | $spl144),
  inference(cnf_transformation,[],[f7703_D])).
fof(f7703_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3)))) | s(bool,X4) != s(bool,X6) | ~p(s(bool,X6)) | s(bool,X4) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2)))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f7664,plain,(
  ( ! [X80,X85,X83,X81,X79,X84,X82] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X79,sK1),X80),s(X79,X81),s(X79,X82)))) | s(bool,X83) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X79,sK1),X80),s(X79,X82),s(X79,X81))) | p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X83) != s(bool,X85)) ) | $spl6),
  inference(resolution,[],[f2431,f125])).
fof(f2431,plain,(
  ( ! [X308,X304,X306,X305,X307] : (~p(s(bool,X308)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X304,sK1),X305),s(X304,X306),s(X304,X307)))) | s(bool,X308) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X304,sK1),X305),s(X304,X307),s(X304,X306)))) ) | $spl6),
  inference(superposition,[],[f147,f2341])).
fof(f7660,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X56,sK1),X57),s(X56,X58),s(X56,X59)))) | s(bool,X60) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X56,sK1),X57),s(X56,X59),s(X56,X58))) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X60) != s(bool,X62) | s(bool,X61) != s(bool,X63)) ) | $spl6),
  inference(resolution,[],[f2431,f89])).
fof(f7658,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X40,sK1),X41),s(X40,X42),s(X40,X43)))) | s(bool,X44) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X40,sK1),X41),s(X40,X43),s(X40,X42))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X44) != s(bool,X46)) ) | $spl6),
  inference(resolution,[],[f2431,f88])).
fof(f7656,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X24,sK1),X25),s(X24,X26),s(X24,X27)))) | s(bool,X28) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X24,sK1),X25),s(X24,X27),s(X24,X26))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X28) != s(bool,X30)) ) | $spl6),
  inference(resolution,[],[f2431,f87])).
fof(f7655,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X16,sK1),X17),s(X16,X18),s(X16,X19)))) | s(bool,X20) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X16,sK1),X17),s(X16,X19),s(X16,X18))) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl6),
  inference(resolution,[],[f2431,f85])).
fof(f7653,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3)))) | s(bool,X4) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X4) != s(bool,X6)) ) | $spl6),
  inference(resolution,[],[f2431,f83])).
fof(f7648,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,X3)) | s(bool,X7) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,X4))) | ~p(s(bool,X7)) | s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,X5)))) ) | $spl142),
  inference(cnf_transformation,[],[f7648_D])).
fof(f7648_D,plain,(
  ( ! [X4,X7,X5,X3] : (p(s(bool,X3)) | s(bool,X7) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,X4))) | ~p(s(bool,X7)) | s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,X5)))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f7625,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X61))),s(sK1,X62))) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X62))),s(sK1,X61)))) ) | $spl6),
  inference(resolution,[],[f2011,f125])).
fof(f7621,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,X45)) | s(bool,X45) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X46))),s(sK1,X47))) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X47))),s(sK1,X46))) | s(bool,X48) != s(bool,X50)) ) | $spl6),
  inference(resolution,[],[f2011,f89])).
fof(f7619,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (p(s(bool,X33)) | s(bool,X33) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X34))),s(sK1,X35))) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X35))),s(sK1,X34)))) ) | $spl6),
  inference(resolution,[],[f2011,f88])).
fof(f7617,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (p(s(bool,X21)) | s(bool,X21) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,X23))) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X23))),s(sK1,X22)))) ) | $spl6),
  inference(resolution,[],[f2011,f87])).
fof(f7616,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (p(s(bool,X15)) | s(bool,X15) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X16))),s(sK1,X17))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X17))),s(sK1,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl6),
  inference(resolution,[],[f2011,f85])).
fof(f7614,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,X5))) | p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,X4)))) ) | $spl6),
  inference(resolution,[],[f2011,f83])).
fof(f2553,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X4) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,X4))) ) | $spl84),
  inference(cnf_transformation,[],[f2553_D])).
fof(f2553_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X4) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,X4))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f7543,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X72,sK1),X73),s(X72,X74),s(X72,X75)))) | p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X77) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X72,sK1),X73),s(X72,X75),s(X72,X74)))) ) | $spl6),
  inference(resolution,[],[f7530,f125])).
fof(f7539,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X53,sK1),X54),s(X53,X55),s(X53,X56)))) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X53,sK1),X54),s(X53,X56),s(X53,X55))) | s(bool,X57) != s(bool,X59)) ) | $spl6),
  inference(resolution,[],[f7530,f89])).
fof(f7537,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X39,sK1),X40),s(X39,X41),s(X39,X42)))) | p(s(bool,X43)) | ~p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X39,sK1),X40),s(X39,X42),s(X39,X41)))) ) | $spl6),
  inference(resolution,[],[f7530,f88])).
fof(f7535,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X25,sK1),X26),s(X25,X27),s(X25,X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X25,sK1),X26),s(X25,X28),s(X25,X27)))) ) | $spl6),
  inference(resolution,[],[f7530,f87])).
fof(f7534,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X18,sK1),X19),s(X18,X20),s(X18,X21)))) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X18,sK1),X19),s(X18,X21),s(X18,X20))) | s(bool,X22) != s(bool,X24)) ) | $spl6),
  inference(resolution,[],[f7530,f85])).
fof(f7532,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X4,sK1),X5),s(X4,X6),s(X4,X7)))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X4,sK1),X5),s(X4,X7),s(X4,X6)))) ) | $spl6),
  inference(resolution,[],[f7530,f83])).
fof(f7362,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X108,sK1),X109),s(X108,X111),s(X108,X110)))) ) | $spl140),
  inference(cnf_transformation,[],[f7362_D])).
fof(f7362_D,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X108,sK1),X109),s(X108,X111),s(X108,X110)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f7361,plain,(
  ( ! [X111,X109,X114,X112,X110,X108,X115,X113] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X112,fun(X112,bool)),X113))) | s(bool,happ(s(fun(X112,bool),happ(s(fun(X112,fun(X112,bool)),X113),s(X112,X114))),s(X112,X115))) = s(bool,happ(s(fun(X112,bool),happ(s(fun(X112,fun(X112,bool)),X113),s(X112,X115))),s(X112,X114))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X108,sK1),X109),s(X108,X111),s(X108,X110)))) ) | ($spl24 | $spl50)),
  inference(subsumption_resolution,[],[f7339,f2434])).
fof(f7339,plain,(
  ( ! [X111,X109,X114,X112,X110,X108,X115,X113] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X112,fun(X112,bool)),X113))) | s(bool,happ(s(fun(X112,bool),happ(s(fun(X112,fun(X112,bool)),X113),s(X112,X114))),s(X112,X115))) = s(bool,happ(s(fun(X112,bool),happ(s(fun(X112,fun(X112,bool)),X113),s(X112,X115))),s(X112,X114))) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X108,sK1),X109),s(X108,X110),s(X108,X111)))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X108,sK1),X109),s(X108,X111),s(X108,X110)))) ) | $spl50),
  inference(superposition,[],[f1285,f2343])).
fof(f7359,plain,(
  ( ! [X83,X82] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X83))),s(sK1,X82)))) ) | $spl138),
  inference(cnf_transformation,[],[f7359_D])).
fof(f7359_D,plain,(
  ( ! [X83,X82] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X83))),s(sK1,X82)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f7358,plain,(
  ( ! [X87,X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X84,fun(X84,bool)),X85))) | s(bool,happ(s(fun(X84,bool),happ(s(fun(X84,fun(X84,bool)),X85),s(X84,X86))),s(X84,X87))) = s(bool,happ(s(fun(X84,bool),happ(s(fun(X84,fun(X84,bool)),X85),s(X84,X87))),s(X84,X86))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X83))),s(sK1,X82)))) ) | ($spl24 | $spl50)),
  inference(subsumption_resolution,[],[f7335,f2013])).
fof(f2013,plain,(
  ( ! [X225,X224] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X225))),s(sK1,X224)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X224))),s(sK1,X225)))) ) | $spl24),
  inference(superposition,[],[f432,f127])).
fof(f7335,plain,(
  ( ! [X87,X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X84,fun(X84,bool)),X85))) | s(bool,happ(s(fun(X84,bool),happ(s(fun(X84,fun(X84,bool)),X85),s(X84,X86))),s(X84,X87))) = s(bool,happ(s(fun(X84,bool),happ(s(fun(X84,fun(X84,bool)),X85),s(X84,X87))),s(X84,X86))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X82))),s(sK1,X83)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X83))),s(sK1,X82)))) ) | $spl50),
  inference(superposition,[],[f1285,f1951])).
fof(f7349,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),X2))) | s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X4))),s(X1,X3))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl50)),
  inference(subsumption_resolution,[],[f7324,f432])).
fof(f7324,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(X1,fun(X1,bool)),X2))) | s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X3))),s(X1,X4))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X4))),s(X1,X3))) | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl50),
  inference(superposition,[],[f1285,f122])).
fof(f3216,plain,(
  ( ! [X6,X4,X7,X5,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X6))),s(X4,X7))) | s(bool,X9) != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5))) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X7))),s(X4,X6)))) ) | $spl98),
  inference(cnf_transformation,[],[f3216_D])).
fof(f3216_D,plain,(
  ( ! [X6,X4,X7,X5,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X6))),s(X4,X7))) | s(bool,X9) != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5))) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X7))),s(X4,X6)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f7225,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (s(bool,'T_0') = s(bool,happ(s(fun(X72,bool),happ(s(fun(X72,fun(X72,bool)),X73),s(X72,X74))),s(X72,X75))) | s(bool,'F_0') = s(bool,happ(s(fun(X72,bool),happ(s(fun(X72,fun(X72,bool)),X73),s(X72,X75))),s(X72,X74))) | p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X77) != s(bool,'const.relation.symmetric_1'(s(fun(X72,fun(X72,bool)),X73)))) )),
  inference(resolution,[],[f1123,f125])).
fof(f1123,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X8,fun(X8,bool)),X9)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X9),s(X8,X11))),s(X8,X10))) | s(bool,'F_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X9),s(X8,X10))),s(X8,X11)))) )),
  inference(superposition,[],[f118,f122])).
fof(f7221,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,happ(s(fun(X53,bool),happ(s(fun(X53,fun(X53,bool)),X54),s(X53,X55))),s(X53,X56))) | s(bool,'F_0') = s(bool,happ(s(fun(X53,bool),happ(s(fun(X53,fun(X53,bool)),X54),s(X53,X56))),s(X53,X55))) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.symmetric_1'(s(fun(X53,fun(X53,bool)),X54))) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f1123,f89])).
fof(f7219,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),X40),s(X39,X41))),s(X39,X42))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),X40),s(X39,X42))),s(X39,X41))) | p(s(bool,X43)) | ~p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.relation.symmetric_1'(s(fun(X39,fun(X39,bool)),X40)))) )),
  inference(resolution,[],[f1123,f88])).
fof(f7217,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),happ(s(fun(X25,fun(X25,bool)),X26),s(X25,X27))),s(X25,X28))) | s(bool,'F_0') = s(bool,happ(s(fun(X25,bool),happ(s(fun(X25,fun(X25,bool)),X26),s(X25,X28))),s(X25,X27))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.symmetric_1'(s(fun(X25,fun(X25,bool)),X26)))) )),
  inference(resolution,[],[f1123,f87])).
fof(f7216,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(X18,bool),happ(s(fun(X18,fun(X18,bool)),X19),s(X18,X20))),s(X18,X21))) | s(bool,'F_0') = s(bool,happ(s(fun(X18,bool),happ(s(fun(X18,fun(X18,bool)),X19),s(X18,X21))),s(X18,X20))) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.relation.symmetric_1'(s(fun(X18,fun(X18,bool)),X19))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f1123,f85])).
fof(f7214,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X6))),s(X4,X7))) | s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X7))),s(X4,X6))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5)))) )),
  inference(resolution,[],[f1123,f83])).
fof(f4039,plain,(
  ( ! [X520,X516,X514,X519,X513,X515] : (s(bool,X519) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X513,sK1),X514),s(X513,X515),s(X513,X516))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X513,sK1),X514),s(X513,X516),s(X513,X515))) | s(bool,'F_0') = s(bool,X520) | s(bool,X519) != s(bool,X520)) ) | $spl128),
  inference(cnf_transformation,[],[f4039_D])).
fof(f4039_D,plain,(
  ( ! [X520,X516,X514,X519,X513,X515] : (s(bool,X519) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X513,sK1),X514),s(X513,X515),s(X513,X516))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X513,sK1),X514),s(X513,X516),s(X513,X515))) | s(bool,'F_0') = s(bool,X520) | s(bool,X519) != s(bool,X520)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7212,plain,(
  ( ! [X509,X511,X505,X507,X512,X508,X510,X506,X513] : (s(bool,X509) != s(bool,X510) | s(bool,X511) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X505,sK1),X506),s(X505,X507),s(X505,X508))) | s(bool,X511) != s(bool,X512) | s(bool,'F_0') = s(bool,X512) | s(bool,X510) != s(bool,X513) | s(bool,'T_0') = s(bool,X513) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X505,sK1),X506),s(X505,X508),s(X505,X507)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f7103,f789])).
fof(f7103,plain,(
  ( ! [X509,X511,X505,X507,X512,X508,X510,X506,X513] : (s(bool,'F_0') != s(bool,X509) | s(bool,X509) != s(bool,X510) | s(bool,X511) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X505,sK1),X506),s(X505,X507),s(X505,X508))) | s(bool,X511) != s(bool,X512) | s(bool,'F_0') = s(bool,X512) | s(bool,X510) != s(bool,X513) | s(bool,'T_0') = s(bool,X513) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X505,sK1),X506),s(X505,X508),s(X505,X507)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f2343])).
fof(f4036,plain,(
  ( ! [X486,X480,X482,X481,X483] : (s(bool,X486) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X480,sK1),X481),s(X480,X482),s(X480,X483))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X480,sK1),X481),s(X480,X483),s(X480,X482))) | s(bool,'F_0') = s(bool,X486)) ) | $spl126),
  inference(cnf_transformation,[],[f4036_D])).
fof(f4036_D,plain,(
  ( ! [X486,X480,X482,X481,X483] : (s(bool,X486) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X480,sK1),X481),s(X480,X482),s(X480,X483))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X480,sK1),X481),s(X480,X483),s(X480,X482))) | s(bool,'F_0') = s(bool,X486)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f7210,plain,(
  ( ! [X476,X478,X472,X474,X475,X477,X479,X473] : (s(bool,X476) != s(bool,X477) | s(bool,X478) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X472,sK1),X473),s(X472,X474),s(X472,X475))) | s(bool,'F_0') = s(bool,X478) | s(bool,X477) != s(bool,X479) | s(bool,'T_0') = s(bool,X479) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X472,sK1),X473),s(X472,X475),s(X472,X474)))) ) | ($spl34 | $spl38 | $spl56)),
  inference(subsumption_resolution,[],[f7099,f789])).
fof(f7099,plain,(
  ( ! [X476,X478,X472,X474,X475,X477,X479,X473] : (s(bool,'F_0') != s(bool,X476) | s(bool,X476) != s(bool,X477) | s(bool,X478) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X472,sK1),X473),s(X472,X474),s(X472,X475))) | s(bool,'F_0') = s(bool,X478) | s(bool,X477) != s(bool,X479) | s(bool,'T_0') = s(bool,X479) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X472,sK1),X473),s(X472,X475),s(X472,X474)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f2343])).
fof(f2733,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3))) | s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2))) | ~p(s(bool,X5))) ) | $spl88),
  inference(cnf_transformation,[],[f2733_D])).
fof(f2733_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3))) | s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2))) | ~p(s(bool,X5))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f7196,plain,(
  ( ! [X302,X304,X306,X303,X305,X307] : (p(s(bool,X306)) | ~p(s(bool,X307)) | s(bool,X307) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X302,sK1),X303),s(X302,X304),s(X302,X305))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X302,sK1),X303),s(X302,X305),s(X302,X304)))) )),
  inference(subsumption_resolution,[],[f7068,f115])).
fof(f7068,plain,(
  ( ! [X302,X304,X306,X303,X305,X307] : (p(s(bool,'F_0')) | p(s(bool,X306)) | ~p(s(bool,X307)) | s(bool,X307) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X302,sK1),X303),s(X302,X304),s(X302,X305))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X302,sK1),X303),s(X302,X305),s(X302,X304)))) )),
  inference(superposition,[],[f125,f2343])).
fof(f7194,plain,(
  ( ! [X243,X241,X246,X244,X242,X247,X245] : (~p(s(bool,X245)) | ~p(s(bool,X246)) | s(bool,X246) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X241,sK1),X242),s(X241,X243),s(X241,X244))) | s(bool,X245) != s(bool,X247) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X241,sK1),X242),s(X241,X244),s(X241,X243)))) )),
  inference(subsumption_resolution,[],[f7057,f115])).
fof(f7057,plain,(
  ( ! [X243,X241,X246,X244,X242,X247,X245] : (p(s(bool,'F_0')) | ~p(s(bool,X245)) | ~p(s(bool,X246)) | s(bool,X246) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X241,sK1),X242),s(X241,X243),s(X241,X244))) | s(bool,X245) != s(bool,X247) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X241,sK1),X242),s(X241,X244),s(X241,X243)))) )),
  inference(superposition,[],[f89,f2343])).
fof(f7191,plain,(
  ( ! [X191,X189,X187,X185,X190,X188,X186] : (p(s(bool,X189)) | ~p(s(bool,X190)) | ~p(s(bool,X191)) | s(bool,X190) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X185,sK1),X186),s(X185,X187),s(X185,X188))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X185,sK1),X186),s(X185,X188),s(X185,X187)))) )),
  inference(subsumption_resolution,[],[f7049,f115])).
fof(f7049,plain,(
  ( ! [X191,X189,X187,X185,X190,X188,X186] : (p(s(bool,'F_0')) | p(s(bool,X189)) | ~p(s(bool,X190)) | ~p(s(bool,X191)) | s(bool,X190) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X185,sK1),X186),s(X185,X187),s(X185,X188))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X185,sK1),X186),s(X185,X188),s(X185,X187)))) )),
  inference(superposition,[],[f88,f2343])).
fof(f7189,plain,(
  ( ! [X156,X154,X152,X151,X155,X153,X150] : (p(s(bool,X154)) | ~p(s(bool,X155)) | s(bool,X154) != s(bool,X156) | s(bool,X155) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X150,sK1),X151),s(X150,X152),s(X150,X153))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X150,sK1),X151),s(X150,X153),s(X150,X152)))) )),
  inference(subsumption_resolution,[],[f7044,f115])).
fof(f7044,plain,(
  ( ! [X156,X154,X152,X151,X155,X153,X150] : (p(s(bool,'F_0')) | p(s(bool,X154)) | ~p(s(bool,X155)) | s(bool,X154) != s(bool,X156) | s(bool,X155) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X150,sK1),X151),s(X150,X152),s(X150,X153))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X150,sK1),X151),s(X150,X153),s(X150,X152)))) )),
  inference(superposition,[],[f87,f2343])).
fof(f7186,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (~p(s(bool,X105)) | ~p(s(bool,X106)) | s(bool,X106) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X101,sK1),X102),s(X101,X103),s(X101,X104))) | s(bool,X105) != s(bool,X107) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X101,sK1),X102),s(X101,X104),s(X101,X103)))) )),
  inference(subsumption_resolution,[],[f7037,f115])).
fof(f7037,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (p(s(bool,'F_0')) | ~p(s(bool,X105)) | ~p(s(bool,X106)) | s(bool,X106) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X101,sK1),X102),s(X101,X103),s(X101,X104))) | s(bool,X105) != s(bool,X107) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X101,sK1),X102),s(X101,X104),s(X101,X103)))) )),
  inference(superposition,[],[f85,f2343])).
fof(f7183,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X56) != s(bool,X58) | s(bool,X57) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X52,sK1),X53),s(X52,X54),s(X52,X55))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X52,sK1),X53),s(X52,X55),s(X52,X54)))) )),
  inference(subsumption_resolution,[],[f7030,f115])).
fof(f7030,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,'F_0')) | p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X56) != s(bool,X58) | s(bool,X57) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X52,sK1),X53),s(X52,X54),s(X52,X55))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X52,sK1),X53),s(X52,X55),s(X52,X54)))) )),
  inference(superposition,[],[f83,f2343])).
fof(f3989,plain,(
  ( ! [X428,X430,X432,X429,X427,X433] : (s(bool,X432) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X427,sK1),X428),s(X427,X429),s(X427,X430))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X427,sK1),X428),s(X427,X430),s(X427,X429))) | s(bool,'T_0') = s(bool,X433) | s(bool,X432) != s(bool,X433)) ) | $spl124),
  inference(cnf_transformation,[],[f3989_D])).
fof(f3989_D,plain,(
  ( ! [X428,X430,X432,X429,X427,X433] : (s(bool,X432) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X427,sK1),X428),s(X427,X429),s(X427,X430))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X427,sK1),X428),s(X427,X430),s(X427,X429))) | s(bool,'T_0') = s(bool,X433) | s(bool,X432) != s(bool,X433)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f7180,plain,(
  ( ! [X499,X500,X502,X496,X498,X504,X501,X503,X497] : (s(bool,X501) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X496,sK1),X497),s(X496,X498),s(X496,X499))) | s(bool,X500) != s(bool,X502) | s(bool,X502) != s(bool,X503) | s(bool,'F_0') = s(bool,X503) | s(bool,X501) != s(bool,X504) | s(bool,'T_0') = s(bool,X504) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X496,sK1),X497),s(X496,X499),s(X496,X498)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f6999,f789])).
fof(f6999,plain,(
  ( ! [X499,X500,X502,X496,X498,X504,X501,X503,X497] : (s(bool,'T_0') != s(bool,X500) | s(bool,X501) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X496,sK1),X497),s(X496,X498),s(X496,X499))) | s(bool,X500) != s(bool,X502) | s(bool,X502) != s(bool,X503) | s(bool,'F_0') = s(bool,X503) | s(bool,X501) != s(bool,X504) | s(bool,'T_0') = s(bool,X504) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X496,sK1),X497),s(X496,X499),s(X496,X498)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f2343])).
fof(f7179,plain,(
  ( ! [X466,X469,X471,X465,X467,X468,X470,X464] : (s(bool,X469) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X464,sK1),X465),s(X464,X466),s(X464,X467))) | s(bool,X468) != s(bool,X470) | s(bool,'F_0') = s(bool,X470) | s(bool,X469) != s(bool,X471) | s(bool,'T_0') = s(bool,X471) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X464,sK1),X465),s(X464,X467),s(X464,X466)))) ) | ($spl6 | $spl34 | $spl56)),
  inference(subsumption_resolution,[],[f6995,f589])).
fof(f6995,plain,(
  ( ! [X466,X469,X471,X465,X467,X468,X470,X464] : (s(bool,'T_0') != s(bool,X468) | s(bool,X469) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X464,sK1),X465),s(X464,X466),s(X464,X467))) | s(bool,X468) != s(bool,X470) | s(bool,'F_0') = s(bool,X470) | s(bool,X469) != s(bool,X471) | s(bool,'T_0') = s(bool,X471) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X464,sK1),X465),s(X464,X467),s(X464,X466)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f2343])).
fof(f7177,plain,(
  ( ! [X420,X422,X424,X421,X423,X419,X425] : (s(bool,X424) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X419,sK1),X420),s(X419,X421),s(X419,X422))) | s(bool,'F_0') = s(bool,X423) | s(bool,X424) != s(bool,X425) | s(bool,'T_0') = s(bool,X425) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X419,sK1),X420),s(X419,X422),s(X419,X421)))) ) | ($spl34 | $spl40)),
  inference(subsumption_resolution,[],[f6988,f122])).
fof(f6988,plain,(
  ( ! [X420,X422,X424,X421,X423,X419,X425] : (s(bool,'T_0') != s(bool,X423) | s(bool,X424) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X419,sK1),X420),s(X419,X421),s(X419,X422))) | s(bool,'F_0') = s(bool,X423) | s(bool,X424) != s(bool,X425) | s(bool,'T_0') = s(bool,X425) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X419,sK1),X420),s(X419,X422),s(X419,X421)))) ) | ($spl34 | $spl40)),
  inference(superposition,[],[f1342,f2343])).
fof(f3987,plain,(
  ( ! [X389,X391,X390,X392,X394] : (s(bool,X394) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X389,sK1),X390),s(X389,X391),s(X389,X392))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X389,sK1),X390),s(X389,X392),s(X389,X391))) | s(bool,'T_0') = s(bool,X394)) ) | $spl122),
  inference(cnf_transformation,[],[f3987_D])).
fof(f3987_D,plain,(
  ( ! [X389,X391,X390,X392,X394] : (s(bool,X394) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X389,sK1),X390),s(X389,X391),s(X389,X392))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X389,sK1),X390),s(X389,X392),s(X389,X391))) | s(bool,'T_0') = s(bool,X394)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f7176,plain,(
  ( ! [X383,X385,X382,X384,X386,X381] : (s(bool,X386) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X381,sK1),X382),s(X381,X383),s(X381,X384))) | s(bool,'F_0') = s(bool,X385) | s(bool,'T_0') = s(bool,X386) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X381,sK1),X382),s(X381,X384),s(X381,X383)))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6981,f122])).
fof(f6981,plain,(
  ( ! [X383,X385,X382,X384,X386,X381] : (s(bool,'T_0') != s(bool,X385) | s(bool,X386) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X381,sK1),X382),s(X381,X383),s(X381,X384))) | s(bool,'F_0') = s(bool,X385) | s(bool,'T_0') = s(bool,X386) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X381,sK1),X382),s(X381,X384),s(X381,X383)))) ) | ($spl34 | $spl38)),
  inference(superposition,[],[f789,f2343])).
fof(f2826,plain,(
  ( ! [X6,X4,X7,X5,X9] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X4,sK1),X5),s(X4,X6),s(X4,X7))) | s(bool,X9) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X4,sK1),X5),s(X4,X7),s(X4,X6))) | p(s(bool,X9))) ) | $spl92),
  inference(cnf_transformation,[],[f2826_D])).
fof(f2826_D,plain,(
  ( ! [X6,X4,X7,X5,X9] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X4,sK1),X5),s(X4,X6),s(X4,X7))) | s(bool,X9) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X4,sK1),X5),s(X4,X7),s(X4,X6))) | p(s(bool,X9))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f7172,plain,(
  ( ! [X315,X316,X318,X314,X317,X319] : (p(s(bool,X318)) | p(s(bool,X319)) | s(bool,X319) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X314,sK1),X315),s(X314,X316),s(X314,X317))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X314,sK1),X315),s(X314,X317),s(X314,X316)))) )),
  inference(subsumption_resolution,[],[f6967,f126])).
fof(f6967,plain,(
  ( ! [X315,X316,X318,X314,X317,X319] : (~p(s(bool,'T_0')) | p(s(bool,X318)) | p(s(bool,X319)) | s(bool,X319) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X314,sK1),X315),s(X314,X316),s(X314,X317))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X314,sK1),X315),s(X314,X317),s(X314,X316)))) )),
  inference(superposition,[],[f125,f2343])).
fof(f7170,plain,(
  ( ! [X255,X261,X257,X259,X260,X256,X258] : (p(s(bool,X259)) | ~p(s(bool,X260)) | s(bool,X259) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X255,sK1),X256),s(X255,X257),s(X255,X258))) | s(bool,X260) != s(bool,X261) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X255,sK1),X256),s(X255,X258),s(X255,X257)))) )),
  inference(subsumption_resolution,[],[f6956,f126])).
fof(f6956,plain,(
  ( ! [X255,X261,X257,X259,X260,X256,X258] : (~p(s(bool,'T_0')) | p(s(bool,X259)) | ~p(s(bool,X260)) | s(bool,X259) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X255,sK1),X256),s(X255,X257),s(X255,X258))) | s(bool,X260) != s(bool,X261) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X255,sK1),X256),s(X255,X258),s(X255,X257)))) )),
  inference(superposition,[],[f89,f2343])).
fof(f7166,plain,(
  ( ! [X204,X202,X200,X199,X205,X203,X201] : (p(s(bool,X203)) | ~p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X199,sK1),X200),s(X199,X201),s(X199,X202))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X199,sK1),X200),s(X199,X202),s(X199,X201)))) )),
  inference(subsumption_resolution,[],[f6948,f126])).
fof(f6948,plain,(
  ( ! [X204,X202,X200,X199,X205,X203,X201] : (~p(s(bool,'T_0')) | p(s(bool,X203)) | ~p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X199,sK1),X200),s(X199,X201),s(X199,X202))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X199,sK1),X200),s(X199,X202),s(X199,X201)))) )),
  inference(superposition,[],[f88,f2343])).
fof(f7164,plain,(
  ( ! [X167,X165,X169,X166,X164,X170,X168] : (p(s(bool,X168)) | p(s(bool,X169)) | s(bool,X168) != s(bool,X170) | s(bool,X169) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X164,sK1),X165),s(X164,X166),s(X164,X167))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X164,sK1),X165),s(X164,X167),s(X164,X166)))) )),
  inference(subsumption_resolution,[],[f6943,f126])).
fof(f6943,plain,(
  ( ! [X167,X165,X169,X166,X164,X170,X168] : (~p(s(bool,'T_0')) | p(s(bool,X168)) | p(s(bool,X169)) | s(bool,X168) != s(bool,X170) | s(bool,X169) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X164,sK1),X165),s(X164,X166),s(X164,X167))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X164,sK1),X165),s(X164,X167),s(X164,X166)))) )),
  inference(superposition,[],[f87,f2343])).
fof(f7161,plain,(
  ( ! [X121,X118,X116,X120,X119,X117,X115] : (p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X115,sK1),X116),s(X115,X117),s(X115,X118))) | s(bool,X120) != s(bool,X121) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X115,sK1),X116),s(X115,X118),s(X115,X117)))) )),
  inference(subsumption_resolution,[],[f6936,f126])).
fof(f6936,plain,(
  ( ! [X121,X118,X116,X120,X119,X117,X115] : (~p(s(bool,'T_0')) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X115,sK1),X116),s(X115,X117),s(X115,X118))) | s(bool,X120) != s(bool,X121) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X115,sK1),X116),s(X115,X118),s(X115,X117)))) )),
  inference(superposition,[],[f85,f2343])).
fof(f7157,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (p(s(bool,X70)) | p(s(bool,X71)) | s(bool,X70) != s(bool,X72) | s(bool,X71) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X66,sK1),X67),s(X66,X68),s(X66,X69))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X66,sK1),X67),s(X66,X69),s(X66,X68)))) )),
  inference(subsumption_resolution,[],[f6929,f126])).
fof(f6929,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X70)) | p(s(bool,X71)) | s(bool,X70) != s(bool,X72) | s(bool,X71) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X66,sK1),X67),s(X66,X68),s(X66,X69))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X66,sK1),X67),s(X66,X69),s(X66,X68)))) )),
  inference(superposition,[],[f83,f2343])).
fof(f3569,plain,(
  ( ! [X358,X354,X357,X353] : (s(bool,X357) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X353))),s(sK1,X354))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X354))),s(sK1,X353))) | s(bool,'F_0') = s(bool,X358) | s(bool,X357) != s(bool,X358)) ) | $spl110),
  inference(cnf_transformation,[],[f3569_D])).
fof(f3569_D,plain,(
  ( ! [X358,X354,X357,X353] : (s(bool,X357) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X353))),s(sK1,X354))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X354))),s(sK1,X353))) | s(bool,'F_0') = s(bool,X358) | s(bool,X357) != s(bool,X358)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f6812,plain,(
  ( ! [X350,X352,X354,X349,X351,X353,X355] : (s(bool,X351) != s(bool,X352) | s(bool,X353) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X349))),s(sK1,X350))) | s(bool,X353) != s(bool,X354) | s(bool,'F_0') = s(bool,X354) | s(bool,X352) != s(bool,X355) | s(bool,'T_0') = s(bool,X355) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X350))),s(sK1,X349)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f6713,f789])).
fof(f6713,plain,(
  ( ! [X350,X352,X354,X349,X351,X353,X355] : (s(bool,'F_0') != s(bool,X351) | s(bool,X351) != s(bool,X352) | s(bool,X353) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X349))),s(sK1,X350))) | s(bool,X353) != s(bool,X354) | s(bool,'F_0') = s(bool,X354) | s(bool,X352) != s(bool,X355) | s(bool,'T_0') = s(bool,X355) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X350))),s(sK1,X349)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f1951])).
fof(f3565,plain,(
  ( ! [X329,X332,X328] : (s(bool,X332) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X328))),s(sK1,X329))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X329))),s(sK1,X328))) | s(bool,'F_0') = s(bool,X332)) ) | $spl106),
  inference(cnf_transformation,[],[f3565_D])).
fof(f3565_D,plain,(
  ( ! [X329,X332,X328] : (s(bool,X332) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X328))),s(sK1,X329))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X329))),s(sK1,X328))) | s(bool,'F_0') = s(bool,X332)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f6810,plain,(
  ( ! [X325,X327,X329,X324,X326,X328] : (s(bool,X326) != s(bool,X327) | s(bool,X328) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X324))),s(sK1,X325))) | s(bool,'F_0') = s(bool,X328) | s(bool,X327) != s(bool,X329) | s(bool,'T_0') = s(bool,X329) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X325))),s(sK1,X324)))) ) | ($spl34 | $spl38 | $spl56)),
  inference(subsumption_resolution,[],[f6709,f789])).
fof(f6709,plain,(
  ( ! [X325,X327,X329,X324,X326,X328] : (s(bool,'F_0') != s(bool,X326) | s(bool,X326) != s(bool,X327) | s(bool,X328) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X324))),s(sK1,X325))) | s(bool,'F_0') = s(bool,X328) | s(bool,X327) != s(bool,X329) | s(bool,'T_0') = s(bool,X329) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X325))),s(sK1,X324)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f1951])).
fof(f2308,plain,(
  ( ! [X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1))) | s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0))) | ~p(s(bool,X3))) ) | $spl80),
  inference(cnf_transformation,[],[f2308_D])).
fof(f2308_D,plain,(
  ( ! [X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1))) | s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0))) | ~p(s(bool,X3))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f6796,plain,(
  ( ! [X218,X216,X219,X217] : (p(s(bool,X218)) | ~p(s(bool,X219)) | s(bool,X219) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X216))),s(sK1,X217))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X217))),s(sK1,X216)))) )),
  inference(subsumption_resolution,[],[f6678,f115])).
fof(f6678,plain,(
  ( ! [X218,X216,X219,X217] : (p(s(bool,'F_0')) | p(s(bool,X218)) | ~p(s(bool,X219)) | s(bool,X219) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X216))),s(sK1,X217))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X217))),s(sK1,X216)))) )),
  inference(superposition,[],[f125,f1951])).
fof(f6794,plain,(
  ( ! [X182,X180,X183,X181,X179] : (~p(s(bool,X181)) | ~p(s(bool,X182)) | s(bool,X182) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X179))),s(sK1,X180))) | s(bool,X181) != s(bool,X183) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X180))),s(sK1,X179)))) )),
  inference(subsumption_resolution,[],[f6667,f115])).
fof(f6667,plain,(
  ( ! [X182,X180,X183,X181,X179] : (p(s(bool,'F_0')) | ~p(s(bool,X181)) | ~p(s(bool,X182)) | s(bool,X182) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X179))),s(sK1,X180))) | s(bool,X181) != s(bool,X183) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X180))),s(sK1,X179)))) )),
  inference(superposition,[],[f89,f1951])).
fof(f6791,plain,(
  ( ! [X142,X140,X143,X141,X139] : (p(s(bool,X141)) | ~p(s(bool,X142)) | ~p(s(bool,X143)) | s(bool,X142) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X139))),s(sK1,X140))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X140))),s(sK1,X139)))) )),
  inference(subsumption_resolution,[],[f6659,f115])).
fof(f6659,plain,(
  ( ! [X142,X140,X143,X141,X139] : (p(s(bool,'F_0')) | p(s(bool,X141)) | ~p(s(bool,X142)) | ~p(s(bool,X143)) | s(bool,X142) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X139))),s(sK1,X140))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X140))),s(sK1,X139)))) )),
  inference(superposition,[],[f88,f1951])).
fof(f6789,plain,(
  ( ! [X118,X116,X114,X117,X115] : (p(s(bool,X116)) | ~p(s(bool,X117)) | s(bool,X116) != s(bool,X118) | s(bool,X117) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X114))),s(sK1,X115))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X115))),s(sK1,X114)))) )),
  inference(subsumption_resolution,[],[f6654,f115])).
fof(f6654,plain,(
  ( ! [X118,X116,X114,X117,X115] : (p(s(bool,'F_0')) | p(s(bool,X116)) | ~p(s(bool,X117)) | s(bool,X116) != s(bool,X118) | s(bool,X117) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X114))),s(sK1,X115))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X115))),s(sK1,X114)))) )),
  inference(superposition,[],[f87,f1951])).
fof(f6786,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X79))),s(sK1,X80))) | s(bool,X81) != s(bool,X83) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X80))),s(sK1,X79)))) )),
  inference(subsumption_resolution,[],[f6647,f115])).
fof(f6647,plain,(
  ( ! [X80,X83,X81,X79,X82] : (p(s(bool,'F_0')) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X79))),s(sK1,X80))) | s(bool,X81) != s(bool,X83) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X80))),s(sK1,X79)))) )),
  inference(superposition,[],[f85,f1951])).
fof(f6783,plain,(
  ( ! [X47,X45,X48,X46,X44] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X44))),s(sK1,X45))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X45))),s(sK1,X44)))) )),
  inference(subsumption_resolution,[],[f6640,f115])).
fof(f6640,plain,(
  ( ! [X47,X45,X48,X46,X44] : (p(s(bool,'F_0')) | p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X44))),s(sK1,X45))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X45))),s(sK1,X44)))) )),
  inference(superposition,[],[f83,f1951])).
fof(f3515,plain,(
  ( ! [X292,X294,X295,X291] : (s(bool,X294) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X291))),s(sK1,X292))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X292))),s(sK1,X291))) | s(bool,'T_0') = s(bool,X295) | s(bool,X294) != s(bool,X295)) ) | $spl102),
  inference(cnf_transformation,[],[f3515_D])).
fof(f3515_D,plain,(
  ( ! [X292,X294,X295,X291] : (s(bool,X294) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X291))),s(sK1,X292))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X292))),s(sK1,X291))) | s(bool,'T_0') = s(bool,X295) | s(bool,X294) != s(bool,X295)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f6778,plain,(
  ( ! [X344,X346,X343,X345,X347,X342,X348] : (s(bool,X345) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X342))),s(sK1,X343))) | s(bool,X344) != s(bool,X346) | s(bool,X346) != s(bool,X347) | s(bool,'F_0') = s(bool,X347) | s(bool,X345) != s(bool,X348) | s(bool,'T_0') = s(bool,X348) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X343))),s(sK1,X342)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f6601,f789])).
fof(f6601,plain,(
  ( ! [X344,X346,X343,X345,X347,X342,X348] : (s(bool,'T_0') != s(bool,X344) | s(bool,X345) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X342))),s(sK1,X343))) | s(bool,X344) != s(bool,X346) | s(bool,X346) != s(bool,X347) | s(bool,'F_0') = s(bool,X347) | s(bool,X345) != s(bool,X348) | s(bool,'T_0') = s(bool,X348) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X343))),s(sK1,X342)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f1951])).
fof(f6777,plain,(
  ( ! [X321,X323,X318,X320,X322,X319] : (s(bool,X321) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X318))),s(sK1,X319))) | s(bool,X320) != s(bool,X322) | s(bool,'F_0') = s(bool,X322) | s(bool,X321) != s(bool,X323) | s(bool,'T_0') = s(bool,X323) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X319))),s(sK1,X318)))) ) | ($spl6 | $spl34 | $spl56)),
  inference(subsumption_resolution,[],[f6597,f589])).
fof(f6597,plain,(
  ( ! [X321,X323,X318,X320,X322,X319] : (s(bool,'T_0') != s(bool,X320) | s(bool,X321) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X318))),s(sK1,X319))) | s(bool,X320) != s(bool,X322) | s(bool,'F_0') = s(bool,X322) | s(bool,X321) != s(bool,X323) | s(bool,'T_0') = s(bool,X323) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X319))),s(sK1,X318)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f1951])).
fof(f6775,plain,(
  ( ! [X288,X290,X287,X289,X291] : (s(bool,X290) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X287))),s(sK1,X288))) | s(bool,'F_0') = s(bool,X289) | s(bool,X290) != s(bool,X291) | s(bool,'T_0') = s(bool,X291) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X288))),s(sK1,X287)))) ) | ($spl34 | $spl40)),
  inference(subsumption_resolution,[],[f6590,f122])).
fof(f6590,plain,(
  ( ! [X288,X290,X287,X289,X291] : (s(bool,'T_0') != s(bool,X289) | s(bool,X290) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X287))),s(sK1,X288))) | s(bool,'F_0') = s(bool,X289) | s(bool,X290) != s(bool,X291) | s(bool,'T_0') = s(bool,X291) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X288))),s(sK1,X287)))) ) | ($spl34 | $spl40)),
  inference(superposition,[],[f1342,f1951])).
fof(f3513,plain,(
  ( ! [X267,X268,X270] : (s(bool,X270) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X267))),s(sK1,X268))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X268))),s(sK1,X267))) | s(bool,'T_0') = s(bool,X270)) ) | $spl100),
  inference(cnf_transformation,[],[f3513_D])).
fof(f3513_D,plain,(
  ( ! [X267,X268,X270] : (s(bool,X270) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X267))),s(sK1,X268))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X268))),s(sK1,X267))) | s(bool,'T_0') = s(bool,X270)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f6774,plain,(
  ( ! [X263,X265,X264,X266] : (s(bool,X266) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X263))),s(sK1,X264))) | s(bool,'F_0') = s(bool,X265) | s(bool,'T_0') = s(bool,X266) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X264))),s(sK1,X263)))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6583,f122])).
fof(f6583,plain,(
  ( ! [X263,X265,X264,X266] : (s(bool,'T_0') != s(bool,X265) | s(bool,X266) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X263))),s(sK1,X264))) | s(bool,'F_0') = s(bool,X265) | s(bool,'T_0') = s(bool,X266) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X264))),s(sK1,X263)))) ) | ($spl34 | $spl38)),
  inference(superposition,[],[f789,f1951])).
fof(f2769,plain,(
  ( ! [X2,X5,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3))) | s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3))),s(sK1,X2))) | p(s(bool,X5))) ) | $spl90),
  inference(cnf_transformation,[],[f2769_D])).
fof(f2769_D,plain,(
  ( ! [X2,X5,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3))) | s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3))),s(sK1,X2))) | p(s(bool,X5))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f6770,plain,(
  ( ! [X227,X225,X226,X224] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X224))),s(sK1,X225))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X225))),s(sK1,X224)))) )),
  inference(subsumption_resolution,[],[f6569,f126])).
fof(f6569,plain,(
  ( ! [X227,X225,X226,X224] : (~p(s(bool,'T_0')) | p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X227) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X224))),s(sK1,X225))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X225))),s(sK1,X224)))) )),
  inference(superposition,[],[f125,f1951])).
fof(f6768,plain,(
  ( ! [X191,X189,X192,X190,X193] : (p(s(bool,X191)) | ~p(s(bool,X192)) | s(bool,X191) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X189))),s(sK1,X190))) | s(bool,X192) != s(bool,X193) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X190))),s(sK1,X189)))) )),
  inference(subsumption_resolution,[],[f6558,f126])).
fof(f6558,plain,(
  ( ! [X191,X189,X192,X190,X193] : (~p(s(bool,'T_0')) | p(s(bool,X191)) | ~p(s(bool,X192)) | s(bool,X191) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X189))),s(sK1,X190))) | s(bool,X192) != s(bool,X193) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X190))),s(sK1,X189)))) )),
  inference(superposition,[],[f89,f1951])).
fof(f6764,plain,(
  ( ! [X152,X151,X149,X153,X150] : (p(s(bool,X151)) | ~p(s(bool,X152)) | p(s(bool,X153)) | s(bool,X153) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X149))),s(sK1,X150))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X150))),s(sK1,X149)))) )),
  inference(subsumption_resolution,[],[f6550,f126])).
fof(f6550,plain,(
  ( ! [X152,X151,X149,X153,X150] : (~p(s(bool,'T_0')) | p(s(bool,X151)) | ~p(s(bool,X152)) | p(s(bool,X153)) | s(bool,X153) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X149))),s(sK1,X150))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X150))),s(sK1,X149)))) )),
  inference(superposition,[],[f88,f1951])).
fof(f6762,plain,(
  ( ! [X127,X125,X128,X126,X124] : (p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X126) != s(bool,X128) | s(bool,X127) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X124))),s(sK1,X125))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X125))),s(sK1,X124)))) )),
  inference(subsumption_resolution,[],[f6545,f126])).
fof(f6545,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,'T_0')) | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X126) != s(bool,X128) | s(bool,X127) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X124))),s(sK1,X125))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X125))),s(sK1,X124)))) )),
  inference(superposition,[],[f87,f1951])).
fof(f6759,plain,(
  ( ! [X92,X90,X93,X91,X89] : (p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X89))),s(sK1,X90))) | s(bool,X92) != s(bool,X93) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X90))),s(sK1,X89)))) )),
  inference(subsumption_resolution,[],[f6538,f126])).
fof(f6538,plain,(
  ( ! [X92,X90,X93,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X89))),s(sK1,X90))) | s(bool,X92) != s(bool,X93) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X90))),s(sK1,X89)))) )),
  inference(superposition,[],[f85,f1951])).
fof(f6755,plain,(
  ( ! [X57,X54,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X56) != s(bool,X58) | s(bool,X57) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X54))),s(sK1,X55))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X55))),s(sK1,X54)))) )),
  inference(subsumption_resolution,[],[f6531,f126])).
fof(f6531,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X56) != s(bool,X58) | s(bool,X57) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X54))),s(sK1,X55))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X55))),s(sK1,X54)))) )),
  inference(superposition,[],[f83,f1951])).
fof(f3130,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,X47) != s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X46))),s(X43,X45))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X43,fun(X43,bool)),X44))) | p(s(bool,X47)) | ~p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X45))),s(X43,X46))))) ) | $spl96),
  inference(cnf_transformation,[],[f3130_D])).
fof(f3130_D,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,X47) != s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X46))),s(X43,X45))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X43,fun(X43,bool)),X44))) | p(s(bool,X47)) | ~p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X45))),s(X43,X46))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f6410,plain,(
  ( ! [X327,X329,X331,X326,X328,X330] : (~p(s(bool,happ(s(fun(X326,bool),happ(s(fun(X326,fun(X326,bool)),X327),s(X326,X329))),s(X326,X328)))) | p(s(bool,X330)) | p(s(bool,X331)) | s(bool,X331) != s(bool,happ(s(fun(X326,bool),happ(s(fun(X326,fun(X326,bool)),X327),s(X326,X328))),s(X326,X329))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X326,fun(X326,bool)),X327)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f1288,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X6))),s(X4,X7))) = s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X7))),s(X4,X6))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5)))) )),
  inference(subsumption_resolution,[],[f1155,f126])).
fof(f1155,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X6))),s(X4,X7))) = s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X7))),s(X4,X6))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5)))) )),
  inference(superposition,[],[f118,f122])).
fof(f3129,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,X40) != s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X39))),s(X36,X38))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X36,fun(X36,bool)),X37))) | p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X38))),s(X36,X39)))) | ~p(s(bool,X40))) ) | $spl94),
  inference(cnf_transformation,[],[f3129_D])).
fof(f3129_D,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,X40) != s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X39))),s(X36,X38))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X36,fun(X36,bool)),X37))) | p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X38))),s(X36,X39)))) | ~p(s(bool,X40))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f6408,plain,(
  ( ! [X315,X316,X318,X314,X317,X319] : (p(s(bool,happ(s(fun(X314,bool),happ(s(fun(X314,fun(X314,bool)),X315),s(X314,X317))),s(X314,X316)))) | p(s(bool,X318)) | ~p(s(bool,X319)) | s(bool,X319) != s(bool,happ(s(fun(X314,bool),happ(s(fun(X314,fun(X314,bool)),X315),s(X314,X316))),s(X314,X317))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X314,fun(X314,bool)),X315)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f6407,plain,(
  ( ! [X313,X308,X310,X312,X309,X311] : (s(bool,X312) != s(bool,happ(s(fun(X308,bool),happ(s(fun(X308,fun(X308,bool)),X309),s(X308,X311))),s(X308,X310))) | p(s(bool,X313)) | ~p(s(bool,X312)) | p(s(bool,happ(s(fun(X308,bool),happ(s(fun(X308,fun(X308,bool)),X309),s(X308,X310))),s(X308,X311)))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X308,fun(X308,bool)),X309)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f6406,plain,(
  ( ! [X302,X304,X306,X303,X305,X307] : (s(bool,X306) != s(bool,happ(s(fun(X302,bool),happ(s(fun(X302,fun(X302,bool)),X303),s(X302,X305))),s(X302,X304))) | p(s(bool,X307)) | ~p(s(bool,happ(s(fun(X302,bool),happ(s(fun(X302,fun(X302,bool)),X303),s(X302,X304))),s(X302,X305)))) | p(s(bool,X306)) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X302,fun(X302,bool)),X303)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f6398,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272] : (~p(s(bool,happ(s(fun(X267,bool),happ(s(fun(X267,fun(X267,bool)),X268),s(X267,X270))),s(X267,X269)))) | p(s(bool,X271)) | ~p(s(bool,X272)) | s(bool,X271) != s(bool,happ(s(fun(X267,bool),happ(s(fun(X267,fun(X267,bool)),X268),s(X267,X269))),s(X267,X270))) | s(bool,X272) != s(bool,X273) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X267,fun(X267,bool)),X268)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f6396,plain,(
  ( ! [X255,X253,X257,X259,X254,X256,X258] : (p(s(bool,happ(s(fun(X253,bool),happ(s(fun(X253,fun(X253,bool)),X254),s(X253,X256))),s(X253,X255)))) | ~p(s(bool,X257)) | ~p(s(bool,X258)) | s(bool,X258) != s(bool,happ(s(fun(X253,bool),happ(s(fun(X253,fun(X253,bool)),X254),s(X253,X255))),s(X253,X256))) | s(bool,X257) != s(bool,X259) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X253,fun(X253,bool)),X254)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f6395,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247] : (s(bool,X250) != s(bool,happ(s(fun(X246,bool),happ(s(fun(X246,fun(X246,bool)),X247),s(X246,X249))),s(X246,X248))) | p(s(bool,X250)) | ~p(s(bool,X251)) | ~p(s(bool,happ(s(fun(X246,bool),happ(s(fun(X246,fun(X246,bool)),X247),s(X246,X248))),s(X246,X249)))) | s(bool,X251) != s(bool,X252) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X246,fun(X246,bool)),X247)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f6394,plain,(
  ( ! [X243,X241,X239,X244,X242,X240,X245] : (s(bool,X243) != s(bool,happ(s(fun(X239,bool),happ(s(fun(X239,fun(X239,bool)),X240),s(X239,X242))),s(X239,X241))) | p(s(bool,happ(s(fun(X239,bool),happ(s(fun(X239,fun(X239,bool)),X240),s(X239,X241))),s(X239,X242)))) | ~p(s(bool,X244)) | ~p(s(bool,X243)) | s(bool,X244) != s(bool,X245) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X239,fun(X239,bool)),X240)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f6390,plain,(
  ( ! [X212,X216,X215,X213,X211,X217,X214] : (~p(s(bool,happ(s(fun(X211,bool),happ(s(fun(X211,fun(X211,bool)),X212),s(X211,X214))),s(X211,X213)))) | p(s(bool,X215)) | ~p(s(bool,X216)) | p(s(bool,X217)) | s(bool,X217) != s(bool,happ(s(fun(X211,bool),happ(s(fun(X211,fun(X211,bool)),X212),s(X211,X213))),s(X211,X214))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X211,fun(X211,bool)),X212)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f6388,plain,(
  ( ! [X198,X202,X200,X199,X197,X203,X201] : (p(s(bool,happ(s(fun(X197,bool),happ(s(fun(X197,fun(X197,bool)),X198),s(X197,X200))),s(X197,X199)))) | p(s(bool,X201)) | ~p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X202) != s(bool,happ(s(fun(X197,bool),happ(s(fun(X197,fun(X197,bool)),X198),s(X197,X199))),s(X197,X200))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X197,fun(X197,bool)),X198)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f6387,plain,(
  ( ! [X191,X196,X194,X192,X190,X195,X193] : (s(bool,X194) != s(bool,happ(s(fun(X190,bool),happ(s(fun(X190,fun(X190,bool)),X191),s(X190,X193))),s(X190,X192))) | p(s(bool,X195)) | ~p(s(bool,X194)) | ~p(s(bool,X196)) | p(s(bool,happ(s(fun(X190,bool),happ(s(fun(X190,fun(X190,bool)),X191),s(X190,X192))),s(X190,X193)))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X190,fun(X190,bool)),X191)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f6386,plain,(
  ( ! [X189,X187,X185,X188,X186,X184,X183] : (s(bool,X187) != s(bool,happ(s(fun(X183,bool),happ(s(fun(X183,fun(X183,bool)),X184),s(X183,X186))),s(X183,X185))) | p(s(bool,X188)) | ~p(s(bool,happ(s(fun(X183,bool),happ(s(fun(X183,fun(X183,bool)),X184),s(X183,X185))),s(X183,X186)))) | ~p(s(bool,X189)) | p(s(bool,X187)) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X183,fun(X183,bool)),X184)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f6385,plain,(
  ( ! [X177,X182,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X176,bool),happ(s(fun(X176,fun(X176,bool)),X177),s(X176,X179))),s(X176,X178)))) | p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X182) | s(bool,X181) != s(bool,happ(s(fun(X176,bool),happ(s(fun(X176,fun(X176,bool)),X177),s(X176,X178))),s(X176,X179))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X176,fun(X176,bool)),X177)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f6383,plain,(
  ( ! [X167,X165,X163,X166,X164,X162,X168] : (p(s(bool,happ(s(fun(X162,bool),happ(s(fun(X162,fun(X162,bool)),X163),s(X162,X165))),s(X162,X164)))) | p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,X168) | s(bool,X167) != s(bool,happ(s(fun(X162,bool),happ(s(fun(X162,fun(X162,bool)),X163),s(X162,X164))),s(X162,X165))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X162,fun(X162,bool)),X163)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f6380,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141] : (s(bool,X145) != s(bool,happ(s(fun(X141,bool),happ(s(fun(X141,fun(X141,bool)),X142),s(X141,X144))),s(X141,X143))) | p(s(bool,X146)) | ~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),happ(s(fun(X141,fun(X141,bool)),X142),s(X141,X143))),s(X141,X144)))) | s(bool,X146) != s(bool,X147) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X141,fun(X141,bool)),X142)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f6379,plain,(
  ( ! [X134,X140,X138,X136,X135,X139,X137] : (s(bool,X138) != s(bool,happ(s(fun(X134,bool),happ(s(fun(X134,fun(X134,bool)),X135),s(X134,X137))),s(X134,X136))) | p(s(bool,X139)) | ~p(s(bool,happ(s(fun(X134,bool),happ(s(fun(X134,fun(X134,bool)),X135),s(X134,X136))),s(X134,X137)))) | p(s(bool,X138)) | s(bool,X139) != s(bool,X140) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X134,fun(X134,bool)),X135)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f6378,plain,(
  ( ! [X127,X132,X130,X128,X133,X131,X129] : (~p(s(bool,happ(s(fun(X127,bool),happ(s(fun(X127,fun(X127,bool)),X128),s(X127,X130))),s(X127,X129)))) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X131) != s(bool,happ(s(fun(X127,bool),happ(s(fun(X127,fun(X127,bool)),X128),s(X127,X129))),s(X127,X130))) | s(bool,X132) != s(bool,X133) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X127,fun(X127,bool)),X128)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f6376,plain,(
  ( ! [X118,X116,X114,X119,X117,X115,X113] : (p(s(bool,happ(s(fun(X113,bool),happ(s(fun(X113,fun(X113,bool)),X114),s(X113,X116))),s(X113,X115)))) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X118) != s(bool,happ(s(fun(X113,bool),happ(s(fun(X113,fun(X113,bool)),X114),s(X113,X115))),s(X113,X116))) | s(bool,X117) != s(bool,X119) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X113,fun(X113,bool)),X114)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f6375,plain,(
  ( ! [X111,X109,X107,X112,X110,X108,X106] : (s(bool,X110) != s(bool,happ(s(fun(X106,bool),happ(s(fun(X106,fun(X106,bool)),X107),s(X106,X109))),s(X106,X108))) | p(s(bool,happ(s(fun(X106,bool),happ(s(fun(X106,fun(X106,bool)),X107),s(X106,X108))),s(X106,X109)))) | ~p(s(bool,X111)) | ~p(s(bool,X110)) | s(bool,X111) != s(bool,X112) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X106,fun(X106,bool)),X107)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f6374,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X104] : (s(bool,X103) != s(bool,happ(s(fun(X99,bool),happ(s(fun(X99,fun(X99,bool)),X100),s(X99,X102))),s(X99,X101))) | p(s(bool,X103)) | ~p(s(bool,X104)) | ~p(s(bool,happ(s(fun(X99,bool),happ(s(fun(X99,fun(X99,bool)),X100),s(X99,X101))),s(X99,X102)))) | s(bool,X104) != s(bool,X105) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X99,fun(X99,bool)),X100)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f6371,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (~p(s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X81))),s(X78,X80)))) | p(s(bool,X82)) | p(s(bool,X83)) | s(bool,X82) != s(bool,X84) | s(bool,X83) != s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X80))),s(X78,X81))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X78,fun(X78,bool)),X79)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f6369,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),happ(s(fun(X64,fun(X64,bool)),X65),s(X64,X67))),s(X64,X66)))) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,happ(s(fun(X64,bool),happ(s(fun(X64,fun(X64,bool)),X65),s(X64,X66))),s(X64,X67))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X64,fun(X64,bool)),X65)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f6366,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,X47) != s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X46))),s(X43,X45))) | p(s(bool,X48)) | ~p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X45))),s(X43,X46)))) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X43,fun(X43,bool)),X44)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f6365,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,X40) != s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X39))),s(X36,X38))) | p(s(bool,X41)) | ~p(s(bool,X40)) | p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X38))),s(X36,X39)))) | s(bool,X41) != s(bool,X42) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X36,fun(X36,bool)),X37)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f6305,plain,(
  ( ! [X327,X329,X331,X326,X328,X330] : (~p(s(bool,happ(s(fun(X326,bool),happ(s(fun(X326,fun(X326,bool)),X327),s(X326,X329))),s(X326,X328)))) | p(s(bool,X330)) | p(s(bool,X331)) | s(bool,X331) != s(bool,happ(s(fun(X326,bool),happ(s(fun(X326,fun(X326,bool)),X327),s(X326,X328))),s(X326,X329))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X326,fun(X326,bool)),X327)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f6303,plain,(
  ( ! [X315,X316,X318,X314,X317,X319] : (p(s(bool,happ(s(fun(X314,bool),happ(s(fun(X314,fun(X314,bool)),X315),s(X314,X317))),s(X314,X316)))) | p(s(bool,X318)) | ~p(s(bool,X319)) | s(bool,X319) != s(bool,happ(s(fun(X314,bool),happ(s(fun(X314,fun(X314,bool)),X315),s(X314,X316))),s(X314,X317))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X314,fun(X314,bool)),X315)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f6302,plain,(
  ( ! [X313,X308,X310,X312,X309,X311] : (s(bool,X312) != s(bool,happ(s(fun(X308,bool),happ(s(fun(X308,fun(X308,bool)),X309),s(X308,X311))),s(X308,X310))) | p(s(bool,X313)) | ~p(s(bool,X312)) | p(s(bool,happ(s(fun(X308,bool),happ(s(fun(X308,fun(X308,bool)),X309),s(X308,X310))),s(X308,X311)))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X308,fun(X308,bool)),X309)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f6301,plain,(
  ( ! [X302,X304,X306,X303,X305,X307] : (s(bool,X306) != s(bool,happ(s(fun(X302,bool),happ(s(fun(X302,fun(X302,bool)),X303),s(X302,X305))),s(X302,X304))) | p(s(bool,X307)) | ~p(s(bool,happ(s(fun(X302,bool),happ(s(fun(X302,fun(X302,bool)),X303),s(X302,X304))),s(X302,X305)))) | p(s(bool,X306)) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X302,fun(X302,bool)),X303)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f6293,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272] : (~p(s(bool,happ(s(fun(X267,bool),happ(s(fun(X267,fun(X267,bool)),X268),s(X267,X270))),s(X267,X269)))) | p(s(bool,X271)) | ~p(s(bool,X272)) | s(bool,X271) != s(bool,happ(s(fun(X267,bool),happ(s(fun(X267,fun(X267,bool)),X268),s(X267,X269))),s(X267,X270))) | s(bool,X272) != s(bool,X273) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X267,fun(X267,bool)),X268)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f6291,plain,(
  ( ! [X255,X253,X257,X259,X254,X256,X258] : (p(s(bool,happ(s(fun(X253,bool),happ(s(fun(X253,fun(X253,bool)),X254),s(X253,X256))),s(X253,X255)))) | ~p(s(bool,X257)) | ~p(s(bool,X258)) | s(bool,X258) != s(bool,happ(s(fun(X253,bool),happ(s(fun(X253,fun(X253,bool)),X254),s(X253,X255))),s(X253,X256))) | s(bool,X257) != s(bool,X259) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X253,fun(X253,bool)),X254)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f6290,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247] : (s(bool,X250) != s(bool,happ(s(fun(X246,bool),happ(s(fun(X246,fun(X246,bool)),X247),s(X246,X249))),s(X246,X248))) | p(s(bool,X250)) | ~p(s(bool,X251)) | ~p(s(bool,happ(s(fun(X246,bool),happ(s(fun(X246,fun(X246,bool)),X247),s(X246,X248))),s(X246,X249)))) | s(bool,X251) != s(bool,X252) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X246,fun(X246,bool)),X247)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f6289,plain,(
  ( ! [X243,X241,X239,X244,X242,X240,X245] : (s(bool,X243) != s(bool,happ(s(fun(X239,bool),happ(s(fun(X239,fun(X239,bool)),X240),s(X239,X242))),s(X239,X241))) | p(s(bool,happ(s(fun(X239,bool),happ(s(fun(X239,fun(X239,bool)),X240),s(X239,X241))),s(X239,X242)))) | ~p(s(bool,X244)) | ~p(s(bool,X243)) | s(bool,X244) != s(bool,X245) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X239,fun(X239,bool)),X240)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f6285,plain,(
  ( ! [X212,X216,X215,X213,X211,X217,X214] : (~p(s(bool,happ(s(fun(X211,bool),happ(s(fun(X211,fun(X211,bool)),X212),s(X211,X214))),s(X211,X213)))) | p(s(bool,X215)) | ~p(s(bool,X216)) | p(s(bool,X217)) | s(bool,X217) != s(bool,happ(s(fun(X211,bool),happ(s(fun(X211,fun(X211,bool)),X212),s(X211,X213))),s(X211,X214))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X211,fun(X211,bool)),X212)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f6283,plain,(
  ( ! [X198,X202,X200,X199,X197,X203,X201] : (p(s(bool,happ(s(fun(X197,bool),happ(s(fun(X197,fun(X197,bool)),X198),s(X197,X200))),s(X197,X199)))) | p(s(bool,X201)) | ~p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X202) != s(bool,happ(s(fun(X197,bool),happ(s(fun(X197,fun(X197,bool)),X198),s(X197,X199))),s(X197,X200))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X197,fun(X197,bool)),X198)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f6282,plain,(
  ( ! [X191,X196,X194,X192,X190,X195,X193] : (s(bool,X194) != s(bool,happ(s(fun(X190,bool),happ(s(fun(X190,fun(X190,bool)),X191),s(X190,X193))),s(X190,X192))) | p(s(bool,X195)) | ~p(s(bool,X194)) | ~p(s(bool,X196)) | p(s(bool,happ(s(fun(X190,bool),happ(s(fun(X190,fun(X190,bool)),X191),s(X190,X192))),s(X190,X193)))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X190,fun(X190,bool)),X191)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f6281,plain,(
  ( ! [X189,X187,X185,X188,X186,X184,X183] : (s(bool,X187) != s(bool,happ(s(fun(X183,bool),happ(s(fun(X183,fun(X183,bool)),X184),s(X183,X186))),s(X183,X185))) | p(s(bool,X188)) | ~p(s(bool,happ(s(fun(X183,bool),happ(s(fun(X183,fun(X183,bool)),X184),s(X183,X185))),s(X183,X186)))) | ~p(s(bool,X189)) | p(s(bool,X187)) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X183,fun(X183,bool)),X184)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f6280,plain,(
  ( ! [X177,X182,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X176,bool),happ(s(fun(X176,fun(X176,bool)),X177),s(X176,X179))),s(X176,X178)))) | p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X182) | s(bool,X181) != s(bool,happ(s(fun(X176,bool),happ(s(fun(X176,fun(X176,bool)),X177),s(X176,X178))),s(X176,X179))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X176,fun(X176,bool)),X177)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f6278,plain,(
  ( ! [X167,X165,X163,X166,X164,X162,X168] : (p(s(bool,happ(s(fun(X162,bool),happ(s(fun(X162,fun(X162,bool)),X163),s(X162,X165))),s(X162,X164)))) | p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,X168) | s(bool,X167) != s(bool,happ(s(fun(X162,bool),happ(s(fun(X162,fun(X162,bool)),X163),s(X162,X164))),s(X162,X165))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X162,fun(X162,bool)),X163)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f6275,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141] : (s(bool,X145) != s(bool,happ(s(fun(X141,bool),happ(s(fun(X141,fun(X141,bool)),X142),s(X141,X144))),s(X141,X143))) | p(s(bool,X146)) | ~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),happ(s(fun(X141,fun(X141,bool)),X142),s(X141,X143))),s(X141,X144)))) | s(bool,X146) != s(bool,X147) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X141,fun(X141,bool)),X142)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f6274,plain,(
  ( ! [X134,X140,X138,X136,X135,X139,X137] : (s(bool,X138) != s(bool,happ(s(fun(X134,bool),happ(s(fun(X134,fun(X134,bool)),X135),s(X134,X137))),s(X134,X136))) | p(s(bool,X139)) | ~p(s(bool,happ(s(fun(X134,bool),happ(s(fun(X134,fun(X134,bool)),X135),s(X134,X136))),s(X134,X137)))) | p(s(bool,X138)) | s(bool,X139) != s(bool,X140) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X134,fun(X134,bool)),X135)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f6273,plain,(
  ( ! [X127,X132,X130,X128,X133,X131,X129] : (~p(s(bool,happ(s(fun(X127,bool),happ(s(fun(X127,fun(X127,bool)),X128),s(X127,X130))),s(X127,X129)))) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X131) != s(bool,happ(s(fun(X127,bool),happ(s(fun(X127,fun(X127,bool)),X128),s(X127,X129))),s(X127,X130))) | s(bool,X132) != s(bool,X133) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X127,fun(X127,bool)),X128)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f6271,plain,(
  ( ! [X118,X116,X114,X119,X117,X115,X113] : (p(s(bool,happ(s(fun(X113,bool),happ(s(fun(X113,fun(X113,bool)),X114),s(X113,X116))),s(X113,X115)))) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X118) != s(bool,happ(s(fun(X113,bool),happ(s(fun(X113,fun(X113,bool)),X114),s(X113,X115))),s(X113,X116))) | s(bool,X117) != s(bool,X119) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X113,fun(X113,bool)),X114)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f6270,plain,(
  ( ! [X111,X109,X107,X112,X110,X108,X106] : (s(bool,X110) != s(bool,happ(s(fun(X106,bool),happ(s(fun(X106,fun(X106,bool)),X107),s(X106,X109))),s(X106,X108))) | p(s(bool,happ(s(fun(X106,bool),happ(s(fun(X106,fun(X106,bool)),X107),s(X106,X108))),s(X106,X109)))) | ~p(s(bool,X111)) | ~p(s(bool,X110)) | s(bool,X111) != s(bool,X112) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X106,fun(X106,bool)),X107)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f6269,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X104] : (s(bool,X103) != s(bool,happ(s(fun(X99,bool),happ(s(fun(X99,fun(X99,bool)),X100),s(X99,X102))),s(X99,X101))) | p(s(bool,X103)) | ~p(s(bool,X104)) | ~p(s(bool,happ(s(fun(X99,bool),happ(s(fun(X99,fun(X99,bool)),X100),s(X99,X101))),s(X99,X102)))) | s(bool,X104) != s(bool,X105) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X99,fun(X99,bool)),X100)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f6266,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (~p(s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X81))),s(X78,X80)))) | p(s(bool,X82)) | p(s(bool,X83)) | s(bool,X82) != s(bool,X84) | s(bool,X83) != s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X80))),s(X78,X81))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X78,fun(X78,bool)),X79)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f6264,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),happ(s(fun(X64,fun(X64,bool)),X65),s(X64,X67))),s(X64,X66)))) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,happ(s(fun(X64,bool),happ(s(fun(X64,fun(X64,bool)),X65),s(X64,X66))),s(X64,X67))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X64,fun(X64,bool)),X65)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f6261,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,X47) != s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X46))),s(X43,X45))) | p(s(bool,X48)) | ~p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X45))),s(X43,X46)))) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X43,fun(X43,bool)),X44)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f6260,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,X40) != s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X39))),s(X36,X38))) | p(s(bool,X41)) | ~p(s(bool,X40)) | p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X38))),s(X36,X39)))) | s(bool,X41) != s(bool,X42) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X36,fun(X36,bool)),X37)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f6121,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X68,sK1),X69),s(X68,X70),s(X68,X71))) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X73) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X68,sK1),X69),s(X68,X71),s(X68,X70)))) ) | $spl24),
  inference(resolution,[],[f2434,f125])).
fof(f6118,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X53,sK1),X54),s(X53,X55),s(X53,X56))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X53,sK1),X54),s(X53,X56),s(X53,X55))) | s(bool,X58) != s(bool,X59)) ) | $spl24),
  inference(resolution,[],[f2434,f89])).
fof(f6115,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X32,sK1),X33),s(X32,X34),s(X32,X35))) | p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X32,sK1),X33),s(X32,X35),s(X32,X34)))) ) | $spl24),
  inference(resolution,[],[f2434,f88])).
fof(f6114,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X25,sK1),X26),s(X25,X27),s(X25,X28))) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X25,sK1),X26),s(X25,X28),s(X25,X27)))) ) | $spl24),
  inference(resolution,[],[f2434,f87])).
fof(f6113,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X18,sK1),X19),s(X18,X20),s(X18,X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X18,sK1),X19),s(X18,X21),s(X18,X20))) | s(bool,X23) != s(bool,X24)) ) | $spl24),
  inference(resolution,[],[f2434,f85])).
fof(f6111,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X4,sK1),X5),s(X4,X6),s(X4,X7))) | p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X4,sK1),X5),s(X4,X7),s(X4,X6)))) ) | $spl24),
  inference(resolution,[],[f2434,f83])).
fof(f6095,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X46))),s(sK1,X47))) | p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X49) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X47))),s(sK1,X46)))) ) | $spl24),
  inference(resolution,[],[f2013,f125])).
fof(f6092,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X37))),s(sK1,X38))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X38))),s(sK1,X37))) | s(bool,X40) != s(bool,X41)) ) | $spl24),
  inference(resolution,[],[f2013,f89])).
fof(f6089,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,X23))) | p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X23))),s(sK1,X22)))) ) | $spl24),
  inference(resolution,[],[f2013,f88])).
fof(f6088,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X17))),s(sK1,X18))) | p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X18))),s(sK1,X17)))) ) | $spl24),
  inference(resolution,[],[f2013,f87])).
fof(f6087,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X12))),s(sK1,X13))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X13))),s(sK1,X12))) | s(bool,X15) != s(bool,X16)) ) | $spl24),
  inference(resolution,[],[f2013,f85])).
fof(f6085,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3))),s(sK1,X2)))) ) | $spl24),
  inference(resolution,[],[f2013,f83])).
fof(f2339,plain,(
  ( ! [X6,X10,X8,X7,X13,X11,X9] : (s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X10),s(X8,X11))) = s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X11),s(X8,X10))) | s(bool,X13) != s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7))) | ~p(s(bool,X13))) ) | $spl82),
  inference(cnf_transformation,[],[f2339_D])).
fof(f2339_D,plain,(
  ( ! [X6,X10,X8,X7,X13,X11,X9] : (s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X10),s(X8,X11))) = s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X11),s(X8,X10))) | s(bool,X13) != s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7))) | ~p(s(bool,X13))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f6068,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (s(bool,'const.relation.inv_image_4'(s(fun(X96,fun(X96,bool)),X97),s(fun(X98,X96),X99),s(X98,X100),s(X98,X101))) = s(bool,'const.relation.inv_image_4'(s(fun(X96,fun(X96,bool)),X97),s(fun(X98,X96),X99),s(X98,X101),s(X98,X100))) | p(s(bool,X102)) | ~p(s(bool,X103)) | s(bool,X103) != s(bool,'const.relation.symmetric_1'(s(fun(X96,fun(X96,bool)),X97)))) )),
  inference(resolution,[],[f1728,f125])).
fof(f6064,plain,(
  ( ! [X70,X76,X74,X72,X71,X69,X77,X75,X73] : (s(bool,'const.relation.inv_image_4'(s(fun(X69,fun(X69,bool)),X70),s(fun(X71,X69),X72),s(X71,X73),s(X71,X74))) = s(bool,'const.relation.inv_image_4'(s(fun(X69,fun(X69,bool)),X70),s(fun(X71,X69),X72),s(X71,X74),s(X71,X73))) | ~p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X76) != s(bool,'const.relation.symmetric_1'(s(fun(X69,fun(X69,bool)),X70))) | s(bool,X75) != s(bool,X77)) )),
  inference(resolution,[],[f1728,f89])).
fof(f6062,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53,X51] : (s(bool,'const.relation.inv_image_4'(s(fun(X51,fun(X51,bool)),X52),s(fun(X53,X51),X54),s(X53,X55),s(X53,X56))) = s(bool,'const.relation.inv_image_4'(s(fun(X51,fun(X51,bool)),X52),s(fun(X53,X51),X54),s(X53,X56),s(X53,X55))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,'const.relation.symmetric_1'(s(fun(X51,fun(X51,bool)),X52)))) )),
  inference(resolution,[],[f1728,f88])).
fof(f6060,plain,(
  ( ! [X39,X37,X35,X33,X41,X38,X36,X34,X40] : (s(bool,'const.relation.inv_image_4'(s(fun(X33,fun(X33,bool)),X34),s(fun(X35,X33),X36),s(X35,X37),s(X35,X38))) = s(bool,'const.relation.inv_image_4'(s(fun(X33,fun(X33,bool)),X34),s(fun(X35,X33),X36),s(X35,X38),s(X35,X37))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.relation.symmetric_1'(s(fun(X33,fun(X33,bool)),X34)))) )),
  inference(resolution,[],[f1728,f87])).
fof(f6059,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25,X32] : (s(bool,'const.relation.inv_image_4'(s(fun(X24,fun(X24,bool)),X25),s(fun(X26,X24),X27),s(X26,X28),s(X26,X29))) = s(bool,'const.relation.inv_image_4'(s(fun(X24,fun(X24,bool)),X25),s(fun(X26,X24),X27),s(X26,X29),s(X26,X28))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.relation.symmetric_1'(s(fun(X24,fun(X24,bool)),X25))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f1728,f85])).
fof(f6057,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X13,X11,X9] : (s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X10),s(X8,X11))) = s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X11),s(X8,X10))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) )),
  inference(resolution,[],[f1728,f83])).
fof(f5998,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X68,sK1),X69),s(X68,X70),s(X68,X71))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X68,sK1),X69),s(X68,X71),s(X68,X70)))) ) | $spl34),
  inference(resolution,[],[f2438,f125])).
fof(f5994,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X49,sK1),X50),s(X49,X51),s(X49,X52))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X49,sK1),X50),s(X49,X52),s(X49,X51))) | s(bool,X53) != s(bool,X55)) ) | $spl34),
  inference(resolution,[],[f2438,f89])).
fof(f5992,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X35,sK1),X36),s(X35,X37),s(X35,X38))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X35,sK1),X36),s(X35,X38),s(X35,X37)))) ) | $spl34),
  inference(resolution,[],[f2438,f88])).
fof(f5990,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X21,sK1),X22),s(X21,X23),s(X21,X24))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X21,sK1),X22),s(X21,X24),s(X21,X23)))) ) | $spl34),
  inference(resolution,[],[f2438,f87])).
fof(f5989,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X14,sK1),X15),s(X14,X16),s(X14,X17))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X14,sK1),X15),s(X14,X17),s(X14,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl34),
  inference(resolution,[],[f2438,f85])).
fof(f5987,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2)))) ) | $spl34),
  inference(resolution,[],[f2438,f83])).
fof(f5865,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X46))),s(sK1,X47))) | p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X47))),s(sK1,X46)))) ) | $spl34),
  inference(resolution,[],[f2017,f125])).
fof(f5861,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X35))),s(sK1,X36))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X36))),s(sK1,X35))) | s(bool,X37) != s(bool,X39)) ) | $spl34),
  inference(resolution,[],[f2017,f89])).
fof(f5859,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X25))),s(sK1,X26))) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X26))),s(sK1,X25)))) ) | $spl34),
  inference(resolution,[],[f2017,f88])).
fof(f5857,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X15))),s(sK1,X16))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X16))),s(sK1,X15)))) ) | $spl34),
  inference(resolution,[],[f2017,f87])).
fof(f5856,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X10))),s(sK1,X11))) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X11))),s(sK1,X10))) | s(bool,X12) != s(bool,X14)) ) | $spl34),
  inference(resolution,[],[f2017,f85])).
fof(f5854,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0)))) ) | $spl34),
  inference(resolution,[],[f2017,f83])).
fof(f2219,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,X7)) | s(bool,X1) != s(bool,X7) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) | $spl76),
  inference(cnf_transformation,[],[f2219_D])).
fof(f2219_D,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,X7)) | s(bool,X1) != s(bool,X7) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5804,plain,(
  ( ! [X90,X88,X87,X85,X91,X89,X86,X84] : (s(bool,X84) != s(bool,X85) | s(bool,X84) != s(bool,X86) | s(bool,X87) != s(bool,X88) | s(bool,X88) != s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,X86) != s(bool,X87) | p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X85) != s(bool,X91)) ) | $spl66),
  inference(resolution,[],[f1852,f125])).
fof(f5801,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (s(bool,X63) != s(bool,X64) | s(bool,X63) != s(bool,X65) | s(bool,X66) != s(bool,X67) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X65) != s(bool,X66) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X64) != s(bool,X69) | s(bool,X70) != s(bool,X71)) ) | $spl66),
  inference(resolution,[],[f1852,f89])).
fof(f5798,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X44,X42,X40] : (s(bool,X36) != s(bool,X37) | s(bool,X36) != s(bool,X38) | s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,X38) != s(bool,X39) | p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X37) != s(bool,X44)) ) | $spl66),
  inference(resolution,[],[f1852,f88])).
fof(f5797,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (s(bool,X27) != s(bool,X28) | s(bool,X27) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X29) != s(bool,X30) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X28) != s(bool,X34)) ) | $spl66),
  inference(resolution,[],[f1852,f87])).
fof(f5796,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X19) != s(bool,X24) | s(bool,X25) != s(bool,X26)) ) | $spl66),
  inference(resolution,[],[f1852,f85])).
fof(f5794,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X2) != s(bool,X3) | p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X1) != s(bool,X7)) ) | $spl66),
  inference(resolution,[],[f1852,f83])).
fof(f2554,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X10),s(X7,X9))) | p(s(bool,X11)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X9),s(X7,X10))))) ) | $spl86),
  inference(cnf_transformation,[],[f2554_D])).
fof(f2554_D,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X10),s(X7,X9))) | p(s(bool,X11)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X9),s(X7,X10))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5735,plain,(
  ( ! [X292,X294,X290,X293,X295,X291] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X290,sK1),X291),s(X290,X293),s(X290,X292)))) | p(s(bool,X294)) | p(s(bool,X295)) | s(bool,X295) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X290,sK1),X291),s(X290,X292),s(X290,X293)))) )),
  inference(superposition,[],[f125,f2341])).
fof(f5733,plain,(
  ( ! [X280,X282,X279,X281,X283,X278] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X278,sK1),X279),s(X278,X281),s(X278,X280)))) | p(s(bool,X282)) | ~p(s(bool,X283)) | s(bool,X283) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X278,sK1),X279),s(X278,X280),s(X278,X281)))) )),
  inference(superposition,[],[f125,f2341])).
fof(f5732,plain,(
  ( ! [X277,X273,X275,X276,X272,X274] : (s(bool,X276) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X272,sK1),X273),s(X272,X275),s(X272,X274))) | p(s(bool,X277)) | ~p(s(bool,X276)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X272,sK1),X273),s(X272,X274),s(X272,X275))))) )),
  inference(superposition,[],[f125,f2341])).
fof(f5731,plain,(
  ( ! [X269,X271,X267,X268,X270,X266] : (s(bool,X270) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X266,sK1),X267),s(X266,X269),s(X266,X268))) | p(s(bool,X271)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X266,sK1),X267),s(X266,X268),s(X266,X269)))) | p(s(bool,X270))) )),
  inference(superposition,[],[f125,f2341])).
fof(f5723,plain,(
  ( ! [X231,X237,X235,X233,X236,X234,X232] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X231,sK1),X232),s(X231,X234),s(X231,X233)))) | p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X231,sK1),X232),s(X231,X233),s(X231,X234))) | s(bool,X236) != s(bool,X237)) )),
  inference(superposition,[],[f89,f2341])).
fof(f5721,plain,(
  ( ! [X218,X222,X220,X219,X217,X223,X221] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X217,sK1),X218),s(X217,X220),s(X217,X219)))) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X217,sK1),X218),s(X217,X219),s(X217,X220))) | s(bool,X221) != s(bool,X223)) )),
  inference(superposition,[],[f89,f2341])).
fof(f5720,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X214] : (s(bool,X214) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X210,sK1),X211),s(X210,X213),s(X210,X212))) | p(s(bool,X214)) | ~p(s(bool,X215)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X210,sK1),X211),s(X210,X212),s(X210,X213)))) | s(bool,X215) != s(bool,X216)) )),
  inference(superposition,[],[f89,f2341])).
fof(f5719,plain,(
  ( ! [X208,X206,X204,X209,X207,X205,X203] : (s(bool,X207) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X203,sK1),X204),s(X203,X206),s(X203,X205))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X203,sK1),X204),s(X203,X205),s(X203,X206)))) | ~p(s(bool,X208)) | ~p(s(bool,X207)) | s(bool,X208) != s(bool,X209)) )),
  inference(superposition,[],[f89,f2341])).
fof(f5715,plain,(
  ( ! [X177,X175,X180,X178,X176,X181,X179] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X175,sK1),X176),s(X175,X178),s(X175,X177)))) | p(s(bool,X179)) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X175,sK1),X176),s(X175,X177),s(X175,X178)))) )),
  inference(superposition,[],[f88,f2341])).
fof(f5713,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X161,sK1),X162),s(X161,X164),s(X161,X163)))) | p(s(bool,X165)) | ~p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X161,sK1),X162),s(X161,X163),s(X161,X164)))) )),
  inference(superposition,[],[f88,f2341])).
fof(f5712,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X160] : (s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X154,sK1),X155),s(X154,X157),s(X154,X156))) | p(s(bool,X159)) | ~p(s(bool,X158)) | ~p(s(bool,X160)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X154,sK1),X155),s(X154,X156),s(X154,X157))))) )),
  inference(superposition,[],[f88,f2341])).
fof(f5711,plain,(
  ( ! [X152,X151,X149,X147,X153,X150,X148] : (s(bool,X151) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X147,sK1),X148),s(X147,X150),s(X147,X149))) | p(s(bool,X152)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X147,sK1),X148),s(X147,X149),s(X147,X150)))) | ~p(s(bool,X153)) | p(s(bool,X151))) )),
  inference(superposition,[],[f88,f2341])).
fof(f5710,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X140,sK1),X141),s(X140,X143),s(X140,X142)))) | p(s(bool,X144)) | p(s(bool,X145)) | s(bool,X144) != s(bool,X146) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X140,sK1),X141),s(X140,X142),s(X140,X143)))) )),
  inference(superposition,[],[f87,f2341])).
fof(f5708,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X126,sK1),X127),s(X126,X129),s(X126,X128)))) | p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,X132) | s(bool,X131) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X126,sK1),X127),s(X126,X128),s(X126,X129)))) )),
  inference(superposition,[],[f87,f2341])).
fof(f5705,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(bool,X109) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X105,sK1),X106),s(X105,X108),s(X105,X107))) | p(s(bool,X110)) | ~p(s(bool,X109)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X105,sK1),X106),s(X105,X107),s(X105,X108)))) | s(bool,X110) != s(bool,X111)) )),
  inference(superposition,[],[f87,f2341])).
fof(f5704,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (s(bool,X102) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X98,sK1),X99),s(X98,X101),s(X98,X100))) | p(s(bool,X103)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X98,sK1),X99),s(X98,X100),s(X98,X101)))) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104)) )),
  inference(superposition,[],[f87,f2341])).
fof(f5703,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X91,sK1),X92),s(X91,X94),s(X91,X93)))) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X91,sK1),X92),s(X91,X93),s(X91,X94))) | s(bool,X96) != s(bool,X97)) )),
  inference(superposition,[],[f85,f2341])).
fof(f5701,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X77,sK1),X78),s(X77,X80),s(X77,X79)))) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X77,sK1),X78),s(X77,X79),s(X77,X80))) | s(bool,X81) != s(bool,X83)) )),
  inference(superposition,[],[f85,f2341])).
fof(f5700,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,X74) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X70,sK1),X71),s(X70,X73),s(X70,X72))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X70,sK1),X71),s(X70,X72),s(X70,X73)))) | ~p(s(bool,X75)) | ~p(s(bool,X74)) | s(bool,X75) != s(bool,X76)) )),
  inference(superposition,[],[f85,f2341])).
fof(f5699,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,X67) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X63,sK1),X64),s(X63,X66),s(X63,X65))) | p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X63,sK1),X64),s(X63,X65),s(X63,X66)))) | s(bool,X68) != s(bool,X69)) )),
  inference(superposition,[],[f85,f2341])).
fof(f5696,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X42,sK1),X43),s(X42,X45),s(X42,X44)))) | p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X42,sK1),X43),s(X42,X44),s(X42,X45)))) )),
  inference(superposition,[],[f83,f2341])).
fof(f5694,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X28,sK1),X29),s(X28,X31),s(X28,X30)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X28,sK1),X29),s(X28,X30),s(X28,X31)))) )),
  inference(superposition,[],[f83,f2341])).
fof(f5691,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X10),s(X7,X9))) | p(s(bool,X12)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X9),s(X7,X10)))) | p(s(bool,X11)) | s(bool,X12) != s(bool,X13)) )),
  inference(superposition,[],[f83,f2341])).
fof(f5690,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X4) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2))) | p(s(bool,X5)) | ~p(s(bool,X4)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3)))) | s(bool,X5) != s(bool,X6)) )),
  inference(superposition,[],[f83,f2341])).
fof(f5651,plain,(
  ( ! [X292,X294,X290,X293,X295,X291] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X290,sK1),X291),s(X290,X293),s(X290,X292)))) | p(s(bool,X294)) | p(s(bool,X295)) | s(bool,X295) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X290,sK1),X291),s(X290,X292),s(X290,X293)))) )),
  inference(superposition,[],[f125,f2341])).
fof(f5649,plain,(
  ( ! [X280,X282,X279,X281,X283,X278] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X278,sK1),X279),s(X278,X281),s(X278,X280)))) | p(s(bool,X282)) | ~p(s(bool,X283)) | s(bool,X283) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X278,sK1),X279),s(X278,X280),s(X278,X281)))) )),
  inference(superposition,[],[f125,f2341])).
fof(f5648,plain,(
  ( ! [X277,X273,X275,X276,X272,X274] : (s(bool,X276) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X272,sK1),X273),s(X272,X275),s(X272,X274))) | p(s(bool,X277)) | ~p(s(bool,X276)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X272,sK1),X273),s(X272,X274),s(X272,X275))))) )),
  inference(superposition,[],[f125,f2341])).
fof(f5647,plain,(
  ( ! [X269,X271,X267,X268,X270,X266] : (s(bool,X270) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X266,sK1),X267),s(X266,X269),s(X266,X268))) | p(s(bool,X271)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X266,sK1),X267),s(X266,X268),s(X266,X269)))) | p(s(bool,X270))) )),
  inference(superposition,[],[f125,f2341])).
fof(f5639,plain,(
  ( ! [X231,X237,X235,X233,X236,X234,X232] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X231,sK1),X232),s(X231,X234),s(X231,X233)))) | p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X231,sK1),X232),s(X231,X233),s(X231,X234))) | s(bool,X236) != s(bool,X237)) )),
  inference(superposition,[],[f89,f2341])).
fof(f5637,plain,(
  ( ! [X218,X222,X220,X219,X217,X223,X221] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X217,sK1),X218),s(X217,X220),s(X217,X219)))) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X217,sK1),X218),s(X217,X219),s(X217,X220))) | s(bool,X221) != s(bool,X223)) )),
  inference(superposition,[],[f89,f2341])).
fof(f5636,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X214] : (s(bool,X214) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X210,sK1),X211),s(X210,X213),s(X210,X212))) | p(s(bool,X214)) | ~p(s(bool,X215)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X210,sK1),X211),s(X210,X212),s(X210,X213)))) | s(bool,X215) != s(bool,X216)) )),
  inference(superposition,[],[f89,f2341])).
fof(f5635,plain,(
  ( ! [X208,X206,X204,X209,X207,X205,X203] : (s(bool,X207) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X203,sK1),X204),s(X203,X206),s(X203,X205))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X203,sK1),X204),s(X203,X205),s(X203,X206)))) | ~p(s(bool,X208)) | ~p(s(bool,X207)) | s(bool,X208) != s(bool,X209)) )),
  inference(superposition,[],[f89,f2341])).
fof(f5631,plain,(
  ( ! [X177,X175,X180,X178,X176,X181,X179] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X175,sK1),X176),s(X175,X178),s(X175,X177)))) | p(s(bool,X179)) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X175,sK1),X176),s(X175,X177),s(X175,X178)))) )),
  inference(superposition,[],[f88,f2341])).
fof(f5629,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X161,sK1),X162),s(X161,X164),s(X161,X163)))) | p(s(bool,X165)) | ~p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X161,sK1),X162),s(X161,X163),s(X161,X164)))) )),
  inference(superposition,[],[f88,f2341])).
fof(f5628,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X160] : (s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X154,sK1),X155),s(X154,X157),s(X154,X156))) | p(s(bool,X159)) | ~p(s(bool,X158)) | ~p(s(bool,X160)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X154,sK1),X155),s(X154,X156),s(X154,X157))))) )),
  inference(superposition,[],[f88,f2341])).
fof(f5627,plain,(
  ( ! [X152,X151,X149,X147,X153,X150,X148] : (s(bool,X151) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X147,sK1),X148),s(X147,X150),s(X147,X149))) | p(s(bool,X152)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X147,sK1),X148),s(X147,X149),s(X147,X150)))) | ~p(s(bool,X153)) | p(s(bool,X151))) )),
  inference(superposition,[],[f88,f2341])).
fof(f5626,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X140,sK1),X141),s(X140,X143),s(X140,X142)))) | p(s(bool,X144)) | p(s(bool,X145)) | s(bool,X144) != s(bool,X146) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X140,sK1),X141),s(X140,X142),s(X140,X143)))) )),
  inference(superposition,[],[f87,f2341])).
fof(f5624,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X126,sK1),X127),s(X126,X129),s(X126,X128)))) | p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,X132) | s(bool,X131) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X126,sK1),X127),s(X126,X128),s(X126,X129)))) )),
  inference(superposition,[],[f87,f2341])).
fof(f5621,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(bool,X109) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X105,sK1),X106),s(X105,X108),s(X105,X107))) | p(s(bool,X110)) | ~p(s(bool,X109)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X105,sK1),X106),s(X105,X107),s(X105,X108)))) | s(bool,X110) != s(bool,X111)) )),
  inference(superposition,[],[f87,f2341])).
fof(f5620,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (s(bool,X102) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X98,sK1),X99),s(X98,X101),s(X98,X100))) | p(s(bool,X103)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X98,sK1),X99),s(X98,X100),s(X98,X101)))) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104)) )),
  inference(superposition,[],[f87,f2341])).
fof(f5619,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X91,sK1),X92),s(X91,X94),s(X91,X93)))) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X91,sK1),X92),s(X91,X93),s(X91,X94))) | s(bool,X96) != s(bool,X97)) )),
  inference(superposition,[],[f85,f2341])).
fof(f5617,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X77,sK1),X78),s(X77,X80),s(X77,X79)))) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X77,sK1),X78),s(X77,X79),s(X77,X80))) | s(bool,X81) != s(bool,X83)) )),
  inference(superposition,[],[f85,f2341])).
fof(f5616,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,X74) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X70,sK1),X71),s(X70,X73),s(X70,X72))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X70,sK1),X71),s(X70,X72),s(X70,X73)))) | ~p(s(bool,X75)) | ~p(s(bool,X74)) | s(bool,X75) != s(bool,X76)) )),
  inference(superposition,[],[f85,f2341])).
fof(f5615,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,X67) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X63,sK1),X64),s(X63,X66),s(X63,X65))) | p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X63,sK1),X64),s(X63,X65),s(X63,X66)))) | s(bool,X68) != s(bool,X69)) )),
  inference(superposition,[],[f85,f2341])).
fof(f5612,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X42,sK1),X43),s(X42,X45),s(X42,X44)))) | p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X42,sK1),X43),s(X42,X44),s(X42,X45)))) )),
  inference(superposition,[],[f83,f2341])).
fof(f5610,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X28,sK1),X29),s(X28,X31),s(X28,X30)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X28,sK1),X29),s(X28,X30),s(X28,X31)))) )),
  inference(superposition,[],[f83,f2341])).
fof(f5607,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X10),s(X7,X9))) | p(s(bool,X12)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X9),s(X7,X10)))) | p(s(bool,X11)) | s(bool,X12) != s(bool,X13)) )),
  inference(superposition,[],[f83,f2341])).
fof(f5606,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X4) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2))) | p(s(bool,X5)) | ~p(s(bool,X4)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3)))) | s(bool,X5) != s(bool,X6)) )),
  inference(superposition,[],[f83,f2341])).
fof(f1898,plain,(
  ( ! [X4,X7,X5] : (s(bool,X7) != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5))) | s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK4(X5,X4)))),s(X4,sK5(X5,X4)))) != s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK5(X5,X4)))),s(X4,sK4(X5,X4)))) | p(s(bool,X7))) ) | $spl70),
  inference(cnf_transformation,[],[f1898_D])).
fof(f1898_D,plain,(
  ( ! [X4,X7,X5] : (s(bool,X7) != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5))) | s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK4(X5,X4)))),s(X4,sK5(X5,X4)))) != s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK5(X5,X4)))),s(X4,sK4(X5,X4)))) | p(s(bool,X7))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f5526,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,happ(s(fun(X48,bool),happ(s(fun(X48,fun(X48,bool)),X49),s(X48,sK4(X49,X48)))),s(X48,sK5(X49,X48)))) != s(bool,happ(s(fun(X48,bool),happ(s(fun(X48,fun(X48,bool)),X49),s(X48,sK5(X49,X48)))),s(X48,sK4(X49,X48)))) | p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.relation.symmetric_1'(s(fun(X48,fun(X48,bool)),X49)))) )),
  inference(resolution,[],[f119,f125])).
fof(f5523,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),X40),s(X39,sK4(X40,X39)))),s(X39,sK5(X40,X39)))) != s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),X40),s(X39,sK5(X40,X39)))),s(X39,sK4(X40,X39)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.relation.symmetric_1'(s(fun(X39,fun(X39,bool)),X40))) | s(bool,X42) != s(bool,X43)) )),
  inference(resolution,[],[f119,f89])).
fof(f5520,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,happ(s(fun(X24,bool),happ(s(fun(X24,fun(X24,bool)),X25),s(X24,sK4(X25,X24)))),s(X24,sK5(X25,X24)))) != s(bool,happ(s(fun(X24,bool),happ(s(fun(X24,fun(X24,bool)),X25),s(X24,sK5(X25,X24)))),s(X24,sK4(X25,X24)))) | p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.symmetric_1'(s(fun(X24,fun(X24,bool)),X25)))) )),
  inference(resolution,[],[f119,f88])).
fof(f5519,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,happ(s(fun(X19,bool),happ(s(fun(X19,fun(X19,bool)),X20),s(X19,sK4(X20,X19)))),s(X19,sK5(X20,X19)))) != s(bool,happ(s(fun(X19,bool),happ(s(fun(X19,fun(X19,bool)),X20),s(X19,sK5(X20,X19)))),s(X19,sK4(X20,X19)))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.relation.symmetric_1'(s(fun(X19,fun(X19,bool)),X20)))) )),
  inference(resolution,[],[f119,f87])).
fof(f5518,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,sK4(X15,X14)))),s(X14,sK5(X15,X14)))) != s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,sK5(X15,X14)))),s(X14,sK4(X15,X14)))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.relation.symmetric_1'(s(fun(X14,fun(X14,bool)),X15))) | s(bool,X17) != s(bool,X18)) )),
  inference(resolution,[],[f119,f85])).
fof(f5516,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK4(X5,X4)))),s(X4,sK5(X5,X4)))) != s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK5(X5,X4)))),s(X4,sK4(X5,X4)))) | p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5)))) )),
  inference(resolution,[],[f119,f83])).
fof(f2134,plain,(
  ( ! [X14,X15,X13] : (s(bool,X15) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X14))),s(sK1,X13))) | p(s(bool,X15)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X13))),s(sK1,X14))))) ) | $spl74),
  inference(cnf_transformation,[],[f2134_D])).
fof(f2134_D,plain,(
  ( ! [X14,X15,X13] : (s(bool,X15) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X14))),s(sK1,X13))) | p(s(bool,X15)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X13))),s(sK1,X14))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f5454,plain,(
  ( ! [X210,X208,X211,X209] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X209))),s(sK1,X208)))) | p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X208))),s(sK1,X209)))) )),
  inference(superposition,[],[f125,f127])).
fof(f2133,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X9))),s(sK1,X8))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9)))) | ~p(s(bool,X10))) ) | $spl72),
  inference(cnf_transformation,[],[f2133_D])).
fof(f2133_D,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X9))),s(sK1,X8))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9)))) | ~p(s(bool,X10))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f5452,plain,(
  ( ! [X202,X200,X203,X201] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X201))),s(sK1,X200)))) | p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X200))),s(sK1,X201)))) )),
  inference(superposition,[],[f125,f127])).
fof(f5451,plain,(
  ( ! [X198,X196,X199,X197] : (s(bool,X198) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X197))),s(sK1,X196))) | p(s(bool,X199)) | ~p(s(bool,X198)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X196))),s(sK1,X197))))) )),
  inference(superposition,[],[f125,f127])).
fof(f5450,plain,(
  ( ! [X194,X192,X195,X193] : (s(bool,X194) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X193))),s(sK1,X192))) | p(s(bool,X195)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X192))),s(sK1,X193)))) | p(s(bool,X194))) )),
  inference(superposition,[],[f125,f127])).
fof(f5442,plain,(
  ( ! [X177,X175,X173,X176,X174] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X174))),s(sK1,X173)))) | p(s(bool,X175)) | ~p(s(bool,X176)) | s(bool,X175) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X173))),s(sK1,X174))) | s(bool,X176) != s(bool,X177)) )),
  inference(superposition,[],[f89,f127])).
fof(f5440,plain,(
  ( ! [X167,X165,X163,X166,X164] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X164))),s(sK1,X163)))) | ~p(s(bool,X165)) | ~p(s(bool,X166)) | s(bool,X166) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X163))),s(sK1,X164))) | s(bool,X165) != s(bool,X167)) )),
  inference(superposition,[],[f89,f127])).
fof(f5439,plain,(
  ( ! [X158,X161,X159,X162,X160] : (s(bool,X160) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X159))),s(sK1,X158))) | p(s(bool,X160)) | ~p(s(bool,X161)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X158))),s(sK1,X159)))) | s(bool,X161) != s(bool,X162)) )),
  inference(superposition,[],[f89,f127])).
fof(f5438,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,X155) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X154))),s(sK1,X153))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X153))),s(sK1,X154)))) | ~p(s(bool,X156)) | ~p(s(bool,X155)) | s(bool,X156) != s(bool,X157)) )),
  inference(superposition,[],[f89,f127])).
fof(f5434,plain,(
  ( ! [X134,X136,X135,X133,X137] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X134))),s(sK1,X133)))) | p(s(bool,X135)) | ~p(s(bool,X136)) | p(s(bool,X137)) | s(bool,X137) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X133))),s(sK1,X134)))) )),
  inference(superposition,[],[f88,f127])).
fof(f5432,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X124))),s(sK1,X123)))) | p(s(bool,X125)) | ~p(s(bool,X126)) | ~p(s(bool,X127)) | s(bool,X126) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X123))),s(sK1,X124)))) )),
  inference(superposition,[],[f88,f127])).
fof(f5431,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,X120) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X119))),s(sK1,X118))) | p(s(bool,X121)) | ~p(s(bool,X120)) | ~p(s(bool,X122)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X118))),s(sK1,X119))))) )),
  inference(superposition,[],[f88,f127])).
fof(f5430,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,X115) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X114))),s(sK1,X113))) | p(s(bool,X116)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X113))),s(sK1,X114)))) | ~p(s(bool,X117)) | p(s(bool,X115))) )),
  inference(superposition,[],[f88,f127])).
fof(f5429,plain,(
  ( ! [X111,X109,X112,X110,X108] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X109))),s(sK1,X108)))) | p(s(bool,X110)) | p(s(bool,X111)) | s(bool,X110) != s(bool,X112) | s(bool,X111) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X108))),s(sK1,X109)))) )),
  inference(superposition,[],[f87,f127])).
fof(f5427,plain,(
  ( ! [X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X99))),s(sK1,X98)))) | p(s(bool,X100)) | ~p(s(bool,X101)) | s(bool,X100) != s(bool,X102) | s(bool,X101) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X98))),s(sK1,X99)))) )),
  inference(superposition,[],[f87,f127])).
fof(f5424,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,X85) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X84))),s(sK1,X83))) | p(s(bool,X86)) | ~p(s(bool,X85)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X83))),s(sK1,X84)))) | s(bool,X86) != s(bool,X87)) )),
  inference(superposition,[],[f87,f127])).
fof(f5423,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,X80) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X79))),s(sK1,X78))) | p(s(bool,X81)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X78))),s(sK1,X79)))) | p(s(bool,X80)) | s(bool,X81) != s(bool,X82)) )),
  inference(superposition,[],[f87,f127])).
fof(f5422,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X74))),s(sK1,X73)))) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X75) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X73))),s(sK1,X74))) | s(bool,X76) != s(bool,X77)) )),
  inference(superposition,[],[f85,f127])).
fof(f5420,plain,(
  ( ! [X66,X64,X67,X65,X63] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X64))),s(sK1,X63)))) | ~p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X66) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X63))),s(sK1,X64))) | s(bool,X65) != s(bool,X67)) )),
  inference(superposition,[],[f85,f127])).
fof(f5419,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,X60) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X59))),s(sK1,X58))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X58))),s(sK1,X59)))) | ~p(s(bool,X61)) | ~p(s(bool,X60)) | s(bool,X61) != s(bool,X62)) )),
  inference(superposition,[],[f85,f127])).
fof(f5418,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X55) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X54))),s(sK1,X53))) | p(s(bool,X55)) | ~p(s(bool,X56)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X53))),s(sK1,X54)))) | s(bool,X56) != s(bool,X57)) )),
  inference(superposition,[],[f85,f127])).
fof(f5415,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X39))),s(sK1,X38)))) | p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X38))),s(sK1,X39)))) )),
  inference(superposition,[],[f83,f127])).
fof(f5413,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,X28)))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X28))),s(sK1,X29)))) )),
  inference(superposition,[],[f83,f127])).
fof(f5410,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,X15) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X14))),s(sK1,X13))) | p(s(bool,X16)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X13))),s(sK1,X14)))) | p(s(bool,X15)) | s(bool,X16) != s(bool,X17)) )),
  inference(superposition,[],[f83,f127])).
fof(f5409,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X9))),s(sK1,X8))) | p(s(bool,X11)) | ~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9)))) | s(bool,X11) != s(bool,X12)) )),
  inference(superposition,[],[f83,f127])).
fof(f5367,plain,(
  ( ! [X210,X208,X211,X209] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X209))),s(sK1,X208)))) | p(s(bool,X210)) | p(s(bool,X211)) | s(bool,X211) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X208))),s(sK1,X209)))) )),
  inference(superposition,[],[f125,f127])).
fof(f5365,plain,(
  ( ! [X202,X200,X203,X201] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X201))),s(sK1,X200)))) | p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X200))),s(sK1,X201)))) )),
  inference(superposition,[],[f125,f127])).
fof(f5364,plain,(
  ( ! [X198,X196,X199,X197] : (s(bool,X198) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X197))),s(sK1,X196))) | p(s(bool,X199)) | ~p(s(bool,X198)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X196))),s(sK1,X197))))) )),
  inference(superposition,[],[f125,f127])).
fof(f5363,plain,(
  ( ! [X194,X192,X195,X193] : (s(bool,X194) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X193))),s(sK1,X192))) | p(s(bool,X195)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X192))),s(sK1,X193)))) | p(s(bool,X194))) )),
  inference(superposition,[],[f125,f127])).
fof(f5355,plain,(
  ( ! [X177,X175,X173,X176,X174] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X174))),s(sK1,X173)))) | p(s(bool,X175)) | ~p(s(bool,X176)) | s(bool,X175) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X173))),s(sK1,X174))) | s(bool,X176) != s(bool,X177)) )),
  inference(superposition,[],[f89,f127])).
fof(f5353,plain,(
  ( ! [X167,X165,X163,X166,X164] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X164))),s(sK1,X163)))) | ~p(s(bool,X165)) | ~p(s(bool,X166)) | s(bool,X166) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X163))),s(sK1,X164))) | s(bool,X165) != s(bool,X167)) )),
  inference(superposition,[],[f89,f127])).
fof(f5352,plain,(
  ( ! [X158,X161,X159,X162,X160] : (s(bool,X160) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X159))),s(sK1,X158))) | p(s(bool,X160)) | ~p(s(bool,X161)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X158))),s(sK1,X159)))) | s(bool,X161) != s(bool,X162)) )),
  inference(superposition,[],[f89,f127])).
fof(f5351,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,X155) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X154))),s(sK1,X153))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X153))),s(sK1,X154)))) | ~p(s(bool,X156)) | ~p(s(bool,X155)) | s(bool,X156) != s(bool,X157)) )),
  inference(superposition,[],[f89,f127])).
fof(f5347,plain,(
  ( ! [X134,X136,X135,X133,X137] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X134))),s(sK1,X133)))) | p(s(bool,X135)) | ~p(s(bool,X136)) | p(s(bool,X137)) | s(bool,X137) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X133))),s(sK1,X134)))) )),
  inference(superposition,[],[f88,f127])).
fof(f5345,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X124))),s(sK1,X123)))) | p(s(bool,X125)) | ~p(s(bool,X126)) | ~p(s(bool,X127)) | s(bool,X126) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X123))),s(sK1,X124)))) )),
  inference(superposition,[],[f88,f127])).
fof(f5344,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,X120) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X119))),s(sK1,X118))) | p(s(bool,X121)) | ~p(s(bool,X120)) | ~p(s(bool,X122)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X118))),s(sK1,X119))))) )),
  inference(superposition,[],[f88,f127])).
fof(f5343,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,X115) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X114))),s(sK1,X113))) | p(s(bool,X116)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X113))),s(sK1,X114)))) | ~p(s(bool,X117)) | p(s(bool,X115))) )),
  inference(superposition,[],[f88,f127])).
fof(f5342,plain,(
  ( ! [X111,X109,X112,X110,X108] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X109))),s(sK1,X108)))) | p(s(bool,X110)) | p(s(bool,X111)) | s(bool,X110) != s(bool,X112) | s(bool,X111) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X108))),s(sK1,X109)))) )),
  inference(superposition,[],[f87,f127])).
fof(f5340,plain,(
  ( ! [X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X99))),s(sK1,X98)))) | p(s(bool,X100)) | ~p(s(bool,X101)) | s(bool,X100) != s(bool,X102) | s(bool,X101) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X98))),s(sK1,X99)))) )),
  inference(superposition,[],[f87,f127])).
fof(f5337,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,X85) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X84))),s(sK1,X83))) | p(s(bool,X86)) | ~p(s(bool,X85)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X83))),s(sK1,X84)))) | s(bool,X86) != s(bool,X87)) )),
  inference(superposition,[],[f87,f127])).
fof(f5336,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,X80) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X79))),s(sK1,X78))) | p(s(bool,X81)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X78))),s(sK1,X79)))) | p(s(bool,X80)) | s(bool,X81) != s(bool,X82)) )),
  inference(superposition,[],[f87,f127])).
fof(f5335,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X74))),s(sK1,X73)))) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X75) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X73))),s(sK1,X74))) | s(bool,X76) != s(bool,X77)) )),
  inference(superposition,[],[f85,f127])).
fof(f5333,plain,(
  ( ! [X66,X64,X67,X65,X63] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X64))),s(sK1,X63)))) | ~p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X66) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X63))),s(sK1,X64))) | s(bool,X65) != s(bool,X67)) )),
  inference(superposition,[],[f85,f127])).
fof(f5332,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,X60) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X59))),s(sK1,X58))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X58))),s(sK1,X59)))) | ~p(s(bool,X61)) | ~p(s(bool,X60)) | s(bool,X61) != s(bool,X62)) )),
  inference(superposition,[],[f85,f127])).
fof(f5331,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X55) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X54))),s(sK1,X53))) | p(s(bool,X55)) | ~p(s(bool,X56)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X53))),s(sK1,X54)))) | s(bool,X56) != s(bool,X57)) )),
  inference(superposition,[],[f85,f127])).
fof(f5328,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X39))),s(sK1,X38)))) | p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X38))),s(sK1,X39)))) )),
  inference(superposition,[],[f83,f127])).
fof(f5326,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,X28)))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X28))),s(sK1,X29)))) )),
  inference(superposition,[],[f83,f127])).
fof(f5323,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,X15) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X14))),s(sK1,X13))) | p(s(bool,X16)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X13))),s(sK1,X14)))) | p(s(bool,X15)) | s(bool,X16) != s(bool,X17)) )),
  inference(superposition,[],[f83,f127])).
fof(f5322,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X9))),s(sK1,X8))) | p(s(bool,X11)) | ~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9)))) | s(bool,X11) != s(bool,X12)) )),
  inference(superposition,[],[f83,f127])).
fof(f1852,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,X6)) | s(bool,X1) != s(bool,X6) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) != s(bool,X2)) ) | $spl66),
  inference(cnf_transformation,[],[f1852_D])).
fof(f1852_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,X6)) | s(bool,X1) != s(bool,X6) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f5070,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (s(bool,X74) != s(bool,X75) | s(bool,X74) != s(bool,X76) | s(bool,X77) != s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,X76) != s(bool,X77) | p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X75) != s(bool,X80)) ) | $spl56),
  inference(resolution,[],[f1356,f125])).
fof(f5067,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X59) != s(bool,X60) | s(bool,'F_0') = s(bool,X60) | s(bool,X58) != s(bool,X59) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X57) != s(bool,X61) | s(bool,X62) != s(bool,X63)) ) | $spl56),
  inference(resolution,[],[f1356,f89])).
fof(f5064,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,X32) != s(bool,X34) | s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) != s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X33) != s(bool,X39)) ) | $spl56),
  inference(resolution,[],[f1356,f88])).
fof(f5063,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X26) != s(bool,X27) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X25) != s(bool,X30)) ) | $spl56),
  inference(resolution,[],[f1356,f87])).
fof(f5062,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,X19) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | s(bool,X18) != s(bool,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X17) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl56),
  inference(resolution,[],[f1356,f85])).
fof(f5060,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X2) != s(bool,X3) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X1) != s(bool,X6)) ) | $spl56),
  inference(resolution,[],[f1356,f83])).
fof(f1356,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X5)) | s(bool,X1) != s(bool,X5) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl56),
  inference(cnf_transformation,[],[f1356_D])).
fof(f1356_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X5)) | s(bool,X1) != s(bool,X5) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4855,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(bool,X64) != s(bool,X65) | s(bool,X64) != s(bool,X66) | s(bool,'F_0') = s(bool,X67) | s(bool,X66) != s(bool,X67) | p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X65) != s(bool,X69)) ) | $spl40),
  inference(resolution,[],[f805,f125])).
fof(f4852,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X52) | s(bool,X51) != s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl40),
  inference(resolution,[],[f805,f89])).
fof(f4849,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X31) | s(bool,X30) != s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) ) | $spl40),
  inference(resolution,[],[f805,f88])).
fof(f4848,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) != s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl40),
  inference(resolution,[],[f805,f87])).
fof(f4847,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,'F_0') = s(bool,X17) | s(bool,X16) != s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl40),
  inference(resolution,[],[f805,f85])).
fof(f4845,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl40),
  inference(resolution,[],[f805,f83])).
fof(f4844,plain,(
  ( ! [X158,X159] : (s(bool,X158) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,'F_0') = s(bool,X159) | s(bool,X158) != s(bool,X159)) ) | $spl134),
  inference(cnf_transformation,[],[f4844_D])).
fof(f4844_D,plain,(
  ( ! [X158,X159] : (s(bool,X158) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,'F_0') = s(bool,X159) | s(bool,X158) != s(bool,X159)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f4843,plain,(
  ( ! [X158,X156,X159,X157,X160] : (s(bool,X156) != s(bool,X157) | s(bool,X158) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,X158) != s(bool,X159) | s(bool,'F_0') = s(bool,X159) | s(bool,X157) != s(bool,X160) | s(bool,'T_0') = s(bool,X160)) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f4794,f789])).
fof(f4794,plain,(
  ( ! [X158,X156,X159,X157,X160] : (s(bool,'F_0') != s(bool,X156) | s(bool,X156) != s(bool,X157) | s(bool,X158) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,X158) != s(bool,X159) | s(bool,'F_0') = s(bool,X159) | s(bool,X157) != s(bool,X160) | s(bool,'T_0') = s(bool,X160)) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f180])).
fof(f4842,plain,(
  ( ! [X146] : (s(bool,X146) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,'F_0') = s(bool,X146)) ) | $spl132),
  inference(cnf_transformation,[],[f4842_D])).
fof(f4842_D,plain,(
  ( ! [X146] : (s(bool,X146) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,'F_0') = s(bool,X146)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f4841,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,X144) != s(bool,X145) | s(bool,X146) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,'F_0') = s(bool,X146) | s(bool,X145) != s(bool,X147) | s(bool,'T_0') = s(bool,X147)) ) | ($spl34 | $spl38 | $spl56)),
  inference(subsumption_resolution,[],[f4791,f789])).
fof(f4791,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,'F_0') != s(bool,X144) | s(bool,X144) != s(bool,X145) | s(bool,X146) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,'F_0') = s(bool,X146) | s(bool,X145) != s(bool,X147) | s(bool,'T_0') = s(bool,X147)) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f180])).
fof(f805,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl40),
  inference(cnf_transformation,[],[f805_D])).
fof(f805_D,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4702,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,X54) != s(bool,X55) | s(bool,X54) != s(bool,X56) | s(bool,'F_0') = s(bool,X56) | p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X55) != s(bool,X58)) ) | $spl38),
  inference(resolution,[],[f745,f125])).
fof(f4699,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl38),
  inference(resolution,[],[f745,f89])).
fof(f4696,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl38),
  inference(resolution,[],[f745,f88])).
fof(f4695,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl38),
  inference(resolution,[],[f745,f87])).
fof(f4694,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl38),
  inference(resolution,[],[f745,f85])).
fof(f4692,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl38),
  inference(resolution,[],[f745,f83])).
fof(f1290,plain,(
  ( ! [X10,X8,X7,X11,X9] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X7,fun(X7,bool)),X8)))) | s(bool,X11) != s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X10))),s(X7,X9))) | p(s(bool,X11)) | ~p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X9))),s(X7,X10))))) ) | $spl54),
  inference(cnf_transformation,[],[f1290_D])).
fof(f1290_D,plain,(
  ( ! [X10,X8,X7,X11,X9] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X7,fun(X7,bool)),X8)))) | s(bool,X11) != s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X10))),s(X7,X9))) | p(s(bool,X11)) | ~p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X9))),s(X7,X10))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4633,plain,(
  ( ! [X292,X294,X290,X293,X295,X291] : (~p(s(bool,happ(s(fun(X290,bool),happ(s(fun(X290,fun(X290,bool)),X291),s(X290,X293))),s(X290,X292)))) | p(s(bool,X294)) | p(s(bool,X295)) | s(bool,X295) != s(bool,happ(s(fun(X290,bool),happ(s(fun(X290,fun(X290,bool)),X291),s(X290,X292))),s(X290,X293))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X290,fun(X290,bool)),X291))))) )),
  inference(superposition,[],[f125,f118])).
fof(f1289,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(bool,X4) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,X4))) ) | $spl52),
  inference(cnf_transformation,[],[f1289_D])).
fof(f1289_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(bool,X4) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,X4))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4631,plain,(
  ( ! [X280,X282,X279,X281,X283,X278] : (p(s(bool,happ(s(fun(X278,bool),happ(s(fun(X278,fun(X278,bool)),X279),s(X278,X281))),s(X278,X280)))) | p(s(bool,X282)) | ~p(s(bool,X283)) | s(bool,X283) != s(bool,happ(s(fun(X278,bool),happ(s(fun(X278,fun(X278,bool)),X279),s(X278,X280))),s(X278,X281))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X278,fun(X278,bool)),X279))))) )),
  inference(superposition,[],[f125,f118])).
fof(f4630,plain,(
  ( ! [X277,X273,X275,X276,X272,X274] : (s(bool,X276) != s(bool,happ(s(fun(X272,bool),happ(s(fun(X272,fun(X272,bool)),X273),s(X272,X275))),s(X272,X274))) | p(s(bool,X277)) | ~p(s(bool,X276)) | p(s(bool,happ(s(fun(X272,bool),happ(s(fun(X272,fun(X272,bool)),X273),s(X272,X274))),s(X272,X275)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X272,fun(X272,bool)),X273))))) )),
  inference(superposition,[],[f125,f118])).
fof(f4629,plain,(
  ( ! [X269,X271,X267,X268,X270,X266] : (s(bool,X270) != s(bool,happ(s(fun(X266,bool),happ(s(fun(X266,fun(X266,bool)),X267),s(X266,X269))),s(X266,X268))) | p(s(bool,X271)) | ~p(s(bool,happ(s(fun(X266,bool),happ(s(fun(X266,fun(X266,bool)),X267),s(X266,X268))),s(X266,X269)))) | p(s(bool,X270)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X266,fun(X266,bool)),X267))))) )),
  inference(superposition,[],[f125,f118])).
fof(f4621,plain,(
  ( ! [X231,X237,X235,X233,X236,X234,X232] : (~p(s(bool,happ(s(fun(X231,bool),happ(s(fun(X231,fun(X231,bool)),X232),s(X231,X234))),s(X231,X233)))) | p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,happ(s(fun(X231,bool),happ(s(fun(X231,fun(X231,bool)),X232),s(X231,X233))),s(X231,X234))) | s(bool,X236) != s(bool,X237) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X231,fun(X231,bool)),X232))))) )),
  inference(superposition,[],[f89,f118])).
fof(f4619,plain,(
  ( ! [X218,X222,X220,X219,X217,X223,X221] : (p(s(bool,happ(s(fun(X217,bool),happ(s(fun(X217,fun(X217,bool)),X218),s(X217,X220))),s(X217,X219)))) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,happ(s(fun(X217,bool),happ(s(fun(X217,fun(X217,bool)),X218),s(X217,X219))),s(X217,X220))) | s(bool,X221) != s(bool,X223) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X217,fun(X217,bool)),X218))))) )),
  inference(superposition,[],[f89,f118])).
fof(f4618,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X214] : (s(bool,X214) != s(bool,happ(s(fun(X210,bool),happ(s(fun(X210,fun(X210,bool)),X211),s(X210,X213))),s(X210,X212))) | p(s(bool,X214)) | ~p(s(bool,X215)) | ~p(s(bool,happ(s(fun(X210,bool),happ(s(fun(X210,fun(X210,bool)),X211),s(X210,X212))),s(X210,X213)))) | s(bool,X215) != s(bool,X216) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X210,fun(X210,bool)),X211))))) )),
  inference(superposition,[],[f89,f118])).
fof(f4617,plain,(
  ( ! [X208,X206,X204,X209,X207,X205,X203] : (s(bool,X207) != s(bool,happ(s(fun(X203,bool),happ(s(fun(X203,fun(X203,bool)),X204),s(X203,X206))),s(X203,X205))) | p(s(bool,happ(s(fun(X203,bool),happ(s(fun(X203,fun(X203,bool)),X204),s(X203,X205))),s(X203,X206)))) | ~p(s(bool,X208)) | ~p(s(bool,X207)) | s(bool,X208) != s(bool,X209) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X203,fun(X203,bool)),X204))))) )),
  inference(superposition,[],[f89,f118])).
fof(f4613,plain,(
  ( ! [X177,X175,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X175,bool),happ(s(fun(X175,fun(X175,bool)),X176),s(X175,X178))),s(X175,X177)))) | p(s(bool,X179)) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,happ(s(fun(X175,bool),happ(s(fun(X175,fun(X175,bool)),X176),s(X175,X177))),s(X175,X178))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X175,fun(X175,bool)),X176))))) )),
  inference(superposition,[],[f88,f118])).
fof(f4611,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162] : (p(s(bool,happ(s(fun(X161,bool),happ(s(fun(X161,fun(X161,bool)),X162),s(X161,X164))),s(X161,X163)))) | p(s(bool,X165)) | ~p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,happ(s(fun(X161,bool),happ(s(fun(X161,fun(X161,bool)),X162),s(X161,X163))),s(X161,X164))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X161,fun(X161,bool)),X162))))) )),
  inference(superposition,[],[f88,f118])).
fof(f4610,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X160] : (s(bool,X158) != s(bool,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),X155),s(X154,X157))),s(X154,X156))) | p(s(bool,X159)) | ~p(s(bool,X158)) | ~p(s(bool,X160)) | p(s(bool,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),X155),s(X154,X156))),s(X154,X157)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X154,fun(X154,bool)),X155))))) )),
  inference(superposition,[],[f88,f118])).
fof(f4609,plain,(
  ( ! [X152,X151,X149,X147,X153,X150,X148] : (s(bool,X151) != s(bool,happ(s(fun(X147,bool),happ(s(fun(X147,fun(X147,bool)),X148),s(X147,X150))),s(X147,X149))) | p(s(bool,X152)) | ~p(s(bool,happ(s(fun(X147,bool),happ(s(fun(X147,fun(X147,bool)),X148),s(X147,X149))),s(X147,X150)))) | ~p(s(bool,X153)) | p(s(bool,X151)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X147,fun(X147,bool)),X148))))) )),
  inference(superposition,[],[f88,f118])).
fof(f4608,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X140,bool),happ(s(fun(X140,fun(X140,bool)),X141),s(X140,X143))),s(X140,X142)))) | p(s(bool,X144)) | p(s(bool,X145)) | s(bool,X144) != s(bool,X146) | s(bool,X145) != s(bool,happ(s(fun(X140,bool),happ(s(fun(X140,fun(X140,bool)),X141),s(X140,X142))),s(X140,X143))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X140,fun(X140,bool)),X141))))) )),
  inference(superposition,[],[f87,f118])).
fof(f4606,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (p(s(bool,happ(s(fun(X126,bool),happ(s(fun(X126,fun(X126,bool)),X127),s(X126,X129))),s(X126,X128)))) | p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,X132) | s(bool,X131) != s(bool,happ(s(fun(X126,bool),happ(s(fun(X126,fun(X126,bool)),X127),s(X126,X128))),s(X126,X129))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X126,fun(X126,bool)),X127))))) )),
  inference(superposition,[],[f87,f118])).
fof(f4603,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(bool,X109) != s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X108))),s(X105,X107))) | p(s(bool,X110)) | ~p(s(bool,X109)) | p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108)))) | s(bool,X110) != s(bool,X111) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X105,fun(X105,bool)),X106))))) )),
  inference(superposition,[],[f87,f118])).
fof(f4602,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (s(bool,X102) != s(bool,happ(s(fun(X98,bool),happ(s(fun(X98,fun(X98,bool)),X99),s(X98,X101))),s(X98,X100))) | p(s(bool,X103)) | ~p(s(bool,happ(s(fun(X98,bool),happ(s(fun(X98,fun(X98,bool)),X99),s(X98,X100))),s(X98,X101)))) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X98,fun(X98,bool)),X99))))) )),
  inference(superposition,[],[f87,f118])).
fof(f4601,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (~p(s(bool,happ(s(fun(X91,bool),happ(s(fun(X91,fun(X91,bool)),X92),s(X91,X94))),s(X91,X93)))) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,happ(s(fun(X91,bool),happ(s(fun(X91,fun(X91,bool)),X92),s(X91,X93))),s(X91,X94))) | s(bool,X96) != s(bool,X97) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X91,fun(X91,bool)),X92))))) )),
  inference(superposition,[],[f85,f118])).
fof(f4599,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X77,bool),happ(s(fun(X77,fun(X77,bool)),X78),s(X77,X80))),s(X77,X79)))) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X77,bool),happ(s(fun(X77,fun(X77,bool)),X78),s(X77,X79))),s(X77,X80))) | s(bool,X81) != s(bool,X83) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X77,fun(X77,bool)),X78))))) )),
  inference(superposition,[],[f85,f118])).
fof(f4598,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,X74) != s(bool,happ(s(fun(X70,bool),happ(s(fun(X70,fun(X70,bool)),X71),s(X70,X73))),s(X70,X72))) | p(s(bool,happ(s(fun(X70,bool),happ(s(fun(X70,fun(X70,bool)),X71),s(X70,X72))),s(X70,X73)))) | ~p(s(bool,X75)) | ~p(s(bool,X74)) | s(bool,X75) != s(bool,X76) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X70,fun(X70,bool)),X71))))) )),
  inference(superposition,[],[f85,f118])).
fof(f4597,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,X67) != s(bool,happ(s(fun(X63,bool),happ(s(fun(X63,fun(X63,bool)),X64),s(X63,X66))),s(X63,X65))) | p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,happ(s(fun(X63,bool),happ(s(fun(X63,fun(X63,bool)),X64),s(X63,X65))),s(X63,X66)))) | s(bool,X68) != s(bool,X69) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X63,fun(X63,bool)),X64))))) )),
  inference(superposition,[],[f85,f118])).
fof(f4594,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X45))),s(X42,X44)))) | p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X42,fun(X42,bool)),X43))))) )),
  inference(superposition,[],[f83,f118])).
fof(f4592,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),happ(s(fun(X28,fun(X28,bool)),X29),s(X28,X31))),s(X28,X30)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,happ(s(fun(X28,bool),happ(s(fun(X28,fun(X28,bool)),X29),s(X28,X30))),s(X28,X31))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X28,fun(X28,bool)),X29))))) )),
  inference(superposition,[],[f83,f118])).
fof(f4589,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,X11) != s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X10))),s(X7,X9))) | p(s(bool,X12)) | ~p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X9))),s(X7,X10)))) | p(s(bool,X11)) | s(bool,X12) != s(bool,X13) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X7,fun(X7,bool)),X8))))) )),
  inference(superposition,[],[f83,f118])).
fof(f4588,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X4) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,X5)) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(bool,X5) != s(bool,X6) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(superposition,[],[f83,f118])).
fof(f4560,plain,(
  ( ! [X292,X294,X290,X293,X295,X291] : (~p(s(bool,happ(s(fun(X290,bool),happ(s(fun(X290,fun(X290,bool)),X291),s(X290,X293))),s(X290,X292)))) | p(s(bool,X294)) | p(s(bool,X295)) | s(bool,X295) != s(bool,happ(s(fun(X290,bool),happ(s(fun(X290,fun(X290,bool)),X291),s(X290,X292))),s(X290,X293))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X290,fun(X290,bool)),X291))))) )),
  inference(superposition,[],[f125,f118])).
fof(f4558,plain,(
  ( ! [X280,X282,X279,X281,X283,X278] : (p(s(bool,happ(s(fun(X278,bool),happ(s(fun(X278,fun(X278,bool)),X279),s(X278,X281))),s(X278,X280)))) | p(s(bool,X282)) | ~p(s(bool,X283)) | s(bool,X283) != s(bool,happ(s(fun(X278,bool),happ(s(fun(X278,fun(X278,bool)),X279),s(X278,X280))),s(X278,X281))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X278,fun(X278,bool)),X279))))) )),
  inference(superposition,[],[f125,f118])).
fof(f4557,plain,(
  ( ! [X277,X273,X275,X276,X272,X274] : (s(bool,X276) != s(bool,happ(s(fun(X272,bool),happ(s(fun(X272,fun(X272,bool)),X273),s(X272,X275))),s(X272,X274))) | p(s(bool,X277)) | ~p(s(bool,X276)) | p(s(bool,happ(s(fun(X272,bool),happ(s(fun(X272,fun(X272,bool)),X273),s(X272,X274))),s(X272,X275)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X272,fun(X272,bool)),X273))))) )),
  inference(superposition,[],[f125,f118])).
fof(f4556,plain,(
  ( ! [X269,X271,X267,X268,X270,X266] : (s(bool,X270) != s(bool,happ(s(fun(X266,bool),happ(s(fun(X266,fun(X266,bool)),X267),s(X266,X269))),s(X266,X268))) | p(s(bool,X271)) | ~p(s(bool,happ(s(fun(X266,bool),happ(s(fun(X266,fun(X266,bool)),X267),s(X266,X268))),s(X266,X269)))) | p(s(bool,X270)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X266,fun(X266,bool)),X267))))) )),
  inference(superposition,[],[f125,f118])).
fof(f4548,plain,(
  ( ! [X231,X237,X235,X233,X236,X234,X232] : (~p(s(bool,happ(s(fun(X231,bool),happ(s(fun(X231,fun(X231,bool)),X232),s(X231,X234))),s(X231,X233)))) | p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,happ(s(fun(X231,bool),happ(s(fun(X231,fun(X231,bool)),X232),s(X231,X233))),s(X231,X234))) | s(bool,X236) != s(bool,X237) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X231,fun(X231,bool)),X232))))) )),
  inference(superposition,[],[f89,f118])).
fof(f4546,plain,(
  ( ! [X218,X222,X220,X219,X217,X223,X221] : (p(s(bool,happ(s(fun(X217,bool),happ(s(fun(X217,fun(X217,bool)),X218),s(X217,X220))),s(X217,X219)))) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,happ(s(fun(X217,bool),happ(s(fun(X217,fun(X217,bool)),X218),s(X217,X219))),s(X217,X220))) | s(bool,X221) != s(bool,X223) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X217,fun(X217,bool)),X218))))) )),
  inference(superposition,[],[f89,f118])).
fof(f4545,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X214] : (s(bool,X214) != s(bool,happ(s(fun(X210,bool),happ(s(fun(X210,fun(X210,bool)),X211),s(X210,X213))),s(X210,X212))) | p(s(bool,X214)) | ~p(s(bool,X215)) | ~p(s(bool,happ(s(fun(X210,bool),happ(s(fun(X210,fun(X210,bool)),X211),s(X210,X212))),s(X210,X213)))) | s(bool,X215) != s(bool,X216) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X210,fun(X210,bool)),X211))))) )),
  inference(superposition,[],[f89,f118])).
fof(f4544,plain,(
  ( ! [X208,X206,X204,X209,X207,X205,X203] : (s(bool,X207) != s(bool,happ(s(fun(X203,bool),happ(s(fun(X203,fun(X203,bool)),X204),s(X203,X206))),s(X203,X205))) | p(s(bool,happ(s(fun(X203,bool),happ(s(fun(X203,fun(X203,bool)),X204),s(X203,X205))),s(X203,X206)))) | ~p(s(bool,X208)) | ~p(s(bool,X207)) | s(bool,X208) != s(bool,X209) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X203,fun(X203,bool)),X204))))) )),
  inference(superposition,[],[f89,f118])).
fof(f4540,plain,(
  ( ! [X177,X175,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X175,bool),happ(s(fun(X175,fun(X175,bool)),X176),s(X175,X178))),s(X175,X177)))) | p(s(bool,X179)) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,happ(s(fun(X175,bool),happ(s(fun(X175,fun(X175,bool)),X176),s(X175,X177))),s(X175,X178))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X175,fun(X175,bool)),X176))))) )),
  inference(superposition,[],[f88,f118])).
fof(f4538,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162] : (p(s(bool,happ(s(fun(X161,bool),happ(s(fun(X161,fun(X161,bool)),X162),s(X161,X164))),s(X161,X163)))) | p(s(bool,X165)) | ~p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,happ(s(fun(X161,bool),happ(s(fun(X161,fun(X161,bool)),X162),s(X161,X163))),s(X161,X164))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X161,fun(X161,bool)),X162))))) )),
  inference(superposition,[],[f88,f118])).
fof(f4537,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X160] : (s(bool,X158) != s(bool,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),X155),s(X154,X157))),s(X154,X156))) | p(s(bool,X159)) | ~p(s(bool,X158)) | ~p(s(bool,X160)) | p(s(bool,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),X155),s(X154,X156))),s(X154,X157)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X154,fun(X154,bool)),X155))))) )),
  inference(superposition,[],[f88,f118])).
fof(f4536,plain,(
  ( ! [X152,X151,X149,X147,X153,X150,X148] : (s(bool,X151) != s(bool,happ(s(fun(X147,bool),happ(s(fun(X147,fun(X147,bool)),X148),s(X147,X150))),s(X147,X149))) | p(s(bool,X152)) | ~p(s(bool,happ(s(fun(X147,bool),happ(s(fun(X147,fun(X147,bool)),X148),s(X147,X149))),s(X147,X150)))) | ~p(s(bool,X153)) | p(s(bool,X151)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X147,fun(X147,bool)),X148))))) )),
  inference(superposition,[],[f88,f118])).
fof(f4535,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X140,bool),happ(s(fun(X140,fun(X140,bool)),X141),s(X140,X143))),s(X140,X142)))) | p(s(bool,X144)) | p(s(bool,X145)) | s(bool,X144) != s(bool,X146) | s(bool,X145) != s(bool,happ(s(fun(X140,bool),happ(s(fun(X140,fun(X140,bool)),X141),s(X140,X142))),s(X140,X143))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X140,fun(X140,bool)),X141))))) )),
  inference(superposition,[],[f87,f118])).
fof(f4533,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (p(s(bool,happ(s(fun(X126,bool),happ(s(fun(X126,fun(X126,bool)),X127),s(X126,X129))),s(X126,X128)))) | p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,X132) | s(bool,X131) != s(bool,happ(s(fun(X126,bool),happ(s(fun(X126,fun(X126,bool)),X127),s(X126,X128))),s(X126,X129))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X126,fun(X126,bool)),X127))))) )),
  inference(superposition,[],[f87,f118])).
fof(f4530,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(bool,X109) != s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X108))),s(X105,X107))) | p(s(bool,X110)) | ~p(s(bool,X109)) | p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108)))) | s(bool,X110) != s(bool,X111) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X105,fun(X105,bool)),X106))))) )),
  inference(superposition,[],[f87,f118])).
fof(f4529,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (s(bool,X102) != s(bool,happ(s(fun(X98,bool),happ(s(fun(X98,fun(X98,bool)),X99),s(X98,X101))),s(X98,X100))) | p(s(bool,X103)) | ~p(s(bool,happ(s(fun(X98,bool),happ(s(fun(X98,fun(X98,bool)),X99),s(X98,X100))),s(X98,X101)))) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X98,fun(X98,bool)),X99))))) )),
  inference(superposition,[],[f87,f118])).
fof(f4528,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (~p(s(bool,happ(s(fun(X91,bool),happ(s(fun(X91,fun(X91,bool)),X92),s(X91,X94))),s(X91,X93)))) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,happ(s(fun(X91,bool),happ(s(fun(X91,fun(X91,bool)),X92),s(X91,X93))),s(X91,X94))) | s(bool,X96) != s(bool,X97) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X91,fun(X91,bool)),X92))))) )),
  inference(superposition,[],[f85,f118])).
fof(f4526,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X77,bool),happ(s(fun(X77,fun(X77,bool)),X78),s(X77,X80))),s(X77,X79)))) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X77,bool),happ(s(fun(X77,fun(X77,bool)),X78),s(X77,X79))),s(X77,X80))) | s(bool,X81) != s(bool,X83) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X77,fun(X77,bool)),X78))))) )),
  inference(superposition,[],[f85,f118])).
fof(f4525,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,X74) != s(bool,happ(s(fun(X70,bool),happ(s(fun(X70,fun(X70,bool)),X71),s(X70,X73))),s(X70,X72))) | p(s(bool,happ(s(fun(X70,bool),happ(s(fun(X70,fun(X70,bool)),X71),s(X70,X72))),s(X70,X73)))) | ~p(s(bool,X75)) | ~p(s(bool,X74)) | s(bool,X75) != s(bool,X76) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X70,fun(X70,bool)),X71))))) )),
  inference(superposition,[],[f85,f118])).
fof(f4524,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,X67) != s(bool,happ(s(fun(X63,bool),happ(s(fun(X63,fun(X63,bool)),X64),s(X63,X66))),s(X63,X65))) | p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,happ(s(fun(X63,bool),happ(s(fun(X63,fun(X63,bool)),X64),s(X63,X65))),s(X63,X66)))) | s(bool,X68) != s(bool,X69) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X63,fun(X63,bool)),X64))))) )),
  inference(superposition,[],[f85,f118])).
fof(f4521,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X45))),s(X42,X44)))) | p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X42,fun(X42,bool)),X43))))) )),
  inference(superposition,[],[f83,f118])).
fof(f4519,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),happ(s(fun(X28,fun(X28,bool)),X29),s(X28,X31))),s(X28,X30)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,happ(s(fun(X28,bool),happ(s(fun(X28,fun(X28,bool)),X29),s(X28,X30))),s(X28,X31))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X28,fun(X28,bool)),X29))))) )),
  inference(superposition,[],[f83,f118])).
fof(f4516,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,X11) != s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X10))),s(X7,X9))) | p(s(bool,X12)) | ~p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X9))),s(X7,X10)))) | p(s(bool,X11)) | s(bool,X12) != s(bool,X13) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X7,fun(X7,bool)),X8))))) )),
  inference(superposition,[],[f83,f118])).
fof(f4515,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X4) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,X5)) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(bool,X5) != s(bool,X6) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(superposition,[],[f83,f118])).
fof(f1285,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X5) != s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))) | s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | ~p(s(bool,X5))) ) | $spl50),
  inference(cnf_transformation,[],[f1285_D])).
fof(f1285_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(bool,X5) != s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))) | s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | ~p(s(bool,X5))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4478,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (s(bool,happ(s(fun(X68,bool),happ(s(fun(X68,fun(X68,bool)),X69),s(X68,X70))),s(X68,X71))) = s(bool,happ(s(fun(X68,bool),happ(s(fun(X68,fun(X68,bool)),X69),s(X68,X71))),s(X68,X70))) | p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.relation.symmetric_1'(s(fun(X68,fun(X68,bool)),X69)))) )),
  inference(resolution,[],[f118,f125])).
fof(f4474,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,happ(s(fun(X49,bool),happ(s(fun(X49,fun(X49,bool)),X50),s(X49,X51))),s(X49,X52))) = s(bool,happ(s(fun(X49,bool),happ(s(fun(X49,fun(X49,bool)),X50),s(X49,X52))),s(X49,X51))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.relation.symmetric_1'(s(fun(X49,fun(X49,bool)),X50))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f118,f89])).
fof(f4472,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38))) = s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X38))),s(X35,X37))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.relation.symmetric_1'(s(fun(X35,fun(X35,bool)),X36)))) )),
  inference(resolution,[],[f118,f88])).
fof(f4470,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,happ(s(fun(X21,bool),happ(s(fun(X21,fun(X21,bool)),X22),s(X21,X23))),s(X21,X24))) = s(bool,happ(s(fun(X21,bool),happ(s(fun(X21,fun(X21,bool)),X22),s(X21,X24))),s(X21,X23))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(X21,fun(X21,bool)),X22)))) )),
  inference(resolution,[],[f118,f87])).
fof(f4469,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,X16))),s(X14,X17))) = s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,X17))),s(X14,X16))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.symmetric_1'(s(fun(X14,fun(X14,bool)),X15))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f118,f85])).
fof(f4467,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) )),
  inference(resolution,[],[f118,f83])).
fof(f745,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl38),
  inference(cnf_transformation,[],[f745_D])).
fof(f745_D,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4459,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X45) != s(bool,X47)) ) | $spl6),
  inference(resolution,[],[f504,f125])).
fof(f4456,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl6),
  inference(resolution,[],[f504,f89])).
fof(f4453,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl6),
  inference(resolution,[],[f504,f88])).
fof(f4452,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl6),
  inference(resolution,[],[f504,f87])).
fof(f4451,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl6),
  inference(resolution,[],[f504,f85])).
fof(f4449,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl6),
  inference(resolution,[],[f504,f83])).
fof(f3610,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4)) ) | $spl114),
  inference(cnf_transformation,[],[f3610_D])).
fof(f3610_D,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f4446,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl118),
  inference(trivial_inequality_removal,[],[f4443])).
fof(f4443,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X5) | s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl118),
  inference(superposition,[],[f3615,f122])).
fof(f4435,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl112),
  inference(duplicate_literal_removal,[],[f4425])).
fof(f4425,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl112),
  inference(superposition,[],[f3609,f122])).
fof(f240,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5)) ) | $spl20),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4436,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl112),
  inference(subsumption_resolution,[],[f4424,f3609])).
fof(f4424,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0)) ) | $spl112),
  inference(superposition,[],[f3609,f122])).
fof(f470,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl28),
  inference(cnf_transformation,[],[f470_D])).
fof(f470_D,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4415,plain,(
  ( ! [X35,X36,X34] : (s(bool,'F_0') = s(bool,X34) | p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X34) != s(bool,X36)) ) | $spl24),
  inference(resolution,[],[f432,f125])).
fof(f4412,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl24),
  inference(resolution,[],[f432,f89])).
fof(f4409,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl24),
  inference(resolution,[],[f432,f88])).
fof(f4408,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl24),
  inference(resolution,[],[f432,f87])).
fof(f4407,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl24),
  inference(resolution,[],[f432,f85])).
fof(f4405,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl24),
  inference(resolution,[],[f432,f83])).
fof(f496,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X3))) ) | $spl32),
  inference(cnf_transformation,[],[f496_D])).
fof(f496_D,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X3))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4389,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,X25)) )),
  inference(resolution,[],[f115,f125])).
fof(f4385,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f115,f89])).
fof(f4383,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) )),
  inference(resolution,[],[f115,f88])).
fof(f4381,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) )),
  inference(resolution,[],[f115,f87])).
fof(f4380,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f115,f85])).
fof(f4378,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) )),
  inference(resolution,[],[f115,f83])).
fof(f4377,plain,(
  ~p(s(bool,'F_0')) | $spl131),
  inference(cnf_transformation,[],[f4377_D])).
fof(f4377_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f101,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X10)))),
  file('thm.relation.symmetric_inv_image','thm.bool.IMP_CLAUSES|split|4')).
fof(f4375,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(subsumption_resolution,[],[f105,f126])).
fof(f105,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.relation.symmetric_inv_image','thm.bool.IMP_CLAUSES|split|2')).
fof(f4305,plain,(
  $false | ($spl32 | $spl114)),
  inference(subsumption_resolution,[],[f4304,f117])).
fof(f117,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.symmetric_inv_image',aHL_TRUTH)).
fof(f4304,plain,(
  ~p(s(bool,'T_0')) | ($spl32 | $spl114)),
  inference(forward_demodulation,[],[f4303,f3610])).
fof(f4303,plain,(
  ( ! [X325,X323,X324,X322] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X322,sK1),X323),s(X322,X324),s(X322,X325))))) ) | ($spl32 | $spl114)),
  inference(subsumption_resolution,[],[f4092,f3610])).
fof(f4092,plain,(
  ( ! [X325,X323,X324,X322] : (s(bool,'T_0') != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X322,sK1),X323),s(X322,X325),s(X322,X324))) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X322,sK1),X323),s(X322,X324),s(X322,X325))))) ) | ($spl32 | $spl114)),
  inference(backward_demodulation,[],[f3610,f2435])).
fof(f4302,plain,(
  $false | ($spl32 | $spl114)),
  inference(subsumption_resolution,[],[f4301,f117])).
fof(f4301,plain,(
  ~p(s(bool,'T_0')) | ($spl32 | $spl114)),
  inference(forward_demodulation,[],[f4300,f3610])).
fof(f4300,plain,(
  ( ! [X227,X226] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X226))),s(sK1,X227))))) ) | ($spl32 | $spl114)),
  inference(subsumption_resolution,[],[f4089,f3610])).
fof(f4089,plain,(
  ( ! [X227,X226] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X227))),s(sK1,X226))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X226))),s(sK1,X227))))) ) | ($spl32 | $spl114)),
  inference(backward_demodulation,[],[f3610,f2014])).
fof(f2014,plain,(
  ( ! [X227,X226] : (s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X227))),s(sK1,X226))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X226))),s(sK1,X227))))) ) | $spl32),
  inference(superposition,[],[f496,f127])).
fof(f4298,plain,(
  $false | $spl114),
  inference(subsumption_resolution,[],[f4071,f117])).
fof(f4071,plain,(
  ~p(s(bool,'T_0')) | $spl114),
  inference(backward_demodulation,[],[f3610,f115])).
fof(f4281,plain,(
  $false | ($spl27 | $spl114)),
  inference(trivial_inequality_removal,[],[f4270])).
fof(f4270,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl27 | $spl114)),
  inference(superposition,[],[f435,f3610])).
fof(f4280,plain,(
  $false | ($spl61 | $spl114)),
  inference(trivial_inequality_removal,[],[f4275])).
fof(f4275,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl61 | $spl114)),
  inference(superposition,[],[f1558,f3610])).
fof(f4065,plain,(
  $false | $spl18),
  inference(trivial_inequality_removal,[],[f4047])).
fof(f4047,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl18),
  inference(superposition,[],[f237,f180])).
fof(f4062,plain,(
  $false | $spl18),
  inference(equality_resolution,[],[f237])).
fof(f3615,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5)) ) | $spl118),
  inference(cnf_transformation,[],[f3615_D])).
fof(f3615_D,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f3607,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) != s(bool,X5)) ) | $spl108),
  inference(duplicate_literal_removal,[],[f3589])).
fof(f3589,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl108),
  inference(superposition,[],[f3566,f122])).
fof(f3609,plain,(
  ( ! [X5,X3] : (s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl112),
  inference(cnf_transformation,[],[f3609_D])).
fof(f3609_D,plain,(
  ( ! [X5,X3] : (s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f3608,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X4) | s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl108),
  inference(subsumption_resolution,[],[f3571,f122])).
fof(f3571,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl108),
  inference(superposition,[],[f3566,f122])).
fof(f3611,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'F_0') != s(bool,X11) | s(bool,'T_0') = s(bool,X9)) ) | $spl108),
  inference(subsumption_resolution,[],[f3573,f3566])).
fof(f3573,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,'F_0') != s(bool,X11) | s(bool,'T_0') = s(bool,X9)) ) | $spl108),
  inference(superposition,[],[f3566,f122])).
fof(f4038,plain,(
  ( ! [X520,X516,X518,X514,X521,X517,X519,X513,X515] : (s(bool,X517) != s(bool,X518) | s(bool,X519) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X513,sK1),X514),s(X513,X515),s(X513,X516))) | s(bool,X519) != s(bool,X520) | s(bool,'F_0') = s(bool,X520) | s(bool,X518) != s(bool,X521) | s(bool,'T_0') = s(bool,X521) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X513,sK1),X514),s(X513,X516),s(X513,X515)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f3912,f789])).
fof(f3912,plain,(
  ( ! [X520,X516,X518,X514,X521,X517,X519,X513,X515] : (s(bool,'F_0') != s(bool,X517) | s(bool,X517) != s(bool,X518) | s(bool,X519) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X513,sK1),X514),s(X513,X515),s(X513,X516))) | s(bool,X519) != s(bool,X520) | s(bool,'F_0') = s(bool,X520) | s(bool,X518) != s(bool,X521) | s(bool,'T_0') = s(bool,X521) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X513,sK1),X514),s(X513,X516),s(X513,X515)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f2343])).
fof(f4035,plain,(
  ( ! [X484,X486,X480,X482,X485,X487,X481,X483] : (s(bool,X484) != s(bool,X485) | s(bool,X486) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X480,sK1),X481),s(X480,X482),s(X480,X483))) | s(bool,'F_0') = s(bool,X486) | s(bool,X485) != s(bool,X487) | s(bool,'T_0') = s(bool,X487) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X480,sK1),X481),s(X480,X483),s(X480,X482)))) ) | ($spl34 | $spl38 | $spl56)),
  inference(subsumption_resolution,[],[f3908,f789])).
fof(f3908,plain,(
  ( ! [X484,X486,X480,X482,X485,X487,X481,X483] : (s(bool,'F_0') != s(bool,X484) | s(bool,X484) != s(bool,X485) | s(bool,X486) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X480,sK1),X481),s(X480,X482),s(X480,X483))) | s(bool,'F_0') = s(bool,X486) | s(bool,X485) != s(bool,X487) | s(bool,'T_0') = s(bool,X487) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X480,sK1),X481),s(X480,X483),s(X480,X482)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f2343])).
fof(f4018,plain,(
  ( ! [X313,X315,X310,X312,X314,X311] : (p(s(bool,X314)) | ~p(s(bool,X315)) | s(bool,X315) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X310,sK1),X311),s(X310,X312),s(X310,X313))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X310,sK1),X311),s(X310,X313),s(X310,X312)))) )),
  inference(subsumption_resolution,[],[f3877,f101])).
fof(f3877,plain,(
  ( ! [X313,X315,X310,X312,X314,X311] : (p(s(bool,'F_0')) | p(s(bool,X314)) | ~p(s(bool,X315)) | s(bool,X315) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X310,sK1),X311),s(X310,X312),s(X310,X313))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X310,sK1),X311),s(X310,X313),s(X310,X312)))) )),
  inference(superposition,[],[f125,f2343])).
fof(f4014,plain,(
  ( ! [X243,X241,X246,X244,X242,X247,X245] : (~p(s(bool,X245)) | ~p(s(bool,X246)) | s(bool,X246) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X241,sK1),X242),s(X241,X243),s(X241,X244))) | s(bool,X245) != s(bool,X247) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X241,sK1),X242),s(X241,X244),s(X241,X243)))) )),
  inference(subsumption_resolution,[],[f3864,f101])).
fof(f3864,plain,(
  ( ! [X243,X241,X246,X244,X242,X247,X245] : (p(s(bool,'F_0')) | ~p(s(bool,X245)) | ~p(s(bool,X246)) | s(bool,X246) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X241,sK1),X242),s(X241,X243),s(X241,X244))) | s(bool,X245) != s(bool,X247) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X241,sK1),X242),s(X241,X244),s(X241,X243)))) )),
  inference(superposition,[],[f89,f2343])).
fof(f4009,plain,(
  ( ! [X191,X189,X187,X185,X190,X188,X186] : (p(s(bool,X189)) | ~p(s(bool,X190)) | ~p(s(bool,X191)) | s(bool,X190) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X185,sK1),X186),s(X185,X187),s(X185,X188))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X185,sK1),X186),s(X185,X188),s(X185,X187)))) )),
  inference(subsumption_resolution,[],[f3856,f101])).
fof(f3856,plain,(
  ( ! [X191,X189,X187,X185,X190,X188,X186] : (p(s(bool,'F_0')) | p(s(bool,X189)) | ~p(s(bool,X190)) | ~p(s(bool,X191)) | s(bool,X190) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X185,sK1),X186),s(X185,X187),s(X185,X188))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X185,sK1),X186),s(X185,X188),s(X185,X187)))) )),
  inference(superposition,[],[f88,f2343])).
fof(f4004,plain,(
  ( ! [X156,X154,X152,X151,X155,X153,X150] : (p(s(bool,X154)) | ~p(s(bool,X155)) | s(bool,X154) != s(bool,X156) | s(bool,X155) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X150,sK1),X151),s(X150,X152),s(X150,X153))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X150,sK1),X151),s(X150,X153),s(X150,X152)))) )),
  inference(subsumption_resolution,[],[f3851,f101])).
fof(f3851,plain,(
  ( ! [X156,X154,X152,X151,X155,X153,X150] : (p(s(bool,'F_0')) | p(s(bool,X154)) | ~p(s(bool,X155)) | s(bool,X154) != s(bool,X156) | s(bool,X155) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X150,sK1),X151),s(X150,X152),s(X150,X153))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X150,sK1),X151),s(X150,X153),s(X150,X152)))) )),
  inference(superposition,[],[f87,f2343])).
fof(f4001,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (~p(s(bool,X105)) | ~p(s(bool,X106)) | s(bool,X106) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X101,sK1),X102),s(X101,X103),s(X101,X104))) | s(bool,X105) != s(bool,X107) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X101,sK1),X102),s(X101,X104),s(X101,X103)))) )),
  inference(subsumption_resolution,[],[f3844,f101])).
fof(f3844,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (p(s(bool,'F_0')) | ~p(s(bool,X105)) | ~p(s(bool,X106)) | s(bool,X106) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X101,sK1),X102),s(X101,X103),s(X101,X104))) | s(bool,X105) != s(bool,X107) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X101,sK1),X102),s(X101,X104),s(X101,X103)))) )),
  inference(superposition,[],[f85,f2343])).
fof(f3995,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X56) != s(bool,X58) | s(bool,X57) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X52,sK1),X53),s(X52,X54),s(X52,X55))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X52,sK1),X53),s(X52,X55),s(X52,X54)))) )),
  inference(subsumption_resolution,[],[f3837,f101])).
fof(f3837,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,'F_0')) | p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X56) != s(bool,X58) | s(bool,X57) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X52,sK1),X53),s(X52,X54),s(X52,X55))) | s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X52,sK1),X53),s(X52,X55),s(X52,X54)))) )),
  inference(superposition,[],[f83,f2343])).
fof(f3992,plain,(
  ( ! [X509,X511,X505,X507,X512,X508,X510,X504,X506] : (s(bool,X509) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X504,sK1),X505),s(X504,X506),s(X504,X507))) | s(bool,X508) != s(bool,X510) | s(bool,X510) != s(bool,X511) | s(bool,'F_0') = s(bool,X511) | s(bool,X509) != s(bool,X512) | s(bool,'T_0') = s(bool,X512) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X504,sK1),X505),s(X504,X507),s(X504,X506)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f3806,f789])).
fof(f3806,plain,(
  ( ! [X509,X511,X505,X507,X512,X508,X510,X504,X506] : (s(bool,'T_0') != s(bool,X508) | s(bool,X509) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X504,sK1),X505),s(X504,X506),s(X504,X507))) | s(bool,X508) != s(bool,X510) | s(bool,X510) != s(bool,X511) | s(bool,'F_0') = s(bool,X511) | s(bool,X509) != s(bool,X512) | s(bool,'T_0') = s(bool,X512) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X504,sK1),X505),s(X504,X507),s(X504,X506)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f2343])).
fof(f3991,plain,(
  ( ! [X476,X478,X472,X474,X475,X477,X479,X473] : (s(bool,X477) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X472,sK1),X473),s(X472,X474),s(X472,X475))) | s(bool,X476) != s(bool,X478) | s(bool,'F_0') = s(bool,X478) | s(bool,X477) != s(bool,X479) | s(bool,'T_0') = s(bool,X479) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X472,sK1),X473),s(X472,X475),s(X472,X474)))) ) | ($spl6 | $spl34 | $spl56)),
  inference(subsumption_resolution,[],[f3802,f589])).
fof(f3802,plain,(
  ( ! [X476,X478,X472,X474,X475,X477,X479,X473] : (s(bool,'T_0') != s(bool,X476) | s(bool,X477) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X472,sK1),X473),s(X472,X474),s(X472,X475))) | s(bool,X476) != s(bool,X478) | s(bool,'F_0') = s(bool,X478) | s(bool,X477) != s(bool,X479) | s(bool,'T_0') = s(bool,X479) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X472,sK1),X473),s(X472,X475),s(X472,X474)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f2343])).
fof(f3988,plain,(
  ( ! [X428,X430,X432,X429,X431,X427,X433] : (s(bool,X432) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X427,sK1),X428),s(X427,X429),s(X427,X430))) | s(bool,'F_0') = s(bool,X431) | s(bool,X432) != s(bool,X433) | s(bool,'T_0') = s(bool,X433) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X427,sK1),X428),s(X427,X430),s(X427,X429)))) ) | ($spl34 | $spl40)),
  inference(subsumption_resolution,[],[f3795,f122])).
fof(f3795,plain,(
  ( ! [X428,X430,X432,X429,X431,X427,X433] : (s(bool,'T_0') != s(bool,X431) | s(bool,X432) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X427,sK1),X428),s(X427,X429),s(X427,X430))) | s(bool,'F_0') = s(bool,X431) | s(bool,X432) != s(bool,X433) | s(bool,'T_0') = s(bool,X433) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X427,sK1),X428),s(X427,X430),s(X427,X429)))) ) | ($spl34 | $spl40)),
  inference(superposition,[],[f1342,f2343])).
fof(f3986,plain,(
  ( ! [X389,X391,X393,X390,X392,X394] : (s(bool,X394) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X389,sK1),X390),s(X389,X391),s(X389,X392))) | s(bool,'F_0') = s(bool,X393) | s(bool,'T_0') = s(bool,X394) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X389,sK1),X390),s(X389,X392),s(X389,X391)))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f3788,f122])).
fof(f3788,plain,(
  ( ! [X389,X391,X393,X390,X392,X394] : (s(bool,'T_0') != s(bool,X393) | s(bool,X394) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X389,sK1),X390),s(X389,X391),s(X389,X392))) | s(bool,'F_0') = s(bool,X393) | s(bool,'T_0') = s(bool,X394) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X389,sK1),X390),s(X389,X392),s(X389,X391)))) ) | ($spl34 | $spl38)),
  inference(superposition,[],[f789,f2343])).
fof(f3982,plain,(
  ( ! [X325,X327,X323,X324,X326,X322] : (p(s(bool,X326)) | p(s(bool,X327)) | s(bool,X327) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X322,sK1),X323),s(X322,X324),s(X322,X325))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X322,sK1),X323),s(X322,X325),s(X322,X324)))) )),
  inference(subsumption_resolution,[],[f3774,f126])).
fof(f3774,plain,(
  ( ! [X325,X327,X323,X324,X326,X322] : (~p(s(bool,'T_0')) | p(s(bool,X326)) | p(s(bool,X327)) | s(bool,X327) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X322,sK1),X323),s(X322,X324),s(X322,X325))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X322,sK1),X323),s(X322,X325),s(X322,X324)))) )),
  inference(superposition,[],[f125,f2343])).
fof(f3980,plain,(
  ( ! [X255,X261,X257,X259,X260,X256,X258] : (p(s(bool,X259)) | ~p(s(bool,X260)) | s(bool,X259) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X255,sK1),X256),s(X255,X257),s(X255,X258))) | s(bool,X260) != s(bool,X261) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X255,sK1),X256),s(X255,X258),s(X255,X257)))) )),
  inference(subsumption_resolution,[],[f3761,f126])).
fof(f3761,plain,(
  ( ! [X255,X261,X257,X259,X260,X256,X258] : (~p(s(bool,'T_0')) | p(s(bool,X259)) | ~p(s(bool,X260)) | s(bool,X259) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X255,sK1),X256),s(X255,X257),s(X255,X258))) | s(bool,X260) != s(bool,X261) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X255,sK1),X256),s(X255,X258),s(X255,X257)))) )),
  inference(superposition,[],[f89,f2343])).
fof(f3976,plain,(
  ( ! [X204,X202,X200,X199,X205,X203,X201] : (p(s(bool,X203)) | ~p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X199,sK1),X200),s(X199,X201),s(X199,X202))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X199,sK1),X200),s(X199,X202),s(X199,X201)))) )),
  inference(subsumption_resolution,[],[f3753,f126])).
fof(f3753,plain,(
  ( ! [X204,X202,X200,X199,X205,X203,X201] : (~p(s(bool,'T_0')) | p(s(bool,X203)) | ~p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X199,sK1),X200),s(X199,X201),s(X199,X202))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X199,sK1),X200),s(X199,X202),s(X199,X201)))) )),
  inference(superposition,[],[f88,f2343])).
fof(f3974,plain,(
  ( ! [X167,X165,X169,X166,X164,X170,X168] : (p(s(bool,X168)) | p(s(bool,X169)) | s(bool,X168) != s(bool,X170) | s(bool,X169) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X164,sK1),X165),s(X164,X166),s(X164,X167))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X164,sK1),X165),s(X164,X167),s(X164,X166)))) )),
  inference(subsumption_resolution,[],[f3748,f126])).
fof(f3748,plain,(
  ( ! [X167,X165,X169,X166,X164,X170,X168] : (~p(s(bool,'T_0')) | p(s(bool,X168)) | p(s(bool,X169)) | s(bool,X168) != s(bool,X170) | s(bool,X169) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X164,sK1),X165),s(X164,X166),s(X164,X167))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X164,sK1),X165),s(X164,X167),s(X164,X166)))) )),
  inference(superposition,[],[f87,f2343])).
fof(f3971,plain,(
  ( ! [X121,X118,X116,X120,X119,X117,X115] : (p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X115,sK1),X116),s(X115,X117),s(X115,X118))) | s(bool,X120) != s(bool,X121) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X115,sK1),X116),s(X115,X118),s(X115,X117)))) )),
  inference(subsumption_resolution,[],[f3741,f126])).
fof(f3741,plain,(
  ( ! [X121,X118,X116,X120,X119,X117,X115] : (~p(s(bool,'T_0')) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X115,sK1),X116),s(X115,X117),s(X115,X118))) | s(bool,X120) != s(bool,X121) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X115,sK1),X116),s(X115,X118),s(X115,X117)))) )),
  inference(superposition,[],[f85,f2343])).
fof(f3967,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (p(s(bool,X70)) | p(s(bool,X71)) | s(bool,X70) != s(bool,X72) | s(bool,X71) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X66,sK1),X67),s(X66,X68),s(X66,X69))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X66,sK1),X67),s(X66,X69),s(X66,X68)))) )),
  inference(subsumption_resolution,[],[f3734,f126])).
fof(f3734,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X70)) | p(s(bool,X71)) | s(bool,X70) != s(bool,X72) | s(bool,X71) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X66,sK1),X67),s(X66,X68),s(X66,X69))) | s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X66,sK1),X67),s(X66,X69),s(X66,X68)))) )),
  inference(superposition,[],[f83,f2343])).
fof(f3617,plain,(
  ( ! [X56,X55] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X55))),s(sK1,X56))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X56))),s(sK1,X55)))) ) | $spl120),
  inference(cnf_transformation,[],[f3617_D])).
fof(f3617_D,plain,(
  ( ! [X56,X55] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X55))),s(sK1,X56))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X56))),s(sK1,X55)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f3600,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'F_0') != s(bool,X57) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X55))),s(sK1,X56))) | s(bool,X57) != s(bool,X58) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X56))),s(sK1,X55)))) ) | $spl108),
  inference(superposition,[],[f3566,f1951])).
fof(f3614,plain,(
  ( ! [X59,X57,X56] : (s(bool,X59) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X56))),s(sK1,X57))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X57))),s(sK1,X56)))) ) | $spl116),
  inference(cnf_transformation,[],[f3614_D])).
fof(f3614_D,plain,(
  ( ! [X59,X57,X56] : (s(bool,X59) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X56))),s(sK1,X57))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X57))),s(sK1,X56)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f3613,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') = s(bool,X58) | s(bool,X59) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X56))),s(sK1,X57))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X57))),s(sK1,X56)))) ) | $spl108),
  inference(subsumption_resolution,[],[f3582,f122])).
fof(f3582,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'F_0') != s(bool,X58) | s(bool,'T_0') = s(bool,X58) | s(bool,X59) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X56))),s(sK1,X57))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X57))),s(sK1,X56)))) ) | $spl108),
  inference(superposition,[],[f3566,f1951])).
fof(f178,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))) ) | $spl16),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3612,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X14) | s(bool,X15) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))) ) | $spl108),
  inference(subsumption_resolution,[],[f3575,f122])).
fof(f3575,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))) ) | $spl108),
  inference(superposition,[],[f3566,f180])).
fof(f3568,plain,(
  ( ! [X356,X358,X354,X357,X359,X353,X355] : (s(bool,X355) != s(bool,X356) | s(bool,X357) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X353))),s(sK1,X354))) | s(bool,X357) != s(bool,X358) | s(bool,'F_0') = s(bool,X358) | s(bool,X356) != s(bool,X359) | s(bool,'T_0') = s(bool,X359) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X354))),s(sK1,X353)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f3450,f789])).
fof(f3450,plain,(
  ( ! [X356,X358,X354,X357,X359,X353,X355] : (s(bool,'F_0') != s(bool,X355) | s(bool,X355) != s(bool,X356) | s(bool,X357) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X353))),s(sK1,X354))) | s(bool,X357) != s(bool,X358) | s(bool,'F_0') = s(bool,X358) | s(bool,X356) != s(bool,X359) | s(bool,'T_0') = s(bool,X359) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X354))),s(sK1,X353)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f1951])).
fof(f3564,plain,(
  ( ! [X333,X329,X331,X332,X328,X330] : (s(bool,X330) != s(bool,X331) | s(bool,X332) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X328))),s(sK1,X329))) | s(bool,'F_0') = s(bool,X332) | s(bool,X331) != s(bool,X333) | s(bool,'T_0') = s(bool,X333) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X329))),s(sK1,X328)))) ) | ($spl34 | $spl38 | $spl56)),
  inference(subsumption_resolution,[],[f3446,f789])).
fof(f3446,plain,(
  ( ! [X333,X329,X331,X332,X328,X330] : (s(bool,'F_0') != s(bool,X330) | s(bool,X330) != s(bool,X331) | s(bool,X332) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X328))),s(sK1,X329))) | s(bool,'F_0') = s(bool,X332) | s(bool,X331) != s(bool,X333) | s(bool,'T_0') = s(bool,X333) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X329))),s(sK1,X328)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f1951])).
fof(f3547,plain,(
  ( ! [X222,X220,X223,X221] : (p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X223) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X220))),s(sK1,X221))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X221))),s(sK1,X220)))) )),
  inference(subsumption_resolution,[],[f3415,f101])).
fof(f3415,plain,(
  ( ! [X222,X220,X223,X221] : (p(s(bool,'F_0')) | p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X223) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X220))),s(sK1,X221))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X221))),s(sK1,X220)))) )),
  inference(superposition,[],[f125,f1951])).
fof(f3543,plain,(
  ( ! [X182,X180,X183,X181,X179] : (~p(s(bool,X181)) | ~p(s(bool,X182)) | s(bool,X182) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X179))),s(sK1,X180))) | s(bool,X181) != s(bool,X183) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X180))),s(sK1,X179)))) )),
  inference(subsumption_resolution,[],[f3402,f101])).
fof(f3402,plain,(
  ( ! [X182,X180,X183,X181,X179] : (p(s(bool,'F_0')) | ~p(s(bool,X181)) | ~p(s(bool,X182)) | s(bool,X182) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X179))),s(sK1,X180))) | s(bool,X181) != s(bool,X183) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X180))),s(sK1,X179)))) )),
  inference(superposition,[],[f89,f1951])).
fof(f3538,plain,(
  ( ! [X142,X140,X143,X141,X139] : (p(s(bool,X141)) | ~p(s(bool,X142)) | ~p(s(bool,X143)) | s(bool,X142) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X139))),s(sK1,X140))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X140))),s(sK1,X139)))) )),
  inference(subsumption_resolution,[],[f3394,f101])).
fof(f3394,plain,(
  ( ! [X142,X140,X143,X141,X139] : (p(s(bool,'F_0')) | p(s(bool,X141)) | ~p(s(bool,X142)) | ~p(s(bool,X143)) | s(bool,X142) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X139))),s(sK1,X140))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X140))),s(sK1,X139)))) )),
  inference(superposition,[],[f88,f1951])).
fof(f3533,plain,(
  ( ! [X118,X116,X114,X117,X115] : (p(s(bool,X116)) | ~p(s(bool,X117)) | s(bool,X116) != s(bool,X118) | s(bool,X117) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X114))),s(sK1,X115))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X115))),s(sK1,X114)))) )),
  inference(subsumption_resolution,[],[f3389,f101])).
fof(f3389,plain,(
  ( ! [X118,X116,X114,X117,X115] : (p(s(bool,'F_0')) | p(s(bool,X116)) | ~p(s(bool,X117)) | s(bool,X116) != s(bool,X118) | s(bool,X117) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X114))),s(sK1,X115))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X115))),s(sK1,X114)))) )),
  inference(superposition,[],[f87,f1951])).
fof(f3530,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X79))),s(sK1,X80))) | s(bool,X81) != s(bool,X83) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X80))),s(sK1,X79)))) )),
  inference(subsumption_resolution,[],[f3382,f101])).
fof(f3382,plain,(
  ( ! [X80,X83,X81,X79,X82] : (p(s(bool,'F_0')) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X79))),s(sK1,X80))) | s(bool,X81) != s(bool,X83) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X80))),s(sK1,X79)))) )),
  inference(superposition,[],[f85,f1951])).
fof(f3524,plain,(
  ( ! [X47,X45,X48,X46,X44] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X44))),s(sK1,X45))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X45))),s(sK1,X44)))) )),
  inference(subsumption_resolution,[],[f3375,f101])).
fof(f3375,plain,(
  ( ! [X47,X45,X48,X46,X44] : (p(s(bool,'F_0')) | p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X44))),s(sK1,X45))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X45))),s(sK1,X44)))) )),
  inference(superposition,[],[f83,f1951])).
fof(f3518,plain,(
  ( ! [X350,X346,X352,X349,X351,X347,X348] : (s(bool,X349) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X346))),s(sK1,X347))) | s(bool,X348) != s(bool,X350) | s(bool,X350) != s(bool,X351) | s(bool,'F_0') = s(bool,X351) | s(bool,X349) != s(bool,X352) | s(bool,'T_0') = s(bool,X352) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X347))),s(sK1,X346)))) ) | ($spl34 | $spl38 | $spl66)),
  inference(subsumption_resolution,[],[f3336,f789])).
fof(f3336,plain,(
  ( ! [X350,X346,X352,X349,X351,X347,X348] : (s(bool,'T_0') != s(bool,X348) | s(bool,X349) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X346))),s(sK1,X347))) | s(bool,X348) != s(bool,X350) | s(bool,X350) != s(bool,X351) | s(bool,'F_0') = s(bool,X351) | s(bool,X349) != s(bool,X352) | s(bool,'T_0') = s(bool,X352) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X347))),s(sK1,X346)))) ) | ($spl34 | $spl66)),
  inference(superposition,[],[f2201,f1951])).
fof(f3517,plain,(
  ( ! [X325,X327,X323,X324,X326,X322] : (s(bool,X325) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X322))),s(sK1,X323))) | s(bool,X324) != s(bool,X326) | s(bool,'F_0') = s(bool,X326) | s(bool,X325) != s(bool,X327) | s(bool,'T_0') = s(bool,X327) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X323))),s(sK1,X322)))) ) | ($spl6 | $spl34 | $spl56)),
  inference(subsumption_resolution,[],[f3332,f589])).
fof(f3332,plain,(
  ( ! [X325,X327,X323,X324,X326,X322] : (s(bool,'T_0') != s(bool,X324) | s(bool,X325) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X322))),s(sK1,X323))) | s(bool,X324) != s(bool,X326) | s(bool,'F_0') = s(bool,X326) | s(bool,X325) != s(bool,X327) | s(bool,'T_0') = s(bool,X327) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X323))),s(sK1,X322)))) ) | ($spl34 | $spl56)),
  inference(superposition,[],[f1836,f1951])).
fof(f3514,plain,(
  ( ! [X292,X294,X293,X295,X291] : (s(bool,X294) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X291))),s(sK1,X292))) | s(bool,'F_0') = s(bool,X293) | s(bool,X294) != s(bool,X295) | s(bool,'T_0') = s(bool,X295) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X292))),s(sK1,X291)))) ) | ($spl34 | $spl40)),
  inference(subsumption_resolution,[],[f3325,f122])).
fof(f3325,plain,(
  ( ! [X292,X294,X293,X295,X291] : (s(bool,'T_0') != s(bool,X293) | s(bool,X294) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X291))),s(sK1,X292))) | s(bool,'F_0') = s(bool,X293) | s(bool,X294) != s(bool,X295) | s(bool,'T_0') = s(bool,X295) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X292))),s(sK1,X291)))) ) | ($spl34 | $spl40)),
  inference(superposition,[],[f1342,f1951])).
fof(f3512,plain,(
  ( ! [X269,X267,X268,X270] : (s(bool,X270) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X267))),s(sK1,X268))) | s(bool,'F_0') = s(bool,X269) | s(bool,'T_0') = s(bool,X270) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X268))),s(sK1,X267)))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f3318,f122])).
fof(f3318,plain,(
  ( ! [X269,X267,X268,X270] : (s(bool,'T_0') != s(bool,X269) | s(bool,X270) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X267))),s(sK1,X268))) | s(bool,'F_0') = s(bool,X269) | s(bool,'T_0') = s(bool,X270) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X268))),s(sK1,X267)))) ) | ($spl34 | $spl38)),
  inference(superposition,[],[f789,f1951])).
fof(f3508,plain,(
  ( ! [X231,X229,X230,X228] : (p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X231) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X228))),s(sK1,X229))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X229))),s(sK1,X228)))) )),
  inference(subsumption_resolution,[],[f3304,f126])).
fof(f3304,plain,(
  ( ! [X231,X229,X230,X228] : (~p(s(bool,'T_0')) | p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X231) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X228))),s(sK1,X229))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X229))),s(sK1,X228)))) )),
  inference(superposition,[],[f125,f1951])).
fof(f3506,plain,(
  ( ! [X191,X189,X192,X190,X193] : (p(s(bool,X191)) | ~p(s(bool,X192)) | s(bool,X191) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X189))),s(sK1,X190))) | s(bool,X192) != s(bool,X193) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X190))),s(sK1,X189)))) )),
  inference(subsumption_resolution,[],[f3291,f126])).
fof(f3291,plain,(
  ( ! [X191,X189,X192,X190,X193] : (~p(s(bool,'T_0')) | p(s(bool,X191)) | ~p(s(bool,X192)) | s(bool,X191) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X189))),s(sK1,X190))) | s(bool,X192) != s(bool,X193) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X190))),s(sK1,X189)))) )),
  inference(superposition,[],[f89,f1951])).
fof(f3502,plain,(
  ( ! [X152,X151,X149,X153,X150] : (p(s(bool,X151)) | ~p(s(bool,X152)) | p(s(bool,X153)) | s(bool,X153) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X149))),s(sK1,X150))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X150))),s(sK1,X149)))) )),
  inference(subsumption_resolution,[],[f3283,f126])).
fof(f3283,plain,(
  ( ! [X152,X151,X149,X153,X150] : (~p(s(bool,'T_0')) | p(s(bool,X151)) | ~p(s(bool,X152)) | p(s(bool,X153)) | s(bool,X153) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X149))),s(sK1,X150))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X150))),s(sK1,X149)))) )),
  inference(superposition,[],[f88,f1951])).
fof(f3500,plain,(
  ( ! [X127,X125,X128,X126,X124] : (p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X126) != s(bool,X128) | s(bool,X127) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X124))),s(sK1,X125))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X125))),s(sK1,X124)))) )),
  inference(subsumption_resolution,[],[f3278,f126])).
fof(f3278,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,'T_0')) | p(s(bool,X126)) | p(s(bool,X127)) | s(bool,X126) != s(bool,X128) | s(bool,X127) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X124))),s(sK1,X125))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X125))),s(sK1,X124)))) )),
  inference(superposition,[],[f87,f1951])).
fof(f3497,plain,(
  ( ! [X92,X90,X93,X91,X89] : (p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X89))),s(sK1,X90))) | s(bool,X92) != s(bool,X93) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X90))),s(sK1,X89)))) )),
  inference(subsumption_resolution,[],[f3271,f126])).
fof(f3271,plain,(
  ( ! [X92,X90,X93,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X89))),s(sK1,X90))) | s(bool,X92) != s(bool,X93) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X90))),s(sK1,X89)))) )),
  inference(superposition,[],[f85,f1951])).
fof(f3493,plain,(
  ( ! [X57,X54,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X56) != s(bool,X58) | s(bool,X57) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X54))),s(sK1,X55))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X55))),s(sK1,X54)))) )),
  inference(subsumption_resolution,[],[f3264,f126])).
fof(f3264,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X56) != s(bool,X58) | s(bool,X57) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X54))),s(sK1,X55))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X55))),s(sK1,X54)))) )),
  inference(superposition,[],[f83,f1951])).
fof(f3202,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (s(bool,'T_0') = s(bool,happ(s(fun(X76,bool),happ(s(fun(X76,fun(X76,bool)),X77),s(X76,X78))),s(X76,X79))) | s(bool,'F_0') = s(bool,happ(s(fun(X76,bool),happ(s(fun(X76,fun(X76,bool)),X77),s(X76,X79))),s(X76,X78))) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,'const.relation.symmetric_1'(s(fun(X76,fun(X76,bool)),X77)))) )),
  inference(resolution,[],[f1123,f125])).
fof(f3197,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,happ(s(fun(X53,bool),happ(s(fun(X53,fun(X53,bool)),X54),s(X53,X55))),s(X53,X56))) | s(bool,'F_0') = s(bool,happ(s(fun(X53,bool),happ(s(fun(X53,fun(X53,bool)),X54),s(X53,X56))),s(X53,X55))) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.symmetric_1'(s(fun(X53,fun(X53,bool)),X54))) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f1123,f89])).
fof(f3195,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),X40),s(X39,X41))),s(X39,X42))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),X40),s(X39,X42))),s(X39,X41))) | p(s(bool,X43)) | ~p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.relation.symmetric_1'(s(fun(X39,fun(X39,bool)),X40)))) )),
  inference(resolution,[],[f1123,f88])).
fof(f3193,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),happ(s(fun(X25,fun(X25,bool)),X26),s(X25,X27))),s(X25,X28))) | s(bool,'F_0') = s(bool,happ(s(fun(X25,bool),happ(s(fun(X25,fun(X25,bool)),X26),s(X25,X28))),s(X25,X27))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.symmetric_1'(s(fun(X25,fun(X25,bool)),X26)))) )),
  inference(resolution,[],[f1123,f87])).
fof(f3192,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(X18,bool),happ(s(fun(X18,fun(X18,bool)),X19),s(X18,X20))),s(X18,X21))) | s(bool,'F_0') = s(bool,happ(s(fun(X18,bool),happ(s(fun(X18,fun(X18,bool)),X19),s(X18,X21))),s(X18,X20))) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.relation.symmetric_1'(s(fun(X18,fun(X18,bool)),X19))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f1123,f85])).
fof(f3190,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X6))),s(X4,X7))) | s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,X7))),s(X4,X6))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5)))) )),
  inference(resolution,[],[f1123,f83])).
fof(f3076,plain,(
  ( ! [X335,X337,X339,X334,X336,X338] : (~p(s(bool,happ(s(fun(X334,bool),happ(s(fun(X334,fun(X334,bool)),X335),s(X334,X337))),s(X334,X336)))) | p(s(bool,X338)) | p(s(bool,X339)) | s(bool,X339) != s(bool,happ(s(fun(X334,bool),happ(s(fun(X334,fun(X334,bool)),X335),s(X334,X336))),s(X334,X337))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X334,fun(X334,bool)),X335)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f3074,plain,(
  ( ! [X325,X327,X323,X324,X326,X322] : (p(s(bool,happ(s(fun(X322,bool),happ(s(fun(X322,fun(X322,bool)),X323),s(X322,X325))),s(X322,X324)))) | p(s(bool,X326)) | ~p(s(bool,X327)) | s(bool,X327) != s(bool,happ(s(fun(X322,bool),happ(s(fun(X322,fun(X322,bool)),X323),s(X322,X324))),s(X322,X325))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X322,fun(X322,bool)),X323)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f3073,plain,(
  ( ! [X321,X316,X318,X320,X317,X319] : (s(bool,X320) != s(bool,happ(s(fun(X316,bool),happ(s(fun(X316,fun(X316,bool)),X317),s(X316,X319))),s(X316,X318))) | p(s(bool,X321)) | ~p(s(bool,X320)) | p(s(bool,happ(s(fun(X316,bool),happ(s(fun(X316,fun(X316,bool)),X317),s(X316,X318))),s(X316,X319)))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X316,fun(X316,bool)),X317)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f3072,plain,(
  ( ! [X313,X315,X310,X312,X314,X311] : (s(bool,X314) != s(bool,happ(s(fun(X310,bool),happ(s(fun(X310,fun(X310,bool)),X311),s(X310,X313))),s(X310,X312))) | p(s(bool,X315)) | ~p(s(bool,happ(s(fun(X310,bool),happ(s(fun(X310,fun(X310,bool)),X311),s(X310,X312))),s(X310,X313)))) | p(s(bool,X314)) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X310,fun(X310,bool)),X311)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f3062,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272] : (~p(s(bool,happ(s(fun(X267,bool),happ(s(fun(X267,fun(X267,bool)),X268),s(X267,X270))),s(X267,X269)))) | p(s(bool,X271)) | ~p(s(bool,X272)) | s(bool,X271) != s(bool,happ(s(fun(X267,bool),happ(s(fun(X267,fun(X267,bool)),X268),s(X267,X269))),s(X267,X270))) | s(bool,X272) != s(bool,X273) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X267,fun(X267,bool)),X268)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f3060,plain,(
  ( ! [X255,X253,X257,X259,X254,X256,X258] : (p(s(bool,happ(s(fun(X253,bool),happ(s(fun(X253,fun(X253,bool)),X254),s(X253,X256))),s(X253,X255)))) | ~p(s(bool,X257)) | ~p(s(bool,X258)) | s(bool,X258) != s(bool,happ(s(fun(X253,bool),happ(s(fun(X253,fun(X253,bool)),X254),s(X253,X255))),s(X253,X256))) | s(bool,X257) != s(bool,X259) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X253,fun(X253,bool)),X254)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f3059,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247] : (s(bool,X250) != s(bool,happ(s(fun(X246,bool),happ(s(fun(X246,fun(X246,bool)),X247),s(X246,X249))),s(X246,X248))) | p(s(bool,X250)) | ~p(s(bool,X251)) | ~p(s(bool,happ(s(fun(X246,bool),happ(s(fun(X246,fun(X246,bool)),X247),s(X246,X248))),s(X246,X249)))) | s(bool,X251) != s(bool,X252) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X246,fun(X246,bool)),X247)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f3058,plain,(
  ( ! [X243,X241,X239,X244,X242,X240,X245] : (s(bool,X243) != s(bool,happ(s(fun(X239,bool),happ(s(fun(X239,fun(X239,bool)),X240),s(X239,X242))),s(X239,X241))) | p(s(bool,happ(s(fun(X239,bool),happ(s(fun(X239,fun(X239,bool)),X240),s(X239,X241))),s(X239,X242)))) | ~p(s(bool,X244)) | ~p(s(bool,X243)) | s(bool,X244) != s(bool,X245) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X239,fun(X239,bool)),X240)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f3054,plain,(
  ( ! [X212,X216,X215,X213,X211,X217,X214] : (~p(s(bool,happ(s(fun(X211,bool),happ(s(fun(X211,fun(X211,bool)),X212),s(X211,X214))),s(X211,X213)))) | p(s(bool,X215)) | ~p(s(bool,X216)) | p(s(bool,X217)) | s(bool,X217) != s(bool,happ(s(fun(X211,bool),happ(s(fun(X211,fun(X211,bool)),X212),s(X211,X213))),s(X211,X214))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X211,fun(X211,bool)),X212)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f3052,plain,(
  ( ! [X198,X202,X200,X199,X197,X203,X201] : (p(s(bool,happ(s(fun(X197,bool),happ(s(fun(X197,fun(X197,bool)),X198),s(X197,X200))),s(X197,X199)))) | p(s(bool,X201)) | ~p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X202) != s(bool,happ(s(fun(X197,bool),happ(s(fun(X197,fun(X197,bool)),X198),s(X197,X199))),s(X197,X200))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X197,fun(X197,bool)),X198)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f3051,plain,(
  ( ! [X191,X196,X194,X192,X190,X195,X193] : (s(bool,X194) != s(bool,happ(s(fun(X190,bool),happ(s(fun(X190,fun(X190,bool)),X191),s(X190,X193))),s(X190,X192))) | p(s(bool,X195)) | ~p(s(bool,X194)) | ~p(s(bool,X196)) | p(s(bool,happ(s(fun(X190,bool),happ(s(fun(X190,fun(X190,bool)),X191),s(X190,X192))),s(X190,X193)))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X190,fun(X190,bool)),X191)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f3050,plain,(
  ( ! [X189,X187,X185,X188,X186,X184,X183] : (s(bool,X187) != s(bool,happ(s(fun(X183,bool),happ(s(fun(X183,fun(X183,bool)),X184),s(X183,X186))),s(X183,X185))) | p(s(bool,X188)) | ~p(s(bool,happ(s(fun(X183,bool),happ(s(fun(X183,fun(X183,bool)),X184),s(X183,X185))),s(X183,X186)))) | ~p(s(bool,X189)) | p(s(bool,X187)) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X183,fun(X183,bool)),X184)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f3049,plain,(
  ( ! [X177,X182,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X176,bool),happ(s(fun(X176,fun(X176,bool)),X177),s(X176,X179))),s(X176,X178)))) | p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X182) | s(bool,X181) != s(bool,happ(s(fun(X176,bool),happ(s(fun(X176,fun(X176,bool)),X177),s(X176,X178))),s(X176,X179))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X176,fun(X176,bool)),X177)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f3047,plain,(
  ( ! [X167,X165,X163,X166,X164,X162,X168] : (p(s(bool,happ(s(fun(X162,bool),happ(s(fun(X162,fun(X162,bool)),X163),s(X162,X165))),s(X162,X164)))) | p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,X168) | s(bool,X167) != s(bool,happ(s(fun(X162,bool),happ(s(fun(X162,fun(X162,bool)),X163),s(X162,X164))),s(X162,X165))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X162,fun(X162,bool)),X163)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f3044,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141] : (s(bool,X145) != s(bool,happ(s(fun(X141,bool),happ(s(fun(X141,fun(X141,bool)),X142),s(X141,X144))),s(X141,X143))) | p(s(bool,X146)) | ~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),happ(s(fun(X141,fun(X141,bool)),X142),s(X141,X143))),s(X141,X144)))) | s(bool,X146) != s(bool,X147) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X141,fun(X141,bool)),X142)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f3043,plain,(
  ( ! [X134,X140,X138,X136,X135,X139,X137] : (s(bool,X138) != s(bool,happ(s(fun(X134,bool),happ(s(fun(X134,fun(X134,bool)),X135),s(X134,X137))),s(X134,X136))) | p(s(bool,X139)) | ~p(s(bool,happ(s(fun(X134,bool),happ(s(fun(X134,fun(X134,bool)),X135),s(X134,X136))),s(X134,X137)))) | p(s(bool,X138)) | s(bool,X139) != s(bool,X140) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X134,fun(X134,bool)),X135)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f3042,plain,(
  ( ! [X127,X132,X130,X128,X133,X131,X129] : (~p(s(bool,happ(s(fun(X127,bool),happ(s(fun(X127,fun(X127,bool)),X128),s(X127,X130))),s(X127,X129)))) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X131) != s(bool,happ(s(fun(X127,bool),happ(s(fun(X127,fun(X127,bool)),X128),s(X127,X129))),s(X127,X130))) | s(bool,X132) != s(bool,X133) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X127,fun(X127,bool)),X128)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f3040,plain,(
  ( ! [X118,X116,X114,X119,X117,X115,X113] : (p(s(bool,happ(s(fun(X113,bool),happ(s(fun(X113,fun(X113,bool)),X114),s(X113,X116))),s(X113,X115)))) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X118) != s(bool,happ(s(fun(X113,bool),happ(s(fun(X113,fun(X113,bool)),X114),s(X113,X115))),s(X113,X116))) | s(bool,X117) != s(bool,X119) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X113,fun(X113,bool)),X114)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f3039,plain,(
  ( ! [X111,X109,X107,X112,X110,X108,X106] : (s(bool,X110) != s(bool,happ(s(fun(X106,bool),happ(s(fun(X106,fun(X106,bool)),X107),s(X106,X109))),s(X106,X108))) | p(s(bool,happ(s(fun(X106,bool),happ(s(fun(X106,fun(X106,bool)),X107),s(X106,X108))),s(X106,X109)))) | ~p(s(bool,X111)) | ~p(s(bool,X110)) | s(bool,X111) != s(bool,X112) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X106,fun(X106,bool)),X107)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f3038,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X104] : (s(bool,X103) != s(bool,happ(s(fun(X99,bool),happ(s(fun(X99,fun(X99,bool)),X100),s(X99,X102))),s(X99,X101))) | p(s(bool,X103)) | ~p(s(bool,X104)) | ~p(s(bool,happ(s(fun(X99,bool),happ(s(fun(X99,fun(X99,bool)),X100),s(X99,X101))),s(X99,X102)))) | s(bool,X104) != s(bool,X105) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X99,fun(X99,bool)),X100)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f3035,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (~p(s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X81))),s(X78,X80)))) | p(s(bool,X82)) | p(s(bool,X83)) | s(bool,X82) != s(bool,X84) | s(bool,X83) != s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X80))),s(X78,X81))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X78,fun(X78,bool)),X79)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f3033,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),happ(s(fun(X64,fun(X64,bool)),X65),s(X64,X67))),s(X64,X66)))) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,happ(s(fun(X64,bool),happ(s(fun(X64,fun(X64,bool)),X65),s(X64,X66))),s(X64,X67))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X64,fun(X64,bool)),X65)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f3030,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,X47) != s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X46))),s(X43,X45))) | p(s(bool,X48)) | ~p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X45))),s(X43,X46)))) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X43,fun(X43,bool)),X44)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f3029,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,X40) != s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X39))),s(X36,X38))) | p(s(bool,X41)) | ~p(s(bool,X40)) | p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X38))),s(X36,X39)))) | s(bool,X41) != s(bool,X42) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X36,fun(X36,bool)),X37)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f2971,plain,(
  ( ! [X335,X337,X339,X334,X336,X338] : (~p(s(bool,happ(s(fun(X334,bool),happ(s(fun(X334,fun(X334,bool)),X335),s(X334,X337))),s(X334,X336)))) | p(s(bool,X338)) | p(s(bool,X339)) | s(bool,X339) != s(bool,happ(s(fun(X334,bool),happ(s(fun(X334,fun(X334,bool)),X335),s(X334,X336))),s(X334,X337))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X334,fun(X334,bool)),X335)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f2969,plain,(
  ( ! [X325,X327,X323,X324,X326,X322] : (p(s(bool,happ(s(fun(X322,bool),happ(s(fun(X322,fun(X322,bool)),X323),s(X322,X325))),s(X322,X324)))) | p(s(bool,X326)) | ~p(s(bool,X327)) | s(bool,X327) != s(bool,happ(s(fun(X322,bool),happ(s(fun(X322,fun(X322,bool)),X323),s(X322,X324))),s(X322,X325))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X322,fun(X322,bool)),X323)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f2968,plain,(
  ( ! [X321,X316,X318,X320,X317,X319] : (s(bool,X320) != s(bool,happ(s(fun(X316,bool),happ(s(fun(X316,fun(X316,bool)),X317),s(X316,X319))),s(X316,X318))) | p(s(bool,X321)) | ~p(s(bool,X320)) | p(s(bool,happ(s(fun(X316,bool),happ(s(fun(X316,fun(X316,bool)),X317),s(X316,X318))),s(X316,X319)))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X316,fun(X316,bool)),X317)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f2967,plain,(
  ( ! [X313,X315,X310,X312,X314,X311] : (s(bool,X314) != s(bool,happ(s(fun(X310,bool),happ(s(fun(X310,fun(X310,bool)),X311),s(X310,X313))),s(X310,X312))) | p(s(bool,X315)) | ~p(s(bool,happ(s(fun(X310,bool),happ(s(fun(X310,fun(X310,bool)),X311),s(X310,X312))),s(X310,X313)))) | p(s(bool,X314)) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X310,fun(X310,bool)),X311)))) )),
  inference(superposition,[],[f125,f1288])).
fof(f2957,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272] : (~p(s(bool,happ(s(fun(X267,bool),happ(s(fun(X267,fun(X267,bool)),X268),s(X267,X270))),s(X267,X269)))) | p(s(bool,X271)) | ~p(s(bool,X272)) | s(bool,X271) != s(bool,happ(s(fun(X267,bool),happ(s(fun(X267,fun(X267,bool)),X268),s(X267,X269))),s(X267,X270))) | s(bool,X272) != s(bool,X273) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X267,fun(X267,bool)),X268)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f2955,plain,(
  ( ! [X255,X253,X257,X259,X254,X256,X258] : (p(s(bool,happ(s(fun(X253,bool),happ(s(fun(X253,fun(X253,bool)),X254),s(X253,X256))),s(X253,X255)))) | ~p(s(bool,X257)) | ~p(s(bool,X258)) | s(bool,X258) != s(bool,happ(s(fun(X253,bool),happ(s(fun(X253,fun(X253,bool)),X254),s(X253,X255))),s(X253,X256))) | s(bool,X257) != s(bool,X259) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X253,fun(X253,bool)),X254)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f2954,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247] : (s(bool,X250) != s(bool,happ(s(fun(X246,bool),happ(s(fun(X246,fun(X246,bool)),X247),s(X246,X249))),s(X246,X248))) | p(s(bool,X250)) | ~p(s(bool,X251)) | ~p(s(bool,happ(s(fun(X246,bool),happ(s(fun(X246,fun(X246,bool)),X247),s(X246,X248))),s(X246,X249)))) | s(bool,X251) != s(bool,X252) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X246,fun(X246,bool)),X247)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f2953,plain,(
  ( ! [X243,X241,X239,X244,X242,X240,X245] : (s(bool,X243) != s(bool,happ(s(fun(X239,bool),happ(s(fun(X239,fun(X239,bool)),X240),s(X239,X242))),s(X239,X241))) | p(s(bool,happ(s(fun(X239,bool),happ(s(fun(X239,fun(X239,bool)),X240),s(X239,X241))),s(X239,X242)))) | ~p(s(bool,X244)) | ~p(s(bool,X243)) | s(bool,X244) != s(bool,X245) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X239,fun(X239,bool)),X240)))) )),
  inference(superposition,[],[f89,f1288])).
fof(f2949,plain,(
  ( ! [X212,X216,X215,X213,X211,X217,X214] : (~p(s(bool,happ(s(fun(X211,bool),happ(s(fun(X211,fun(X211,bool)),X212),s(X211,X214))),s(X211,X213)))) | p(s(bool,X215)) | ~p(s(bool,X216)) | p(s(bool,X217)) | s(bool,X217) != s(bool,happ(s(fun(X211,bool),happ(s(fun(X211,fun(X211,bool)),X212),s(X211,X213))),s(X211,X214))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X211,fun(X211,bool)),X212)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f2947,plain,(
  ( ! [X198,X202,X200,X199,X197,X203,X201] : (p(s(bool,happ(s(fun(X197,bool),happ(s(fun(X197,fun(X197,bool)),X198),s(X197,X200))),s(X197,X199)))) | p(s(bool,X201)) | ~p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X202) != s(bool,happ(s(fun(X197,bool),happ(s(fun(X197,fun(X197,bool)),X198),s(X197,X199))),s(X197,X200))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X197,fun(X197,bool)),X198)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f2946,plain,(
  ( ! [X191,X196,X194,X192,X190,X195,X193] : (s(bool,X194) != s(bool,happ(s(fun(X190,bool),happ(s(fun(X190,fun(X190,bool)),X191),s(X190,X193))),s(X190,X192))) | p(s(bool,X195)) | ~p(s(bool,X194)) | ~p(s(bool,X196)) | p(s(bool,happ(s(fun(X190,bool),happ(s(fun(X190,fun(X190,bool)),X191),s(X190,X192))),s(X190,X193)))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X190,fun(X190,bool)),X191)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f2945,plain,(
  ( ! [X189,X187,X185,X188,X186,X184,X183] : (s(bool,X187) != s(bool,happ(s(fun(X183,bool),happ(s(fun(X183,fun(X183,bool)),X184),s(X183,X186))),s(X183,X185))) | p(s(bool,X188)) | ~p(s(bool,happ(s(fun(X183,bool),happ(s(fun(X183,fun(X183,bool)),X184),s(X183,X185))),s(X183,X186)))) | ~p(s(bool,X189)) | p(s(bool,X187)) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X183,fun(X183,bool)),X184)))) )),
  inference(superposition,[],[f88,f1288])).
fof(f2944,plain,(
  ( ! [X177,X182,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X176,bool),happ(s(fun(X176,fun(X176,bool)),X177),s(X176,X179))),s(X176,X178)))) | p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X180) != s(bool,X182) | s(bool,X181) != s(bool,happ(s(fun(X176,bool),happ(s(fun(X176,fun(X176,bool)),X177),s(X176,X178))),s(X176,X179))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X176,fun(X176,bool)),X177)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f2942,plain,(
  ( ! [X167,X165,X163,X166,X164,X162,X168] : (p(s(bool,happ(s(fun(X162,bool),happ(s(fun(X162,fun(X162,bool)),X163),s(X162,X165))),s(X162,X164)))) | p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,X168) | s(bool,X167) != s(bool,happ(s(fun(X162,bool),happ(s(fun(X162,fun(X162,bool)),X163),s(X162,X164))),s(X162,X165))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X162,fun(X162,bool)),X163)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f2939,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141] : (s(bool,X145) != s(bool,happ(s(fun(X141,bool),happ(s(fun(X141,fun(X141,bool)),X142),s(X141,X144))),s(X141,X143))) | p(s(bool,X146)) | ~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),happ(s(fun(X141,fun(X141,bool)),X142),s(X141,X143))),s(X141,X144)))) | s(bool,X146) != s(bool,X147) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X141,fun(X141,bool)),X142)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f2938,plain,(
  ( ! [X134,X140,X138,X136,X135,X139,X137] : (s(bool,X138) != s(bool,happ(s(fun(X134,bool),happ(s(fun(X134,fun(X134,bool)),X135),s(X134,X137))),s(X134,X136))) | p(s(bool,X139)) | ~p(s(bool,happ(s(fun(X134,bool),happ(s(fun(X134,fun(X134,bool)),X135),s(X134,X136))),s(X134,X137)))) | p(s(bool,X138)) | s(bool,X139) != s(bool,X140) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X134,fun(X134,bool)),X135)))) )),
  inference(superposition,[],[f87,f1288])).
fof(f2937,plain,(
  ( ! [X127,X132,X130,X128,X133,X131,X129] : (~p(s(bool,happ(s(fun(X127,bool),happ(s(fun(X127,fun(X127,bool)),X128),s(X127,X130))),s(X127,X129)))) | p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X131) != s(bool,happ(s(fun(X127,bool),happ(s(fun(X127,fun(X127,bool)),X128),s(X127,X129))),s(X127,X130))) | s(bool,X132) != s(bool,X133) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X127,fun(X127,bool)),X128)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f2935,plain,(
  ( ! [X118,X116,X114,X119,X117,X115,X113] : (p(s(bool,happ(s(fun(X113,bool),happ(s(fun(X113,fun(X113,bool)),X114),s(X113,X116))),s(X113,X115)))) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X118) != s(bool,happ(s(fun(X113,bool),happ(s(fun(X113,fun(X113,bool)),X114),s(X113,X115))),s(X113,X116))) | s(bool,X117) != s(bool,X119) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X113,fun(X113,bool)),X114)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f2934,plain,(
  ( ! [X111,X109,X107,X112,X110,X108,X106] : (s(bool,X110) != s(bool,happ(s(fun(X106,bool),happ(s(fun(X106,fun(X106,bool)),X107),s(X106,X109))),s(X106,X108))) | p(s(bool,happ(s(fun(X106,bool),happ(s(fun(X106,fun(X106,bool)),X107),s(X106,X108))),s(X106,X109)))) | ~p(s(bool,X111)) | ~p(s(bool,X110)) | s(bool,X111) != s(bool,X112) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X106,fun(X106,bool)),X107)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f2933,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X104] : (s(bool,X103) != s(bool,happ(s(fun(X99,bool),happ(s(fun(X99,fun(X99,bool)),X100),s(X99,X102))),s(X99,X101))) | p(s(bool,X103)) | ~p(s(bool,X104)) | ~p(s(bool,happ(s(fun(X99,bool),happ(s(fun(X99,fun(X99,bool)),X100),s(X99,X101))),s(X99,X102)))) | s(bool,X104) != s(bool,X105) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X99,fun(X99,bool)),X100)))) )),
  inference(superposition,[],[f85,f1288])).
fof(f2930,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (~p(s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X81))),s(X78,X80)))) | p(s(bool,X82)) | p(s(bool,X83)) | s(bool,X82) != s(bool,X84) | s(bool,X83) != s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X80))),s(X78,X81))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X78,fun(X78,bool)),X79)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f2928,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),happ(s(fun(X64,fun(X64,bool)),X65),s(X64,X67))),s(X64,X66)))) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,happ(s(fun(X64,bool),happ(s(fun(X64,fun(X64,bool)),X65),s(X64,X66))),s(X64,X67))) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X64,fun(X64,bool)),X65)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f2925,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,X47) != s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X46))),s(X43,X45))) | p(s(bool,X48)) | ~p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),X44),s(X43,X45))),s(X43,X46)))) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X43,fun(X43,bool)),X44)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f2924,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,X40) != s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X39))),s(X36,X38))) | p(s(bool,X41)) | ~p(s(bool,X40)) | p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),X37),s(X36,X38))),s(X36,X39)))) | s(bool,X41) != s(bool,X42) | s(bool,'F_0') = s(bool,'const.relation.symmetric_1'(s(fun(X36,fun(X36,bool)),X37)))) )),
  inference(superposition,[],[f83,f1288])).
fof(f2784,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X72,sK1),X73),s(X72,X74),s(X72,X75))) | p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X72,sK1),X73),s(X72,X75),s(X72,X74)))) ) | $spl24),
  inference(resolution,[],[f2434,f125])).
fof(f2780,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X53,sK1),X54),s(X53,X55),s(X53,X56))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X53,sK1),X54),s(X53,X56),s(X53,X55))) | s(bool,X58) != s(bool,X59)) ) | $spl24),
  inference(resolution,[],[f2434,f89])).
fof(f2777,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X32,sK1),X33),s(X32,X34),s(X32,X35))) | p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X32,sK1),X33),s(X32,X35),s(X32,X34)))) ) | $spl24),
  inference(resolution,[],[f2434,f88])).
fof(f2776,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X25,sK1),X26),s(X25,X27),s(X25,X28))) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X25,sK1),X26),s(X25,X28),s(X25,X27)))) ) | $spl24),
  inference(resolution,[],[f2434,f87])).
fof(f2775,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X18,sK1),X19),s(X18,X20),s(X18,X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X18,sK1),X19),s(X18,X21),s(X18,X20))) | s(bool,X23) != s(bool,X24)) ) | $spl24),
  inference(resolution,[],[f2434,f85])).
fof(f2773,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s(bool,'F_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X4,sK1),X5),s(X4,X6),s(X4,X7))) | p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X4,sK1),X5),s(X4,X7),s(X4,X6)))) ) | $spl24),
  inference(resolution,[],[f2434,f83])).
fof(f2756,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X48))),s(sK1,X49))) | p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X49))),s(sK1,X48)))) ) | $spl24),
  inference(resolution,[],[f2013,f125])).
fof(f2752,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X37))),s(sK1,X38))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X38))),s(sK1,X37))) | s(bool,X40) != s(bool,X41)) ) | $spl24),
  inference(resolution,[],[f2013,f89])).
fof(f2749,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,X23))) | p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X23))),s(sK1,X22)))) ) | $spl24),
  inference(resolution,[],[f2013,f88])).
fof(f2748,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X17))),s(sK1,X18))) | p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X18))),s(sK1,X17)))) ) | $spl24),
  inference(resolution,[],[f2013,f87])).
fof(f2747,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X12))),s(sK1,X13))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X13))),s(sK1,X12))) | s(bool,X15) != s(bool,X16)) ) | $spl24),
  inference(resolution,[],[f2013,f85])).
fof(f2745,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X3))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X3))),s(sK1,X2)))) ) | $spl24),
  inference(resolution,[],[f2013,f83])).
fof(f2685,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X72,sK1),X73),s(X72,X74),s(X72,X75))) | p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X77) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X72,sK1),X73),s(X72,X75),s(X72,X74)))) ) | $spl34),
  inference(resolution,[],[f2438,f125])).
fof(f2680,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X49,sK1),X50),s(X49,X51),s(X49,X52))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X49,sK1),X50),s(X49,X52),s(X49,X51))) | s(bool,X53) != s(bool,X55)) ) | $spl34),
  inference(resolution,[],[f2438,f89])).
fof(f2678,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X35,sK1),X36),s(X35,X37),s(X35,X38))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X35,sK1),X36),s(X35,X38),s(X35,X37)))) ) | $spl34),
  inference(resolution,[],[f2438,f88])).
fof(f2676,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X21,sK1),X22),s(X21,X23),s(X21,X24))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X21,sK1),X22),s(X21,X24),s(X21,X23)))) ) | $spl34),
  inference(resolution,[],[f2438,f87])).
fof(f2675,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X14,sK1),X15),s(X14,X16),s(X14,X17))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X14,sK1),X15),s(X14,X17),s(X14,X16))) | s(bool,X18) != s(bool,X20)) ) | $spl34),
  inference(resolution,[],[f2438,f85])).
fof(f2673,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2)))) ) | $spl34),
  inference(resolution,[],[f2438,f83])).
fof(f2515,plain,(
  ( ! [X300,X302,X298,X301,X303,X299] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X298,sK1),X299),s(X298,X301),s(X298,X300)))) | p(s(bool,X302)) | p(s(bool,X303)) | s(bool,X303) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X298,sK1),X299),s(X298,X300),s(X298,X301)))) )),
  inference(superposition,[],[f125,f2341])).
fof(f2513,plain,(
  ( ! [X288,X290,X287,X289,X291,X286] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X286,sK1),X287),s(X286,X289),s(X286,X288)))) | p(s(bool,X290)) | ~p(s(bool,X291)) | s(bool,X291) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X286,sK1),X287),s(X286,X288),s(X286,X289)))) )),
  inference(superposition,[],[f125,f2341])).
fof(f2512,plain,(
  ( ! [X280,X282,X285,X281,X283,X284] : (s(bool,X284) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X280,sK1),X281),s(X280,X283),s(X280,X282))) | p(s(bool,X285)) | ~p(s(bool,X284)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X280,sK1),X281),s(X280,X282),s(X280,X283))))) )),
  inference(superposition,[],[f125,f2341])).
fof(f2511,plain,(
  ( ! [X277,X279,X275,X276,X278,X274] : (s(bool,X278) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X274,sK1),X275),s(X274,X277),s(X274,X276))) | p(s(bool,X279)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X274,sK1),X275),s(X274,X276),s(X274,X277)))) | p(s(bool,X278))) )),
  inference(superposition,[],[f125,f2341])).
fof(f2501,plain,(
  ( ! [X231,X237,X235,X233,X236,X234,X232] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X231,sK1),X232),s(X231,X234),s(X231,X233)))) | p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X231,sK1),X232),s(X231,X233),s(X231,X234))) | s(bool,X236) != s(bool,X237)) )),
  inference(superposition,[],[f89,f2341])).
fof(f2499,plain,(
  ( ! [X218,X222,X220,X219,X217,X223,X221] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X217,sK1),X218),s(X217,X220),s(X217,X219)))) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X217,sK1),X218),s(X217,X219),s(X217,X220))) | s(bool,X221) != s(bool,X223)) )),
  inference(superposition,[],[f89,f2341])).
fof(f2498,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X214] : (s(bool,X214) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X210,sK1),X211),s(X210,X213),s(X210,X212))) | p(s(bool,X214)) | ~p(s(bool,X215)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X210,sK1),X211),s(X210,X212),s(X210,X213)))) | s(bool,X215) != s(bool,X216)) )),
  inference(superposition,[],[f89,f2341])).
fof(f2497,plain,(
  ( ! [X208,X206,X204,X209,X207,X205,X203] : (s(bool,X207) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X203,sK1),X204),s(X203,X206),s(X203,X205))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X203,sK1),X204),s(X203,X205),s(X203,X206)))) | ~p(s(bool,X208)) | ~p(s(bool,X207)) | s(bool,X208) != s(bool,X209)) )),
  inference(superposition,[],[f89,f2341])).
fof(f2493,plain,(
  ( ! [X177,X175,X180,X178,X176,X181,X179] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X175,sK1),X176),s(X175,X178),s(X175,X177)))) | p(s(bool,X179)) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X175,sK1),X176),s(X175,X177),s(X175,X178)))) )),
  inference(superposition,[],[f88,f2341])).
fof(f2491,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X161,sK1),X162),s(X161,X164),s(X161,X163)))) | p(s(bool,X165)) | ~p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X161,sK1),X162),s(X161,X163),s(X161,X164)))) )),
  inference(superposition,[],[f88,f2341])).
fof(f2490,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X160] : (s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X154,sK1),X155),s(X154,X157),s(X154,X156))) | p(s(bool,X159)) | ~p(s(bool,X158)) | ~p(s(bool,X160)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X154,sK1),X155),s(X154,X156),s(X154,X157))))) )),
  inference(superposition,[],[f88,f2341])).
fof(f2489,plain,(
  ( ! [X152,X151,X149,X147,X153,X150,X148] : (s(bool,X151) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X147,sK1),X148),s(X147,X150),s(X147,X149))) | p(s(bool,X152)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X147,sK1),X148),s(X147,X149),s(X147,X150)))) | ~p(s(bool,X153)) | p(s(bool,X151))) )),
  inference(superposition,[],[f88,f2341])).
fof(f2488,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X140,sK1),X141),s(X140,X143),s(X140,X142)))) | p(s(bool,X144)) | p(s(bool,X145)) | s(bool,X144) != s(bool,X146) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X140,sK1),X141),s(X140,X142),s(X140,X143)))) )),
  inference(superposition,[],[f87,f2341])).
fof(f2486,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X126,sK1),X127),s(X126,X129),s(X126,X128)))) | p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,X132) | s(bool,X131) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X126,sK1),X127),s(X126,X128),s(X126,X129)))) )),
  inference(superposition,[],[f87,f2341])).
fof(f2483,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(bool,X109) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X105,sK1),X106),s(X105,X108),s(X105,X107))) | p(s(bool,X110)) | ~p(s(bool,X109)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X105,sK1),X106),s(X105,X107),s(X105,X108)))) | s(bool,X110) != s(bool,X111)) )),
  inference(superposition,[],[f87,f2341])).
fof(f2482,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (s(bool,X102) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X98,sK1),X99),s(X98,X101),s(X98,X100))) | p(s(bool,X103)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X98,sK1),X99),s(X98,X100),s(X98,X101)))) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104)) )),
  inference(superposition,[],[f87,f2341])).
fof(f2481,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X91,sK1),X92),s(X91,X94),s(X91,X93)))) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X91,sK1),X92),s(X91,X93),s(X91,X94))) | s(bool,X96) != s(bool,X97)) )),
  inference(superposition,[],[f85,f2341])).
fof(f2479,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X77,sK1),X78),s(X77,X80),s(X77,X79)))) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X77,sK1),X78),s(X77,X79),s(X77,X80))) | s(bool,X81) != s(bool,X83)) )),
  inference(superposition,[],[f85,f2341])).
fof(f2478,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,X74) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X70,sK1),X71),s(X70,X73),s(X70,X72))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X70,sK1),X71),s(X70,X72),s(X70,X73)))) | ~p(s(bool,X75)) | ~p(s(bool,X74)) | s(bool,X75) != s(bool,X76)) )),
  inference(superposition,[],[f85,f2341])).
fof(f2477,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,X67) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X63,sK1),X64),s(X63,X66),s(X63,X65))) | p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X63,sK1),X64),s(X63,X65),s(X63,X66)))) | s(bool,X68) != s(bool,X69)) )),
  inference(superposition,[],[f85,f2341])).
fof(f2474,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X42,sK1),X43),s(X42,X45),s(X42,X44)))) | p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X42,sK1),X43),s(X42,X44),s(X42,X45)))) )),
  inference(superposition,[],[f83,f2341])).
fof(f2472,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X28,sK1),X29),s(X28,X31),s(X28,X30)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X28,sK1),X29),s(X28,X30),s(X28,X31)))) )),
  inference(superposition,[],[f83,f2341])).
fof(f2469,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X10),s(X7,X9))) | p(s(bool,X12)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X9),s(X7,X10)))) | p(s(bool,X11)) | s(bool,X12) != s(bool,X13)) )),
  inference(superposition,[],[f83,f2341])).
fof(f2468,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X4) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2))) | p(s(bool,X5)) | ~p(s(bool,X4)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3)))) | s(bool,X5) != s(bool,X6)) )),
  inference(superposition,[],[f83,f2341])).
fof(f2430,plain,(
  ( ! [X300,X302,X298,X301,X303,X299] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X298,sK1),X299),s(X298,X301),s(X298,X300)))) | p(s(bool,X302)) | p(s(bool,X303)) | s(bool,X303) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X298,sK1),X299),s(X298,X300),s(X298,X301)))) )),
  inference(superposition,[],[f125,f2341])).
fof(f2428,plain,(
  ( ! [X288,X290,X287,X289,X291,X286] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X286,sK1),X287),s(X286,X289),s(X286,X288)))) | p(s(bool,X290)) | ~p(s(bool,X291)) | s(bool,X291) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X286,sK1),X287),s(X286,X288),s(X286,X289)))) )),
  inference(superposition,[],[f125,f2341])).
fof(f2427,plain,(
  ( ! [X280,X282,X285,X281,X283,X284] : (s(bool,X284) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X280,sK1),X281),s(X280,X283),s(X280,X282))) | p(s(bool,X285)) | ~p(s(bool,X284)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X280,sK1),X281),s(X280,X282),s(X280,X283))))) )),
  inference(superposition,[],[f125,f2341])).
fof(f2426,plain,(
  ( ! [X277,X279,X275,X276,X278,X274] : (s(bool,X278) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X274,sK1),X275),s(X274,X277),s(X274,X276))) | p(s(bool,X279)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X274,sK1),X275),s(X274,X276),s(X274,X277)))) | p(s(bool,X278))) )),
  inference(superposition,[],[f125,f2341])).
fof(f2416,plain,(
  ( ! [X231,X237,X235,X233,X236,X234,X232] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X231,sK1),X232),s(X231,X234),s(X231,X233)))) | p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X231,sK1),X232),s(X231,X233),s(X231,X234))) | s(bool,X236) != s(bool,X237)) )),
  inference(superposition,[],[f89,f2341])).
fof(f2414,plain,(
  ( ! [X218,X222,X220,X219,X217,X223,X221] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X217,sK1),X218),s(X217,X220),s(X217,X219)))) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X217,sK1),X218),s(X217,X219),s(X217,X220))) | s(bool,X221) != s(bool,X223)) )),
  inference(superposition,[],[f89,f2341])).
fof(f2413,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X214] : (s(bool,X214) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X210,sK1),X211),s(X210,X213),s(X210,X212))) | p(s(bool,X214)) | ~p(s(bool,X215)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X210,sK1),X211),s(X210,X212),s(X210,X213)))) | s(bool,X215) != s(bool,X216)) )),
  inference(superposition,[],[f89,f2341])).
fof(f2412,plain,(
  ( ! [X208,X206,X204,X209,X207,X205,X203] : (s(bool,X207) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X203,sK1),X204),s(X203,X206),s(X203,X205))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X203,sK1),X204),s(X203,X205),s(X203,X206)))) | ~p(s(bool,X208)) | ~p(s(bool,X207)) | s(bool,X208) != s(bool,X209)) )),
  inference(superposition,[],[f89,f2341])).
fof(f2408,plain,(
  ( ! [X177,X175,X180,X178,X176,X181,X179] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X175,sK1),X176),s(X175,X178),s(X175,X177)))) | p(s(bool,X179)) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X175,sK1),X176),s(X175,X177),s(X175,X178)))) )),
  inference(superposition,[],[f88,f2341])).
fof(f2406,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X161,sK1),X162),s(X161,X164),s(X161,X163)))) | p(s(bool,X165)) | ~p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X161,sK1),X162),s(X161,X163),s(X161,X164)))) )),
  inference(superposition,[],[f88,f2341])).
fof(f2405,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X160] : (s(bool,X158) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X154,sK1),X155),s(X154,X157),s(X154,X156))) | p(s(bool,X159)) | ~p(s(bool,X158)) | ~p(s(bool,X160)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X154,sK1),X155),s(X154,X156),s(X154,X157))))) )),
  inference(superposition,[],[f88,f2341])).
fof(f2404,plain,(
  ( ! [X152,X151,X149,X147,X153,X150,X148] : (s(bool,X151) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X147,sK1),X148),s(X147,X150),s(X147,X149))) | p(s(bool,X152)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X147,sK1),X148),s(X147,X149),s(X147,X150)))) | ~p(s(bool,X153)) | p(s(bool,X151))) )),
  inference(superposition,[],[f88,f2341])).
fof(f2403,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X140,sK1),X141),s(X140,X143),s(X140,X142)))) | p(s(bool,X144)) | p(s(bool,X145)) | s(bool,X144) != s(bool,X146) | s(bool,X145) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X140,sK1),X141),s(X140,X142),s(X140,X143)))) )),
  inference(superposition,[],[f87,f2341])).
fof(f2401,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X126,sK1),X127),s(X126,X129),s(X126,X128)))) | p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,X132) | s(bool,X131) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X126,sK1),X127),s(X126,X128),s(X126,X129)))) )),
  inference(superposition,[],[f87,f2341])).
fof(f2398,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(bool,X109) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X105,sK1),X106),s(X105,X108),s(X105,X107))) | p(s(bool,X110)) | ~p(s(bool,X109)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X105,sK1),X106),s(X105,X107),s(X105,X108)))) | s(bool,X110) != s(bool,X111)) )),
  inference(superposition,[],[f87,f2341])).
fof(f2397,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (s(bool,X102) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X98,sK1),X99),s(X98,X101),s(X98,X100))) | p(s(bool,X103)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X98,sK1),X99),s(X98,X100),s(X98,X101)))) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104)) )),
  inference(superposition,[],[f87,f2341])).
fof(f2396,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X91,sK1),X92),s(X91,X94),s(X91,X93)))) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X91,sK1),X92),s(X91,X93),s(X91,X94))) | s(bool,X96) != s(bool,X97)) )),
  inference(superposition,[],[f85,f2341])).
fof(f2394,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X77,sK1),X78),s(X77,X80),s(X77,X79)))) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X77,sK1),X78),s(X77,X79),s(X77,X80))) | s(bool,X81) != s(bool,X83)) )),
  inference(superposition,[],[f85,f2341])).
fof(f2393,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,X74) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X70,sK1),X71),s(X70,X73),s(X70,X72))) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X70,sK1),X71),s(X70,X72),s(X70,X73)))) | ~p(s(bool,X75)) | ~p(s(bool,X74)) | s(bool,X75) != s(bool,X76)) )),
  inference(superposition,[],[f85,f2341])).
fof(f2392,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,X67) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X63,sK1),X64),s(X63,X66),s(X63,X65))) | p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X63,sK1),X64),s(X63,X65),s(X63,X66)))) | s(bool,X68) != s(bool,X69)) )),
  inference(superposition,[],[f85,f2341])).
fof(f2389,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X42,sK1),X43),s(X42,X45),s(X42,X44)))) | p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X42,sK1),X43),s(X42,X44),s(X42,X45)))) )),
  inference(superposition,[],[f83,f2341])).
fof(f2387,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X28,sK1),X29),s(X28,X31),s(X28,X30)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X28,sK1),X29),s(X28,X30),s(X28,X31)))) )),
  inference(superposition,[],[f83,f2341])).
fof(f2384,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,X11) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X10),s(X7,X9))) | p(s(bool,X12)) | ~p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X7,sK1),X8),s(X7,X9),s(X7,X10)))) | p(s(bool,X11)) | s(bool,X12) != s(bool,X13)) )),
  inference(superposition,[],[f83,f2341])).
fof(f2383,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X4) != s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X3),s(X0,X2))) | p(s(bool,X5)) | ~p(s(bool,X4)) | p(s(bool,'const.relation.inv_image_4'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(X0,sK1),X1),s(X0,X2),s(X0,X3)))) | s(bool,X5) != s(bool,X6)) )),
  inference(superposition,[],[f83,f2341])).
fof(f2325,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (s(bool,'const.relation.inv_image_4'(s(fun(X102,fun(X102,bool)),X103),s(fun(X104,X102),X105),s(X104,X106),s(X104,X107))) = s(bool,'const.relation.inv_image_4'(s(fun(X102,fun(X102,bool)),X103),s(fun(X104,X102),X105),s(X104,X107),s(X104,X106))) | p(s(bool,X108)) | ~p(s(bool,X109)) | s(bool,X109) != s(bool,'const.relation.symmetric_1'(s(fun(X102,fun(X102,bool)),X103)))) )),
  inference(resolution,[],[f1728,f125])).
fof(f2320,plain,(
  ( ! [X70,X76,X74,X72,X71,X69,X77,X75,X73] : (s(bool,'const.relation.inv_image_4'(s(fun(X69,fun(X69,bool)),X70),s(fun(X71,X69),X72),s(X71,X73),s(X71,X74))) = s(bool,'const.relation.inv_image_4'(s(fun(X69,fun(X69,bool)),X70),s(fun(X71,X69),X72),s(X71,X74),s(X71,X73))) | ~p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X76) != s(bool,'const.relation.symmetric_1'(s(fun(X69,fun(X69,bool)),X70))) | s(bool,X75) != s(bool,X77)) )),
  inference(resolution,[],[f1728,f89])).
fof(f2318,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53,X51] : (s(bool,'const.relation.inv_image_4'(s(fun(X51,fun(X51,bool)),X52),s(fun(X53,X51),X54),s(X53,X55),s(X53,X56))) = s(bool,'const.relation.inv_image_4'(s(fun(X51,fun(X51,bool)),X52),s(fun(X53,X51),X54),s(X53,X56),s(X53,X55))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,'const.relation.symmetric_1'(s(fun(X51,fun(X51,bool)),X52)))) )),
  inference(resolution,[],[f1728,f88])).
fof(f2316,plain,(
  ( ! [X39,X37,X35,X33,X41,X38,X36,X34,X40] : (s(bool,'const.relation.inv_image_4'(s(fun(X33,fun(X33,bool)),X34),s(fun(X35,X33),X36),s(X35,X37),s(X35,X38))) = s(bool,'const.relation.inv_image_4'(s(fun(X33,fun(X33,bool)),X34),s(fun(X35,X33),X36),s(X35,X38),s(X35,X37))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.relation.symmetric_1'(s(fun(X33,fun(X33,bool)),X34)))) )),
  inference(resolution,[],[f1728,f87])).
fof(f2315,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25,X32] : (s(bool,'const.relation.inv_image_4'(s(fun(X24,fun(X24,bool)),X25),s(fun(X26,X24),X27),s(X26,X28),s(X26,X29))) = s(bool,'const.relation.inv_image_4'(s(fun(X24,fun(X24,bool)),X25),s(fun(X26,X24),X27),s(X26,X29),s(X26,X28))) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.relation.symmetric_1'(s(fun(X24,fun(X24,bool)),X25))) | s(bool,X30) != s(bool,X32)) )),
  inference(resolution,[],[f1728,f85])).
fof(f2313,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X13,X11,X9] : (s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X10),s(X8,X11))) = s(bool,'const.relation.inv_image_4'(s(fun(X6,fun(X6,bool)),X7),s(fun(X8,X6),X9),s(X8,X11),s(X8,X10))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) )),
  inference(resolution,[],[f1728,f83])).
fof(f2289,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X48))),s(sK1,X49))) | p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X51) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X49))),s(sK1,X48)))) ) | $spl34),
  inference(resolution,[],[f2017,f125])).
fof(f2284,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X35))),s(sK1,X36))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X36))),s(sK1,X35))) | s(bool,X37) != s(bool,X39)) ) | $spl34),
  inference(resolution,[],[f2017,f89])).
fof(f2282,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X25))),s(sK1,X26))) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X26))),s(sK1,X25)))) ) | $spl34),
  inference(resolution,[],[f2017,f88])).
fof(f2280,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X15))),s(sK1,X16))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X16))),s(sK1,X15)))) ) | $spl34),
  inference(resolution,[],[f2017,f87])).
fof(f2279,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X10))),s(sK1,X11))) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X11))),s(sK1,X10))) | s(bool,X12) != s(bool,X14)) ) | $spl34),
  inference(resolution,[],[f2017,f85])).
fof(f2277,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1))) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0)))) ) | $spl34),
  inference(resolution,[],[f2017,f83])).
fof(f2222,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'F_0') != s(bool,X115) | s(bool,X116) != s(bool,X117) | s(bool,X117) != s(bool,X118) | s(bool,'F_0') = s(bool,X119) | s(bool,X118) != s(bool,X119) | s(bool,X115) != s(bool,X116)) ) | $spl78),
  inference(cnf_transformation,[],[f2222_D])).
fof(f2222_D,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'F_0') != s(bool,X115) | s(bool,X116) != s(bool,X117) | s(bool,X117) != s(bool,X118) | s(bool,'F_0') = s(bool,X119) | s(bool,X118) != s(bool,X119) | s(bool,X115) != s(bool,X116)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2224,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21)) ) | ($spl24 | $spl66)),
  inference(subsumption_resolution,[],[f2210,f385])).
fof(f2210,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X18)) ) | $spl66),
  inference(superposition,[],[f1852,f122])).
fof(f2205,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (s(bool,'F_0') != s(bool,X121) | s(bool,X121) != s(bool,X122) | s(bool,X123) != s(bool,X124) | s(bool,X124) != s(bool,X125) | s(bool,'F_0') = s(bool,X125) | s(bool,X122) != s(bool,X123) | ~p(s(bool,X126))) ) | $spl66),
  inference(resolution,[],[f1852,f101])).
fof(f2221,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(bool,'F_0') != s(bool,X115) | s(bool,X115) != s(bool,X116) | s(bool,X117) != s(bool,X118) | s(bool,X118) != s(bool,X119) | s(bool,'F_0') = s(bool,X119) | s(bool,X116) != s(bool,X117) | p(s(bool,X120))) ) | $spl66),
  inference(subsumption_resolution,[],[f2204,f126])).
fof(f2204,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(bool,'F_0') != s(bool,X115) | s(bool,X115) != s(bool,X116) | s(bool,X117) != s(bool,X118) | s(bool,X118) != s(bool,X119) | s(bool,'F_0') = s(bool,X119) | s(bool,X116) != s(bool,X117) | p(s(bool,X120)) | ~p(s(bool,'T_0'))) ) | $spl66),
  inference(resolution,[],[f1852,f105])).
fof(f2200,plain,(
  ( ! [X94,X92,X90,X97,X95,X93,X91,X96] : (s(bool,X90) != s(bool,X91) | s(bool,X90) != s(bool,X92) | s(bool,X93) != s(bool,X94) | s(bool,X94) != s(bool,X95) | s(bool,'F_0') = s(bool,X95) | s(bool,X92) != s(bool,X93) | p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X91) != s(bool,X97)) ) | $spl66),
  inference(resolution,[],[f1852,f125])).
fof(f2196,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (s(bool,X63) != s(bool,X64) | s(bool,X63) != s(bool,X65) | s(bool,X66) != s(bool,X67) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X65) != s(bool,X66) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X64) != s(bool,X69) | s(bool,X70) != s(bool,X71)) ) | $spl66),
  inference(resolution,[],[f1852,f89])).
fof(f2193,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X44,X42,X40] : (s(bool,X36) != s(bool,X37) | s(bool,X36) != s(bool,X38) | s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,X38) != s(bool,X39) | p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X37) != s(bool,X44)) ) | $spl66),
  inference(resolution,[],[f1852,f88])).
fof(f2192,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (s(bool,X27) != s(bool,X28) | s(bool,X27) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X29) != s(bool,X30) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X28) != s(bool,X34)) ) | $spl66),
  inference(resolution,[],[f1852,f87])).
fof(f2191,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X19) != s(bool,X24) | s(bool,X25) != s(bool,X26)) ) | $spl66),
  inference(resolution,[],[f1852,f85])).
fof(f2189,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X2) != s(bool,X3) | p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X1) != s(bool,X7)) ) | $spl66),
  inference(resolution,[],[f1852,f83])).
fof(f2094,plain,(
  ( ! [X212,X215,X213,X214] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X213))),s(sK1,X212)))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X215) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X212))),s(sK1,X213)))) )),
  inference(superposition,[],[f125,f127])).
fof(f2092,plain,(
  ( ! [X206,X204,X207,X205] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X205))),s(sK1,X204)))) | p(s(bool,X206)) | ~p(s(bool,X207)) | s(bool,X207) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X204))),s(sK1,X205)))) )),
  inference(superposition,[],[f125,f127])).
fof(f2091,plain,(
  ( ! [X202,X200,X203,X201] : (s(bool,X202) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X201))),s(sK1,X200))) | p(s(bool,X203)) | ~p(s(bool,X202)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X200))),s(sK1,X201))))) )),
  inference(superposition,[],[f125,f127])).
fof(f2090,plain,(
  ( ! [X198,X196,X199,X197] : (s(bool,X198) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X197))),s(sK1,X196))) | p(s(bool,X199)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X196))),s(sK1,X197)))) | p(s(bool,X198))) )),
  inference(superposition,[],[f125,f127])).
fof(f2080,plain,(
  ( ! [X177,X175,X173,X176,X174] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X174))),s(sK1,X173)))) | p(s(bool,X175)) | ~p(s(bool,X176)) | s(bool,X175) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X173))),s(sK1,X174))) | s(bool,X176) != s(bool,X177)) )),
  inference(superposition,[],[f89,f127])).
fof(f2078,plain,(
  ( ! [X167,X165,X163,X166,X164] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X164))),s(sK1,X163)))) | ~p(s(bool,X165)) | ~p(s(bool,X166)) | s(bool,X166) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X163))),s(sK1,X164))) | s(bool,X165) != s(bool,X167)) )),
  inference(superposition,[],[f89,f127])).
fof(f2077,plain,(
  ( ! [X158,X161,X159,X162,X160] : (s(bool,X160) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X159))),s(sK1,X158))) | p(s(bool,X160)) | ~p(s(bool,X161)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X158))),s(sK1,X159)))) | s(bool,X161) != s(bool,X162)) )),
  inference(superposition,[],[f89,f127])).
fof(f2076,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,X155) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X154))),s(sK1,X153))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X153))),s(sK1,X154)))) | ~p(s(bool,X156)) | ~p(s(bool,X155)) | s(bool,X156) != s(bool,X157)) )),
  inference(superposition,[],[f89,f127])).
fof(f2072,plain,(
  ( ! [X134,X136,X135,X133,X137] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X134))),s(sK1,X133)))) | p(s(bool,X135)) | ~p(s(bool,X136)) | p(s(bool,X137)) | s(bool,X137) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X133))),s(sK1,X134)))) )),
  inference(superposition,[],[f88,f127])).
fof(f2070,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X124))),s(sK1,X123)))) | p(s(bool,X125)) | ~p(s(bool,X126)) | ~p(s(bool,X127)) | s(bool,X126) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X123))),s(sK1,X124)))) )),
  inference(superposition,[],[f88,f127])).
fof(f2069,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,X120) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X119))),s(sK1,X118))) | p(s(bool,X121)) | ~p(s(bool,X120)) | ~p(s(bool,X122)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X118))),s(sK1,X119))))) )),
  inference(superposition,[],[f88,f127])).
fof(f2068,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,X115) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X114))),s(sK1,X113))) | p(s(bool,X116)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X113))),s(sK1,X114)))) | ~p(s(bool,X117)) | p(s(bool,X115))) )),
  inference(superposition,[],[f88,f127])).
fof(f2067,plain,(
  ( ! [X111,X109,X112,X110,X108] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X109))),s(sK1,X108)))) | p(s(bool,X110)) | p(s(bool,X111)) | s(bool,X110) != s(bool,X112) | s(bool,X111) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X108))),s(sK1,X109)))) )),
  inference(superposition,[],[f87,f127])).
fof(f2065,plain,(
  ( ! [X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X99))),s(sK1,X98)))) | p(s(bool,X100)) | ~p(s(bool,X101)) | s(bool,X100) != s(bool,X102) | s(bool,X101) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X98))),s(sK1,X99)))) )),
  inference(superposition,[],[f87,f127])).
fof(f2062,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,X85) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X84))),s(sK1,X83))) | p(s(bool,X86)) | ~p(s(bool,X85)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X83))),s(sK1,X84)))) | s(bool,X86) != s(bool,X87)) )),
  inference(superposition,[],[f87,f127])).
fof(f2061,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,X80) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X79))),s(sK1,X78))) | p(s(bool,X81)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X78))),s(sK1,X79)))) | p(s(bool,X80)) | s(bool,X81) != s(bool,X82)) )),
  inference(superposition,[],[f87,f127])).
fof(f2060,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X74))),s(sK1,X73)))) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X75) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X73))),s(sK1,X74))) | s(bool,X76) != s(bool,X77)) )),
  inference(superposition,[],[f85,f127])).
fof(f2058,plain,(
  ( ! [X66,X64,X67,X65,X63] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X64))),s(sK1,X63)))) | ~p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X66) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X63))),s(sK1,X64))) | s(bool,X65) != s(bool,X67)) )),
  inference(superposition,[],[f85,f127])).
fof(f2057,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,X60) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X59))),s(sK1,X58))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X58))),s(sK1,X59)))) | ~p(s(bool,X61)) | ~p(s(bool,X60)) | s(bool,X61) != s(bool,X62)) )),
  inference(superposition,[],[f85,f127])).
fof(f2056,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X55) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X54))),s(sK1,X53))) | p(s(bool,X55)) | ~p(s(bool,X56)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X53))),s(sK1,X54)))) | s(bool,X56) != s(bool,X57)) )),
  inference(superposition,[],[f85,f127])).
fof(f2053,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X39))),s(sK1,X38)))) | p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X38))),s(sK1,X39)))) )),
  inference(superposition,[],[f83,f127])).
fof(f2051,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,X28)))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X28))),s(sK1,X29)))) )),
  inference(superposition,[],[f83,f127])).
fof(f2048,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,X15) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X14))),s(sK1,X13))) | p(s(bool,X16)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X13))),s(sK1,X14)))) | p(s(bool,X15)) | s(bool,X16) != s(bool,X17)) )),
  inference(superposition,[],[f83,f127])).
fof(f2047,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X9))),s(sK1,X8))) | p(s(bool,X11)) | ~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9)))) | s(bool,X11) != s(bool,X12)) )),
  inference(superposition,[],[f83,f127])).
fof(f2009,plain,(
  ( ! [X212,X215,X213,X214] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X213))),s(sK1,X212)))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X215) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X212))),s(sK1,X213)))) )),
  inference(superposition,[],[f125,f127])).
fof(f2007,plain,(
  ( ! [X206,X204,X207,X205] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X205))),s(sK1,X204)))) | p(s(bool,X206)) | ~p(s(bool,X207)) | s(bool,X207) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X204))),s(sK1,X205)))) )),
  inference(superposition,[],[f125,f127])).
fof(f2006,plain,(
  ( ! [X202,X200,X203,X201] : (s(bool,X202) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X201))),s(sK1,X200))) | p(s(bool,X203)) | ~p(s(bool,X202)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X200))),s(sK1,X201))))) )),
  inference(superposition,[],[f125,f127])).
fof(f2005,plain,(
  ( ! [X198,X196,X199,X197] : (s(bool,X198) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X197))),s(sK1,X196))) | p(s(bool,X199)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X196))),s(sK1,X197)))) | p(s(bool,X198))) )),
  inference(superposition,[],[f125,f127])).
fof(f1995,plain,(
  ( ! [X177,X175,X173,X176,X174] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X174))),s(sK1,X173)))) | p(s(bool,X175)) | ~p(s(bool,X176)) | s(bool,X175) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X173))),s(sK1,X174))) | s(bool,X176) != s(bool,X177)) )),
  inference(superposition,[],[f89,f127])).
fof(f1993,plain,(
  ( ! [X167,X165,X163,X166,X164] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X164))),s(sK1,X163)))) | ~p(s(bool,X165)) | ~p(s(bool,X166)) | s(bool,X166) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X163))),s(sK1,X164))) | s(bool,X165) != s(bool,X167)) )),
  inference(superposition,[],[f89,f127])).
fof(f1992,plain,(
  ( ! [X158,X161,X159,X162,X160] : (s(bool,X160) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X159))),s(sK1,X158))) | p(s(bool,X160)) | ~p(s(bool,X161)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X158))),s(sK1,X159)))) | s(bool,X161) != s(bool,X162)) )),
  inference(superposition,[],[f89,f127])).
fof(f1991,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,X155) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X154))),s(sK1,X153))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X153))),s(sK1,X154)))) | ~p(s(bool,X156)) | ~p(s(bool,X155)) | s(bool,X156) != s(bool,X157)) )),
  inference(superposition,[],[f89,f127])).
fof(f1987,plain,(
  ( ! [X134,X136,X135,X133,X137] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X134))),s(sK1,X133)))) | p(s(bool,X135)) | ~p(s(bool,X136)) | p(s(bool,X137)) | s(bool,X137) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X133))),s(sK1,X134)))) )),
  inference(superposition,[],[f88,f127])).
fof(f1985,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X124))),s(sK1,X123)))) | p(s(bool,X125)) | ~p(s(bool,X126)) | ~p(s(bool,X127)) | s(bool,X126) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X123))),s(sK1,X124)))) )),
  inference(superposition,[],[f88,f127])).
fof(f1984,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,X120) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X119))),s(sK1,X118))) | p(s(bool,X121)) | ~p(s(bool,X120)) | ~p(s(bool,X122)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X118))),s(sK1,X119))))) )),
  inference(superposition,[],[f88,f127])).
fof(f1983,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,X115) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X114))),s(sK1,X113))) | p(s(bool,X116)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X113))),s(sK1,X114)))) | ~p(s(bool,X117)) | p(s(bool,X115))) )),
  inference(superposition,[],[f88,f127])).
fof(f1982,plain,(
  ( ! [X111,X109,X112,X110,X108] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X109))),s(sK1,X108)))) | p(s(bool,X110)) | p(s(bool,X111)) | s(bool,X110) != s(bool,X112) | s(bool,X111) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X108))),s(sK1,X109)))) )),
  inference(superposition,[],[f87,f127])).
fof(f1980,plain,(
  ( ! [X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X99))),s(sK1,X98)))) | p(s(bool,X100)) | ~p(s(bool,X101)) | s(bool,X100) != s(bool,X102) | s(bool,X101) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X98))),s(sK1,X99)))) )),
  inference(superposition,[],[f87,f127])).
fof(f1977,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,X85) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X84))),s(sK1,X83))) | p(s(bool,X86)) | ~p(s(bool,X85)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X83))),s(sK1,X84)))) | s(bool,X86) != s(bool,X87)) )),
  inference(superposition,[],[f87,f127])).
fof(f1976,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,X80) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X79))),s(sK1,X78))) | p(s(bool,X81)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X78))),s(sK1,X79)))) | p(s(bool,X80)) | s(bool,X81) != s(bool,X82)) )),
  inference(superposition,[],[f87,f127])).
fof(f1975,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X74))),s(sK1,X73)))) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X75) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X73))),s(sK1,X74))) | s(bool,X76) != s(bool,X77)) )),
  inference(superposition,[],[f85,f127])).
fof(f1973,plain,(
  ( ! [X66,X64,X67,X65,X63] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X64))),s(sK1,X63)))) | ~p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X66) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X63))),s(sK1,X64))) | s(bool,X65) != s(bool,X67)) )),
  inference(superposition,[],[f85,f127])).
fof(f1972,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,X60) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X59))),s(sK1,X58))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X58))),s(sK1,X59)))) | ~p(s(bool,X61)) | ~p(s(bool,X60)) | s(bool,X61) != s(bool,X62)) )),
  inference(superposition,[],[f85,f127])).
fof(f1971,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X55) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X54))),s(sK1,X53))) | p(s(bool,X55)) | ~p(s(bool,X56)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X53))),s(sK1,X54)))) | s(bool,X56) != s(bool,X57)) )),
  inference(superposition,[],[f85,f127])).
fof(f1968,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X39))),s(sK1,X38)))) | p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X38))),s(sK1,X39)))) )),
  inference(superposition,[],[f83,f127])).
fof(f1966,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,X28)))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X28))),s(sK1,X29)))) )),
  inference(superposition,[],[f83,f127])).
fof(f1963,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,X15) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X14))),s(sK1,X13))) | p(s(bool,X16)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X13))),s(sK1,X14)))) | p(s(bool,X15)) | s(bool,X16) != s(bool,X17)) )),
  inference(superposition,[],[f83,f127])).
fof(f1962,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X9))),s(sK1,X8))) | p(s(bool,X11)) | ~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,X9)))) | s(bool,X11) != s(bool,X12)) )),
  inference(superposition,[],[f83,f127])).
fof(f1873,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,happ(s(fun(X50,bool),happ(s(fun(X50,fun(X50,bool)),X51),s(X50,sK4(X51,X50)))),s(X50,sK5(X51,X50)))) != s(bool,happ(s(fun(X50,bool),happ(s(fun(X50,fun(X50,bool)),X51),s(X50,sK5(X51,X50)))),s(X50,sK4(X51,X50)))) | p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X53) != s(bool,'const.relation.symmetric_1'(s(fun(X50,fun(X50,bool)),X51)))) )),
  inference(resolution,[],[f119,f125])).
fof(f1869,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),X40),s(X39,sK4(X40,X39)))),s(X39,sK5(X40,X39)))) != s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),X40),s(X39,sK5(X40,X39)))),s(X39,sK4(X40,X39)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.relation.symmetric_1'(s(fun(X39,fun(X39,bool)),X40))) | s(bool,X42) != s(bool,X43)) )),
  inference(resolution,[],[f119,f89])).
fof(f1866,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,happ(s(fun(X24,bool),happ(s(fun(X24,fun(X24,bool)),X25),s(X24,sK4(X25,X24)))),s(X24,sK5(X25,X24)))) != s(bool,happ(s(fun(X24,bool),happ(s(fun(X24,fun(X24,bool)),X25),s(X24,sK5(X25,X24)))),s(X24,sK4(X25,X24)))) | p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.symmetric_1'(s(fun(X24,fun(X24,bool)),X25)))) )),
  inference(resolution,[],[f119,f88])).
fof(f1865,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,happ(s(fun(X19,bool),happ(s(fun(X19,fun(X19,bool)),X20),s(X19,sK4(X20,X19)))),s(X19,sK5(X20,X19)))) != s(bool,happ(s(fun(X19,bool),happ(s(fun(X19,fun(X19,bool)),X20),s(X19,sK5(X20,X19)))),s(X19,sK4(X20,X19)))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.relation.symmetric_1'(s(fun(X19,fun(X19,bool)),X20)))) )),
  inference(resolution,[],[f119,f87])).
fof(f1864,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,sK4(X15,X14)))),s(X14,sK5(X15,X14)))) != s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,sK5(X15,X14)))),s(X14,sK4(X15,X14)))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.relation.symmetric_1'(s(fun(X14,fun(X14,bool)),X15))) | s(bool,X17) != s(bool,X18)) )),
  inference(resolution,[],[f119,f85])).
fof(f1862,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK4(X5,X4)))),s(X4,sK5(X5,X4)))) != s(bool,happ(s(fun(X4,bool),happ(s(fun(X4,fun(X4,bool)),X5),s(X4,sK5(X5,X4)))),s(X4,sK4(X5,X4)))) | p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.relation.symmetric_1'(s(fun(X4,fun(X4,bool)),X5)))) )),
  inference(resolution,[],[f119,f83])).
fof(f1855,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,'F_0') != s(bool,X100) | s(bool,X101) != s(bool,X102) | s(bool,X102) != s(bool,X103) | s(bool,'F_0') = s(bool,X103) | s(bool,X100) != s(bool,X101)) ) | $spl68),
  inference(cnf_transformation,[],[f1855_D])).
fof(f1855_D,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,'F_0') != s(bool,X100) | s(bool,X101) != s(bool,X102) | s(bool,X102) != s(bool,X103) | s(bool,'F_0') = s(bool,X103) | s(bool,X100) != s(bool,X101)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f1857,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | ($spl24 | $spl56)),
  inference(subsumption_resolution,[],[f1845,f385])).
fof(f1845,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,X15)) ) | $spl56),
  inference(superposition,[],[f1356,f122])).
fof(f1840,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,'F_0') != s(bool,X105) | s(bool,X105) != s(bool,X106) | s(bool,X107) != s(bool,X108) | s(bool,'F_0') = s(bool,X108) | s(bool,X106) != s(bool,X107) | ~p(s(bool,X109))) ) | $spl56),
  inference(resolution,[],[f1356,f101])).
fof(f1854,plain,(
  ( ! [X103,X101,X102,X100,X104] : (s(bool,'F_0') != s(bool,X100) | s(bool,X100) != s(bool,X101) | s(bool,X102) != s(bool,X103) | s(bool,'F_0') = s(bool,X103) | s(bool,X101) != s(bool,X102) | p(s(bool,X104))) ) | $spl56),
  inference(subsumption_resolution,[],[f1839,f126])).
fof(f1839,plain,(
  ( ! [X103,X101,X102,X100,X104] : (s(bool,'F_0') != s(bool,X100) | s(bool,X100) != s(bool,X101) | s(bool,X102) != s(bool,X103) | s(bool,'F_0') = s(bool,X103) | s(bool,X101) != s(bool,X102) | p(s(bool,X104)) | ~p(s(bool,'T_0'))) ) | $spl56),
  inference(resolution,[],[f1356,f105])).
fof(f1835,plain,(
  ( ! [X80,X85,X83,X81,X79,X84,X82] : (s(bool,X79) != s(bool,X80) | s(bool,X79) != s(bool,X81) | s(bool,X82) != s(bool,X83) | s(bool,'F_0') = s(bool,X83) | s(bool,X81) != s(bool,X82) | p(s(bool,X84)) | p(s(bool,X85)) | s(bool,X80) != s(bool,X85)) ) | $spl56),
  inference(resolution,[],[f1356,f125])).
fof(f1831,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X59) != s(bool,X60) | s(bool,'F_0') = s(bool,X60) | s(bool,X58) != s(bool,X59) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X57) != s(bool,X61) | s(bool,X62) != s(bool,X63)) ) | $spl56),
  inference(resolution,[],[f1356,f89])).
fof(f1828,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,X32) != s(bool,X34) | s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) != s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X33) != s(bool,X39)) ) | $spl56),
  inference(resolution,[],[f1356,f88])).
fof(f1827,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X26) != s(bool,X27) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X25) != s(bool,X30)) ) | $spl56),
  inference(resolution,[],[f1356,f87])).
fof(f1826,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,X19) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | s(bool,X18) != s(bool,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X17) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl56),
  inference(resolution,[],[f1356,f85])).
fof(f1824,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X2) != s(bool,X3) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X1) != s(bool,X6)) ) | $spl56),
  inference(resolution,[],[f1356,f83])).
fof(f1595,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X1) != s(bool,X2)) ) | $spl64),
  inference(cnf_transformation,[],[f1595_D])).
fof(f1595_D,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X1) != s(bool,X2)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1591,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl34 | $spl40)),
  inference(duplicate_literal_removal,[],[f1581])).
fof(f1581,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X0)) ) | ($spl34 | $spl40)),
  inference(superposition,[],[f1342,f122])).
fof(f1594,plain,(
  ( ! [X17,X18] : (s(bool,X17) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,'T_0') = s(bool,X18) | s(bool,X17) != s(bool,X18)) ) | $spl62),
  inference(cnf_transformation,[],[f1594_D])).
fof(f1594_D,plain,(
  ( ! [X17,X18] : (s(bool,X17) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,'T_0') = s(bool,X18) | s(bool,X17) != s(bool,X18)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1593,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,'F_0') = s(bool,X16) | s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,X18)) ) | ($spl34 | $spl40)),
  inference(subsumption_resolution,[],[f1576,f122])).
fof(f1576,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,X16) | s(bool,X17) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,'F_0') = s(bool,X16) | s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,X18)) ) | ($spl34 | $spl40)),
  inference(superposition,[],[f1342,f151])).
fof(f1558,plain,(
  s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | $spl61),
  inference(cnf_transformation,[],[f1558_D])).
fof(f1558_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f1556,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1539])).
fof(f1539,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f748,f122])).
fof(f748,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,'F_0') = s(bool,X9)) ) | $spl6),
  inference(subsumption_resolution,[],[f740,f115])).
fof(f740,plain,(
  ( ! [X9] : (p(s(bool,'F_0')) | s(bool,X9) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,'F_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f504,f180])).
fof(f1359,plain,(
  ( ! [X87,X85,X86] : (s(bool,'F_0') != s(bool,X85) | s(bool,X86) != s(bool,X87) | s(bool,'F_0') = s(bool,X87) | s(bool,X85) != s(bool,X86)) ) | $spl58),
  inference(cnf_transformation,[],[f1359_D])).
fof(f1359_D,plain,(
  ( ! [X87,X85,X86] : (s(bool,'F_0') != s(bool,X85) | s(bool,X86) != s(bool,X87) | s(bool,'F_0') = s(bool,X87) | s(bool,X85) != s(bool,X86)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1361,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s(bool,X13) != s(bool,X14) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) != s(bool,X15)) ) | ($spl24 | $spl40)),
  inference(subsumption_resolution,[],[f1351,f385])).
fof(f1351,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s(bool,X13) != s(bool,X14) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) != s(bool,X15) | s(bool,'T_0') = s(bool,X12)) ) | $spl40),
  inference(superposition,[],[f805,f122])).
fof(f1346,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'F_0') != s(bool,X89) | s(bool,X89) != s(bool,X90) | s(bool,'F_0') = s(bool,X91) | s(bool,X90) != s(bool,X91) | ~p(s(bool,X92))) ) | $spl40),
  inference(resolution,[],[f805,f101])).
fof(f1358,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'F_0') != s(bool,X85) | s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X87) | s(bool,X86) != s(bool,X87) | p(s(bool,X88))) ) | $spl40),
  inference(subsumption_resolution,[],[f1345,f126])).
fof(f1345,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'F_0') != s(bool,X85) | s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X87) | s(bool,X86) != s(bool,X87) | p(s(bool,X88)) | ~p(s(bool,'T_0'))) ) | $spl40),
  inference(resolution,[],[f805,f105])).
fof(f1341,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (s(bool,X68) != s(bool,X69) | s(bool,X68) != s(bool,X70) | s(bool,'F_0') = s(bool,X71) | s(bool,X70) != s(bool,X71) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X69) != s(bool,X73)) ) | $spl40),
  inference(resolution,[],[f805,f125])).
fof(f1337,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X52) | s(bool,X51) != s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl40),
  inference(resolution,[],[f805,f89])).
fof(f1334,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X31) | s(bool,X30) != s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) ) | $spl40),
  inference(resolution,[],[f805,f88])).
fof(f1333,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) != s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl40),
  inference(resolution,[],[f805,f87])).
fof(f1332,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,'F_0') = s(bool,X17) | s(bool,X16) != s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl40),
  inference(resolution,[],[f805,f85])).
fof(f1330,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl40),
  inference(resolution,[],[f805,f83])).
fof(f1268,plain,(
  ( ! [X300,X302,X298,X301,X303,X299] : (~p(s(bool,happ(s(fun(X298,bool),happ(s(fun(X298,fun(X298,bool)),X299),s(X298,X301))),s(X298,X300)))) | p(s(bool,X302)) | p(s(bool,X303)) | s(bool,X303) != s(bool,happ(s(fun(X298,bool),happ(s(fun(X298,fun(X298,bool)),X299),s(X298,X300))),s(X298,X301))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X298,fun(X298,bool)),X299))))) )),
  inference(superposition,[],[f125,f118])).
fof(f1266,plain,(
  ( ! [X288,X290,X287,X289,X291,X286] : (p(s(bool,happ(s(fun(X286,bool),happ(s(fun(X286,fun(X286,bool)),X287),s(X286,X289))),s(X286,X288)))) | p(s(bool,X290)) | ~p(s(bool,X291)) | s(bool,X291) != s(bool,happ(s(fun(X286,bool),happ(s(fun(X286,fun(X286,bool)),X287),s(X286,X288))),s(X286,X289))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X286,fun(X286,bool)),X287))))) )),
  inference(superposition,[],[f125,f118])).
fof(f1265,plain,(
  ( ! [X280,X282,X285,X281,X283,X284] : (s(bool,X284) != s(bool,happ(s(fun(X280,bool),happ(s(fun(X280,fun(X280,bool)),X281),s(X280,X283))),s(X280,X282))) | p(s(bool,X285)) | ~p(s(bool,X284)) | p(s(bool,happ(s(fun(X280,bool),happ(s(fun(X280,fun(X280,bool)),X281),s(X280,X282))),s(X280,X283)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X280,fun(X280,bool)),X281))))) )),
  inference(superposition,[],[f125,f118])).
fof(f1264,plain,(
  ( ! [X277,X279,X275,X276,X278,X274] : (s(bool,X278) != s(bool,happ(s(fun(X274,bool),happ(s(fun(X274,fun(X274,bool)),X275),s(X274,X277))),s(X274,X276))) | p(s(bool,X279)) | ~p(s(bool,happ(s(fun(X274,bool),happ(s(fun(X274,fun(X274,bool)),X275),s(X274,X276))),s(X274,X277)))) | p(s(bool,X278)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X274,fun(X274,bool)),X275))))) )),
  inference(superposition,[],[f125,f118])).
fof(f1254,plain,(
  ( ! [X231,X237,X235,X233,X236,X234,X232] : (~p(s(bool,happ(s(fun(X231,bool),happ(s(fun(X231,fun(X231,bool)),X232),s(X231,X234))),s(X231,X233)))) | p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,happ(s(fun(X231,bool),happ(s(fun(X231,fun(X231,bool)),X232),s(X231,X233))),s(X231,X234))) | s(bool,X236) != s(bool,X237) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X231,fun(X231,bool)),X232))))) )),
  inference(superposition,[],[f89,f118])).
fof(f1252,plain,(
  ( ! [X218,X222,X220,X219,X217,X223,X221] : (p(s(bool,happ(s(fun(X217,bool),happ(s(fun(X217,fun(X217,bool)),X218),s(X217,X220))),s(X217,X219)))) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,happ(s(fun(X217,bool),happ(s(fun(X217,fun(X217,bool)),X218),s(X217,X219))),s(X217,X220))) | s(bool,X221) != s(bool,X223) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X217,fun(X217,bool)),X218))))) )),
  inference(superposition,[],[f89,f118])).
fof(f1251,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X214] : (s(bool,X214) != s(bool,happ(s(fun(X210,bool),happ(s(fun(X210,fun(X210,bool)),X211),s(X210,X213))),s(X210,X212))) | p(s(bool,X214)) | ~p(s(bool,X215)) | ~p(s(bool,happ(s(fun(X210,bool),happ(s(fun(X210,fun(X210,bool)),X211),s(X210,X212))),s(X210,X213)))) | s(bool,X215) != s(bool,X216) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X210,fun(X210,bool)),X211))))) )),
  inference(superposition,[],[f89,f118])).
fof(f1250,plain,(
  ( ! [X208,X206,X204,X209,X207,X205,X203] : (s(bool,X207) != s(bool,happ(s(fun(X203,bool),happ(s(fun(X203,fun(X203,bool)),X204),s(X203,X206))),s(X203,X205))) | p(s(bool,happ(s(fun(X203,bool),happ(s(fun(X203,fun(X203,bool)),X204),s(X203,X205))),s(X203,X206)))) | ~p(s(bool,X208)) | ~p(s(bool,X207)) | s(bool,X208) != s(bool,X209) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X203,fun(X203,bool)),X204))))) )),
  inference(superposition,[],[f89,f118])).
fof(f1246,plain,(
  ( ! [X177,X175,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X175,bool),happ(s(fun(X175,fun(X175,bool)),X176),s(X175,X178))),s(X175,X177)))) | p(s(bool,X179)) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,happ(s(fun(X175,bool),happ(s(fun(X175,fun(X175,bool)),X176),s(X175,X177))),s(X175,X178))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X175,fun(X175,bool)),X176))))) )),
  inference(superposition,[],[f88,f118])).
fof(f1244,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162] : (p(s(bool,happ(s(fun(X161,bool),happ(s(fun(X161,fun(X161,bool)),X162),s(X161,X164))),s(X161,X163)))) | p(s(bool,X165)) | ~p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,happ(s(fun(X161,bool),happ(s(fun(X161,fun(X161,bool)),X162),s(X161,X163))),s(X161,X164))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X161,fun(X161,bool)),X162))))) )),
  inference(superposition,[],[f88,f118])).
fof(f1243,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X160] : (s(bool,X158) != s(bool,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),X155),s(X154,X157))),s(X154,X156))) | p(s(bool,X159)) | ~p(s(bool,X158)) | ~p(s(bool,X160)) | p(s(bool,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),X155),s(X154,X156))),s(X154,X157)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X154,fun(X154,bool)),X155))))) )),
  inference(superposition,[],[f88,f118])).
fof(f1242,plain,(
  ( ! [X152,X151,X149,X147,X153,X150,X148] : (s(bool,X151) != s(bool,happ(s(fun(X147,bool),happ(s(fun(X147,fun(X147,bool)),X148),s(X147,X150))),s(X147,X149))) | p(s(bool,X152)) | ~p(s(bool,happ(s(fun(X147,bool),happ(s(fun(X147,fun(X147,bool)),X148),s(X147,X149))),s(X147,X150)))) | ~p(s(bool,X153)) | p(s(bool,X151)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X147,fun(X147,bool)),X148))))) )),
  inference(superposition,[],[f88,f118])).
fof(f1241,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X140,bool),happ(s(fun(X140,fun(X140,bool)),X141),s(X140,X143))),s(X140,X142)))) | p(s(bool,X144)) | p(s(bool,X145)) | s(bool,X144) != s(bool,X146) | s(bool,X145) != s(bool,happ(s(fun(X140,bool),happ(s(fun(X140,fun(X140,bool)),X141),s(X140,X142))),s(X140,X143))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X140,fun(X140,bool)),X141))))) )),
  inference(superposition,[],[f87,f118])).
fof(f1239,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (p(s(bool,happ(s(fun(X126,bool),happ(s(fun(X126,fun(X126,bool)),X127),s(X126,X129))),s(X126,X128)))) | p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,X132) | s(bool,X131) != s(bool,happ(s(fun(X126,bool),happ(s(fun(X126,fun(X126,bool)),X127),s(X126,X128))),s(X126,X129))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X126,fun(X126,bool)),X127))))) )),
  inference(superposition,[],[f87,f118])).
fof(f1236,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(bool,X109) != s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X108))),s(X105,X107))) | p(s(bool,X110)) | ~p(s(bool,X109)) | p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108)))) | s(bool,X110) != s(bool,X111) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X105,fun(X105,bool)),X106))))) )),
  inference(superposition,[],[f87,f118])).
fof(f1235,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (s(bool,X102) != s(bool,happ(s(fun(X98,bool),happ(s(fun(X98,fun(X98,bool)),X99),s(X98,X101))),s(X98,X100))) | p(s(bool,X103)) | ~p(s(bool,happ(s(fun(X98,bool),happ(s(fun(X98,fun(X98,bool)),X99),s(X98,X100))),s(X98,X101)))) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X98,fun(X98,bool)),X99))))) )),
  inference(superposition,[],[f87,f118])).
fof(f1234,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (~p(s(bool,happ(s(fun(X91,bool),happ(s(fun(X91,fun(X91,bool)),X92),s(X91,X94))),s(X91,X93)))) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,happ(s(fun(X91,bool),happ(s(fun(X91,fun(X91,bool)),X92),s(X91,X93))),s(X91,X94))) | s(bool,X96) != s(bool,X97) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X91,fun(X91,bool)),X92))))) )),
  inference(superposition,[],[f85,f118])).
fof(f1232,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X77,bool),happ(s(fun(X77,fun(X77,bool)),X78),s(X77,X80))),s(X77,X79)))) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X77,bool),happ(s(fun(X77,fun(X77,bool)),X78),s(X77,X79))),s(X77,X80))) | s(bool,X81) != s(bool,X83) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X77,fun(X77,bool)),X78))))) )),
  inference(superposition,[],[f85,f118])).
fof(f1231,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,X74) != s(bool,happ(s(fun(X70,bool),happ(s(fun(X70,fun(X70,bool)),X71),s(X70,X73))),s(X70,X72))) | p(s(bool,happ(s(fun(X70,bool),happ(s(fun(X70,fun(X70,bool)),X71),s(X70,X72))),s(X70,X73)))) | ~p(s(bool,X75)) | ~p(s(bool,X74)) | s(bool,X75) != s(bool,X76) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X70,fun(X70,bool)),X71))))) )),
  inference(superposition,[],[f85,f118])).
fof(f1230,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,X67) != s(bool,happ(s(fun(X63,bool),happ(s(fun(X63,fun(X63,bool)),X64),s(X63,X66))),s(X63,X65))) | p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,happ(s(fun(X63,bool),happ(s(fun(X63,fun(X63,bool)),X64),s(X63,X65))),s(X63,X66)))) | s(bool,X68) != s(bool,X69) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X63,fun(X63,bool)),X64))))) )),
  inference(superposition,[],[f85,f118])).
fof(f1227,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X45))),s(X42,X44)))) | p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X42,fun(X42,bool)),X43))))) )),
  inference(superposition,[],[f83,f118])).
fof(f1225,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),happ(s(fun(X28,fun(X28,bool)),X29),s(X28,X31))),s(X28,X30)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,happ(s(fun(X28,bool),happ(s(fun(X28,fun(X28,bool)),X29),s(X28,X30))),s(X28,X31))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X28,fun(X28,bool)),X29))))) )),
  inference(superposition,[],[f83,f118])).
fof(f1222,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,X11) != s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X10))),s(X7,X9))) | p(s(bool,X12)) | ~p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X9))),s(X7,X10)))) | p(s(bool,X11)) | s(bool,X12) != s(bool,X13) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X7,fun(X7,bool)),X8))))) )),
  inference(superposition,[],[f83,f118])).
fof(f1221,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X4) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,X5)) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(bool,X5) != s(bool,X6) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(superposition,[],[f83,f118])).
fof(f1204,plain,(
  ( ! [X300,X302,X298,X301,X303,X299] : (~p(s(bool,happ(s(fun(X298,bool),happ(s(fun(X298,fun(X298,bool)),X299),s(X298,X301))),s(X298,X300)))) | p(s(bool,X302)) | p(s(bool,X303)) | s(bool,X303) != s(bool,happ(s(fun(X298,bool),happ(s(fun(X298,fun(X298,bool)),X299),s(X298,X300))),s(X298,X301))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X298,fun(X298,bool)),X299))))) )),
  inference(superposition,[],[f125,f118])).
fof(f1202,plain,(
  ( ! [X288,X290,X287,X289,X291,X286] : (p(s(bool,happ(s(fun(X286,bool),happ(s(fun(X286,fun(X286,bool)),X287),s(X286,X289))),s(X286,X288)))) | p(s(bool,X290)) | ~p(s(bool,X291)) | s(bool,X291) != s(bool,happ(s(fun(X286,bool),happ(s(fun(X286,fun(X286,bool)),X287),s(X286,X288))),s(X286,X289))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X286,fun(X286,bool)),X287))))) )),
  inference(superposition,[],[f125,f118])).
fof(f1201,plain,(
  ( ! [X280,X282,X285,X281,X283,X284] : (s(bool,X284) != s(bool,happ(s(fun(X280,bool),happ(s(fun(X280,fun(X280,bool)),X281),s(X280,X283))),s(X280,X282))) | p(s(bool,X285)) | ~p(s(bool,X284)) | p(s(bool,happ(s(fun(X280,bool),happ(s(fun(X280,fun(X280,bool)),X281),s(X280,X282))),s(X280,X283)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X280,fun(X280,bool)),X281))))) )),
  inference(superposition,[],[f125,f118])).
fof(f1200,plain,(
  ( ! [X277,X279,X275,X276,X278,X274] : (s(bool,X278) != s(bool,happ(s(fun(X274,bool),happ(s(fun(X274,fun(X274,bool)),X275),s(X274,X277))),s(X274,X276))) | p(s(bool,X279)) | ~p(s(bool,happ(s(fun(X274,bool),happ(s(fun(X274,fun(X274,bool)),X275),s(X274,X276))),s(X274,X277)))) | p(s(bool,X278)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X274,fun(X274,bool)),X275))))) )),
  inference(superposition,[],[f125,f118])).
fof(f1190,plain,(
  ( ! [X231,X237,X235,X233,X236,X234,X232] : (~p(s(bool,happ(s(fun(X231,bool),happ(s(fun(X231,fun(X231,bool)),X232),s(X231,X234))),s(X231,X233)))) | p(s(bool,X235)) | ~p(s(bool,X236)) | s(bool,X235) != s(bool,happ(s(fun(X231,bool),happ(s(fun(X231,fun(X231,bool)),X232),s(X231,X233))),s(X231,X234))) | s(bool,X236) != s(bool,X237) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X231,fun(X231,bool)),X232))))) )),
  inference(superposition,[],[f89,f118])).
fof(f1188,plain,(
  ( ! [X218,X222,X220,X219,X217,X223,X221] : (p(s(bool,happ(s(fun(X217,bool),happ(s(fun(X217,fun(X217,bool)),X218),s(X217,X220))),s(X217,X219)))) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,happ(s(fun(X217,bool),happ(s(fun(X217,fun(X217,bool)),X218),s(X217,X219))),s(X217,X220))) | s(bool,X221) != s(bool,X223) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X217,fun(X217,bool)),X218))))) )),
  inference(superposition,[],[f89,f118])).
fof(f1187,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X214] : (s(bool,X214) != s(bool,happ(s(fun(X210,bool),happ(s(fun(X210,fun(X210,bool)),X211),s(X210,X213))),s(X210,X212))) | p(s(bool,X214)) | ~p(s(bool,X215)) | ~p(s(bool,happ(s(fun(X210,bool),happ(s(fun(X210,fun(X210,bool)),X211),s(X210,X212))),s(X210,X213)))) | s(bool,X215) != s(bool,X216) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X210,fun(X210,bool)),X211))))) )),
  inference(superposition,[],[f89,f118])).
fof(f1186,plain,(
  ( ! [X208,X206,X204,X209,X207,X205,X203] : (s(bool,X207) != s(bool,happ(s(fun(X203,bool),happ(s(fun(X203,fun(X203,bool)),X204),s(X203,X206))),s(X203,X205))) | p(s(bool,happ(s(fun(X203,bool),happ(s(fun(X203,fun(X203,bool)),X204),s(X203,X205))),s(X203,X206)))) | ~p(s(bool,X208)) | ~p(s(bool,X207)) | s(bool,X208) != s(bool,X209) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X203,fun(X203,bool)),X204))))) )),
  inference(superposition,[],[f89,f118])).
fof(f1182,plain,(
  ( ! [X177,X175,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X175,bool),happ(s(fun(X175,fun(X175,bool)),X176),s(X175,X178))),s(X175,X177)))) | p(s(bool,X179)) | ~p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,happ(s(fun(X175,bool),happ(s(fun(X175,fun(X175,bool)),X176),s(X175,X177))),s(X175,X178))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X175,fun(X175,bool)),X176))))) )),
  inference(superposition,[],[f88,f118])).
fof(f1180,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162] : (p(s(bool,happ(s(fun(X161,bool),happ(s(fun(X161,fun(X161,bool)),X162),s(X161,X164))),s(X161,X163)))) | p(s(bool,X165)) | ~p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X166) != s(bool,happ(s(fun(X161,bool),happ(s(fun(X161,fun(X161,bool)),X162),s(X161,X163))),s(X161,X164))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X161,fun(X161,bool)),X162))))) )),
  inference(superposition,[],[f88,f118])).
fof(f1179,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X160] : (s(bool,X158) != s(bool,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),X155),s(X154,X157))),s(X154,X156))) | p(s(bool,X159)) | ~p(s(bool,X158)) | ~p(s(bool,X160)) | p(s(bool,happ(s(fun(X154,bool),happ(s(fun(X154,fun(X154,bool)),X155),s(X154,X156))),s(X154,X157)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X154,fun(X154,bool)),X155))))) )),
  inference(superposition,[],[f88,f118])).
fof(f1178,plain,(
  ( ! [X152,X151,X149,X147,X153,X150,X148] : (s(bool,X151) != s(bool,happ(s(fun(X147,bool),happ(s(fun(X147,fun(X147,bool)),X148),s(X147,X150))),s(X147,X149))) | p(s(bool,X152)) | ~p(s(bool,happ(s(fun(X147,bool),happ(s(fun(X147,fun(X147,bool)),X148),s(X147,X149))),s(X147,X150)))) | ~p(s(bool,X153)) | p(s(bool,X151)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X147,fun(X147,bool)),X148))))) )),
  inference(superposition,[],[f88,f118])).
fof(f1177,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X140,bool),happ(s(fun(X140,fun(X140,bool)),X141),s(X140,X143))),s(X140,X142)))) | p(s(bool,X144)) | p(s(bool,X145)) | s(bool,X144) != s(bool,X146) | s(bool,X145) != s(bool,happ(s(fun(X140,bool),happ(s(fun(X140,fun(X140,bool)),X141),s(X140,X142))),s(X140,X143))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X140,fun(X140,bool)),X141))))) )),
  inference(superposition,[],[f87,f118])).
fof(f1175,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (p(s(bool,happ(s(fun(X126,bool),happ(s(fun(X126,fun(X126,bool)),X127),s(X126,X129))),s(X126,X128)))) | p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,X132) | s(bool,X131) != s(bool,happ(s(fun(X126,bool),happ(s(fun(X126,fun(X126,bool)),X127),s(X126,X128))),s(X126,X129))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X126,fun(X126,bool)),X127))))) )),
  inference(superposition,[],[f87,f118])).
fof(f1172,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(bool,X109) != s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X108))),s(X105,X107))) | p(s(bool,X110)) | ~p(s(bool,X109)) | p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108)))) | s(bool,X110) != s(bool,X111) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X105,fun(X105,bool)),X106))))) )),
  inference(superposition,[],[f87,f118])).
fof(f1171,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (s(bool,X102) != s(bool,happ(s(fun(X98,bool),happ(s(fun(X98,fun(X98,bool)),X99),s(X98,X101))),s(X98,X100))) | p(s(bool,X103)) | ~p(s(bool,happ(s(fun(X98,bool),happ(s(fun(X98,fun(X98,bool)),X99),s(X98,X100))),s(X98,X101)))) | p(s(bool,X102)) | s(bool,X103) != s(bool,X104) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X98,fun(X98,bool)),X99))))) )),
  inference(superposition,[],[f87,f118])).
fof(f1170,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (~p(s(bool,happ(s(fun(X91,bool),happ(s(fun(X91,fun(X91,bool)),X92),s(X91,X94))),s(X91,X93)))) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,happ(s(fun(X91,bool),happ(s(fun(X91,fun(X91,bool)),X92),s(X91,X93))),s(X91,X94))) | s(bool,X96) != s(bool,X97) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X91,fun(X91,bool)),X92))))) )),
  inference(superposition,[],[f85,f118])).
fof(f1168,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X77,bool),happ(s(fun(X77,fun(X77,bool)),X78),s(X77,X80))),s(X77,X79)))) | ~p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(X77,bool),happ(s(fun(X77,fun(X77,bool)),X78),s(X77,X79))),s(X77,X80))) | s(bool,X81) != s(bool,X83) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X77,fun(X77,bool)),X78))))) )),
  inference(superposition,[],[f85,f118])).
fof(f1167,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,X74) != s(bool,happ(s(fun(X70,bool),happ(s(fun(X70,fun(X70,bool)),X71),s(X70,X73))),s(X70,X72))) | p(s(bool,happ(s(fun(X70,bool),happ(s(fun(X70,fun(X70,bool)),X71),s(X70,X72))),s(X70,X73)))) | ~p(s(bool,X75)) | ~p(s(bool,X74)) | s(bool,X75) != s(bool,X76) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X70,fun(X70,bool)),X71))))) )),
  inference(superposition,[],[f85,f118])).
fof(f1166,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,X67) != s(bool,happ(s(fun(X63,bool),happ(s(fun(X63,fun(X63,bool)),X64),s(X63,X66))),s(X63,X65))) | p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,happ(s(fun(X63,bool),happ(s(fun(X63,fun(X63,bool)),X64),s(X63,X65))),s(X63,X66)))) | s(bool,X68) != s(bool,X69) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X63,fun(X63,bool)),X64))))) )),
  inference(superposition,[],[f85,f118])).
fof(f1163,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X45))),s(X42,X44)))) | p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,happ(s(fun(X42,bool),happ(s(fun(X42,fun(X42,bool)),X43),s(X42,X44))),s(X42,X45))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X42,fun(X42,bool)),X43))))) )),
  inference(superposition,[],[f83,f118])).
fof(f1161,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X28,bool),happ(s(fun(X28,fun(X28,bool)),X29),s(X28,X31))),s(X28,X30)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,happ(s(fun(X28,bool),happ(s(fun(X28,fun(X28,bool)),X29),s(X28,X30))),s(X28,X31))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X28,fun(X28,bool)),X29))))) )),
  inference(superposition,[],[f83,f118])).
fof(f1158,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,X11) != s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X10))),s(X7,X9))) | p(s(bool,X12)) | ~p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X9))),s(X7,X10)))) | p(s(bool,X11)) | s(bool,X12) != s(bool,X13) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X7,fun(X7,bool)),X8))))) )),
  inference(superposition,[],[f83,f118])).
fof(f1157,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X4) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,X5)) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(bool,X5) != s(bool,X6) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(superposition,[],[f83,f118])).
fof(f1116,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (s(bool,happ(s(fun(X72,bool),happ(s(fun(X72,fun(X72,bool)),X73),s(X72,X74))),s(X72,X75))) = s(bool,happ(s(fun(X72,bool),happ(s(fun(X72,fun(X72,bool)),X73),s(X72,X75))),s(X72,X74))) | p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X77) != s(bool,'const.relation.symmetric_1'(s(fun(X72,fun(X72,bool)),X73)))) )),
  inference(resolution,[],[f118,f125])).
fof(f1111,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,happ(s(fun(X49,bool),happ(s(fun(X49,fun(X49,bool)),X50),s(X49,X51))),s(X49,X52))) = s(bool,happ(s(fun(X49,bool),happ(s(fun(X49,fun(X49,bool)),X50),s(X49,X52))),s(X49,X51))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.relation.symmetric_1'(s(fun(X49,fun(X49,bool)),X50))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f118,f89])).
fof(f1109,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X37))),s(X35,X38))) = s(bool,happ(s(fun(X35,bool),happ(s(fun(X35,fun(X35,bool)),X36),s(X35,X38))),s(X35,X37))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.relation.symmetric_1'(s(fun(X35,fun(X35,bool)),X36)))) )),
  inference(resolution,[],[f118,f88])).
fof(f1107,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,happ(s(fun(X21,bool),happ(s(fun(X21,fun(X21,bool)),X22),s(X21,X23))),s(X21,X24))) = s(bool,happ(s(fun(X21,bool),happ(s(fun(X21,fun(X21,bool)),X22),s(X21,X24))),s(X21,X23))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.relation.symmetric_1'(s(fun(X21,fun(X21,bool)),X22)))) )),
  inference(resolution,[],[f118,f87])).
fof(f1106,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,X16))),s(X14,X17))) = s(bool,happ(s(fun(X14,bool),happ(s(fun(X14,fun(X14,bool)),X15),s(X14,X17))),s(X14,X16))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.symmetric_1'(s(fun(X14,fun(X14,bool)),X15))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f118,f85])).
fof(f1104,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) )),
  inference(resolution,[],[f118,f83])).
fof(f1021,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f1020,f855])).
fof(f1020,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK4(X1,X0)))),s(X0,sK5(X1,X0))))) ) | $spl44),
  inference(forward_demodulation,[],[f1019,f855])).
fof(f1019,plain,(
  ( ! [X0,X1] : (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK4(X1,X0)))),s(X0,sK5(X1,X0)))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK5(X1,X0)))),s(X0,sK4(X1,X0))))) ) | $spl44),
  inference(subsumption_resolution,[],[f861,f115])).
fof(f861,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK4(X1,X0)))),s(X0,sK5(X1,X0)))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK5(X1,X0)))),s(X0,sK4(X1,X0))))) ) | $spl44),
  inference(backward_demodulation,[],[f855,f119])).
fof(f857,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) != s(bool,X2)) ) | $spl48),
  inference(cnf_transformation,[],[f857_D])).
fof(f857_D,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) != s(bool,X2)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f853,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X2)) ) | ($spl34 | $spl38)),
  inference(duplicate_literal_removal,[],[f842])).
fof(f842,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | ($spl34 | $spl38)),
  inference(superposition,[],[f789,f122])).
fof(f856,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,'T_0') = s(bool,X13)) ) | $spl46),
  inference(cnf_transformation,[],[f856_D])).
fof(f856_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,'T_0') = s(bool,X13)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f854,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X13)) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f836,f122])).
fof(f836,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X12) | s(bool,X13) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X13)) ) | ($spl34 | $spl38)),
  inference(superposition,[],[f789,f151])).
fof(f808,plain,(
  ( ! [X70,X71] : (s(bool,'F_0') != s(bool,X70) | s(bool,'F_0') = s(bool,X71) | s(bool,X70) != s(bool,X71)) ) | $spl42),
  inference(cnf_transformation,[],[f808_D])).
fof(f808_D,plain,(
  ( ! [X70,X71] : (s(bool,'F_0') != s(bool,X70) | s(bool,'F_0') = s(bool,X71) | s(bool,X70) != s(bool,X71)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f810,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | ($spl24 | $spl38)),
  inference(subsumption_resolution,[],[f798,f385])).
fof(f798,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X9)) ) | $spl38),
  inference(superposition,[],[f745,f122])).
fof(f793,plain,(
  ( ! [X74,X75,X73] : (s(bool,'F_0') != s(bool,X73) | s(bool,X73) != s(bool,X74) | s(bool,'F_0') = s(bool,X74) | ~p(s(bool,X75))) ) | $spl38),
  inference(resolution,[],[f745,f101])).
fof(f807,plain,(
  ( ! [X70,X72,X71] : (s(bool,'F_0') != s(bool,X70) | s(bool,X70) != s(bool,X71) | s(bool,'F_0') = s(bool,X71) | p(s(bool,X72))) ) | $spl38),
  inference(subsumption_resolution,[],[f792,f126])).
fof(f792,plain,(
  ( ! [X70,X72,X71] : (s(bool,'F_0') != s(bool,X70) | s(bool,X70) != s(bool,X71) | s(bool,'F_0') = s(bool,X71) | p(s(bool,X72)) | ~p(s(bool,'T_0'))) ) | $spl38),
  inference(resolution,[],[f745,f105])).
fof(f788,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,X57) != s(bool,X58) | s(bool,X57) != s(bool,X59) | s(bool,'F_0') = s(bool,X59) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X58) != s(bool,X61)) ) | $spl38),
  inference(resolution,[],[f745,f125])).
fof(f784,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl38),
  inference(resolution,[],[f745,f89])).
fof(f781,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl38),
  inference(resolution,[],[f745,f88])).
fof(f780,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl38),
  inference(resolution,[],[f745,f87])).
fof(f779,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl38),
  inference(resolution,[],[f745,f85])).
fof(f777,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl38),
  inference(resolution,[],[f745,f83])).
fof(f728,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X46) != s(bool,X47) | s(bool,'F_0') = s(bool,X46) | p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X47) != s(bool,X49)) ) | $spl6),
  inference(resolution,[],[f504,f125])).
fof(f724,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl6),
  inference(resolution,[],[f504,f89])).
fof(f721,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl6),
  inference(resolution,[],[f504,f88])).
fof(f720,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl6),
  inference(resolution,[],[f504,f87])).
fof(f719,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl6),
  inference(resolution,[],[f504,f85])).
fof(f717,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl6),
  inference(resolution,[],[f504,f83])).
fof(f560,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl36),
  inference(cnf_transformation,[],[f560_D])).
fof(f560_D,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f536,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38)) ) | $spl34),
  inference(resolution,[],[f503,f125])).
fof(f531,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl34),
  inference(resolution,[],[f503,f89])).
fof(f529,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl34),
  inference(resolution,[],[f503,f88])).
fof(f527,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl34),
  inference(resolution,[],[f503,f87])).
fof(f526,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl34),
  inference(resolution,[],[f503,f85])).
fof(f524,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl34),
  inference(resolution,[],[f503,f83])).
fof(f503,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(cnf_transformation,[],[f503_D])).
fof(f503_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f502,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f485,f432])).
fof(f485,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X3)) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f147,f122])).
fof(f495,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X2)) | ~p(s(bool,X3))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f476,f385])).
fof(f476,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f147,f122])).
fof(f453,plain,(
  ( ! [X37,X35,X36] : (s(bool,'F_0') = s(bool,X35) | p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X35) != s(bool,X37)) ) | $spl24),
  inference(resolution,[],[f432,f125])).
fof(f449,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl24),
  inference(resolution,[],[f432,f89])).
fof(f147,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X5)) | ~p(s(bool,X6))) ) | $spl6),
  inference(cnf_transformation,[],[f147_D])).
fof(f147_D,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X5)) | ~p(s(bool,X6))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f448,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl24),
  inference(resolution,[],[f432,f89])).
fof(f446,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl24),
  inference(resolution,[],[f432,f88])).
fof(f445,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl24),
  inference(resolution,[],[f432,f87])).
fof(f444,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl24),
  inference(resolution,[],[f432,f85])).
fof(f443,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl24),
  inference(resolution,[],[f432,f85])).
fof(f442,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl24),
  inference(resolution,[],[f432,f83])).
fof(f435,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl27),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f433,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X1))) ) | $spl24),
  inference(subsumption_resolution,[],[f421,f385])).
fof(f421,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(superposition,[],[f385,f122])).
fof(f385,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X5))) ) | $spl24),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X5))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f417,plain,(
  ( ! [X112,X113] : (s(bool,'T_0') != s(bool,X113) | p(s(bool,X112)) | p(s(bool,X113))) )),
  inference(forward_demodulation,[],[f416,f151])).
fof(f416,plain,(
  ( ! [X112,X113] : (p(s(bool,X112)) | p(s(bool,X113)) | s(bool,X113) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(subsumption_resolution,[],[f369,f126])).
fof(f369,plain,(
  ( ! [X112,X113] : (~p(s(bool,'T_0')) | p(s(bool,X112)) | p(s(bool,X113)) | s(bool,X113) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(superposition,[],[f125,f151])).
fof(f413,plain,(
  ( ! [X103,X101,X102] : (s(bool,'T_0') != s(bool,X101) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X102) != s(bool,X103)) )),
  inference(forward_demodulation,[],[f412,f151])).
fof(f412,plain,(
  ( ! [X103,X101,X102] : (p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,X102) != s(bool,X103)) )),
  inference(subsumption_resolution,[],[f356,f126])).
fof(f356,plain,(
  ( ! [X103,X101,X102] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,X102) != s(bool,X103)) )),
  inference(superposition,[],[f89,f151])).
fof(f411,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') != s(bool,X100) | p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,X99)) )),
  inference(forward_demodulation,[],[f410,f151])).
fof(f410,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,X99) | s(bool,X100) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(subsumption_resolution,[],[f355,f125])).
fof(f355,plain,(
  ( ! [X99,X100,X98] : (~p(s(bool,'T_0')) | p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,X99) | s(bool,X100) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(superposition,[],[f89,f151])).
fof(f407,plain,(
  ( ! [X85,X83,X84] : (s(bool,'T_0') != s(bool,X83) | p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X84) != s(bool,X85)) )),
  inference(subsumption_resolution,[],[f350,f125])).
fof(f350,plain,(
  ( ! [X85,X83,X84] : (s(bool,'T_0') != s(bool,X83) | p(s(bool,X84)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) | ~p(s(bool,X85)) | s(bool,X84) != s(bool,X85)) )),
  inference(superposition,[],[f89,f151])).
fof(f406,plain,(
  ( ! [X78,X79,X77] : (s(bool,'T_0') != s(bool,X79) | p(s(bool,X77)) | ~p(s(bool,X78)) | p(s(bool,X79))) )),
  inference(forward_demodulation,[],[f405,f151])).
fof(f405,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | ~p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X79) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(subsumption_resolution,[],[f348,f126])).
fof(f348,plain,(
  ( ! [X78,X79,X77] : (~p(s(bool,'T_0')) | p(s(bool,X77)) | ~p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X79) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(superposition,[],[f88,f151])).
fof(f402,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') != s(bool,X63) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64)) )),
  inference(forward_demodulation,[],[f401,f151])).
fof(f401,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(subsumption_resolution,[],[f343,f126])).
fof(f343,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,'T_0')) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(superposition,[],[f87,f151])).
fof(f400,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X54)) | p(s(bool,X55)) | s(bool,X54) != s(bool,X55)) )),
  inference(subsumption_resolution,[],[f340,f89])).
fof(f340,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X53)) | ~p(s(bool,X54)) | p(s(bool,X55)) | s(bool,X54) != s(bool,X55)) )),
  inference(superposition,[],[f87,f151])).
fof(f384,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7)) ) | $spl22),
  inference(cnf_transformation,[],[f384_D])).
fof(f384_D,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f398,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') != s(bool,X44) | p(s(bool,X44)) | s(bool,X45) != s(bool,X46)) )),
  inference(subsumption_resolution,[],[f397,f126])).
fof(f397,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X44) | p(s(bool,X44)) | s(bool,X45) != s(bool,X46)) )),
  inference(subsumption_resolution,[],[f396,f89])).
fof(f396,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X44) | p(s(bool,X45)) | p(s(bool,X44)) | s(bool,X45) != s(bool,X46)) )),
  inference(forward_demodulation,[],[f337,f151])).
fof(f337,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') != s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) | p(s(bool,X44)) | s(bool,X45) != s(bool,X46)) )),
  inference(superposition,[],[f87,f151])).
fof(f395,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,X41) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,X43)) )),
  inference(forward_demodulation,[],[f394,f151])).
fof(f394,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,X42) != s(bool,X43)) )),
  inference(subsumption_resolution,[],[f336,f126])).
fof(f336,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'T_0')) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,X42) != s(bool,X43)) )),
  inference(superposition,[],[f85,f151])).
fof(f393,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') != s(bool,X40) | p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X39)) )),
  inference(forward_demodulation,[],[f392,f151])).
fof(f392,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X39) | s(bool,X40) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(subsumption_resolution,[],[f335,f125])).
fof(f335,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,'T_0')) | p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X39) | s(bool,X40) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(superposition,[],[f85,f151])).
fof(f389,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) )),
  inference(subsumption_resolution,[],[f331,f125])).
fof(f331,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) )),
  inference(superposition,[],[f85,f151])).
fof(f388,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22)) )),
  inference(forward_demodulation,[],[f387,f151])).
fof(f387,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(subsumption_resolution,[],[f329,f126])).
fof(f329,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'T_0')) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(superposition,[],[f83,f151])).
fof(f386,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X10)) )),
  inference(subsumption_resolution,[],[f325,f89])).
fof(f325,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,X8)) | ~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X10)) )),
  inference(superposition,[],[f83,f151])).
fof(f383,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X5)) | s(bool,X6) != s(bool,X7)) )),
  inference(subsumption_resolution,[],[f382,f126])).
fof(f382,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X5)) | s(bool,X6) != s(bool,X7)) )),
  inference(subsumption_resolution,[],[f381,f89])).
fof(f381,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6)) | p(s(bool,X5)) | s(bool,X6) != s(bool,X7)) )),
  inference(forward_demodulation,[],[f324,f151])).
fof(f324,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) | p(s(bool,X5)) | s(bool,X6) != s(bool,X7)) )),
  inference(superposition,[],[f83,f151])).
fof(f314,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f146])).
fof(f303,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f284,f149])).
fof(f284,plain,(
  ( ! [X39] : (p(s(bool,X39))) ) | $spl10),
  inference(resolution,[],[f149,f112])).
fof(f112,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.relation.symmetric_inv_image','thm.bool.IMP_CLAUSES|split|3')).
fof(f302,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f283,f149])).
fof(f283,plain,(
  ( ! [X38] : (p(s(bool,X38))) ) | $spl10),
  inference(resolution,[],[f149,f109])).
fof(f109,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.relation.symmetric_inv_image','thm.bool.IMP_CLAUSES|split|1')).
fof(f301,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f282,f149])).
fof(f282,plain,(
  ( ! [X37] : (p(s(bool,X37))) ) | $spl10),
  inference(resolution,[],[f149,f106])).
fof(f106,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10] : ((p(s(bool,'T_0')) => p(s(bool,X10))) <=> p(s(bool,X10)))),
  file('thm.relation.symmetric_inv_image','thm.bool.IMP_CLAUSES|split|0')).
fof(f300,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f280,f149])).
fof(f280,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f149,f103])).
fof(f103,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f67])).
fof(f286,plain,(
  $false | $spl10),
  inference(resolution,[],[f149,f116])).
fof(f116,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.symmetric_inv_image','thm.bool.TRUTH')).
fof(f287,plain,(
  $false | $spl10),
  inference(resolution,[],[f149,f117])).
fof(f288,plain,(
  $false | $spl10),
  inference(resolution,[],[f149,f126])).
fof(f291,plain,(
  $false | $spl10),
  inference(resolution,[],[f149,f79])).
fof(f262,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X2))) ) | $spl8),
  inference(subsumption_resolution,[],[f245,f148])).
fof(f245,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl8),
  inference(superposition,[],[f148,f122])).
fof(f261,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f244,f148])).
fof(f244,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f148,f122])).
fof(f241,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | s(bool,'F_0') != s(bool,X7)) ) | $spl2),
  inference(subsumption_resolution,[],[f230,f145])).
fof(f230,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') = s(bool,X6)) ) | $spl2),
  inference(superposition,[],[f145,f122])).
fof(f239,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5)) ) | $spl2),
  inference(subsumption_resolution,[],[f229,f145])).
fof(f229,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | $spl2),
  inference(superposition,[],[f145,f122])).
fof(f224,plain,(
  ( ! [X56,X55] : (s(bool,'F_0') != s(bool,X55) | ~p(s(bool,X56))) ) | $spl2),
  inference(resolution,[],[f145,f101])).
fof(f236,plain,(
  ( ! [X54,X53] : (s(bool,'F_0') != s(bool,X53) | p(s(bool,X54))) ) | $spl2),
  inference(subsumption_resolution,[],[f223,f126])).
fof(f223,plain,(
  ( ! [X54,X53] : (s(bool,'F_0') != s(bool,X53) | p(s(bool,X54)) | ~p(s(bool,'T_0'))) ) | $spl2),
  inference(resolution,[],[f145,f105])).
fof(f208,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f196,f150])).
fof(f196,plain,(
  ( ! [X36] : (~p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f150,f101])).
fof(f194,plain,(
  $false | $spl12),
  inference(resolution,[],[f150,f115])).
fof(f198,plain,(
  $false | $spl12),
  inference(resolution,[],[f150,f80])).
fof(f177,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))) ) | $spl14),
  inference(cnf_transformation,[],[f177_D])).
fof(f177_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f165,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))) )),
  inference(resolution,[],[f80,f125])).
fof(f160,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f80,f89])).
fof(f158,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))) )),
  inference(resolution,[],[f80,f88])).
fof(f157,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f80,f87])).
fof(f156,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))) )),
  inference(resolution,[],[f80,f87])).
fof(f155,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f80,f85])).
fof(f154,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3))))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f80,f83])).
fof(f153,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_image_2'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK0,sK1),sK3)))))) )),
  inference(resolution,[],[f80,f83])).
fof(f144,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) ) | $spl0),
  inference(cnf_transformation,[],[f144_D])).
fof(f144_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f139,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(resolution,[],[f79,f125])).
fof(f135,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,X24) != s(bool,X25)) )),
  inference(resolution,[],[f79,f89])).
fof(f146,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) ) | $spl4),
  inference(cnf_transformation,[],[f146_D])).
fof(f146_D,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f134,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X21) | s(bool,X22) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(resolution,[],[f79,f89])).
fof(f132,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(resolution,[],[f79,f88])).
fof(f131,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(resolution,[],[f79,f87])).
fof(f130,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2))) | s(bool,X9) != s(bool,X10)) )),
  inference(resolution,[],[f79,f85])).
fof(f129,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X7) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(resolution,[],[f79,f85])).
fof(f128,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.relation.symmetric_1'(s(fun(sK1,fun(sK1,bool)),sK2)))) )),
  inference(resolution,[],[f79,f83])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 279
% Passive clauses: 8573
% Generated clauses: 23632
% Final active clauses: 209
% Final passive clauses: 7819
% Input formulas: 19
% Initial clauses: 46
% 
% Selected by SInE selection: 19
% SInE iterations: 1
% 
% Duplicate literals: 145
% Trivial inequalities: 250
% Fw subsumption resolutions: 1480
% Fw demodulations: 4154
% Bw demodulations: 47
% 
% Simple tautologies: 238
% Equational tautologies: 90
% Forward subsumptions: 7562
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 2680
% Forward superposition: 7489
% Backward superposition: 6929
% Self superposition: 100
% Equality factoring: 4
% Equality resolution: 39
% 
% Split clauses: 1378
% Split components: 2965
% SAT solver clauses: 1401
% SAT solver unit clauses: 16
% SAT solver binary clauses: 637
% 
% Sat splits: 1378
% Sat splitting refutations: 22
% 
% Memory used [KB]: 15607
% Time elapsed: 1.457 s
% ------------------------------
% ------------------------------
% Success in time 13.155 s
