Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (64966)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 640
% Passive clauses: 10155
% Generated clauses: 18603
% Final active clauses: 572
% Final passive clauses: 8996
% Input formulas: 40
% Initial clauses: 178
% 
% Duplicate literals: 278
% Trivial inequalities: 95
% Fw subsumption resolutions: 1620
% Fw demodulations: 117
% Bw demodulations: 67
% 
% Simple tautologies: 206
% Equational tautologies: 65
% Forward subsumptions: 5867
% 
% Binary resolution: 10923
% Factoring: 984
% Forward superposition: 1646
% Backward superposition: 2443
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 31
% 
% Split clauses: 286
% Split components: 574
% SAT solver clauses: 603
% SAT solver unit clauses: 18
% SAT solver binary clauses: 135
% 
% Sat splits: 286
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 601
% TWLsolver calls for satisfiability: 349
% 
% Memory used [KB]: 8699
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (64978)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 384
% Passive clauses: 2609
% Generated clauses: 14504
% Final active clauses: 292
% Final passive clauses: 1341
% Input formulas: 40
% Initial clauses: 178
% 
% Duplicate literals: 335
% Trivial inequalities: 63
% Fw subsumption resolutions: 1160
% Fw demodulations: 1309
% Bw demodulations: 906
% Condensations: 373
% Global subsumptions: 13
% 
% Simple tautologies: 281
% Equational tautologies: 20
% Forward subsumptions: 4875
% 
% Binary resolution: 2938
% Factoring: 146
% Forward superposition: 1237
% Backward superposition: 2077
% Self superposition: 32
% Equality factoring: 35
% Equality resolution: 12
% 
% Split clauses: 596
% Split components: 1202
% SAT solver clauses: 3319
% SAT solver unit clauses: 180
% SAT solver binary clauses: 101
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 11
% 
% Sat splits: 596
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 3266
% TWLsolver calls for satisfiability: 1763
% 
% Memory used [KB]: 11769
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f6398,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f236,f426,f424,f422,f235,f425,f234,f423,f421,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2378,f2284,f2380,f2288,f2382,f2291,f2385,f2295,f2387,f2298,f2389,f2301,f2391,f2304,f2393,f2307,f2395,f2310,f2397,f2313,f2399,f2316,f2401,f2319,f2403,f2322,f2405,f2325,f2407,f2328,f2409,f2331,f2411,f2334,f2413,f2285,f2337,f2514,f2515,f2518,f2525,f2527,f2529,f2531,f2533,f2535,f2537,f2539,f2541,f2543,f2545,f2547,f2549,f2551,f2553,f2555,f2557,f2559,f2561,f2563,f4413,f4414,f4421,f4418,f4415,f4416,f4424,f4417,f4426,f4419,f4423,f4611,f4612,f4624,f4622,f5681,f6384,f6389,f6396,f6397])).
fof(f6397,plain,(
  $false | ($spl0 | $spl5)),
  inference(global_subsumption,[],[f426,f236,f235,f234,f237,f244,f254,f253,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1311,f1310,f1763,f1778,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2029,f2028,f2027,f2026,f2025,f2024,f1867,f1866,f1865,f1864,f1863,f1862,f2372,f2371,f2370,f2369,f2368,f2367,f2366,f2365,f2364,f2363,f292,f291,f290,f289,f288,f287,f286,f285,f281,f280,f2362,f2361,f2360,f1728,f1540,f1405,f1393,f1341,f1340,f1339,f1338,f1337,f1336,f1335,f1334,f1333,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1325,f1324,f1323,f1322,f1321,f1319,f1316,f1315,f295,f294,f293,f282,f1861,f2023,f1460,f1448,f1430,f1412,f1381,f1380,f1318,f1317,f1312,f1294,f1282,f427,f258,f257,f3660,f3699,f3724,f3898,f3949,f3975,f4125,f4140,f4285,f4288,f4289,f4290,f4291,f4292,f4293,f4294,f4295,f4296,f4297,f4298,f4299,f4300,f4301,f4302,f4303,f4304,f4305,f4306,f4307,f4308,f4309,f4310,f4311,f4312,f4313,f4314,f4315,f4316,f4368,f4350,f4399,f4400,f4413,f4414,f4415,f4416,f4417,f4501,f4502,f4547,f4548,f4573,f4574,f4593,f4594,f4611,f4612,f4626,f4734,f4732,f4735,f4772,f4770,f4773,f4877,f4879,f4881,f4887,f4908,f4910,f4912,f4918,f4967,f4952,f4968,f421,f6295,f6294,f6291,f6290,f6289,f6288,f6285,f6282,f6279,f6278,f6255,f6253,f6251,f6250,f6249,f6248,f6243,f6236,f6235,f6231,f6230,f6207,f6206,f6205,f6200,f6195,f6194,f6189,f6180,f6175,f6173,f6172,f6168,f6167,f6166,f6165,f6162,f6160,f6155,f6154,f6150,f6149,f6133,f6132,f6131,f6128,f6121,f6120,f6119,f6118,f6114,f6113,f6107,f6103,f6102,f6093,f6091,f6090,f6087,f6086,f6085,f6084,f6082,f6081,f6080,f6078,f6072,f6070,f6068,f6067,f6065,f6057,f6056,f6047,f6044,f6039,f6034,f6029,f6025,f6022,f6021,f6020,f6019,f6018,f6016,f6386,f6011,f6010])).
fof(f6010,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f1405])).
fof(f6011,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f1393])).
fof(f6386,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15)))))) | $spl0),
  inference(forward_demodulation,[],[f6013,f282])).
fof(f6013,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),sK15))),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f421,f295])).
fof(f6016,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f237,f421,f1321])).
fof(f6018,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f1322])).
fof(f6019,plain,(
  ( ! [X0] : (sP0('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f1323])).
fof(f6020,plain,(
  sP1('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f1324])).
fof(f6021,plain,(
  sP1('F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f1325])).
fof(f6022,plain,(
  ~sP5('T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f1326])).
fof(f6025,plain,(
  sP5('T_0','T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f448,f440,f435,f421,f1327])).
fof(f6029,plain,(
  ~sP4('T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f1328])).
fof(f6034,plain,(
  ~sP3('T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f1329])).
fof(f6039,plain,(
  ~sP6('T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f1330])).
fof(f6044,plain,(
  ~sP7('F_0','T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f1331])).
fof(f6047,plain,(
  ~sP8('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f427,f427,f421,f1332])).
fof(f6056,plain,(
  ~sP9('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f427,f427,f421,f1335])).
fof(f6057,plain,(
  ~sP10('T_0','F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f237,f421,f1336])).
fof(f6065,plain,(
  ~sP11('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f1339])).
fof(f6067,plain,(
  ( ! [X0] : (sP11('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f466,f421,f1340])).
fof(f6068,plain,(
  ( ! [X0] : (sP11('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f467,f421,f1340])).
fof(f6070,plain,(
  ( ! [X0] : (sP9('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0,'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f459,f421,f1341])).
fof(f6072,plain,(
  ( ! [X0] : (sP9('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f462,f421,f1341])).
fof(f6078,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f237,f421,f299])).
fof(f6080,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f305])).
fof(f6081,plain,(
  ( ! [X0] : (sP0('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f306])).
fof(f6082,plain,(
  ~sP1('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f421,f320])).
fof(f6084,plain,(
  ~sP1('T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f237,f421,f320])).
fof(f6085,plain,(
  ~sP1('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f421,f320])).
fof(f6086,plain,(
  sP1('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f321])).
fof(f6087,plain,(
  sP1('F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f324])).
fof(f6090,plain,(
  ~sP5('T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f341])).
fof(f6091,plain,(
  ~sP5('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f421,f427,f421,f341])).
fof(f6093,plain,(
  sP5('T_0','T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f448,f440,f435,f421,f342])).
fof(f6102,plain,(
  ~sP4('T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f345])).
fof(f6103,plain,(
  ~sP4('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f421,f427,f421,f345])).
fof(f6107,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f346])).
fof(f6113,plain,(
  ~sP3('T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f349])).
fof(f6114,plain,(
  ~sP3('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f421,f427,f421,f349])).
fof(f6118,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f350])).
fof(f6119,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f354])).
fof(f6120,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f355])).
fof(f6121,plain,(
  ( ! [X0,X1] : (sP2('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f356])).
fof(f6128,plain,(
  ( ! [X2,X0,X1] : (sP5('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)),'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f282,f421,f358])).
fof(f6131,plain,(
  ~sP5('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f471,f421,f359])).
fof(f6132,plain,(
  ( ! [X0] : (~sP7('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f362])).
fof(f6133,plain,(
  ( ! [X0] : (~sP7('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f363])).
fof(f6149,plain,(
  ~sP6('T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f368])).
fof(f6150,plain,(
  ~sP6('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f421,f427,f421,f368])).
fof(f6154,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f369])).
fof(f6155,plain,(
  sP7('T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f237,f421,f374])).
fof(f6160,plain,(
  sP7('T_0','T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f237,f237,f421,f374])).
fof(f6162,plain,(
  sP7('T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f237,f421,f421,f374])).
fof(f6165,plain,(
  sP7('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'T_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f237,f421,f374])).
fof(f6166,plain,(
  sP7('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f421,f237,f421,f374])).
fof(f6167,plain,(
  sP7('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f237,f421,f421,f374])).
fof(f6168,plain,(
  sP7('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f421,f421,f374])).
fof(f6172,plain,(
  ~sP7('F_0','T_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f377])).
fof(f6173,plain,(
  ~sP7('F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f427,f421,f377])).
fof(f6175,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f378])).
fof(f6180,plain,(
  ( ! [X0] : (~sP9('F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f379])).
fof(f6189,plain,(
  ~sP8('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f427,f427,f421,f385])).
fof(f6194,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f386])).
fof(f6195,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f387])).
fof(f6200,plain,(
  ( ! [X0] : (sP9('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f421,f390])).
fof(f6205,plain,(
  ( ! [X0] : (sP9('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f421,f391])).
fof(f6206,plain,(
  ~sP9('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f427,f427,f421,f392])).
fof(f6207,plain,(
  ( ! [X0] : (~sP11('F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f396])).
fof(f6230,plain,(
  ~sP10('T_0','F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f237,f421,f402])).
fof(f6231,plain,(
  ~sP10('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f421,f402])).
fof(f6235,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f404])).
fof(f6236,plain,(
  sP11('F_0','F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f427,f421,f406])).
fof(f6243,plain,(
  ( ! [X0] : (sP11('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f421,f408])).
fof(f6248,plain,(
  ~sP11('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f237,f427,f421,f409])).
fof(f6249,plain,(
  ~sP11('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f427,f421,f409])).
fof(f6250,plain,(
  ( ! [X0] : (sP11('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f466,f421,f412])).
fof(f6251,plain,(
  ( ! [X0] : (sP11('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f467,f421,f412])).
fof(f6253,plain,(
  ( ! [X0] : (sP9('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0,'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f459,f421,f413])).
fof(f6255,plain,(
  ( ! [X0] : (sP9('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f462,f421,f413])).
fof(f6278,plain,(
  ~sP5('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f421,f1326])).
fof(f6279,plain,(
  ~sP4('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f421,f1328])).
fof(f6282,plain,(
  ~sP3('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f421,f1329])).
fof(f6285,plain,(
  ~sP6('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f421,f1330])).
fof(f6288,plain,(
  ~sP7('F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f421,f1331])).
fof(f6289,plain,(
  ( ! [X0] : (sP9('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),X0,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f421,f1333])).
fof(f6290,plain,(
  ( ! [X0] : (sP9('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f421,f1334])).
fof(f6291,plain,(
  ~sP10('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),'F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15))) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f421,f1336])).
fof(f6294,plain,(
  ( ! [X0] : (sP11('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f421,f1338])).
fof(f6295,plain,(
  ~sP11('const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))),s(fun(sK12,bool),sK15)),'F_0','const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f427,f421,f421,f1339])).
fof(f4968,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X21,bool),X23),s(fun(X21,bool),X24)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X21,bool),'const.pred_set.UNION_2'(s(fun(X21,bool),X24),s(fun(X21,bool),X22))),s(fun(X21,bool),X23)))) | s(bool,X25) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X21,bool),X22),s(fun(X21,bool),X23))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f4954,f237])).
fof(f4954,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X21,bool),X23),s(fun(X21,bool),X24)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X21,bool),'const.pred_set.UNION_2'(s(fun(X21,bool),X24),s(fun(X21,bool),X22))),s(fun(X21,bool),X23)))) | s(bool,X25) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X21,bool),X22),s(fun(X21,bool),X23))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f1540,f716])).
fof(f4952,plain,(
  ( ! [X14,X15,X13,X16] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X13,bool),X15),s(fun(X13,bool),X14)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X13,bool),X15),s(fun(X13,bool),X16)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X13,bool),'const.pred_set.UNION_2'(s(fun(X13,bool),X16),s(fun(X13,bool),X14))),s(fun(X13,bool),X15))))) )),
  inference(superposition,[],[f1540,f282])).
fof(f4967,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X5,bool),X7),s(fun(X5,bool),X8)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X5,bool),'const.pred_set.UNION_2'(s(fun(X5,bool),X8),s(fun(X5,bool),X6))),s(fun(X5,bool),X7)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X5,bool),X7),s(fun(X5,bool),X6)))) )),
  inference(subsumption_resolution,[],[f4950,f237])).
fof(f4950,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X5,bool),X7),s(fun(X5,bool),X8)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X5,bool),'const.pred_set.UNION_2'(s(fun(X5,bool),X8),s(fun(X5,bool),X6))),s(fun(X5,bool),X7)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X5,bool),X7),s(fun(X5,bool),X6)))) )),
  inference(superposition,[],[f1540,f1312])).
fof(f4918,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X1,X0,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK19(X1,X0,bool)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f4913,f237])).
fof(f4913,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X1,X0,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK19(X1,X0,bool)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1460,f716])).
fof(f4912,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X18)))) | ~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X16,X17,bool)) | s(bool,'T_0') = s(bool,sK19(X16,X17,bool))) )),
  inference(superposition,[],[f1460,f244])).
fof(f4910,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X11)))) | ~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X8,X9,bool)) | s(bool,'T_0') = s(bool,sK19(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f1460,f716])).
fof(f4908,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X0,X1,bool)) | s(bool,X3) = s(bool,sK19(X0,X1,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1460,f716])).
fof(f4887,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X1,X0,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK19(X1,X0,bool)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f4882,f237])).
fof(f4882,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X1,X0,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK19(X1,X0,bool)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1448,f716])).
fof(f4881,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | ~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X16,X17,bool)) | s(bool,'T_0') = s(bool,sK19(X16,X17,bool))) )),
  inference(superposition,[],[f1448,f244])).
fof(f4879,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X11)))) | ~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X8,X9,bool)) | s(bool,'T_0') = s(bool,sK19(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f1448,f716])).
fof(f4877,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X0,X1,bool)) | s(bool,X3) = s(bool,sK19(X0,X1,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1448,f716])).
fof(f4773,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X1,X0,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK18(X1,X0,bool)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f4765,f237])).
fof(f4765,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X1,X0,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK18(X1,X0,bool)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1430,f716])).
fof(f4770,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X18)))) | ~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X16,X17,bool))) )),
  inference(duplicate_literal_removal,[],[f4764])).
fof(f4764,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X18)))) | ~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X16,X17,bool)) | s(bool,'T_0') = s(bool,sK18(X16,X17,bool))) )),
  inference(superposition,[],[f1430,f244])).
fof(f4772,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X11)))) | ~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(duplicate_literal_removal,[],[f4762])).
fof(f4762,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X11)))) | ~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X8,X9,bool)) | s(bool,'T_0') = s(bool,sK18(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f1430,f716])).
fof(f4735,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X1,X0,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK18(X1,X0,bool)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f4727,f237])).
fof(f4727,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X1,X0,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK18(X1,X0,bool)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1412,f716])).
fof(f4732,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | ~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X16,X17,bool))) )),
  inference(duplicate_literal_removal,[],[f4726])).
fof(f4726,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | ~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X16,X17,bool)) | s(bool,'T_0') = s(bool,sK18(X16,X17,bool))) )),
  inference(superposition,[],[f1412,f244])).
fof(f4734,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X11)))) | ~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(duplicate_literal_removal,[],[f4724])).
fof(f4724,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X11)))) | ~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X8,X9,bool)) | s(bool,'T_0') = s(bool,sK18(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f1412,f716])).
fof(f4626,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK24(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK24(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f4625])).
fof(f4625,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK24(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK24(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f4618,f244])).
fof(f4618,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK24(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK24(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1381,f716])).
fof(f4594,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X16,bool),X19),s(fun(X16,bool),X17)))) | s(bool,X20) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X16,bool),X17),s(fun(X16,bool),'const.pred_set.UNION_2'(s(fun(X16,bool),X18),s(fun(X16,bool),X19))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f4588,f237])).
fof(f4588,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X16,bool),X19),s(fun(X16,bool),X17)))) | s(bool,X20) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X16,bool),X17),s(fun(X16,bool),'const.pred_set.UNION_2'(s(fun(X16,bool),X18),s(fun(X16,bool),X19))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f1405,f716])).
fof(f4593,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f4584,f237])).
fof(f4584,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X1)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f1405,f1312])).
fof(f4574,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X16,bool),X18),s(fun(X16,bool),X17)))) | s(bool,X20) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X16,bool),X17),s(fun(X16,bool),'const.pred_set.UNION_2'(s(fun(X16,bool),X18),s(fun(X16,bool),X19))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f4568,f237])).
fof(f4568,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X16,bool),X18),s(fun(X16,bool),X17)))) | s(bool,X20) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X16,bool),X17),s(fun(X16,bool),'const.pred_set.UNION_2'(s(fun(X16,bool),X18),s(fun(X16,bool),X19))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f1393,f716])).
fof(f4573,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f4564,f237])).
fof(f4564,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f1393,f1312])).
fof(f4548,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP11('const.pred_set.DISJOINT_2'(s(fun(X30,bool),X32),s(fun(X30,bool),X31)),X34,X33) | s(bool,X35) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X30,bool),X31),s(fun(X30,bool),X32))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f4541,f237])).
fof(f4541,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP11('const.pred_set.DISJOINT_2'(s(fun(X30,bool),X32),s(fun(X30,bool),X31)),X34,X33) | s(bool,X35) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X30,bool),X31),s(fun(X30,bool),X32))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f1339,f716])).
fof(f4547,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,X13)) | p(s(bool,X14)) | ~sP11('const.pred_set.DISJOINT_2'(s(fun(X10,bool),X12),s(fun(X10,bool),X11)),X14,X13) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X10,bool),X12),s(fun(X10,bool),X11)))) )),
  inference(subsumption_resolution,[],[f4537,f237])).
fof(f4537,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | p(s(bool,X14)) | ~sP11('const.pred_set.DISJOINT_2'(s(fun(X10,bool),X12),s(fun(X10,bool),X11)),X14,X13) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X10,bool),X12),s(fun(X10,bool),X11)))) )),
  inference(superposition,[],[f1339,f1312])).
fof(f4502,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,X33)) | p(s(bool,X34)) | ~sP9('const.pred_set.DISJOINT_2'(s(fun(X30,bool),X32),s(fun(X30,bool),X31)),X34,X33) | s(bool,X35) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X30,bool),X31),s(fun(X30,bool),X32))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f4495,f237])).
fof(f4495,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X33)) | p(s(bool,X34)) | ~sP9('const.pred_set.DISJOINT_2'(s(fun(X30,bool),X32),s(fun(X30,bool),X31)),X34,X33) | s(bool,X35) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X30,bool),X31),s(fun(X30,bool),X32))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f1335,f716])).
fof(f4501,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X13)) | p(s(bool,X14)) | ~sP9('const.pred_set.DISJOINT_2'(s(fun(X10,bool),X12),s(fun(X10,bool),X11)),X14,X13) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X10,bool),X12),s(fun(X10,bool),X11)))) )),
  inference(subsumption_resolution,[],[f4491,f237])).
fof(f4491,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'T_0')) | p(s(bool,X13)) | p(s(bool,X14)) | ~sP9('const.pred_set.DISJOINT_2'(s(fun(X10,bool),X12),s(fun(X10,bool),X11)),X14,X13) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X10,bool),X12),s(fun(X10,bool),X11)))) )),
  inference(superposition,[],[f1335,f1312])).
fof(f4400,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,X24)) | s(bool,X24) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X21,bool),X23),s(fun(X21,bool),X22))) | s(bool,X25) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X21,bool),X22),s(fun(X21,bool),X23))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f4385,f237])).
fof(f4385,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | s(bool,X24) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X21,bool),X23),s(fun(X21,bool),X22))) | s(bool,X25) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X21,bool),X22),s(fun(X21,bool),X23))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f1321,f716])).
fof(f4399,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,X8)) | s(bool,X8) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X5,bool),X7),s(fun(X5,bool),X6))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X5,bool),X7),s(fun(X5,bool),X6)))) )),
  inference(subsumption_resolution,[],[f4381,f237])).
fof(f4381,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | s(bool,X8) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X5,bool),X7),s(fun(X5,bool),X6))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X5,bool),X7),s(fun(X5,bool),X6)))) )),
  inference(superposition,[],[f1321,f1312])).
fof(f4350,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X65,bool),X67),s(fun(X65,bool),X66)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X68,bool),X69),s(fun(X68,bool),X70)))) | s(bool,'const.pred_set.DISJOINT_2'(s(fun(X65,bool),X66),s(fun(X65,bool),X67))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X68,bool),X70),s(fun(X68,bool),X69)))) )),
  inference(superposition,[],[f1319,f282])).
fof(f4368,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X62,bool),X63),s(fun(X62,bool),X64)))) | s(bool,'const.pred_set.DISJOINT_2'(s(fun(X59,bool),X60),s(fun(X59,bool),X61))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X62,bool),X64),s(fun(X62,bool),X63))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X59,bool),X61),s(fun(X59,bool),X60)))) )),
  inference(subsumption_resolution,[],[f4349,f427])).
fof(f4349,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (p(s(bool,'F_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X62,bool),X63),s(fun(X62,bool),X64)))) | s(bool,'const.pred_set.DISJOINT_2'(s(fun(X59,bool),X60),s(fun(X59,bool),X61))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X62,bool),X64),s(fun(X62,bool),X63))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X59,bool),X61),s(fun(X59,bool),X60)))) )),
  inference(superposition,[],[f1319,f1312])).
fof(f4316,plain,(
  ( ! [X167,X165,X166,X164,X168] : (~sP8('const.pred_set.DISJOINT_2'(s(fun(X164,bool),X165),s(fun(X164,bool),X166)),X167,X168) | sP9('const.pred_set.DISJOINT_2'(s(fun(X164,bool),X165),s(fun(X164,bool),X166)),X167,X168) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X164,bool),X166),s(fun(X164,bool),X165)))) )),
  inference(subsumption_resolution,[],[f4261,f237])).
fof(f4261,plain,(
  ( ! [X167,X165,X166,X164,X168] : (~p(s(bool,'T_0')) | ~sP8('const.pred_set.DISJOINT_2'(s(fun(X164,bool),X165),s(fun(X164,bool),X166)),X167,X168) | sP9('const.pred_set.DISJOINT_2'(s(fun(X164,bool),X165),s(fun(X164,bool),X166)),X167,X168) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X164,bool),X166),s(fun(X164,bool),X165)))) )),
  inference(superposition,[],[f413,f1312])).
fof(f4315,plain,(
  ( ! [X163,X161,X159,X162,X160] : (~sP10('const.pred_set.DISJOINT_2'(s(fun(X159,bool),X160),s(fun(X159,bool),X161)),X162,X163) | sP11('const.pred_set.DISJOINT_2'(s(fun(X159,bool),X160),s(fun(X159,bool),X161)),X162,X163) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X159,bool),X161),s(fun(X159,bool),X160)))) )),
  inference(subsumption_resolution,[],[f4260,f237])).
fof(f4260,plain,(
  ( ! [X163,X161,X159,X162,X160] : (~p(s(bool,'T_0')) | ~sP10('const.pred_set.DISJOINT_2'(s(fun(X159,bool),X160),s(fun(X159,bool),X161)),X162,X163) | sP11('const.pred_set.DISJOINT_2'(s(fun(X159,bool),X160),s(fun(X159,bool),X161)),X162,X163) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X159,bool),X161),s(fun(X159,bool),X160)))) )),
  inference(superposition,[],[f412,f1312])).
fof(f4314,plain,(
  ( ! [X158,X156,X154,X157,X155] : (~p(s(bool,X157)) | p(s(bool,X158)) | ~sP11('const.pred_set.DISJOINT_2'(s(fun(X154,bool),X155),s(fun(X154,bool),X156)),X158,X157) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X154,bool),X156),s(fun(X154,bool),X155)))) )),
  inference(subsumption_resolution,[],[f4259,f237])).
fof(f4259,plain,(
  ( ! [X158,X156,X154,X157,X155] : (~p(s(bool,'T_0')) | ~p(s(bool,X157)) | p(s(bool,X158)) | ~sP11('const.pred_set.DISJOINT_2'(s(fun(X154,bool),X155),s(fun(X154,bool),X156)),X158,X157) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X154,bool),X156),s(fun(X154,bool),X155)))) )),
  inference(superposition,[],[f409,f1312])).
fof(f4313,plain,(
  ( ! [X152,X151,X149,X153,X150] : (~p(s(bool,X152)) | sP11('const.pred_set.DISJOINT_2'(s(fun(X149,bool),X150),s(fun(X149,bool),X151)),X152,X153) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X149,bool),X151),s(fun(X149,bool),X150)))) )),
  inference(subsumption_resolution,[],[f4258,f237])).
fof(f4258,plain,(
  ( ! [X152,X151,X149,X153,X150] : (~p(s(bool,'T_0')) | ~p(s(bool,X152)) | sP11('const.pred_set.DISJOINT_2'(s(fun(X149,bool),X150),s(fun(X149,bool),X151)),X152,X153) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X149,bool),X151),s(fun(X149,bool),X150)))) )),
  inference(superposition,[],[f408,f1312])).
fof(f4312,plain,(
  ( ! [X146,X144,X147,X145,X148] : (p(s(bool,X147)) | sP11('const.pred_set.DISJOINT_2'(s(fun(X144,bool),X145),s(fun(X144,bool),X146)),X148,X147) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X144,bool),X146),s(fun(X144,bool),X145)))) )),
  inference(subsumption_resolution,[],[f4257,f237])).
fof(f4257,plain,(
  ( ! [X146,X144,X147,X145,X148] : (~p(s(bool,'T_0')) | p(s(bool,X147)) | sP11('const.pred_set.DISJOINT_2'(s(fun(X144,bool),X145),s(fun(X144,bool),X146)),X148,X147) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X144,bool),X146),s(fun(X144,bool),X145)))) )),
  inference(superposition,[],[f407,f1312])).
fof(f4311,plain,(
  ( ! [X142,X140,X143,X141,X139] : (p(s(bool,X142)) | ~p(s(bool,X143)) | ~sP10(X143,X142,'const.pred_set.DISJOINT_2'(s(fun(X139,bool),X140),s(fun(X139,bool),X141))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X139,bool),X141),s(fun(X139,bool),X140)))) )),
  inference(subsumption_resolution,[],[f4256,f237])).
fof(f4256,plain,(
  ( ! [X142,X140,X143,X141,X139] : (~p(s(bool,'T_0')) | p(s(bool,X142)) | ~p(s(bool,X143)) | ~sP10(X143,X142,'const.pred_set.DISJOINT_2'(s(fun(X139,bool),X140),s(fun(X139,bool),X141))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X139,bool),X141),s(fun(X139,bool),X140)))) )),
  inference(superposition,[],[f402,f1312])).
fof(f4310,plain,(
  ( ! [X134,X138,X136,X135,X137] : (p(s(bool,X137)) | p(s(bool,X138)) | ~sP9('const.pred_set.DISJOINT_2'(s(fun(X134,bool),X135),s(fun(X134,bool),X136)),X138,X137) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X134,bool),X136),s(fun(X134,bool),X135)))) )),
  inference(subsumption_resolution,[],[f4255,f237])).
fof(f4255,plain,(
  ( ! [X134,X138,X136,X135,X137] : (~p(s(bool,'T_0')) | p(s(bool,X137)) | p(s(bool,X138)) | ~sP9('const.pred_set.DISJOINT_2'(s(fun(X134,bool),X135),s(fun(X134,bool),X136)),X138,X137) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X134,bool),X136),s(fun(X134,bool),X135)))) )),
  inference(superposition,[],[f392,f1312])).
fof(f4309,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,X132)) | sP9('const.pred_set.DISJOINT_2'(s(fun(X129,bool),X130),s(fun(X129,bool),X131)),X132,X133) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X129,bool),X131),s(fun(X129,bool),X130)))) )),
  inference(subsumption_resolution,[],[f4254,f237])).
fof(f4254,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,'T_0')) | ~p(s(bool,X132)) | sP9('const.pred_set.DISJOINT_2'(s(fun(X129,bool),X130),s(fun(X129,bool),X131)),X132,X133) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X129,bool),X131),s(fun(X129,bool),X130)))) )),
  inference(superposition,[],[f391,f1312])).
fof(f4308,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,X127)) | sP9('const.pred_set.DISJOINT_2'(s(fun(X124,bool),X125),s(fun(X124,bool),X126)),X128,X127) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X124,bool),X126),s(fun(X124,bool),X125)))) )),
  inference(subsumption_resolution,[],[f4253,f237])).
fof(f4253,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,'T_0')) | ~p(s(bool,X127)) | sP9('const.pred_set.DISJOINT_2'(s(fun(X124,bool),X125),s(fun(X124,bool),X126)),X128,X127) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X124,bool),X126),s(fun(X124,bool),X125)))) )),
  inference(superposition,[],[f390,f1312])).
fof(f4307,plain,(
  ( ! [X123,X121,X122,X120,X119] : (p(s(bool,X122)) | p(s(bool,X123)) | ~sP8('const.pred_set.DISJOINT_2'(s(fun(X119,bool),X120),s(fun(X119,bool),X121)),X122,X123) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X119,bool),X121),s(fun(X119,bool),X120)))) )),
  inference(subsumption_resolution,[],[f4252,f237])).
fof(f4252,plain,(
  ( ! [X123,X121,X122,X120,X119] : (~p(s(bool,'T_0')) | p(s(bool,X122)) | p(s(bool,X123)) | ~sP8('const.pred_set.DISJOINT_2'(s(fun(X119,bool),X120),s(fun(X119,bool),X121)),X122,X123) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X119,bool),X121),s(fun(X119,bool),X120)))) )),
  inference(superposition,[],[f385,f1312])).
fof(f4306,plain,(
  ( ! [X118,X116,X114,X117,X115] : (~p(s(bool,X117)) | p(s(bool,X118)) | ~sP7(X118,X117,'const.pred_set.DISJOINT_2'(s(fun(X114,bool),X115),s(fun(X114,bool),X116))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X114,bool),X116),s(fun(X114,bool),X115)))) )),
  inference(subsumption_resolution,[],[f4251,f237])).
fof(f4251,plain,(
  ( ! [X118,X116,X114,X117,X115] : (~p(s(bool,'T_0')) | ~p(s(bool,X117)) | p(s(bool,X118)) | ~sP7(X118,X117,'const.pred_set.DISJOINT_2'(s(fun(X114,bool),X115),s(fun(X114,bool),X116))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X114,bool),X116),s(fun(X114,bool),X115)))) )),
  inference(superposition,[],[f377,f1312])).
fof(f4305,plain,(
  ( ! [X111,X109,X112,X110,X113] : (~p(s(bool,X112)) | p(s(bool,X113)) | ~sP6(X112,'const.pred_set.DISJOINT_2'(s(fun(X109,bool),X110),s(fun(X109,bool),X111)),X113) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X109,bool),X111),s(fun(X109,bool),X110)))) )),
  inference(subsumption_resolution,[],[f4250,f237])).
fof(f4250,plain,(
  ( ! [X111,X109,X112,X110,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X112)) | p(s(bool,X113)) | ~sP6(X112,'const.pred_set.DISJOINT_2'(s(fun(X109,bool),X110),s(fun(X109,bool),X111)),X113) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X109,bool),X111),s(fun(X109,bool),X110)))) )),
  inference(superposition,[],[f368,f1312])).
fof(f4304,plain,(
  ( ! [X107,X105,X108,X106,X104] : (~p(s(bool,X107)) | p(s(bool,X108)) | ~sP3(X107,'const.pred_set.DISJOINT_2'(s(fun(X104,bool),X105),s(fun(X104,bool),X106)),X108) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X104,bool),X106),s(fun(X104,bool),X105)))) )),
  inference(subsumption_resolution,[],[f4249,f237])).
fof(f4249,plain,(
  ( ! [X107,X105,X108,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X107)) | p(s(bool,X108)) | ~sP3(X107,'const.pred_set.DISJOINT_2'(s(fun(X104,bool),X105),s(fun(X104,bool),X106)),X108) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X104,bool),X106),s(fun(X104,bool),X105)))) )),
  inference(superposition,[],[f349,f1312])).
fof(f4303,plain,(
  ( ! [X103,X101,X99,X102,X100] : (~p(s(bool,X102)) | p(s(bool,X103)) | ~sP4(X102,'const.pred_set.DISJOINT_2'(s(fun(X99,bool),X100),s(fun(X99,bool),X101)),X103) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X99,bool),X101),s(fun(X99,bool),X100)))) )),
  inference(subsumption_resolution,[],[f4248,f237])).
fof(f4248,plain,(
  ( ! [X103,X101,X99,X102,X100] : (~p(s(bool,'T_0')) | ~p(s(bool,X102)) | p(s(bool,X103)) | ~sP4(X102,'const.pred_set.DISJOINT_2'(s(fun(X99,bool),X100),s(fun(X99,bool),X101)),X103) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X99,bool),X101),s(fun(X99,bool),X100)))) )),
  inference(superposition,[],[f345,f1312])).
fof(f4302,plain,(
  ( ! [X94,X97,X95,X98,X96] : (sP5(X97,X98,'const.pred_set.DISJOINT_2'(s(fun(X94,bool),X95),s(fun(X94,bool),X96))) | ~sP4(X97,'const.pred_set.DISJOINT_2'(s(fun(X94,bool),X95),s(fun(X94,bool),X96)),X98) | ~sP3(X98,'const.pred_set.DISJOINT_2'(s(fun(X94,bool),X95),s(fun(X94,bool),X96)),X97) | ~sP2('const.pred_set.DISJOINT_2'(s(fun(X94,bool),X95),s(fun(X94,bool),X96)),X98,X97) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X94,bool),X96),s(fun(X94,bool),X95)))) )),
  inference(subsumption_resolution,[],[f4247,f237])).
fof(f4247,plain,(
  ( ! [X94,X97,X95,X98,X96] : (~p(s(bool,'T_0')) | sP5(X97,X98,'const.pred_set.DISJOINT_2'(s(fun(X94,bool),X95),s(fun(X94,bool),X96))) | ~sP4(X97,'const.pred_set.DISJOINT_2'(s(fun(X94,bool),X95),s(fun(X94,bool),X96)),X98) | ~sP3(X98,'const.pred_set.DISJOINT_2'(s(fun(X94,bool),X95),s(fun(X94,bool),X96)),X97) | ~sP2('const.pred_set.DISJOINT_2'(s(fun(X94,bool),X95),s(fun(X94,bool),X96)),X98,X97) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X94,bool),X96),s(fun(X94,bool),X95)))) )),
  inference(superposition,[],[f342,f1312])).
fof(f4301,plain,(
  ( ! [X92,X90,X93,X91,X89] : (~p(s(bool,X92)) | p(s(bool,X93)) | ~sP5(X92,'const.pred_set.DISJOINT_2'(s(fun(X89,bool),X90),s(fun(X89,bool),X91)),X93) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X89,bool),X91),s(fun(X89,bool),X90)))) )),
  inference(subsumption_resolution,[],[f4246,f237])).
fof(f4246,plain,(
  ( ! [X92,X90,X93,X91,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X92)) | p(s(bool,X93)) | ~sP5(X92,'const.pred_set.DISJOINT_2'(s(fun(X89,bool),X90),s(fun(X89,bool),X91)),X93) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X89,bool),X91),s(fun(X89,bool),X90)))) )),
  inference(superposition,[],[f341,f1312])).
fof(f4300,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X88)) | sP1(X88,'const.pred_set.DISJOINT_2'(s(fun(X85,bool),X86),s(fun(X85,bool),X87))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X85,bool),X87),s(fun(X85,bool),X86)))) )),
  inference(subsumption_resolution,[],[f4245,f237])).
fof(f4245,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'T_0')) | p(s(bool,X88)) | sP1(X88,'const.pred_set.DISJOINT_2'(s(fun(X85,bool),X86),s(fun(X85,bool),X87))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X85,bool),X87),s(fun(X85,bool),X86)))) )),
  inference(superposition,[],[f324,f1312])).
fof(f4299,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,X84)) | sP1('const.pred_set.DISJOINT_2'(s(fun(X81,bool),X82),s(fun(X81,bool),X83)),X84) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X81,bool),X83),s(fun(X81,bool),X82)))) )),
  inference(subsumption_resolution,[],[f4244,f237])).
fof(f4244,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'T_0')) | p(s(bool,X84)) | sP1('const.pred_set.DISJOINT_2'(s(fun(X81,bool),X82),s(fun(X81,bool),X83)),X84) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X81,bool),X83),s(fun(X81,bool),X82)))) )),
  inference(superposition,[],[f321,f1312])).
fof(f4298,plain,(
  ( ! [X80,X78,X79,X77] : (sP0('const.pred_set.DISJOINT_2'(s(fun(X77,bool),X78),s(fun(X77,bool),X79)),X80) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X77,bool),X79),s(fun(X77,bool),X78)))) )),
  inference(subsumption_resolution,[],[f4243,f237])).
fof(f4243,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,'T_0')) | sP0('const.pred_set.DISJOINT_2'(s(fun(X77,bool),X78),s(fun(X77,bool),X79)),X80) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X77,bool),X79),s(fun(X77,bool),X78)))) )),
  inference(superposition,[],[f306,f1312])).
fof(f4297,plain,(
  ( ! [X76,X74,X75,X73] : (sP0(X76,'const.pred_set.DISJOINT_2'(s(fun(X73,bool),X74),s(fun(X73,bool),X75))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X73,bool),X75),s(fun(X73,bool),X74)))) )),
  inference(subsumption_resolution,[],[f4242,f237])).
fof(f4242,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | sP0(X76,'const.pred_set.DISJOINT_2'(s(fun(X73,bool),X74),s(fun(X73,bool),X75))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X73,bool),X75),s(fun(X73,bool),X74)))) )),
  inference(superposition,[],[f305,f1312])).
fof(f4296,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X72)) | s(bool,X72) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X69,bool),X70),s(fun(X69,bool),X71))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X69,bool),X71),s(fun(X69,bool),X70)))) )),
  inference(subsumption_resolution,[],[f4241,f237])).
fof(f4241,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'T_0')) | ~p(s(bool,X72)) | s(bool,X72) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X69,bool),X70),s(fun(X69,bool),X71))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X69,bool),X71),s(fun(X69,bool),X70)))) )),
  inference(superposition,[],[f299,f1312])).
fof(f4295,plain,(
  ( ! [X57,X56,X55] : (s(bool,'F_0') != s(bool,'const.pred_set.DISJOINT_2'(s(fun(X55,bool),X56),s(fun(X55,bool),X57))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X55,bool),X57),s(fun(X55,bool),X56)))) )),
  inference(subsumption_resolution,[],[f4237,f237])).
fof(f4237,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.pred_set.DISJOINT_2'(s(fun(X55,bool),X56),s(fun(X55,bool),X57))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X55,bool),X57),s(fun(X55,bool),X56)))) )),
  inference(superposition,[],[f257,f1312])).
fof(f4294,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X42,bool),X43),s(fun(X42,bool),X45)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X42,bool),X45),s(fun(X42,bool),'const.pred_set.UNION_2'(s(fun(X42,bool),X43),s(fun(X42,bool),X44)))))) )),
  inference(subsumption_resolution,[],[f4231,f237])).
fof(f4231,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X42,bool),X43),s(fun(X42,bool),X45)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X42,bool),X45),s(fun(X42,bool),'const.pred_set.UNION_2'(s(fun(X42,bool),X43),s(fun(X42,bool),X44)))))) )),
  inference(superposition,[],[f293,f1312])).
fof(f4293,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X38,bool),X40),s(fun(X38,bool),X41)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X38,bool),X41),s(fun(X38,bool),'const.pred_set.UNION_2'(s(fun(X38,bool),X39),s(fun(X38,bool),X40)))))) )),
  inference(subsumption_resolution,[],[f4230,f237])).
fof(f4230,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X38,bool),X40),s(fun(X38,bool),X41)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X38,bool),X41),s(fun(X38,bool),'const.pred_set.UNION_2'(s(fun(X38,bool),X39),s(fun(X38,bool),X40)))))) )),
  inference(superposition,[],[f294,f1312])).
fof(f4292,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X37)) | sP1(X37,'const.pred_set.DISJOINT_2'(s(fun(X34,bool),X36),s(fun(X34,bool),X35))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X34,bool),X36),s(fun(X34,bool),X35)))) )),
  inference(subsumption_resolution,[],[f4229,f237])).
fof(f4229,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'T_0')) | p(s(bool,X37)) | sP1(X37,'const.pred_set.DISJOINT_2'(s(fun(X34,bool),X36),s(fun(X34,bool),X35))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X34,bool),X36),s(fun(X34,bool),X35)))) )),
  inference(superposition,[],[f1325,f1312])).
fof(f4291,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X33)) | sP1('const.pred_set.DISJOINT_2'(s(fun(X30,bool),X32),s(fun(X30,bool),X31)),X33) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X30,bool),X32),s(fun(X30,bool),X31)))) )),
  inference(subsumption_resolution,[],[f4228,f237])).
fof(f4228,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X33)) | sP1('const.pred_set.DISJOINT_2'(s(fun(X30,bool),X32),s(fun(X30,bool),X31)),X33) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X30,bool),X32),s(fun(X30,bool),X31)))) )),
  inference(superposition,[],[f1324,f1312])).
fof(f4290,plain,(
  ( ! [X28,X26,X29,X27] : (sP0('const.pred_set.DISJOINT_2'(s(fun(X26,bool),X28),s(fun(X26,bool),X27)),X29) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X26,bool),X28),s(fun(X26,bool),X27)))) )),
  inference(subsumption_resolution,[],[f4227,f237])).
fof(f4227,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'T_0')) | sP0('const.pred_set.DISJOINT_2'(s(fun(X26,bool),X28),s(fun(X26,bool),X27)),X29) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X26,bool),X28),s(fun(X26,bool),X27)))) )),
  inference(superposition,[],[f1323,f1312])).
fof(f4289,plain,(
  ( ! [X24,X23,X25,X22] : (sP0(X25,'const.pred_set.DISJOINT_2'(s(fun(X22,bool),X24),s(fun(X22,bool),X23))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X22,bool),X24),s(fun(X22,bool),X23)))) )),
  inference(subsumption_resolution,[],[f4226,f237])).
fof(f4226,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'T_0')) | sP0(X25,'const.pred_set.DISJOINT_2'(s(fun(X22,bool),X24),s(fun(X22,bool),X23))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X22,bool),X24),s(fun(X22,bool),X23)))) )),
  inference(superposition,[],[f1322,f1312])).
fof(f4288,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X8)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9))),s(fun(X6,bool),X8)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))) )),
  inference(subsumption_resolution,[],[f4221,f237])).
fof(f4221,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X8)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9))),s(fun(X6,bool),X8)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))) )),
  inference(superposition,[],[f295,f1312])).
fof(f4285,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,X64)) | s(bool,X64) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X61,bool),X62),s(fun(X61,bool),X63))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X61,bool),X63),s(fun(X61,bool),X62)))) )),
  inference(subsumption_resolution,[],[f4185,f427])).
fof(f4185,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,'F_0')) | p(s(bool,X64)) | s(bool,X64) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X61,bool),X62),s(fun(X61,bool),X63))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X61,bool),X63),s(fun(X61,bool),X62)))) )),
  inference(superposition,[],[f296,f1312])).
fof(f4140,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK17(X0,bool)) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f4135,f237])).
fof(f4135,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK17(X0,bool)) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1294,f716])).
fof(f4125,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK16(X0,bool)) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f4120,f237])).
fof(f4120,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK16(X0,bool)) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1282,f716])).
fof(f3975,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,X15)) | sP1(X15,'const.pred_set.DISJOINT_2'(s(fun(X12,bool),X14),s(fun(X12,bool),X13))) | s(bool,X16) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X12,bool),X13),s(fun(X12,bool),X14))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f3968,f237])).
fof(f3968,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'T_0')) | p(s(bool,X15)) | sP1(X15,'const.pred_set.DISJOINT_2'(s(fun(X12,bool),X14),s(fun(X12,bool),X13))) | s(bool,X16) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X12,bool),X13),s(fun(X12,bool),X14))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f1325,f716])).
fof(f3949,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,X15)) | sP1('const.pred_set.DISJOINT_2'(s(fun(X12,bool),X14),s(fun(X12,bool),X13)),X15) | s(bool,X16) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X12,bool),X13),s(fun(X12,bool),X14))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f3942,f237])).
fof(f3942,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,'T_0')) | p(s(bool,X15)) | sP1('const.pred_set.DISJOINT_2'(s(fun(X12,bool),X14),s(fun(X12,bool),X13)),X15) | s(bool,X16) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X12,bool),X13),s(fun(X12,bool),X14))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f1324,f716])).
fof(f3898,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))) | s(bool,X9) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f3885,f237])).
fof(f3885,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))) | s(bool,X9) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1317,f716])).
fof(f3724,plain,(
  ( ! [X14,X12,X15,X13,X11] : (sP0('const.pred_set.DISJOINT_2'(s(fun(X11,bool),X13),s(fun(X11,bool),X12)),X14) | s(bool,X15) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X11,bool),X12),s(fun(X11,bool),X13))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f3713,f237])).
fof(f3713,plain,(
  ( ! [X14,X12,X15,X13,X11] : (~p(s(bool,'T_0')) | sP0('const.pred_set.DISJOINT_2'(s(fun(X11,bool),X13),s(fun(X11,bool),X12)),X14) | s(bool,X15) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X11,bool),X12),s(fun(X11,bool),X13))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f1323,f716])).
fof(f3699,plain,(
  ( ! [X14,X12,X15,X13,X11] : (sP0(X14,'const.pred_set.DISJOINT_2'(s(fun(X11,bool),X13),s(fun(X11,bool),X12))) | s(bool,X15) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X11,bool),X12),s(fun(X11,bool),X13))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f3688,f237])).
fof(f3688,plain,(
  ( ! [X14,X12,X15,X13,X11] : (~p(s(bool,'T_0')) | sP0(X14,'const.pred_set.DISJOINT_2'(s(fun(X11,bool),X13),s(fun(X11,bool),X12))) | s(bool,X15) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X11,bool),X12),s(fun(X11,bool),X13))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f1322,f716])).
fof(f3660,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0','const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)),'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f1311,f340])).
fof(f257,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f258,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f154])).
fof(f427,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f240,f237])).
fof(f240,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f1282,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK16(X0,bool))) )),
  inference(superposition,[],[f280,f244])).
fof(f1294,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK17(X0,bool))) )),
  inference(superposition,[],[f281,f244])).
fof(f1312,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) )),
  inference(superposition,[],[f282,f244])).
fof(f1317,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X9,bool),X11),s(fun(X9,bool),X10)))) | s(bool,'F_0') != s(bool,'const.pred_set.DISJOINT_2'(s(fun(X9,bool),X10),s(fun(X9,bool),X11)))) )),
  inference(superposition,[],[f257,f282])).
fof(f1318,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X12,bool),X14),s(fun(X12,bool),X13)))) | s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X12,bool),X13),s(fun(X12,bool),X14)))) )),
  inference(superposition,[],[f258,f282])).
fof(f1380,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK24(X0,X1,X2,bool))) )),
  inference(superposition,[],[f337,f244])).
fof(f1381,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK24(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK24(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f337,f244])).
fof(f1412,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK18(X0,X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X0,X1,bool))) )),
  inference(superposition,[],[f285,f244])).
fof(f1430,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK18(X0,X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK18(X0,X1,bool))) )),
  inference(superposition,[],[f286,f244])).
fof(f1448,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK19(X0,X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X0,X1,bool))) )),
  inference(superposition,[],[f287,f244])).
fof(f1460,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK19(X0,X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK20(X0,X1,bool))) )),
  inference(superposition,[],[f288,f244])).
fof(f2023,plain,(
  ( ! [X198,X200,X199,X197] : (s(bool,'T_0') != s(bool,f(s(X197,sK24(X198,X199,bool,X197)))) | s(fun(X197,bool),X198) = s(fun(X197,bool),X199) | s(bool,X200) = s(bool,g(s(X197,sK24(X198,X199,bool,X197)))) | s(bool,'T_0') = s(bool,X200)) )),
  inference(superposition,[],[f337,f716])).
fof(f1861,plain,(
  ( ! [X204,X202,X205,X203] : (s(bool,X205) != s(bool,f(s(X202,sK24(X203,X204,bool,X202)))) | s(fun(X202,bool),X203) = s(fun(X202,bool),X204) | s(bool,'T_0') = s(bool,X205) | s(bool,'T_0') = s(bool,g(s(X202,sK24(X203,X204,bool,X202))))) )),
  inference(superposition,[],[f337,f716])).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X7,X8] : s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.pred_set.DISJOINT_SYM')).
fof(f293,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & ((~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))),s(fun(X6,bool),X9)))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.pred_set.DISJOINT_UNION')).
fof(f294,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f181])).
fof(f295,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f181])).
fof(f1315,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.pred_set.DISJOINT_2'(s(fun(X3,bool),X5),s(fun(X3,bool),X4))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X3,bool),X4),s(fun(X3,bool),X5))))) )),
  inference(superposition,[],[f253,f282])).
fof(f1316,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) )),
  inference(superposition,[],[f254,f282])).
fof(f1319,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X15,bool),X17),s(fun(X15,bool),X16)))) | p(s(bool,X18)) | s(bool,X18) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X15,bool),X16),s(fun(X15,bool),X17)))) )),
  inference(superposition,[],[f296,f282])).
fof(f1321,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X23,bool),X25),s(fun(X23,bool),X24)))) | ~p(s(bool,X26)) | s(bool,X26) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X23,bool),X24),s(fun(X23,bool),X25)))) )),
  inference(superposition,[],[f299,f282])).
fof(f1322,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X27,bool),X29),s(fun(X27,bool),X28)))) | sP0(X30,'const.pred_set.DISJOINT_2'(s(fun(X27,bool),X28),s(fun(X27,bool),X29)))) )),
  inference(superposition,[],[f305,f282])).
fof(f1323,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X31,bool),X33),s(fun(X31,bool),X32)))) | sP0('const.pred_set.DISJOINT_2'(s(fun(X31,bool),X32),s(fun(X31,bool),X33)),X34)) )),
  inference(superposition,[],[f306,f282])).
fof(f1324,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X35,bool),X37),s(fun(X35,bool),X36)))) | p(s(bool,X38)) | sP1('const.pred_set.DISJOINT_2'(s(fun(X35,bool),X36),s(fun(X35,bool),X37)),X38)) )),
  inference(superposition,[],[f321,f282])).
fof(f1325,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X39,bool),X41),s(fun(X39,bool),X40)))) | p(s(bool,X42)) | sP1(X42,'const.pred_set.DISJOINT_2'(s(fun(X39,bool),X40),s(fun(X39,bool),X41)))) )),
  inference(superposition,[],[f324,f282])).
fof(f1326,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~sP5(X46,'const.pred_set.DISJOINT_2'(s(fun(X43,bool),X44),s(fun(X43,bool),X45)),X47) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X43,bool),X45),s(fun(X43,bool),X44))))) )),
  inference(superposition,[],[f341,f282])).
fof(f1327,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~sP4(X51,'const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50)),X52) | sP5(X51,X52,'const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X48,bool),X50),s(fun(X48,bool),X49)))) | ~sP3(X52,'const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50)),X51) | ~sP2('const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50)),X52,X51)) )),
  inference(superposition,[],[f342,f282])).
fof(f1328,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~sP4(X56,'const.pred_set.DISJOINT_2'(s(fun(X53,bool),X54),s(fun(X53,bool),X55)),X57) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X53,bool),X55),s(fun(X53,bool),X54))))) )),
  inference(superposition,[],[f345,f282])).
fof(f1329,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~sP3(X61,'const.pred_set.DISJOINT_2'(s(fun(X58,bool),X59),s(fun(X58,bool),X60)),X62) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X58,bool),X60),s(fun(X58,bool),X59))))) )),
  inference(superposition,[],[f349,f282])).
fof(f1330,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~sP6(X66,'const.pred_set.DISJOINT_2'(s(fun(X63,bool),X64),s(fun(X63,bool),X65)),X67) | ~p(s(bool,X66)) | p(s(bool,X67)) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X63,bool),X65),s(fun(X63,bool),X64))))) )),
  inference(superposition,[],[f368,f282])).
fof(f1331,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~sP7(X72,X71,'const.pred_set.DISJOINT_2'(s(fun(X68,bool),X69),s(fun(X68,bool),X70))) | ~p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X68,bool),X70),s(fun(X68,bool),X69))))) )),
  inference(superposition,[],[f377,f282])).
fof(f1332,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~sP8('const.pred_set.DISJOINT_2'(s(fun(X73,bool),X74),s(fun(X73,bool),X75)),X76,X77) | p(s(bool,X76)) | p(s(bool,X77)) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X73,bool),X75),s(fun(X73,bool),X74))))) )),
  inference(superposition,[],[f385,f282])).
fof(f1333,plain,(
  ( ! [X80,X78,X81,X79,X82] : (sP9('const.pred_set.DISJOINT_2'(s(fun(X78,bool),X79),s(fun(X78,bool),X80)),X82,X81) | ~p(s(bool,X81)) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X78,bool),X80),s(fun(X78,bool),X79))))) )),
  inference(superposition,[],[f390,f282])).
fof(f1334,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP9('const.pred_set.DISJOINT_2'(s(fun(X83,bool),X84),s(fun(X83,bool),X85)),X86,X87) | ~p(s(bool,X86)) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X83,bool),X85),s(fun(X83,bool),X84))))) )),
  inference(superposition,[],[f391,f282])).
fof(f1335,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X88,bool),X90),s(fun(X88,bool),X89)))) | p(s(bool,X91)) | p(s(bool,X92)) | ~sP9('const.pred_set.DISJOINT_2'(s(fun(X88,bool),X89),s(fun(X88,bool),X90)),X92,X91)) )),
  inference(superposition,[],[f392,f282])).
fof(f1336,plain,(
  ( ! [X94,X97,X95,X93,X96] : (~sP10(X97,X96,'const.pred_set.DISJOINT_2'(s(fun(X93,bool),X94),s(fun(X93,bool),X95))) | p(s(bool,X96)) | ~p(s(bool,X97)) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X93,bool),X95),s(fun(X93,bool),X94))))) )),
  inference(superposition,[],[f402,f282])).
fof(f1337,plain,(
  ( ! [X101,X99,X102,X100,X98] : (sP11('const.pred_set.DISJOINT_2'(s(fun(X98,bool),X99),s(fun(X98,bool),X100)),X102,X101) | p(s(bool,X101)) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X98,bool),X100),s(fun(X98,bool),X99))))) )),
  inference(superposition,[],[f407,f282])).
fof(f1338,plain,(
  ( ! [X103,X107,X105,X106,X104] : (sP11('const.pred_set.DISJOINT_2'(s(fun(X103,bool),X104),s(fun(X103,bool),X105)),X106,X107) | ~p(s(bool,X106)) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X103,bool),X105),s(fun(X103,bool),X104))))) )),
  inference(superposition,[],[f408,f282])).
fof(f1339,plain,(
  ( ! [X111,X109,X112,X110,X108] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X108,bool),X110),s(fun(X108,bool),X109)))) | ~p(s(bool,X111)) | p(s(bool,X112)) | ~sP11('const.pred_set.DISJOINT_2'(s(fun(X108,bool),X109),s(fun(X108,bool),X110)),X112,X111)) )),
  inference(superposition,[],[f409,f282])).
fof(f1340,plain,(
  ( ! [X116,X114,X117,X115,X113] : (sP11('const.pred_set.DISJOINT_2'(s(fun(X113,bool),X114),s(fun(X113,bool),X115)),X116,X117) | ~sP10('const.pred_set.DISJOINT_2'(s(fun(X113,bool),X114),s(fun(X113,bool),X115)),X116,X117) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X113,bool),X115),s(fun(X113,bool),X114))))) )),
  inference(superposition,[],[f412,f282])).
fof(f1341,plain,(
  ( ! [X121,X118,X122,X120,X119] : (sP9('const.pred_set.DISJOINT_2'(s(fun(X118,bool),X119),s(fun(X118,bool),X120)),X121,X122) | ~sP8('const.pred_set.DISJOINT_2'(s(fun(X118,bool),X119),s(fun(X118,bool),X120)),X121,X122) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X118,bool),X120),s(fun(X118,bool),X119))))) )),
  inference(superposition,[],[f413,f282])).
fof(f1393,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3))))) )),
  inference(superposition,[],[f293,f282])).
fof(f1405,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) )),
  inference(superposition,[],[f294,f282])).
fof(f1540,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X2,bool),X5),s(fun(X2,bool),X4)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X2,bool),X4),s(fun(X2,bool),X3)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X2,bool),'const.pred_set.UNION_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X5))),s(fun(X2,bool),X4))))) )),
  inference(superposition,[],[f295,f282])).
fof(f1728,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X3,bool),X5),s(fun(X3,bool),X4)))) | s(bool,X6) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X3,bool),X4),s(fun(X3,bool),X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f776,f282])).
fof(f2360,plain,(
  ( ! [X109,X107,X105,X108,X106] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X105,bool),X108),s(fun(X105,bool),X107)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X105,bool),'const.pred_set.UNION_2'(s(fun(X105,bool),X106),s(fun(X105,bool),X108))),s(fun(X105,bool),X107)))) | s(bool,X109) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X105,bool),X106),s(fun(X105,bool),X107))) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f2002,f237])).
fof(f2002,plain,(
  ( ! [X109,X107,X105,X108,X106] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X105,bool),X108),s(fun(X105,bool),X107)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X105,bool),'const.pred_set.UNION_2'(s(fun(X105,bool),X106),s(fun(X105,bool),X108))),s(fun(X105,bool),X107)))) | s(bool,X109) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X105,bool),X106),s(fun(X105,bool),X107))) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f295,f716])).
fof(f2361,plain,(
  ( ! [X121,X118,X122,X120,X119] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X118,bool),X120),s(fun(X118,bool),X121)))) | s(bool,X122) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X118,bool),'const.pred_set.UNION_2'(s(fun(X118,bool),X119),s(fun(X118,bool),X120))),s(fun(X118,bool),X121))) | s(bool,'T_0') = s(bool,X122)) )),
  inference(subsumption_resolution,[],[f2005,f237])).
fof(f2005,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X118,bool),X120),s(fun(X118,bool),X121)))) | s(bool,X122) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X118,bool),'const.pred_set.UNION_2'(s(fun(X118,bool),X119),s(fun(X118,bool),X120))),s(fun(X118,bool),X121))) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f294,f716])).
fof(f2362,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X123,bool),X124),s(fun(X123,bool),X126)))) | s(bool,X127) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X123,bool),'const.pred_set.UNION_2'(s(fun(X123,bool),X124),s(fun(X123,bool),X125))),s(fun(X123,bool),X126))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(subsumption_resolution,[],[f2006,f237])).
fof(f2006,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X123,bool),X124),s(fun(X123,bool),X126)))) | s(bool,X127) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X123,bool),'const.pred_set.UNION_2'(s(fun(X123,bool),X124),s(fun(X123,bool),X125))),s(fun(X123,bool),X126))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f293,f716])).
fof(f280,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f169])).
fof(f169,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X14] : (~! [X12] : p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))) <=> ? [X12] : ~p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.bool.NOT_FORALL_THM')).
fof(f281,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f285,plain,(
  ( ! [X2,X0,X7,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0)))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19,sK20])],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : ((? [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X14,X15] : (! [X12] : (p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))) & p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X12))))) <=> (! [X12] : p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))) & ! [X12] : p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X12))))))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.bool.FORALL_AND_THM')).
fof(f286,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f287,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f288,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f289,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK23(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22,sK23])],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X14,X15] : (? [X12] : (p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))) | p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X12))))) <=> (? [X12] : p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))) | ? [X12] : p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X12))))))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.bool.EXISTS_OR_THM')).
fof(f290,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f291,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK23(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f292,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK23(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f2363,plain,(
  ( ! [X134,X132,X130,X133,X131] : (p(s(bool,happ(s(fun(X130,bool),X133),s(X130,sK23(X131,X133,X130))))) | p(s(bool,happ(s(fun(X130,bool),X131),s(X130,sK23(X131,X133,X130))))) | s(bool,X134) = s(bool,happ(s(fun(X130,bool),X131),s(X130,X132))) | s(bool,'T_0') = s(bool,X134)) )),
  inference(subsumption_resolution,[],[f2009,f237])).
fof(f2009,plain,(
  ( ! [X134,X132,X130,X133,X131] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X130,bool),X133),s(X130,sK23(X131,X133,X130))))) | p(s(bool,happ(s(fun(X130,bool),X131),s(X130,sK23(X131,X133,X130))))) | s(bool,X134) = s(bool,happ(s(fun(X130,bool),X131),s(X130,X132))) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f292,f716])).
fof(f2364,plain,(
  ( ! [X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(X135,bool),X136),s(X135,sK23(X138,X136,X135))))) | p(s(bool,happ(s(fun(X135,bool),X138),s(X135,sK23(X138,X136,X135))))) | s(bool,X139) = s(bool,happ(s(fun(X135,bool),X136),s(X135,X137))) | s(bool,'T_0') = s(bool,X139)) )),
  inference(subsumption_resolution,[],[f2010,f237])).
fof(f2010,plain,(
  ( ! [X138,X136,X135,X139,X137] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X135,bool),X136),s(X135,sK23(X138,X136,X135))))) | p(s(bool,happ(s(fun(X135,bool),X138),s(X135,sK23(X138,X136,X135))))) | s(bool,X139) = s(bool,happ(s(fun(X135,bool),X136),s(X135,X137))) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f291,f716])).
fof(f2365,plain,(
  ( ! [X144,X142,X140,X143,X141] : (p(s(bool,happ(s(fun(X140,bool),X143),s(X140,sK21(X141,X143,X140))))) | p(s(bool,happ(s(fun(X140,bool),X141),s(X140,sK22(X141,X143,X140))))) | s(bool,X144) = s(bool,happ(s(fun(X140,bool),X141),s(X140,X142))) | s(bool,'T_0') = s(bool,X144)) )),
  inference(subsumption_resolution,[],[f2011,f237])).
fof(f2011,plain,(
  ( ! [X144,X142,X140,X143,X141] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X140,bool),X143),s(X140,sK21(X141,X143,X140))))) | p(s(bool,happ(s(fun(X140,bool),X141),s(X140,sK22(X141,X143,X140))))) | s(bool,X144) = s(bool,happ(s(fun(X140,bool),X141),s(X140,X142))) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f290,f716])).
fof(f2366,plain,(
  ( ! [X146,X149,X147,X145,X148] : (p(s(bool,happ(s(fun(X145,bool),X146),s(X145,sK21(X148,X146,X145))))) | p(s(bool,happ(s(fun(X145,bool),X148),s(X145,sK22(X148,X146,X145))))) | s(bool,X149) = s(bool,happ(s(fun(X145,bool),X146),s(X145,X147))) | s(bool,'T_0') = s(bool,X149)) )),
  inference(subsumption_resolution,[],[f2012,f237])).
fof(f2012,plain,(
  ( ! [X146,X149,X147,X145,X148] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X145,bool),X146),s(X145,sK21(X148,X146,X145))))) | p(s(bool,happ(s(fun(X145,bool),X148),s(X145,sK22(X148,X146,X145))))) | s(bool,X149) = s(bool,happ(s(fun(X145,bool),X146),s(X145,X147))) | s(bool,'T_0') = s(bool,X149)) )),
  inference(superposition,[],[f289,f716])).
fof(f2367,plain,(
  ( ! [X152,X151,X153,X150] : (p(s(bool,happ(s(fun(X150,bool),X151),s(X150,X152)))) | s(bool,X153) = s(bool,happ(s(fun(X150,bool),X151),s(X150,sK16(X151,X150)))) | s(bool,'T_0') = s(bool,X153)) )),
  inference(subsumption_resolution,[],[f2013,f237])).
fof(f2013,plain,(
  ( ! [X152,X151,X153,X150] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X150,bool),X151),s(X150,X152)))) | s(bool,X153) = s(bool,happ(s(fun(X150,bool),X151),s(X150,sK16(X151,X150)))) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f280,f716])).
fof(f2368,plain,(
  ( ! [X156,X154,X157,X155] : (p(s(bool,happ(s(fun(X154,bool),X155),s(X154,X156)))) | s(bool,X157) = s(bool,happ(s(fun(X154,bool),X155),s(X154,sK17(X155,X154)))) | s(bool,'T_0') = s(bool,X157)) )),
  inference(subsumption_resolution,[],[f2014,f237])).
fof(f2014,plain,(
  ( ! [X156,X154,X157,X155] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X154,bool),X155),s(X154,X156)))) | s(bool,X157) = s(bool,happ(s(fun(X154,bool),X155),s(X154,sK17(X155,X154)))) | s(bool,'T_0') = s(bool,X157)) )),
  inference(superposition,[],[f281,f716])).
fof(f2369,plain,(
  ( ! [X158,X161,X159,X162,X160] : (p(s(bool,happ(s(fun(X158,bool),X159),s(X158,X161)))) | ~p(s(bool,happ(s(fun(X158,bool),X160),s(X158,sK18(X159,X160,X158))))) | s(bool,X162) = s(bool,happ(s(fun(X158,bool),X159),s(X158,sK18(X159,X160,X158)))) | s(bool,'T_0') = s(bool,X162)) )),
  inference(subsumption_resolution,[],[f2015,f237])).
fof(f2015,plain,(
  ( ! [X158,X161,X159,X162,X160] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X158,bool),X159),s(X158,X161)))) | ~p(s(bool,happ(s(fun(X158,bool),X160),s(X158,sK18(X159,X160,X158))))) | s(bool,X162) = s(bool,happ(s(fun(X158,bool),X159),s(X158,sK18(X159,X160,X158)))) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f286,f716])).
fof(f2370,plain,(
  ( ! [X167,X165,X163,X166,X164] : (p(s(bool,happ(s(fun(X163,bool),X165),s(X163,X166)))) | ~p(s(bool,happ(s(fun(X163,bool),X165),s(X163,sK18(X164,X165,X163))))) | s(bool,X167) = s(bool,happ(s(fun(X163,bool),X164),s(X163,sK18(X164,X165,X163)))) | s(bool,'T_0') = s(bool,X167)) )),
  inference(subsumption_resolution,[],[f2016,f237])).
fof(f2016,plain,(
  ( ! [X167,X165,X163,X166,X164] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X163,bool),X165),s(X163,X166)))) | ~p(s(bool,happ(s(fun(X163,bool),X165),s(X163,sK18(X164,X165,X163))))) | s(bool,X167) = s(bool,happ(s(fun(X163,bool),X164),s(X163,sK18(X164,X165,X163)))) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f285,f716])).
fof(f2371,plain,(
  ( ! [X171,X169,X172,X170,X168] : (p(s(bool,happ(s(fun(X168,bool),X169),s(X168,X171)))) | ~p(s(bool,happ(s(fun(X168,bool),X170),s(X168,sK19(X169,X170,X168))))) | s(bool,X172) = s(bool,happ(s(fun(X168,bool),X169),s(X168,sK20(X169,X170,X168)))) | s(bool,'T_0') = s(bool,X172)) )),
  inference(subsumption_resolution,[],[f2017,f237])).
fof(f2017,plain,(
  ( ! [X171,X169,X172,X170,X168] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X168,bool),X169),s(X168,X171)))) | ~p(s(bool,happ(s(fun(X168,bool),X170),s(X168,sK19(X169,X170,X168))))) | s(bool,X172) = s(bool,happ(s(fun(X168,bool),X169),s(X168,sK20(X169,X170,X168)))) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f288,f716])).
fof(f2372,plain,(
  ( ! [X177,X175,X173,X176,X174] : (p(s(bool,happ(s(fun(X173,bool),X175),s(X173,X176)))) | ~p(s(bool,happ(s(fun(X173,bool),X175),s(X173,sK19(X174,X175,X173))))) | s(bool,X177) = s(bool,happ(s(fun(X173,bool),X174),s(X173,sK20(X174,X175,X173)))) | s(bool,'T_0') = s(bool,X177)) )),
  inference(subsumption_resolution,[],[f2018,f237])).
fof(f2018,plain,(
  ( ! [X177,X175,X173,X176,X174] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X173,bool),X175),s(X173,X176)))) | ~p(s(bool,happ(s(fun(X173,bool),X175),s(X173,sK19(X174,X175,X173))))) | s(bool,X177) = s(bool,happ(s(fun(X173,bool),X174),s(X173,sK20(X174,X175,X173)))) | s(bool,'T_0') = s(bool,X177)) )),
  inference(superposition,[],[f287,f716])).
fof(f1862,plain,(
  ( ! [X208,X206,X207] : (~p(s(bool,happ(s(fun(bool,bool),X206),s(bool,X207)))) | p(s(bool,happ(s(fun(bool,bool),X206),s(bool,X208)))) | s(bool,'T_0') = s(bool,X207) | s(bool,'T_0') = s(bool,sK16(X206,bool))) )),
  inference(superposition,[],[f280,f716])).
fof(f1863,plain,(
  ( ! [X210,X211,X209] : (~p(s(bool,happ(s(fun(bool,bool),X209),s(bool,X210)))) | p(s(bool,happ(s(fun(bool,bool),X209),s(bool,X211)))) | s(bool,'T_0') = s(bool,X210) | s(bool,'T_0') = s(bool,sK17(X209,bool))) )),
  inference(superposition,[],[f281,f716])).
fof(f1864,plain,(
  ( ! [X212,X215,X213,X214] : (~p(s(bool,happ(s(fun(bool,bool),X212),s(bool,X214)))) | p(s(bool,happ(s(fun(bool,bool),X212),s(bool,X215)))) | ~p(s(bool,happ(s(fun(bool,bool),X213),s(bool,sK18(X212,X213,bool))))) | s(bool,'T_0') = s(bool,X214) | s(bool,'T_0') = s(bool,sK18(X212,X213,bool))) )),
  inference(superposition,[],[f286,f716])).
fof(f1865,plain,(
  ( ! [X218,X216,X219,X217] : (~p(s(bool,happ(s(fun(bool,bool),X216),s(bool,X218)))) | p(s(bool,happ(s(fun(bool,bool),X217),s(bool,X219)))) | ~p(s(bool,happ(s(fun(bool,bool),X217),s(bool,sK18(X216,X217,bool))))) | s(bool,'T_0') = s(bool,X218) | s(bool,'T_0') = s(bool,sK18(X216,X217,bool))) )),
  inference(superposition,[],[f285,f716])).
fof(f1866,plain,(
  ( ! [X222,X220,X223,X221] : (~p(s(bool,happ(s(fun(bool,bool),X220),s(bool,X222)))) | p(s(bool,happ(s(fun(bool,bool),X220),s(bool,X223)))) | ~p(s(bool,happ(s(fun(bool,bool),X221),s(bool,sK19(X220,X221,bool))))) | s(bool,'T_0') = s(bool,X222) | s(bool,'T_0') = s(bool,sK20(X220,X221,bool))) )),
  inference(superposition,[],[f288,f716])).
fof(f1867,plain,(
  ( ! [X227,X225,X226,X224] : (~p(s(bool,happ(s(fun(bool,bool),X224),s(bool,X226)))) | p(s(bool,happ(s(fun(bool,bool),X225),s(bool,X227)))) | ~p(s(bool,happ(s(fun(bool,bool),X225),s(bool,sK19(X224,X225,bool))))) | s(bool,'T_0') = s(bool,X226) | s(bool,'T_0') = s(bool,sK20(X224,X225,bool))) )),
  inference(superposition,[],[f287,f716])).
fof(f2024,plain,(
  ( ! [X202,X203,X201] : (~p(s(bool,happ(s(fun(bool,bool),X201),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X201),s(bool,X202)))) | s(bool,X203) = s(bool,sK16(X201,bool)) | s(bool,'T_0') = s(bool,X203)) )),
  inference(superposition,[],[f280,f716])).
fof(f2025,plain,(
  ( ! [X206,X204,X205] : (~p(s(bool,happ(s(fun(bool,bool),X204),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X204),s(bool,X205)))) | s(bool,X206) = s(bool,sK17(X204,bool)) | s(bool,'T_0') = s(bool,X206)) )),
  inference(superposition,[],[f281,f716])).
fof(f2026,plain,(
  ( ! [X210,X208,X209,X207] : (~p(s(bool,happ(s(fun(bool,bool),X207),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X207),s(bool,X209)))) | ~p(s(bool,happ(s(fun(bool,bool),X208),s(bool,sK18(X207,X208,bool))))) | s(bool,X210) = s(bool,sK18(X207,X208,bool)) | s(bool,'T_0') = s(bool,X210)) )),
  inference(superposition,[],[f286,f716])).
fof(f2027,plain,(
  ( ! [X212,X213,X211,X214] : (~p(s(bool,happ(s(fun(bool,bool),X211),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X212),s(bool,X213)))) | ~p(s(bool,happ(s(fun(bool,bool),X212),s(bool,sK18(X211,X212,bool))))) | s(bool,X214) = s(bool,sK18(X211,X212,bool)) | s(bool,'T_0') = s(bool,X214)) )),
  inference(superposition,[],[f285,f716])).
fof(f2028,plain,(
  ( ! [X218,X216,X215,X217] : (~p(s(bool,happ(s(fun(bool,bool),X215),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X215),s(bool,X217)))) | ~p(s(bool,happ(s(fun(bool,bool),X216),s(bool,sK19(X215,X216,bool))))) | s(bool,X218) = s(bool,sK20(X215,X216,bool)) | s(bool,'T_0') = s(bool,X218)) )),
  inference(superposition,[],[f288,f716])).
fof(f2029,plain,(
  ( ! [X222,X220,X219,X221] : (~p(s(bool,happ(s(fun(bool,bool),X219),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X220),s(bool,X221)))) | ~p(s(bool,happ(s(fun(bool,bool),X220),s(bool,sK19(X219,X220,bool))))) | s(bool,X222) = s(bool,sK20(X219,X220,bool)) | s(bool,'T_0') = s(bool,X222)) )),
  inference(superposition,[],[f287,f716])).
fof(f2030,plain,(
  ( ! [X225,X223,X226,X224] : (s(X225,f(s(bool,'T_0'))) != s(X225,g(s(bool,'T_0'))) | s(fun(bool,X225),X223) = s(fun(bool,X225),X224) | s(bool,X226) = s(bool,sK24(X223,X224,X225,bool)) | s(bool,'T_0') = s(bool,X226)) )),
  inference(superposition,[],[f337,f716])).
fof(f2359,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8(X91,X92,X93) | sP9(X91,X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1995,f237])).
fof(f1995,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~sP8(X91,X92,X93) | sP9(X91,X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f413,f716])).
fof(f2358,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10(X87,X88,X89) | sP11(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1994,f237])).
fof(f1994,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP10(X87,X88,X89) | sP11(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f412,f716])).
fof(f2357,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11(X83,X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1993,f237])).
fof(f1993,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11(X83,X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f409,f716])).
fof(f2356,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11(X79,X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1992,f237])).
fof(f1992,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | sP11(X79,X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f408,f716])).
fof(f2355,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1991,f237])).
fof(f1991,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X76)) | sP11(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f407,f716])).
fof(f2354,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,X71) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1990,f237])).
fof(f1990,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,X71) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f402,f716])).
fof(f2353,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1989,f237])).
fof(f1989,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f392,f716])).
fof(f2352,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1988,f237])).
fof(f1988,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X64)) | sP9(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f391,f716])).
fof(f2351,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1987,f237])).
fof(f1987,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f390,f716])).
fof(f2350,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1986,f237])).
fof(f1986,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f385,f716])).
fof(f2349,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,X51) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1985,f237])).
fof(f1985,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,X51) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f377,f716])).
fof(f2348,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1984,f237])).
fof(f1984,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f368,f716])).
fof(f2347,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1983,f237])).
fof(f1983,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f349,f716])).
fof(f2346,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1982,f237])).
fof(f1982,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f345,f716])).
fof(f2345,plain,(
  ( ! [X37,X35,X38,X36] : (sP5(X36,X37,X35) | ~sP4(X36,X35,X37) | ~sP3(X37,X35,X36) | ~sP2(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1981,f237])).
fof(f1981,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | sP5(X36,X37,X35) | ~sP4(X36,X35,X37) | ~sP3(X37,X35,X36) | ~sP2(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f342,f716])).
fof(f2344,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1980,f237])).
fof(f1980,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f341,f716])).
fof(f2343,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f1979,f237])).
fof(f1979,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f324,f716])).
fof(f2342,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f1978,f237])).
fof(f1978,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f321,f716])).
fof(f2341,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1977,f237])).
fof(f1977,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f306,f716])).
fof(f2340,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1976,f237])).
fof(f1976,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f305,f716])).
fof(f2339,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1975,f237])).
fof(f1975,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f299,f716])).
fof(f2338,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1971,f237])).
fof(f1971,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f257,f716])).
fof(f1934,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f253,f716])).
fof(f1868,plain,(
  ( ! [X231,X229,X230,X228] : (s(X230,f(s(bool,X231))) != s(X230,g(s(bool,X231))) | s(fun(bool,X230),X228) = s(fun(bool,X230),X229) | s(bool,'T_0') = s(bool,X231) | s(bool,'T_0') = s(bool,sK24(X228,X229,X230,bool))) )),
  inference(superposition,[],[f337,f716])).
fof(f1778,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f716,f716])).
fof(f1763,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X26) = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X23,bool),X25),s(fun(X23,bool),X24))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X23,bool),X24),s(fun(X23,bool),X25)))) )),
  inference(superposition,[],[f716,f282])).
fof(f1310,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)),'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f427,f282,f360])).
fof(f1311,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)),'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f237,f282,f358])).
fof(f1226,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f445,f440,f438,f427,f343])).
fof(f1181,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f451,f440,f435,f237,f342])).
fof(f1129,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1117,f397])).
fof(f1117,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f237,f427,f237,f402])).
fof(f1103,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f470,f237,f427,f401])).
fof(f1029,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1015,f380])).
fof(f1015,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f427,f427,f237,f385])).
fof(f1002,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f465,f427,f427,f384])).
fof(f951,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f939,f361])).
fof(f939,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f427,f237,f237,f368])).
fof(f926,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f454,f237,f237,f364])).
fof(f922,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f909,f338])).
fof(f909,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f427,f427,f427,f353])).
fof(f901,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f889,f339])).
fof(f889,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f427,f237,f237,f349])).
fof(f869,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f427,f237,f237,f345])).
fof(f823,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f237,f471,f359])).
fof(f825,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f237,f471,f359])).
fof(f802,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f427,f471,f357])).
fof(f803,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f427,f471,f357])).
fof(f776,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f771,f427])).
fof(f771,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f296,f244])).
fof(f716,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f244,f244])).
fof(f651,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f237,f237,f408])).
fof(f640,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f427,f237,f407])).
fof(f632,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f237,f427,f396])).
fof(f630,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f427,f427,f395])).
fof(f617,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f237,f237,f391])).
fof(f606,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f237,f390])).
fof(f601,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f237,f427,f379])).
fof(f598,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f427,f378])).
fof(f592,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f427,f427,f373])).
fof(f585,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f427,f427,f372])).
fof(f580,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f237,f427,f363])).
fof(f578,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f237,f427,f362])).
fof(f533,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f427,f427,f428])).
fof(f519,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f427,f237,f324])).
fof(f509,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f427,f237,f321])).
fof(f507,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f237,f237,f320])).
fof(f506,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f427,f427,f319])).
fof(f471,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f427,f253])).
fof(f470,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f427,f405])).
fof(f467,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f237,f404])).
fof(f466,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f427,f403])).
fof(f465,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f427,f388])).
fof(f462,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f237,f387])).
fof(f459,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f386])).
fof(f458,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f427,f371])).
fof(f457,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f427,f370])).
fof(f454,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f369])).
fof(f451,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f237,f356])).
fof(f448,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f237,f355])).
fof(f445,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f354])).
fof(f444,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f427,f352])).
fof(f443,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f427,f351])).
fof(f440,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f350])).
fof(f439,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f427,f348])).
fof(f438,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f427,f347])).
fof(f435,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f346])).
fof(f432,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f237,f306])).
fof(f429,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f237,f305])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f128,f145,f144])).
fof(f144,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f145,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f128,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X17,X18,X19] : ((p(s(bool,X19)) <=> (p(s(bool,X18)) => p(s(bool,X17)))) <=> ((p(s(bool,X19)) | p(s(bool,X18))) & (p(s(bool,X19)) | ~p(s(bool,X17))) & (~p(s(bool,X18)) | p(s(bool,X17)) | ~p(s(bool,X19)))))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.sat.sat_27_prove_ltb')).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f230])).
fof(f230,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f144])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f231])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f231])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f231])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f145])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f228])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f228])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f102,f142,f141])).
fof(f141,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f142,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X17,X18,X19] : ((p(s(bool,X19)) <=> (p(s(bool,X18)) | p(s(bool,X17)))) <=> ((p(s(bool,X19)) | ~p(s(bool,X18))) & (p(s(bool,X19)) | ~p(s(bool,X17))) & (p(s(bool,X18)) | p(s(bool,X17)) | ~p(s(bool,X19)))))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.sat.sat_25_prove_ltb')).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f225])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f225])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f225])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f222])).
fof(f222,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f141])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f223])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f223])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f223])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f219])).
fof(f219,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f142])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f220])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f220])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f381])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f220])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f220])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f220])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f220])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f100,f139,f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f139,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f100,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X17,X18,X19] : ((p(s(bool,X19)) <=> (p(s(bool,X18)) & p(s(bool,X17)))) <=> ((p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,X17))) & (p(s(bool,X18)) | ~p(s(bool,X19))) & (p(s(bool,X17)) | ~p(s(bool,X19)))))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.sat.sat_23_prove_ltb')).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f217])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f217])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f217])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f138])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f215])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f215])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f215])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f139])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f367])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f98,f136,f135,f134,f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f134,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f135,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f136,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f98,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X17,X18,X19] : ((p(s(bool,X19)) <=> s(bool,X17) = s(bool,X18)) <=> ((p(s(bool,X19)) | p(s(bool,X18)) | p(s(bool,X17))) & (p(s(bool,X19)) | ~p(s(bool,X17)) | ~p(s(bool,X18))) & (p(s(bool,X18)) | ~p(s(bool,X17)) | ~p(s(bool,X19))) & (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)))))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.sat.sat_2|49|_prove_ltb')).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f210])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f210])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f133])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f209])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f209])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f206])).
fof(f206,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f134])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f207])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f207])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f207])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f203])).
fof(f203,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f135])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f204])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f204])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f204])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f200])).
fof(f200,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f136])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f201])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f201])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f337,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK24(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK24(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK24(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK24(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24])],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1',aHL_EXT)).
fof(f319,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f191])).
fof(f191,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f320,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f321,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f192])).
fof(f324,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f192])).
fof(f428,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f304,f427])).
fof(f304,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f305,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f185])).
fof(f306,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f185])).
fof(f296,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f299,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f120])).
fof(f253,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f254,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f152])).
fof(f244,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1',aHL_BOOL_CASES)).
fof(f237,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1',aHL_TRUTH)).
fof(f6396,plain,(
  $false | $spl0),
  inference(global_subsumption,[],[f236,f235,f234,f237,f244,f254,f253,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1311,f1310,f1763,f1778,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2029,f2028,f2027,f2026,f2025,f2024,f1867,f1866,f1865,f1864,f1863,f1862,f2372,f2371,f2370,f2369,f2368,f2367,f2366,f2365,f2364,f2363,f292,f291,f290,f289,f288,f287,f286,f285,f281,f280,f2362,f2361,f2360,f1728,f1540,f1405,f1393,f1341,f1340,f1339,f1338,f1337,f1336,f1335,f1334,f1333,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1325,f1324,f1323,f1322,f1321,f1319,f1316,f1315,f295,f294,f293,f282,f1861,f2023,f1460,f1448,f1430,f1412,f1381,f1380,f1318,f1317,f1312,f1294,f1282,f427,f258,f257,f3660,f3699,f3724,f3898,f3949,f3975,f4125,f4140,f4285,f4288,f4289,f4290,f4291,f4292,f4293,f4294,f4295,f4296,f4297,f4298,f4299,f4300,f4301,f4302,f4303,f4304,f4305,f4306,f4307,f4308,f4309,f4310,f4311,f4312,f4313,f4314,f4315,f4316,f4368,f4350,f4399,f4400,f4413,f4414,f4415,f4416,f4417,f4501,f4502,f4547,f4548,f4573,f4574,f4593,f4594,f4611,f4612,f4626,f4734,f4732,f4735,f4772,f4770,f4773,f4877,f4879,f4881,f4887,f4908,f4910,f4912,f4918,f4967,f4952,f4968,f421,f6295,f6294,f6291,f6290,f6289,f6288,f6285,f6282,f6279,f6278,f6255,f6253,f6251,f6250,f6249,f6248,f6243,f6236,f6235,f6231,f6230,f6207,f6206,f6205,f6200,f6195,f6194,f6189,f6180,f6175,f6173,f6172,f6168,f6167,f6166,f6165,f6162,f6160,f6155,f6154,f6150,f6149,f6133,f6132,f6131,f6128,f6121,f6120,f6119,f6118,f6114,f6113,f6107,f6103,f6102,f6093,f6091,f6090,f6087,f6086,f6085,f6084,f6082,f6081,f6080,f6078,f6072,f6070,f6068,f6067,f6065,f6057,f6056,f6047,f6044,f6039,f6034,f6029,f6025,f6022,f6021,f6020,f6019,f6018,f6016,f6386,f6011,f6010])).
fof(f6389,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f6388,f427])).
fof(f6388,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f6011,f5877])).
fof(f5877,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl3),
  inference(unit_resulting_resolution,[],[f471,f424,f776])).
fof(f6384,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f6024,f5691])).
fof(f5691,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X0))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f424,f1393])).
fof(f6024,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f237,f1311,f421,f1326])).
fof(f5681,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f5680,f427])).
fof(f5680,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f5679,f543])).
fof(f543,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f422,f258])).
fof(f5679,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f5666,f237])).
fof(f5666,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | ($spl2 | $spl4)),
  inference(superposition,[],[f1544,f475])).
fof(f475,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f425,f254])).
fof(f1544,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),X0),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X0))))))) ) | $spl2),
  inference(forward_demodulation,[],[f1543,f282])).
fof(f1543,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),X0),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X0))),s(fun(sK12,bool),sK15))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1538,f237])).
fof(f1538,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),X0),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X0))),s(fun(sK12,bool),sK15))))) ) | $spl2),
  inference(superposition,[],[f295,f474])).
fof(f474,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f423,f254])).
fof(f4622,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK24(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK24(X0,X1,bool,bool))))) ) | $spl56),
  inference(cnf_transformation,[],[f4622_D])).
fof(f4622_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK24(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK24(X0,X1,bool,bool))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4624,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl59),
  inference(cnf_transformation,[],[f4624_D])).
fof(f4624_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f4612,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK24(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK24(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1381,f716])).
fof(f4611,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK24(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK24(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1381,f716])).
fof(f4423,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl50),
  inference(cnf_transformation,[],[f4423_D])).
fof(f4423_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4419,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK24(X0,X1,bool,bool))) ) | $spl46),
  inference(cnf_transformation,[],[f4419_D])).
fof(f4419_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK24(X0,X1,bool,bool))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4426,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl55),
  inference(cnf_transformation,[],[f4426_D])).
fof(f4426_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f4417,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK24(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1380,f244])).
fof(f4424,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl52),
  inference(cnf_transformation,[],[f4424_D])).
fof(f4424_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4416,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK24(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1380,f716])).
fof(f4415,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK24(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1380,f716])).
fof(f4418,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl44),
  inference(cnf_transformation,[],[f4418_D])).
fof(f4418_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4421,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl49),
  inference(cnf_transformation,[],[f4421_D])).
fof(f4421_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f4414,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK24(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1380,f716])).
fof(f4413,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK24(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1380,f716])).
fof(f2563,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540,f2542,f2544,f2546,f2548,f2550,f2552,f2554,f2556,f2558,f2560,f2562])).
fof(f2562,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~sP8('const.pred_set.DISJOINT_2'(s(fun(X118,bool),X119),s(fun(X118,bool),X120)),X121,X122) | sP9('const.pred_set.DISJOINT_2'(s(fun(X118,bool),X119),s(fun(X118,bool),X120)),X121,X122)) ) | $spl8),
  inference(subsumption_resolution,[],[f2457,f237])).
fof(f2457,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,'T_0')) | ~sP8('const.pred_set.DISJOINT_2'(s(fun(X118,bool),X119),s(fun(X118,bool),X120)),X121,X122) | sP9('const.pred_set.DISJOINT_2'(s(fun(X118,bool),X119),s(fun(X118,bool),X120)),X121,X122)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1341])).
fof(f2560,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~sP10('const.pred_set.DISJOINT_2'(s(fun(X113,bool),X114),s(fun(X113,bool),X115)),X116,X117) | sP11('const.pred_set.DISJOINT_2'(s(fun(X113,bool),X114),s(fun(X113,bool),X115)),X116,X117)) ) | $spl8),
  inference(subsumption_resolution,[],[f2456,f237])).
fof(f2456,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,'T_0')) | ~sP10('const.pred_set.DISJOINT_2'(s(fun(X113,bool),X114),s(fun(X113,bool),X115)),X116,X117) | sP11('const.pred_set.DISJOINT_2'(s(fun(X113,bool),X114),s(fun(X113,bool),X115)),X116,X117)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1340])).
fof(f2558,plain,(
  ( ! [X111,X109,X112,X110,X108] : (~p(s(bool,X111)) | p(s(bool,X112)) | ~sP11('const.pred_set.DISJOINT_2'(s(fun(X108,bool),X109),s(fun(X108,bool),X110)),X112,X111)) ) | $spl8),
  inference(subsumption_resolution,[],[f2455,f237])).
fof(f2455,plain,(
  ( ! [X111,X109,X112,X110,X108] : (~p(s(bool,'T_0')) | ~p(s(bool,X111)) | p(s(bool,X112)) | ~sP11('const.pred_set.DISJOINT_2'(s(fun(X108,bool),X109),s(fun(X108,bool),X110)),X112,X111)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1339])).
fof(f2556,plain,(
  ( ! [X103,X107,X105,X106,X104] : (~p(s(bool,X106)) | sP11('const.pred_set.DISJOINT_2'(s(fun(X103,bool),X104),s(fun(X103,bool),X105)),X106,X107)) ) | $spl8),
  inference(subsumption_resolution,[],[f2454,f237])).
fof(f2454,plain,(
  ( ! [X103,X107,X105,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X106)) | sP11('const.pred_set.DISJOINT_2'(s(fun(X103,bool),X104),s(fun(X103,bool),X105)),X106,X107)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1338])).
fof(f2554,plain,(
  ( ! [X101,X99,X102,X100,X98] : (p(s(bool,X101)) | sP11('const.pred_set.DISJOINT_2'(s(fun(X98,bool),X99),s(fun(X98,bool),X100)),X102,X101)) ) | $spl8),
  inference(subsumption_resolution,[],[f2453,f237])).
fof(f2453,plain,(
  ( ! [X101,X99,X102,X100,X98] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | sP11('const.pred_set.DISJOINT_2'(s(fun(X98,bool),X99),s(fun(X98,bool),X100)),X102,X101)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1337])).
fof(f2552,plain,(
  ( ! [X94,X97,X95,X93,X96] : (p(s(bool,X96)) | ~p(s(bool,X97)) | ~sP10(X97,X96,'const.pred_set.DISJOINT_2'(s(fun(X93,bool),X94),s(fun(X93,bool),X95)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2452,f237])).
fof(f2452,plain,(
  ( ! [X94,X97,X95,X93,X96] : (~p(s(bool,'T_0')) | p(s(bool,X96)) | ~p(s(bool,X97)) | ~sP10(X97,X96,'const.pred_set.DISJOINT_2'(s(fun(X93,bool),X94),s(fun(X93,bool),X95)))) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1336])).
fof(f2550,plain,(
  ( ! [X92,X90,X88,X91,X89] : (p(s(bool,X91)) | p(s(bool,X92)) | ~sP9('const.pred_set.DISJOINT_2'(s(fun(X88,bool),X89),s(fun(X88,bool),X90)),X92,X91)) ) | $spl8),
  inference(subsumption_resolution,[],[f2451,f237])).
fof(f2451,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,X91)) | p(s(bool,X92)) | ~sP9('const.pred_set.DISJOINT_2'(s(fun(X88,bool),X89),s(fun(X88,bool),X90)),X92,X91)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1335])).
fof(f2548,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,X86)) | sP9('const.pred_set.DISJOINT_2'(s(fun(X83,bool),X84),s(fun(X83,bool),X85)),X86,X87)) ) | $spl8),
  inference(subsumption_resolution,[],[f2450,f237])).
fof(f2450,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,X86)) | sP9('const.pred_set.DISJOINT_2'(s(fun(X83,bool),X84),s(fun(X83,bool),X85)),X86,X87)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1334])).
fof(f2546,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,X81)) | sP9('const.pred_set.DISJOINT_2'(s(fun(X78,bool),X79),s(fun(X78,bool),X80)),X82,X81)) ) | $spl8),
  inference(subsumption_resolution,[],[f2449,f237])).
fof(f2449,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | sP9('const.pred_set.DISJOINT_2'(s(fun(X78,bool),X79),s(fun(X78,bool),X80)),X82,X81)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1333])).
fof(f2544,plain,(
  ( ! [X76,X74,X77,X75,X73] : (p(s(bool,X76)) | p(s(bool,X77)) | ~sP8('const.pred_set.DISJOINT_2'(s(fun(X73,bool),X74),s(fun(X73,bool),X75)),X76,X77)) ) | $spl8),
  inference(subsumption_resolution,[],[f2448,f237])).
fof(f2448,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X76)) | p(s(bool,X77)) | ~sP8('const.pred_set.DISJOINT_2'(s(fun(X73,bool),X74),s(fun(X73,bool),X75)),X76,X77)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1332])).
fof(f2542,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~sP7(X72,X71,'const.pred_set.DISJOINT_2'(s(fun(X68,bool),X69),s(fun(X68,bool),X70)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2447,f237])).
fof(f2447,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,'T_0')) | ~p(s(bool,X71)) | p(s(bool,X72)) | ~sP7(X72,X71,'const.pred_set.DISJOINT_2'(s(fun(X68,bool),X69),s(fun(X68,bool),X70)))) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1331])).
fof(f2540,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~sP6(X66,'const.pred_set.DISJOINT_2'(s(fun(X63,bool),X64),s(fun(X63,bool),X65)),X67)) ) | $spl8),
  inference(subsumption_resolution,[],[f2446,f237])).
fof(f2446,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X66)) | p(s(bool,X67)) | ~sP6(X66,'const.pred_set.DISJOINT_2'(s(fun(X63,bool),X64),s(fun(X63,bool),X65)),X67)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1330])).
fof(f2538,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP3(X61,'const.pred_set.DISJOINT_2'(s(fun(X58,bool),X59),s(fun(X58,bool),X60)),X62)) ) | $spl8),
  inference(subsumption_resolution,[],[f2445,f237])).
fof(f2445,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP3(X61,'const.pred_set.DISJOINT_2'(s(fun(X58,bool),X59),s(fun(X58,bool),X60)),X62)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1329])).
fof(f2536,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP4(X56,'const.pred_set.DISJOINT_2'(s(fun(X53,bool),X54),s(fun(X53,bool),X55)),X57)) ) | $spl8),
  inference(subsumption_resolution,[],[f2444,f237])).
fof(f2444,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP4(X56,'const.pred_set.DISJOINT_2'(s(fun(X53,bool),X54),s(fun(X53,bool),X55)),X57)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1328])).
fof(f2534,plain,(
  ( ! [X52,X50,X48,X51,X49] : (sP5(X51,X52,'const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50))) | ~sP4(X51,'const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50)),X52) | ~sP3(X52,'const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50)),X51) | ~sP2('const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50)),X52,X51)) ) | $spl8),
  inference(subsumption_resolution,[],[f2443,f237])).
fof(f2443,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,'T_0')) | sP5(X51,X52,'const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50))) | ~sP4(X51,'const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50)),X52) | ~sP3(X52,'const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50)),X51) | ~sP2('const.pred_set.DISJOINT_2'(s(fun(X48,bool),X49),s(fun(X48,bool),X50)),X52,X51)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1327])).
fof(f2532,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,X46)) | p(s(bool,X47)) | ~sP5(X46,'const.pred_set.DISJOINT_2'(s(fun(X43,bool),X44),s(fun(X43,bool),X45)),X47)) ) | $spl8),
  inference(subsumption_resolution,[],[f2442,f237])).
fof(f2442,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP5(X46,'const.pred_set.DISJOINT_2'(s(fun(X43,bool),X44),s(fun(X43,bool),X45)),X47)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1326])).
fof(f2530,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X42)) | sP1(X42,'const.pred_set.DISJOINT_2'(s(fun(X39,bool),X40),s(fun(X39,bool),X41)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2441,f237])).
fof(f2441,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X42)) | sP1(X42,'const.pred_set.DISJOINT_2'(s(fun(X39,bool),X40),s(fun(X39,bool),X41)))) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1325])).
fof(f2528,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,X38)) | sP1('const.pred_set.DISJOINT_2'(s(fun(X35,bool),X36),s(fun(X35,bool),X37)),X38)) ) | $spl8),
  inference(subsumption_resolution,[],[f2440,f237])).
fof(f2440,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,X38)) | sP1('const.pred_set.DISJOINT_2'(s(fun(X35,bool),X36),s(fun(X35,bool),X37)),X38)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1324])).
fof(f2526,plain,(
  ( ! [X33,X31,X34,X32] : (sP0('const.pred_set.DISJOINT_2'(s(fun(X31,bool),X32),s(fun(X31,bool),X33)),X34)) ) | $spl8),
  inference(subsumption_resolution,[],[f2439,f237])).
fof(f2439,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | sP0('const.pred_set.DISJOINT_2'(s(fun(X31,bool),X32),s(fun(X31,bool),X33)),X34)) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1323])).
fof(f2524,plain,(
  ( ! [X30,X28,X29,X27] : (sP0(X30,'const.pred_set.DISJOINT_2'(s(fun(X27,bool),X28),s(fun(X27,bool),X29)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2438,f237])).
fof(f2438,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | sP0(X30,'const.pred_set.DISJOINT_2'(s(fun(X27,bool),X28),s(fun(X27,bool),X29)))) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1322])).
fof(f2513,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f2414,f2285])).
fof(f2414,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f257])).
fof(f2561,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540,f2542,f2544,f2546,f2548,f2550,f2552,f2554,f2556,f2558,f2560])).
fof(f2559,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540,f2542,f2544,f2546,f2548,f2550,f2552,f2554,f2556,f2558])).
fof(f2557,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540,f2542,f2544,f2546,f2548,f2550,f2552,f2554,f2556])).
fof(f2555,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540,f2542,f2544,f2546,f2548,f2550,f2552,f2554])).
fof(f2553,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540,f2542,f2544,f2546,f2548,f2550,f2552])).
fof(f2551,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540,f2542,f2544,f2546,f2548,f2550])).
fof(f2549,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540,f2542,f2544,f2546,f2548])).
fof(f2547,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540,f2542,f2544,f2546])).
fof(f2545,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540,f2542,f2544])).
fof(f2543,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540,f2542])).
fof(f2541,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538,f2540])).
fof(f2539,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536,f2538])).
fof(f2537,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534,f2536])).
fof(f2535,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532,f2534])).
fof(f2533,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530,f2532])).
fof(f2531,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528,f2530])).
fof(f2529,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526,f2528])).
fof(f2527,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524,f2526])).
fof(f2525,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513,f2524])).
fof(f2518,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2517,f237])).
fof(f2517,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(forward_demodulation,[],[f2516,f2285])).
fof(f2516,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X9,bool),X11),s(fun(X9,bool),X10))))) ) | $spl8),
  inference(subsumption_resolution,[],[f2421,f2285])).
fof(f2421,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.pred_set.DISJOINT_2'(s(fun(X9,bool),X10),s(fun(X9,bool),X11))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X9,bool),X11),s(fun(X9,bool),X10))))) ) | $spl8),
  inference(backward_demodulation,[],[f2285,f1317])).
fof(f2515,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2416,f237])).
fof(f2416,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f2285,f427])).
fof(f2514,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f236,f235,f234,f237,f427,f244,f254,f253,f258,f257,f281,f280,f282,f288,f287,f286,f285,f292,f291,f290,f289,f295,f294,f293,f299,f296,f306,f305,f428,f324,f321,f320,f319,f337,f344,f343,f342,f341,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f414,f364,f363,f362,f361,f371,f370,f369,f368,f377,f374,f373,f372,f384,f383,f382,f413,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f401,f400,f399,f412,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f429,f432,f435,f438,f439,f440,f443,f444,f445,f448,f451,f454,f457,f458,f459,f462,f465,f466,f467,f470,f471,f506,f507,f509,f519,f533,f578,f580,f585,f592,f598,f601,f606,f617,f630,f632,f640,f651,f716,f776,f803,f802,f825,f823,f869,f889,f901,f909,f922,f926,f939,f951,f1002,f1015,f1029,f1103,f1117,f1129,f1181,f1226,f1282,f1294,f1311,f1310,f1312,f1315,f1316,f1317,f1318,f1319,f1321,f1322,f1323,f1324,f1325,f1326,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1380,f1381,f1393,f1405,f1412,f1430,f1448,f1460,f1540,f1728,f1763,f1778,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1934,f2283,f2287,f2290,f2294,f2297,f2300,f2303,f2306,f2309,f2312,f2315,f2318,f2321,f2324,f2327,f2330,f2333,f2336,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2355,f2356,f2357,f2358,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2369,f2370,f2371,f2372,f2023,f2024,f2025,f2026,f2027,f2028,f2029,f2030,f2378,f2380,f2382,f2385,f2387,f2389,f2391,f2393,f2395,f2397,f2399,f2401,f2403,f2405,f2407,f2409,f2411,f2413,f2285,f2513])).
fof(f2337,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f2337_D])).
fof(f2337_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2285,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f2285_D])).
fof(f2285_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2413,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2412])).
fof(f2412,plain,(
  ( ! [X99,X97,X98,X96] : (~sP8('T_0',X97,X98) | sP9('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f2059,f776])).
fof(f2059,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~sP8('T_0',X97,X98) | sP9('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f413,f716])).
fof(f2334,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl40),
  inference(cnf_transformation,[],[f2334_D])).
fof(f2334_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2411,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2410])).
fof(f2410,plain,(
  ( ! [X94,X92,X95,X93] : (~sP10('T_0',X93,X94) | sP11('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f2058,f776])).
fof(f2058,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP10('T_0',X93,X94) | sP11('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f412,f716])).
fof(f2331,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f2331_D])).
fof(f2331_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2409,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2408])).
fof(f2408,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | ~sP11('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f2057,f776])).
fof(f2057,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~p(s(bool,X89)) | p(s(bool,X90)) | ~sP11('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f409,f716])).
fof(f2328,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f2328_D])).
fof(f2328_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2407,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2406])).
fof(f2406,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | sP11('T_0',X85,X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f2056,f776])).
fof(f2056,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | sP11('T_0',X85,X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f408,f716])).
fof(f2325,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f2325_D])).
fof(f2325_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2405,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2404])).
fof(f2404,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X81)) | sP11('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f2055,f776])).
fof(f2055,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | sP11('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f407,f716])).
fof(f2322,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f2322_D])).
fof(f2322_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2403,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f2402])).
fof(f2402,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | ~p(s(bool,X78)) | ~sP10(X78,X77,'T_0') | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f2054,f776])).
fof(f2054,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~p(s(bool,X78)) | ~sP10(X78,X77,'T_0') | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f402,f716])).
fof(f2319,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f2319_D])).
fof(f2319_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2401,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2400])).
fof(f2400,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f2053,f776])).
fof(f2053,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f392,f716])).
fof(f2316,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f2316_D])).
fof(f2316_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2399,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2398])).
fof(f2398,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f2052,f776])).
fof(f2052,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | ~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f391,f716])).
fof(f2313,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f2313_D])).
fof(f2313_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2397,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2396])).
fof(f2396,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f2051,f776])).
fof(f2051,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f390,f716])).
fof(f2310,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl24),
  inference(cnf_transformation,[],[f2310_D])).
fof(f2310_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2395,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f2394])).
fof(f2394,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f2050,f776])).
fof(f2050,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f385,f716])).
fof(f2307,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f2307_D])).
fof(f2307_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2393,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f2392])).
fof(f2392,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f2049,f776])).
fof(f2049,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f377,f716])).
fof(f2304,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f2304_D])).
fof(f2304_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2391,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f2390])).
fof(f2390,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f2048,f776])).
fof(f2048,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f368,f716])).
fof(f2301,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f2301_D])).
fof(f2301_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2389,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f2388])).
fof(f2388,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f2047,f776])).
fof(f2047,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f349,f716])).
fof(f2298,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f2298_D])).
fof(f2298_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2387,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2386])).
fof(f2386,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f2046,f776])).
fof(f2046,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f345,f716])).
fof(f2295,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) | $spl14),
  inference(cnf_transformation,[],[f2295_D])).
fof(f2295_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2385,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2384,f451])).
fof(f2384,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f2383])).
fof(f2383,plain,(
  ( ! [X43,X41,X42,X40] : (sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f2045,f776])).
fof(f2045,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f342,f716])).
fof(f2291,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f2291_D])).
fof(f2291_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2382,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2381])).
fof(f2381,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f2044,f776])).
fof(f2044,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f341,f716])).
fof(f2288,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl10),
  inference(cnf_transformation,[],[f2288_D])).
fof(f2288_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2380,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2379])).
fof(f2379,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f2043,f776])).
fof(f2043,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f324,f716])).
fof(f2284,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl6),
  inference(cnf_transformation,[],[f2284_D])).
fof(f2284_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2378,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2377])).
fof(f2377,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f2042,f776])).
fof(f2042,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f321,f716])).
fof(f2336,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2335])).
fof(f2335,plain,(
  ( ! [X99,X97,X98,X96] : (~sP8('T_0',X97,X98) | sP9('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f1962,f776])).
fof(f1962,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~sP8('T_0',X97,X98) | sP9('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f413,f716])).
fof(f2333,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2332])).
fof(f2332,plain,(
  ( ! [X94,X92,X95,X93] : (~sP10('T_0',X93,X94) | sP11('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f1961,f776])).
fof(f1961,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP10('T_0',X93,X94) | sP11('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f412,f716])).
fof(f2330,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2329])).
fof(f2329,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | ~sP11('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f1960,f776])).
fof(f1960,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~p(s(bool,X89)) | p(s(bool,X90)) | ~sP11('T_0',X90,X89) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f409,f716])).
fof(f2327,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2326])).
fof(f2326,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | sP11('T_0',X85,X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f1959,f776])).
fof(f1959,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | sP11('T_0',X85,X86) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f408,f716])).
fof(f2324,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2323])).
fof(f2323,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X81)) | sP11('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f1958,f776])).
fof(f1958,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | sP11('T_0',X82,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f407,f716])).
fof(f2321,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f2320])).
fof(f2320,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | ~p(s(bool,X78)) | ~sP10(X78,X77,'T_0') | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f1957,f776])).
fof(f1957,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~p(s(bool,X78)) | ~sP10(X78,X77,'T_0') | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f402,f716])).
fof(f2318,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2317])).
fof(f2317,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f1956,f776])).
fof(f1956,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f392,f716])).
fof(f2315,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2314])).
fof(f2314,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f1955,f776])).
fof(f1955,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | ~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f391,f716])).
fof(f2312,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2311])).
fof(f2311,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f1954,f776])).
fof(f1954,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f390,f716])).
fof(f2309,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f2308])).
fof(f2308,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f1953,f776])).
fof(f1953,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f385,f716])).
fof(f2306,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f2305])).
fof(f2305,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f1952,f776])).
fof(f1952,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f377,f716])).
fof(f2303,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f2302])).
fof(f2302,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f1951,f776])).
fof(f1951,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f368,f716])).
fof(f2300,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f2299])).
fof(f2299,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f1950,f776])).
fof(f1950,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f349,f716])).
fof(f2297,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2296])).
fof(f2296,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f1949,f776])).
fof(f1949,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f345,f716])).
fof(f2294,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2293,f451])).
fof(f2293,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f2292])).
fof(f2292,plain,(
  ( ! [X43,X41,X42,X40] : (sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f1948,f776])).
fof(f1948,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f342,f716])).
fof(f2290,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2289])).
fof(f2289,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f1947,f776])).
fof(f1947,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f341,f716])).
fof(f2287,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2286])).
fof(f2286,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f1946,f776])).
fof(f1946,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f324,f716])).
fof(f2283,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2282])).
fof(f2282,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f1945,f776])).
fof(f1945,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f321,f716])).
fof(f421,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | $spl0),
  inference(cnf_transformation,[],[f421_D])).
fof(f421_D,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f423,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | $spl2),
  inference(cnf_transformation,[],[f423_D])).
fof(f423_D,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f234,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f148])).
fof(f148,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))))),
  inference(nnf_transformation,[],[f105])).
fof(f105,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <~> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X9),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X9),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  file('thm.pred_set.DISJOINT_UNION_BOTH|split|1','thm.pred_set.DISJOINT_UNION_BOTH|split|1')).
fof(f425,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | $spl4),
  inference(cnf_transformation,[],[f425_D])).
fof(f425_D,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f235,plain,(
  p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))),
  inference(cnf_transformation,[],[f149])).
fof(f422,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | $spl1),
  inference(cnf_transformation,[],[f422_D])).
fof(f422_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f424,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | $spl3),
  inference(cnf_transformation,[],[f424_D])).
fof(f424_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f426,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | $spl5),
  inference(cnf_transformation,[],[f426_D])).
fof(f426_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f236,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK12,bool),sK15),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))),
  inference(cnf_transformation,[],[f149])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 524
% Passive clauses: 1050
% Generated clauses: 6183
% Final active clauses: 240
% Final passive clauses: 263
% Input formulas: 40
% Initial clauses: 178
% 
% Duplicate literals: 142
% Trivial inequalities: 28
% Fw subsumption resolutions: 363
% Fw demodulations: 677
% Bw demodulations: 92
% Condensations: 172
% Global subsumptions: 25
% 
% Simple tautologies: 72
% Equational tautologies: 60
% Forward subsumptions: 3520
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 2504
% Factoring: 2
% Forward superposition: 818
% Backward superposition: 1010
% Self superposition: 26
% Equality factoring: 29
% Equality resolution: 1
% 
% Split clauses: 46
% Split components: 98
% SAT solver clauses: 1197
% SAT solver unit clauses: 90
% SAT solver binary clauses: 485
% 
% Sat splits: 46
% Sat splitting refutations: 28
% 
% Memory used [KB]: 4221
% Time elapsed: 0.212 s
% ------------------------------
% ------------------------------
% Success in time 1.528 s
