Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% (81120)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 598
% Passive clauses: 9061
% Generated clauses: 16426
% Final active clauses: 531
% Final passive clauses: 7971
% Input formulas: 46
% Initial clauses: 203
% 
% Duplicate literals: 258
% Trivial inequalities: 86
% Fw subsumption resolutions: 1233
% Fw demodulations: 43
% Bw demodulations: 66
% 
% Simple tautologies: 222
% Equational tautologies: 65
% Forward subsumptions: 5246
% 
% Binary resolution: 10683
% Factoring: 986
% Forward superposition: 1242
% Backward superposition: 1369
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 37
% 
% Split clauses: 283
% Split components: 567
% SAT solver clauses: 600
% SAT solver unit clauses: 16
% SAT solver binary clauses: 156
% 
% Sat splits: 283
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 598
% TWLsolver calls for satisfiability: 335
% 
% Memory used [KB]: 7675
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f4196,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4159,f478,f480,f4077,f4078,f1149,f360,f523,f877,f878,f263,f477,f852,f863,f680,f686,f517,f518,f654,f669,f359,f524,f4072,f4073,f363,f522,f889,f899,f3908,f3953,f3645,f3646,f610,f613,f262,f475,f285,f488,f692,f699,f358,f525,f4036,f4061,f270,f484,f1911,f1915,f332,f506,f4156,f3290,f344,f516,f3324])).
fof(f3324,plain,(
  $false | ($spl2 | $spl1430)),
  inference(resolution,[],[f3290,f477])).
fof(f516,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,X1))) ) | $spl60),
  inference(cnf_transformation,[],[f516_D])).
fof(f516_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,X1))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f344,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f3290,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl1430),
  inference(cnf_transformation,[],[f3290_D])).
fof(f3290_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl1430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1430])])).
fof(f4156,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl60 | $spl1954)),
  inference(resolution,[],[f4078,f516])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) ) | $spl42),
  inference(cnf_transformation,[],[f506_D])).
fof(f506_D,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK23(X1,X0)),s('type.integer.int',X0)))) & (! [X2] : s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23])],[f199])).
fof(f199,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ? [X3] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X0)))) & (! [X2] : s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(rectify,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ? [X2] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0)))) & (! [X2] : s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ? [X2] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X18,X19] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X18),s('type.integer.int',X19)))) <=> ? [X20] : s('type.integer.int',X19) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X18))))),
  file('thm.integer.INT_DIVIDES_|49||split|0','thm.integer.INT_DIVIDES')).
fof(f1915,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))) ) | $spl708),
  inference(cnf_transformation,[],[f1915_D])).
fof(f1915_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))) ) <=> ~$spl708),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl708])])).
fof(f1911,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))) ) | $spl42),
  inference(equality_resolution,[],[f506])).
fof(f484,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl6),
  inference(cnf_transformation,[],[f484_D])).
fof(f484_D,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f270,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : s('type.integer.int',X6) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.integer.INT_DIVIDES_|49||split|0','thm.integer.INT_MUL_RID')).
fof(f4061,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X0))))) ) | $spl1950),
  inference(cnf_transformation,[],[f4061_D])).
fof(f4061_D,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X0))))) ) <=> ~$spl1950),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1950])])).
fof(f4036,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X0))))) ) | ($spl6 | $spl708)),
  inference(superposition,[],[f1915,f484])).
fof(f525,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP1(X0,X1)) ) | $spl70),
  inference(cnf_transformation,[],[f525_D])).
fof(f525_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP1(X0,X1)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f358,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f699,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~sP1(X4,X3) | ~p(s(bool,X5)) | ~sP1(X5,X4)) ) | $spl270),
  inference(cnf_transformation,[],[f699_D])).
fof(f699_D,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~sP1(X4,X3) | ~p(s(bool,X5)) | ~sP1(X5,X4)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f692,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~sP1(X4,X3) | ~p(s(bool,X5)) | ~sP1(X5,X4)) ) | ($spl68 | $spl70)),
  inference(resolution,[],[f525,f524])).
fof(f488,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f488_D])).
fof(f488_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f285,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.integer.INT_DIVIDES_|49||split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f475,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12)))) | $spl1),
  inference(cnf_transformation,[],[f475_D])).
fof(f475_D,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f262,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))))),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f117])).
fof(f117,plain,(
  ? [X0] : ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X0))))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0] : p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X0))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X6))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X6))))),
  file('thm.integer.INT_DIVIDES_|49||split|0','thm.integer.INT_DIVIDES_|49||split|0')).
fof(f613,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))) | $spl220),
  inference(cnf_transformation,[],[f613_D])).
fof(f613_D,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f610,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))) | ($spl1 | $spl14)),
  inference(resolution,[],[f488,f475])).
fof(f3646,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))) | ~p(s(bool,X1)) | ~sP1(X1,X0)) ) | $spl1676),
  inference(cnf_transformation,[],[f3646_D])).
fof(f3646_D,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))) | ~p(s(bool,X1)) | ~sP1(X1,X0)) ) <=> ~$spl1676),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1676])])).
fof(f3645,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))) | ~p(s(bool,X1)) | ~sP1(X1,X0)) ) | ($spl5 | $spl220 | $spl270)),
  inference(subsumption_resolution,[],[f3623,f480])).
fof(f3623,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~sP1(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))) | ~p(s(bool,X1)) | ~sP1(X1,X0)) ) | ($spl220 | $spl270)),
  inference(superposition,[],[f699,f613])).
fof(f3953,plain,(
  ( ! [X66,X67] : (~sP1(X66,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))) | ~sP1(X67,X66) | ~sP0('F_0',X67)) ) | $spl1866),
  inference(cnf_transformation,[],[f3953_D])).
fof(f3953_D,plain,(
  ( ! [X66,X67] : (~sP1(X66,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))) | ~sP1(X67,X66) | ~sP0('F_0',X67)) ) <=> ~$spl1866),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1866])])).
fof(f3908,plain,(
  ( ! [X66,X67] : (~sP1(X66,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))) | ~sP1(X67,X66) | ~sP0('F_0',X67)) ) | ($spl354 | $spl1676)),
  inference(resolution,[],[f3646,f878])).
fof(f899,plain,(
  ( ! [X2,X3] : (~p(s(bool,X3)) | ~sP1(X2,X3) | ~sP0('F_0',X2)) ) | $spl360),
  inference(cnf_transformation,[],[f899_D])).
fof(f899_D,plain,(
  ( ! [X2,X3] : (~p(s(bool,X3)) | ~sP1(X2,X3) | ~sP0('F_0',X2)) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f889,plain,(
  ( ! [X2,X3] : (~sP0('F_0',X2) | ~p(s(bool,X3)) | ~sP1(X2,X3)) ) | ($spl68 | $spl354)),
  inference(resolution,[],[f878,f524])).
fof(f522,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl64),
  inference(cnf_transformation,[],[f522_D])).
fof(f522_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f363,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f213])).
fof(f4073,plain,(
  ( ! [X2,X3] : (~sP1(X2,X3) | ~sP0('F_0',X2)) ) | $spl1952),
  inference(cnf_transformation,[],[f4073_D])).
fof(f4073_D,plain,(
  ( ! [X2,X3] : (~sP1(X2,X3) | ~sP0('F_0',X2)) ) <=> ~$spl1952),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1952])])).
fof(f4072,plain,(
  ( ! [X2,X3] : (~sP1(X2,X3) | ~sP0('F_0',X2)) ) | ($spl64 | $spl360 | $spl1866 | $spl1950)),
  inference(subsumption_resolution,[],[f4071,f4061])).
fof(f4071,plain,(
  ( ! [X2,X3] : (~sP1(X2,X3) | ~sP0('F_0',X2) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))))) ) | ($spl64 | $spl360 | $spl1866)),
  inference(subsumption_resolution,[],[f4067,f899])).
fof(f4067,plain,(
  ( ! [X2,X3] : (~sP1(X2,X3) | ~sP0('F_0',X2) | p(s(bool,X3)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',sK12))))) ) | ($spl64 | $spl1866)),
  inference(resolution,[],[f3953,f522])).
fof(f524,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP1(X0,X1)) ) | $spl68),
  inference(cnf_transformation,[],[f524_D])).
fof(f524_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP1(X0,X1)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f359,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f669,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~sP1('T_0',X0)) ) | $spl240),
  inference(cnf_transformation,[],[f669_D])).
fof(f669_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~sP1('T_0',X0)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f654,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~sP1('T_0',X0)) ) | ($spl2 | $spl68)),
  inference(resolution,[],[f524,f477])).
fof(f518,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~sP0(X0,X1) | p(s(bool,X0))) ) | $spl62),
  inference(cnf_transformation,[],[f518_D])).
fof(f518_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~sP0(X0,X1) | p(s(bool,X0))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f517,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~sP0(X0,X1) | p(s(bool,X1))) )),
  inference(global_subsumption,[],[f343,f266,f263])).
fof(f266,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.integer.INT_DIVIDES_|49||split|0','thm.bool.NOT_CLAUSES|split|1')).
fof(f343,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f206])).
fof(f686,plain,(
  ( ! [X2,X1] : (p(s(bool,X2)) | ~sP0(X2,X1) | ~sP1('T_0',X1)) ) | $spl262),
  inference(cnf_transformation,[],[f686_D])).
fof(f686_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X2)) | ~sP0(X2,X1) | ~sP1('T_0',X1)) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f680,plain,(
  ( ! [X2,X1] : (~sP1('T_0',X1) | ~sP0(X2,X1) | p(s(bool,X2))) ) | ($spl62 | $spl240)),
  inference(resolution,[],[f669,f518])).
fof(f863,plain,(
  ( ! [X0] : (~sP1('T_0',X0) | ~sP0('F_0',X0)) ) | $spl332),
  inference(cnf_transformation,[],[f863_D])).
fof(f863_D,plain,(
  ( ! [X0] : (~sP1('T_0',X0) | ~sP0('F_0',X0)) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f852,plain,(
  ( ! [X0] : (~sP0('F_0',X0) | ~sP1('T_0',X0)) ) | ($spl5 | $spl262)),
  inference(resolution,[],[f686,f480])).
fof(f477,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f477_D])).
fof(f477_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f263,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_DIVIDES_|49||split|0',aHL_TRUTH)).
fof(f878,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP0('F_0',X0)) ) | $spl354),
  inference(cnf_transformation,[],[f878_D])).
fof(f878_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP0('F_0',X0)) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f877,plain,(
  ( ! [X0] : (~sP0('F_0',X0) | p(s(bool,X0))) ) | ($spl2 | $spl66 | $spl332)),
  inference(subsumption_resolution,[],[f875,f477])).
fof(f875,plain,(
  ( ! [X0] : (~sP0('F_0',X0) | p(s(bool,X0)) | ~p(s(bool,'T_0'))) ) | ($spl66 | $spl332)),
  inference(resolution,[],[f863,f523])).
fof(f523,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl66),
  inference(cnf_transformation,[],[f523_D])).
fof(f523_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f360,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f1149,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl462),
  inference(cnf_transformation,[],[f1149_D])).
fof(f1149_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f4078,plain,(
  ( ! [X0] : (~sP0('F_0',X0)) ) | $spl1954),
  inference(cnf_transformation,[],[f4078_D])).
fof(f4078_D,plain,(
  ( ! [X0] : (~sP0('F_0',X0)) ) <=> ~$spl1954),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1954])])).
fof(f4077,plain,(
  ( ! [X0,X1] : (~sP0('F_0',X0) | p(s(bool,X1))) ) | ($spl66 | $spl354 | $spl1952)),
  inference(subsumption_resolution,[],[f4075,f878])).
fof(f4075,plain,(
  ( ! [X0,X1] : (~sP0('F_0',X0) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | ($spl66 | $spl1952)),
  inference(resolution,[],[f4073,f523])).
fof(f480,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f480_D])).
fof(f480_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f478,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f266,f263])).
fof(f4159,plain,(
  $false | ($spl5 | $spl462)),
  inference(resolution,[],[f1149,f480])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 277
% Passive clauses: 1004
% Generated clauses: 3921
% Final active clauses: 273
% Final passive clauses: 725
% Input formulas: 46
% Initial clauses: 203
% 
% Duplicate literals: 86
% Trivial inequalities: 7
% Fw subsumption resolutions: 288
% Fw demodulations: 41
% Bw demodulations: 4
% Condensations: 59
% Global subsumptions: 13
% 
% Simple tautologies: 195
% Equational tautologies: 4
% Forward subsumptions: 1073
% 
% Binary resolution: 1808
% Factoring: 145
% Forward superposition: 205
% Backward superposition: 42
% Self superposition: 3
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 1321
% SAT solver unit clauses: 188
% SAT solver binary clauses: 66
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 10
% 
% Sat splits: 4
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1268
% TWLsolver calls for satisfiability: 1777
% 
% Memory used [KB]: 4477
% Time elapsed: 0.149 s
% ------------------------------
% ------------------------------
% Success in time 0.818 s
