Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.lebesgue
% (57533)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 701
% Passive clauses: 8280
% Generated clauses: 18714
% Final active clauses: 693
% Final passive clauses: 7126
% Input formulas: 41
% Initial clauses: 181
% 
% Duplicate literals: 423
% Trivial inequalities: 17
% Fw subsumption resolutions: 1891
% Fw demodulations: 493
% Bw demodulations: 7
% 
% Simple tautologies: 234
% Equational tautologies: 60
% Forward subsumptions: 6785
% 
% Binary resolution: 9685
% Factoring: 1255
% Forward superposition: 2498
% Backward superposition: 2003
% Self superposition: 28
% Equality factoring: 30
% Equality resolution: 6
% 
% Split clauses: 569
% Split components: 1145
% SAT solver clauses: 994
% SAT solver unit clauses: 24
% SAT solver binary clauses: 167
% 
% Sat splits: 569
% 
% TWLsolver clauses: 993
% TWLsolver calls for satisfiability: 410
% 
% Memory used [KB]: 9466
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.lebesgue
% (57552)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 395
% Passive clauses: 1892
% Generated clauses: 9863
% Final active clauses: 383
% Final passive clauses: 1474
% Input formulas: 41
% Initial clauses: 181
% 
% Duplicate literals: 282
% Trivial inequalities: 129
% Fw subsumption resolutions: 925
% Fw demodulations: 313
% Bw demodulations: 29
% Condensations: 378
% Global subsumptions: 14
% 
% Simple tautologies: 292
% Equational tautologies: 19
% Forward subsumptions: 3318
% 
% Binary resolution: 2409
% Factoring: 161
% Forward superposition: 1516
% Backward superposition: 1529
% Self superposition: 24
% Equality factoring: 28
% Equality resolution: 38
% 
% Split clauses: 64
% Split components: 139
% SAT solver clauses: 2517
% SAT solver unit clauses: 205
% SAT solver binary clauses: 78
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 12
% 
% Sat splits: 64
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 2461
% TWLsolver calls for satisfiability: 1993
% 
% Memory used [KB]: 11513
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.lebesgue
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.lebesgue
% SZS output start Proof for thm.lebesgue
fof(f57851,plain,(
  $false),
  inference(global_subsumption,[],[f238,f237,f236,f235,f239,f246,f247,f255,f254,f288,f285,f298,f297,f428,f316,f313,f312,f311,f321,f327,f326,f325,f420,f323,f322,f419,f334,f333,f332,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f429,f430,f431,f432,f433,f434,f435,f436,f437,f438,f439,f440,f441,f442,f443,f444,f445,f446,f447,f448,f449,f450,f463,f462,f461,f460,f459,f458,f457,f456,f455,f454,f453,f452,f451,f476,f475,f474,f473,f472,f471,f470,f469,f468,f467,f466,f465,f464,f477,f481,f480,f488,f489,f492,f497,f506,f515,f520,f529,f534,f539,f544,f551,f556,f559,f562,f569,f570,f572,f579,f583,f584,f587,f593,f604,f603,f624,f623,f621,f620,f639,f713,f714,f715,f716,f717,f718,f719,f720,f721,f722,f733,f734,f747,f751,f750,f759,f758,f771,f769,f780,f778,f790,f800,f806,f811,f818,f824,f828,f836,f849,f855,f854,f853,f865,f871,f875,f899,f900,f910,f920,f928,f932,f940,f948,f952,f972,f971,f970,f969,f1008,f1009,f1010,f1011,f1012,f1013,f1014,f1015,f1016,f1017,f1018,f1019,f1020,f1021,f1022,f1039,f1040,f1063,f1068,f1072,f1078,f1085,f1086,f1090,f1103,f1117,f1118,f1122,f1359,f1358,f1357,f1356,f1355,f1354,f1353,f1352,f1351,f1350,f1349,f1348,f1347,f1346,f1345,f1344,f1343,f1342,f1341,f1340,f1339,f1334,f1333,f1329,f1326,f1325,f1319,f1318,f1311,f1290,f1289,f1288,f1284,f1283,f1279,f1278,f1277,f1270,f1259,f1254,f1252,f1251,f1248,f1247,f1246,f1245,f1244,f1241,f1239,f1234,f1233,f1227,f1226,f1219,f1217,f1216,f1215,f1214,f1209,f1207,f1204,f1203,f1202,f1201,f1200,f1199,f1198,f1197,f1196,f1195,f1194,f1193,f1187,f1186,f1179,f1178,f1172,f1171,f1164,f1163,f1162,f1158,f1157,f1154,f1152,f1151,f1148,f1137,f1136,f1135,f1134,f1132,f1131,f1130,f1128,f1399,f1636,f1635,f1634,f1633,f1632,f1631,f1630,f1629,f1628,f1627,f1626,f1625,f1624,f1623,f1622,f1621,f1620,f1619,f1618,f1617,f1616,f1615,f1610,f1609,f1605,f1602,f1601,f1595,f1594,f1587,f1566,f1565,f1564,f1560,f1559,f1555,f1554,f1553,f1546,f1535,f1530,f1528,f1527,f1524,f1523,f1522,f1521,f1520,f1517,f1515,f1510,f1509,f1503,f1502,f1495,f1493,f1492,f1491,f1490,f1485,f1483,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1463,f1462,f1455,f1454,f1448,f1447,f1440,f1439,f1438,f1434,f1433,f1430,f1428,f1427,f1424,f1413,f1412,f1411,f1410,f1408,f1407,f1406,f1404,f1678,f1686,f1714,f1710,f1706,f1705,f1701,f1698,f1692,f1691,f1688,f1720,f1722,f1746,f1745,f1742,f1741,f1740,f1784,f1783,f1776,f1773,f1814,f1813,f1812,f1811,f1904,f1912,f1924,f1938,f1950,f1960,f1976,f1988,f2000,f2009,f2014,f2059,f2068,f2066,f2064,f2078,f2083,f2089,f2086,f2098,f2107,f2119,f2306,f2317,f2315,f2313,f2311,f2327,f2332,f2340,f2339,f2335,f2349,f2359,f2372,f2564,f2614,f2615,f2616,f2617,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2648,f2698,f2699,f2700,f2701,f2703,f2704,f2705,f2706,f2707,f2708,f2709,f2710,f2711,f2712,f2713,f2714,f2715,f2716,f2717,f2758,f2829,f2880,f3197,f3201,f3204,f3207,f3211,f3214,f3217,f3220,f3223,f3226,f3229,f3232,f3235,f3238,f3241,f3244,f3247,f3250,f3253,f3255,f3256,f3257,f3258,f3259,f3260,f3261,f3262,f3263,f3264,f3265,f3266,f3267,f3268,f3269,f3270,f3271,f3272,f3273,f3274,f3275,f3276,f3277,f3278,f2959,f3288,f3290,f3292,f3294,f3297,f3299,f3301,f3303,f3305,f3307,f3309,f3311,f3313,f3315,f3317,f3319,f3321,f3323,f3325,f2768,f4417,f4466,f4628,f4642,f4661,f4679,f4691,f4718,f4728,f4734,f4748,f4754,f4771,f4776,f4893,f4907,f4954,f4953,f4952,f4951,f4980,f4979,f4978,f4977,f4975,f4973,f4972,f5052,f5211,f5208,f5238,f5235,f5233,f5230,f5390,f5391,f5392,f5393,f5394,f5426,f5427,f5916,f5997,f6000,f6001,f6002,f8874,f8876,f8877,f8878,f9775,f9776,f9777,f9780,f11681,f11684,f11688,f11689,f11690,f11692,f11962,f11965,f11970,f11971,f11972,f11974,f12021,f12364,f12363,f12362,f12361,f12360,f12359,f12358,f12357,f12356,f12355,f12354,f12353,f12352,f12351,f12350,f12349,f12348,f12347,f12346,f12345,f12344,f12341,f12874,f12873,f12872,f12871,f12870,f12869,f12868,f12867,f12866,f12865,f12864,f12863,f12862,f12861,f12860,f12859,f12858,f12857,f12856,f12855,f12854,f12851,f14727,f15178,f15173,f15204,f15242,f15237,f17912,f15447,f15448,f17758,f15535,f15784,f18224,f18228,f18234,f18236,f18270,f16058,f18272,f18287,f18289,f18290,f16198,f18320,f18321,f18322,f18323,f18324,f18325,f18326,f18327,f18328,f18329,f18330,f18331,f18332,f18333,f18334,f18335,f18336,f18337,f18338,f18339,f18340,f18341,f18342,f18343,f18344,f18345,f16381,f18372,f18373,f18406,f18449,f16881,f20503,f20508,f20588,f20595,f18361,f18360,f18359,f18358,f18357,f18356,f18355,f18354,f18353,f18352,f15562,f15531,f15503,f12759,f11954,f11673,f8869,f5992,f5609,f5365,f5329,f5305,f5277,f5261,f5245,f5223,f5205,f3281,f2743,f1848,f1847,f1846,f1845,f1844,f1843,f1842,f1841,f1840,f1839,f1838,f1837,f1836,f1835,f1834,f1833,f1832,f1831,f1829,f1828,f1827,f1826,f427,f281,f18269,f16346,f16160,f16023,f15749,f5441,f2957,f2827,f18346,f12036,f11973,f11691,f9778,f5933,f5189,f5188,f5187,f5186,f5185,f5184,f5183,f5182,f5181,f5180,f5179,f5178,f5177,f5176,f5175,f5174,f5173,f5172,f5171,f5170,f5169,f5168,f5167,f5165,f1876,f1875,f1816,f426,f259,f258,f30344,f30340,f30317,f30299,f29273,f29276,f29277,f30267,f29313,f30584,f30585,f29775,f29786,f30699,f30069,f30080,f30081,f30109,f18362,f3282,f295,f293,f18351,f18350,f3280,f3279,f284,f283,f2605,f2608,f2689,f2692,f32104,f32120,f4473,f4453,f4424,f4376,f20259,f20257,f20000,f19775,f18854,f18638,f16366,f16365,f16364,f16363,f16362,f16361,f16360,f16359,f16358,f16357,f16180,f16179,f16178,f16177,f16176,f16175,f16174,f16173,f16172,f16171,f16043,f16042,f16041,f16040,f16039,f16038,f16037,f16036,f16035,f16034,f15769,f15768,f15767,f15766,f15765,f15764,f15763,f15762,f15761,f15760,f12017,f9758,f9429,f5527,f5525,f5523,f5499,f5497,f5495,f5065,f5063,f5062,f5061,f5060,f5059,f1815,f4605,f4607,f4609,f4620,f4622,f4624,f5056,f5107,f5109,f15756,f16030,f16167,f16353,f37298,f37225,f39463,f36241,f36238,f36199,f36164,f16185,f16184,f16183,f14696,f14694,f14669,f14667,f14644,f14642,f14618,f14616,f14592,f14590,f14517,f14515,f14442,f14440,f14367,f14365,f14294,f14292,f14205,f12194,f12071,f5934,f5086,f5083,f5082,f5079,f5078,f5077,f5076,f5075,f4592,f4590,f4579,f4577,f4568,f4566,f4553,f4551,f4540,f4538,f4527,f4525,f4516,f4514,f4499,f4486,f34960,f34784,f34677,f34467,f34306,f34194,f18296,f18295,f16194,f16193,f16192,f16191,f16190,f16189,f16188,f16187,f16186,f16170,f14203,f12192,f12069,f11955,f11674,f8870,f5993,f5088,f5087,f5074,f4501,f4488,f40436,f41847,f41846,f39469,f39467,f39466,f39464,f39462,f39460,f39459,f39457,f37297,f37296,f37292,f37224,f37223,f37219,f36357,f36240,f36237,f36236,f36232,f36198,f36197,f36193,f36163,f36162,f36158,f34959,f34958,f34954,f34783,f34782,f34778,f34676,f34675,f34671,f34466,f34465,f34461,f34305,f34304,f34300,f34193,f34192,f34188,f30070,f29776,f20500,f20426,f20419,f19869,f19649,f18726,f18514,f18294,f18293,f18292,f18291,f18288,f16380,f16379,f16378,f16377,f16376,f16375,f16374,f16373,f16372,f16371,f16370,f16369,f16368,f16367,f16356,f16355,f16354,f16352,f16351,f16350,f16349,f16348,f16182,f16181,f16057,f16056,f16055,f16054,f16053,f16052,f16051,f16050,f16049,f16048,f16047,f16046,f16045,f16044,f16033,f16032,f16031,f16029,f16028,f16027,f16026,f16025,f15783,f15782,f15781,f15780,f15779,f15778,f15777,f15776,f15775,f15774,f15773,f15772,f15771,f15770,f15759,f15758,f15757,f15755,f15754,f15753,f15752,f15751,f15563,f15532,f15504,f14695,f14668,f14643,f14617,f14591,f14516,f14441,f14366,f14293,f14204,f12193,f12070,f9259,f9257,f9255,f9229,f9227,f9225,f5054,f4588,f4575,f4564,f4549,f4536,f4523,f4512,f4497,f4484,f4471,f4469,f4451,f4449,f4422,f4420,f4374,f4372,f2958,f2828,f47734,f49141,f48777,f48819,f49243,f49246,f49247,f49248,f49249,f49250,f49251,f49252,f49253,f49254,f49255,f49256,f49257,f49258,f49259,f49260,f49261,f49648,f49633,f49661,f49662,f49663,f49664,f49665,f49666,f49667,f49668,f49669,f49670,f49671,f49672,f49673,f49674,f49675,f49676,f49677,f49678,f49679,f49680,f49681,f49682,f49683,f49684,f49686,f49687,f49688,f49800,f49805,f49807,f49912,f49919,f49921,f51085,f51086,f51088,f51118,f51119,f51121,f51250,f51251,f52276,f52274,f52272,f52506,f52505,f52504,f52503,f52502,f52500,f52499,f52498,f52497,f52496,f52494,f52493,f52492,f52491,f52490,f52488,f52487,f52486,f52485,f52484,f52482,f52481,f52480,f52479,f52555,f52554,f52553,f52552,f52551,f52549,f52548,f52547,f52546,f52545,f52543,f52542,f52541,f52540,f52539,f52537,f52536,f52535,f52534,f52533,f52531,f52530,f52529,f52528,f52663,f52664,f52666,f52700,f52701,f52703,f55592,f55594,f55589,f56126,f56125,f56119,f56118,f56117,f56116,f56110,f56104,f56098,f56097,f56096,f56091,f56090,f56084,f56083,f56082,f56081,f56075,f56069,f56063,f56062,f56061,f56055,f56051,f56045,f56039,f56038,f56037,f56036,f56035,f56029,f56023,f56017,f56016,f56015,f56014,f56008,f56002,f55996,f55995,f55994,f55993,f55992,f55991,f55990,f55989,f55988,f55987,f55986,f55985,f55984,f55983,f55982,f55981,f55980,f55979,f55978,f55977,f55976,f55975,f55974,f55973,f55972,f55971,f55963,f55954,f55953,f55945,f55944,f55897,f55896,f55895,f55888,f55881,f55880,f55869,f55837,f55815,f55813,f55812,f55808,f55807,f55806,f55805,f55802,f55800,f55795,f55794,f55786,f55785,f55772,f55771,f55770,f55769,f55768,f55767,f55766,f55765,f55758,f55756,f55754,f55752,f55749,f55748,f55744,f55740,f55734,f55733,f55732,f55731,f55730,f55724,f55723,f55722,f55721,f55720,f55719,f55718,f55717,f55716,f55708,f55707,f55697,f55689,f55688,f55678,f55676,f55674,f55672,f55670,f55668,f55666,f55665,f55664,f55663,f55662,f55660,f55659,f55658,f55657,f55656,f55654,f55653,f55652,f55651,f55650,f55648,f55647,f55646,f55645,f55644,f55643,f55639,f55638,f55633,f55632,f55618,f55617,f55616,f55615,f55613,f55612,f55611,f55604,f57844,f57846,f57839])).
fof(f57839,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',sK17))))),
  inference(subsumption_resolution,[],[f57818,f239])).
fof(f57818,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',sK17))))),
  inference(superposition,[],[f51118,f1128])).
fof(f57846,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X28)))) | s(bool,X29) = s(bool,X30) | s(bool,X30) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X28),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f57830,f239])).
fof(f57830,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X28)))) | s(bool,X29) = s(bool,X30) | s(bool,X30) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X28),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f51118,f2758])).
fof(f57844,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X1)))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X1),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f57821,f239])).
fof(f57821,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X1)))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X1),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f51118,f639])).
fof(f55604,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',sK17))))),
  inference(unit_resulting_resolution,[],[f238,f55589,f293])).
fof(f55611,plain,(
  ( ! [X0] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f297])).
fof(f55612,plain,(
  ( ! [X0] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f55589,f298])).
fof(f55613,plain,(
  ~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f55589,f312])).
fof(f55615,plain,(
  ~sP1('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f239,f55589,f312])).
fof(f55616,plain,(
  ~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f55589,f55589,f312])).
fof(f55617,plain,(
  sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f55589,f313])).
fof(f55618,plain,(
  sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f426,f55589,f316])).
fof(f55632,plain,(
  ( ! [X0] : (sP2('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f239,f55589,f327])).
fof(f55633,plain,(
  ( ! [X0] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f55589,f327])).
fof(f55638,plain,(
  ( ! [X0] : (sP3(X0,'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f239,f55589,f334])).
fof(f55639,plain,(
  ( ! [X0] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f55589,f334])).
fof(f55643,plain,(
  ~sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f55589,f345])).
fof(f55644,plain,(
  ~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0')),
  inference(unit_resulting_resolution,[],[f55589,f426,f55589,f345])).
fof(f55645,plain,(
  sP7('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f438,f434,f431,f55589,f346])).
fof(f55646,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f438,f467,f431,f55589,f346])).
fof(f55647,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f438,f1193,f431,f55589,f346])).
fof(f55648,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f438,f1469,f431,f55589,f346])).
fof(f55650,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f455,f454,f431,f55589,f346])).
fof(f55651,plain,(
  sP7('T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f469,f434,f466,f55589,f346])).
fof(f55652,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f469,f467,f466,f55589,f346])).
fof(f55653,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f469,f1193,f466,f55589,f346])).
fof(f55654,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f469,f1469,f466,f55589,f346])).
fof(f55656,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f455,f454,f466,f55589,f346])).
fof(f55657,plain,(
  sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f1195,f434,f1178,f55589,f346])).
fof(f55658,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f1195,f467,f1178,f55589,f346])).
fof(f55659,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f1195,f1193,f1178,f55589,f346])).
fof(f55660,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f1195,f1469,f1178,f55589,f346])).
fof(f55662,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f455,f454,f1178,f55589,f346])).
fof(f55663,plain,(
  sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f1471,f434,f1454,f55589,f346])).
fof(f55664,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f1471,f467,f1454,f55589,f346])).
fof(f55665,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f1471,f1193,f1454,f55589,f346])).
fof(f55666,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f1471,f1469,f1454,f55589,f346])).
fof(f55668,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f455,f454,f1454,f55589,f346])).
fof(f55670,plain,(
  ( ! [X0] : (sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f437,f434,f453,f55589,f346])).
fof(f55672,plain,(
  ( ! [X0] : (sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f468,f467,f453,f55589,f346])).
fof(f55674,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f1194,f1193,f453,f55589,f346])).
fof(f55676,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f1470,f1469,f453,f55589,f346])).
fof(f55678,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f455,f454,f453,f55589,f346])).
fof(f55688,plain,(
  ~sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f55589,f349])).
fof(f55689,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0')),
  inference(unit_resulting_resolution,[],[f55589,f426,f55589,f349])).
fof(f55697,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f350])).
fof(f55707,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f55589,f353])).
fof(f55708,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0')),
  inference(unit_resulting_resolution,[],[f55589,f426,f55589,f353])).
fof(f55716,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f354])).
fof(f55717,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f358])).
fof(f55718,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X1)) )),
  inference(unit_resulting_resolution,[],[f55589,f359])).
fof(f55719,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f55589,f360])).
fof(f55720,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f481,f55589,f362])).
fof(f55721,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1128,f55589,f362])).
fof(f55722,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1404,f55589,f362])).
fof(f55723,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f480,f55589,f362])).
fof(f55724,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f281,f55589,f362])).
fof(f55730,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f481,f55589,f362])).
fof(f55731,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1128,f55589,f362])).
fof(f55732,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1404,f55589,f362])).
fof(f55733,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f480,f55589,f362])).
fof(f55734,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f281,f55589,f362])).
fof(f55740,plain,(
  ~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f477,f55589,f363])).
fof(f55744,plain,(
  ~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f477,f55589,f363])).
fof(f55748,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f55589,f366])).
fof(f55749,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f426,f55589,f367])).
fof(f55752,plain,(
  sP9('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f471,f55589,f368])).
fof(f55754,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f1233,f55589,f368])).
fof(f55756,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f1509,f55589,f368])).
fof(f55758,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f458,f55589,f368])).
fof(f55765,plain,(
  sP9('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f239,f471,f55589,f368])).
fof(f55766,plain,(
  sP9('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f55589,f471,f55589,f368])).
fof(f55767,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f239,f1233,f55589,f368])).
fof(f55768,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f55589,f1233,f55589,f368])).
fof(f55769,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f239,f1509,f55589,f368])).
fof(f55770,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f55589,f1509,f55589,f368])).
fof(f55771,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f239,f458,f55589,f368])).
fof(f55772,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f458,f55589,f368])).
fof(f55785,plain,(
  ~sP8('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f55589,f372])).
fof(f55786,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0')),
  inference(unit_resulting_resolution,[],[f55589,f426,f55589,f372])).
fof(f55794,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f373])).
fof(f55795,plain,(
  sP9('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f239,f55589,f378])).
fof(f55800,plain,(
  sP9('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f239,f239,f55589,f378])).
fof(f55802,plain,(
  sP9('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f239,f55589,f55589,f378])).
fof(f55805,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f239,f55589,f378])).
fof(f55806,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0')),
  inference(unit_resulting_resolution,[],[f55589,f239,f55589,f378])).
fof(f55807,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f239,f55589,f55589,f378])).
fof(f55808,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f55589,f55589,f55589,f378])).
fof(f55812,plain,(
  ~sP9('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f239,f426,f55589,f381])).
fof(f55813,plain,(
  ~sP9('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f55589,f426,f55589,f381])).
fof(f55815,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f426,f55589,f382])).
fof(f55837,plain,(
  ( ! [X0] : (~sP11('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f426,f55589,f383])).
fof(f55869,plain,(
  ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f426,f55589,f389])).
fof(f55880,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f390])).
fof(f55881,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X1)) )),
  inference(unit_resulting_resolution,[],[f55589,f391])).
fof(f55888,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f55589,f394])).
fof(f55895,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f55589,f55589,f395])).
fof(f55896,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f426,f55589,f396])).
fof(f55897,plain,(
  ( ! [X0] : (~sP13('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f426,f55589,f400])).
fof(f55944,plain,(
  ~sP12('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f426,f239,f55589,f406])).
fof(f55945,plain,(
  ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f426,f55589,f55589,f406])).
fof(f55953,plain,(
  ( ! [X0,X1] : (sP12(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X1)) )),
  inference(unit_resulting_resolution,[],[f55589,f408])).
fof(f55954,plain,(
  sP13('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f426,f426,f55589,f410])).
fof(f55963,plain,(
  ( ! [X0] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f55589,f55589,f412])).
fof(f55971,plain,(
  ~sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f55589,f413])).
fof(f55972,plain,(
  ~sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f55589,f426,f55589,f413])).
fof(f55973,plain,(
  ( ! [X0] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f446,f55589,f416])).
fof(f55974,plain,(
  ( ! [X0] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f447,f55589,f416])).
fof(f55975,plain,(
  ( ! [X0] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f474,f55589,f416])).
fof(f55976,plain,(
  ( ! [X0] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f1325,f55589,f416])).
fof(f55977,plain,(
  ( ! [X0] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f1601,f55589,f416])).
fof(f55978,plain,(
  ( ! [X0,X1] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f461,f55589,f416])).
fof(f55979,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f443,f55589,f417])).
fof(f55980,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f472,f55589,f417])).
fof(f55981,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f1277,f55589,f417])).
fof(f55982,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f1553,f55589,f417])).
fof(f55983,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f459,f55589,f417])).
fof(f55984,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f444,f55589,f417])).
fof(f55985,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f473,f55589,f417])).
fof(f55986,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f1278,f55589,f417])).
fof(f55987,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f1554,f55589,f417])).
fof(f55988,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f460,f55589,f417])).
fof(f55989,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f55589,f419])).
fof(f55990,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X1)) )),
  inference(unit_resulting_resolution,[],[f55589,f420])).
fof(f55991,plain,(
  ( ! [X0] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f55589,f715])).
fof(f55992,plain,(
  ( ! [X0] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f55589,f716])).
fof(f55993,plain,(
  ( ! [X0] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f717])).
fof(f55994,plain,(
  ( ! [X0] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f55589,f718])).
fof(f55995,plain,(
  ( ! [X0] : (sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f55589,f720])).
fof(f55996,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f55589,f806])).
fof(f56002,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f55589,f824])).
fof(f56008,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f55589,f871])).
fof(f56014,plain,(
  ~sP9('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f426,f55589,f899])).
fof(f56015,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f55589,f1010])).
fof(f56016,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f55589,f1011])).
fof(f56017,plain,(
  ( ! [X0] : (~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f55589,f1012])).
fof(f56023,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f55589,f1013])).
fof(f56029,plain,(
  ( ! [X0] : (~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f55589,f1014])).
fof(f56035,plain,(
  ( ! [X0] : (~sP9('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f426,f55589,f1015])).
fof(f56036,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f1017])).
fof(f56037,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X1)) )),
  inference(unit_resulting_resolution,[],[f55589,f1018])).
fof(f56038,plain,(
  ( ! [X0,X1] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X1)) )),
  inference(unit_resulting_resolution,[],[f55589,f1020])).
fof(f56039,plain,(
  ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f426,f55589,f1085])).
fof(f56045,plain,(
  ( ! [X0] : (~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f426,f55589,f1086])).
fof(f56051,plain,(
  ~sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f426,f55589,f1117])).
fof(f56055,plain,(
  ( ! [X0] : (~sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f426,f55589,f1118])).
fof(f56061,plain,(
  ( ! [X0] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f55589,f2616])).
fof(f56062,plain,(
  ( ! [X0] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f55589,f2617])).
fof(f56063,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f55589,f2620])).
fof(f56069,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f55589,f2621])).
fof(f56075,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f55589,f2622])).
fof(f56081,plain,(
  ~sP9('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f55589,f2623])).
fof(f56082,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f2625])).
fof(f56083,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f55589,f2626])).
fof(f56084,plain,(
  ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f55589,f2628])).
fof(f56090,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f55589,f2630])).
fof(f56091,plain,(
  ~sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f426,f55589,f2631])).
fof(f56096,plain,(
  ( ! [X0] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f55589,f2700])).
fof(f56097,plain,(
  ( ! [X0] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f55589,f2701])).
fof(f56098,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f55589,f2704])).
fof(f56104,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f55589,f2705])).
fof(f56110,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f55589,f2706])).
fof(f56116,plain,(
  ~sP9('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f55589,f2707])).
fof(f56117,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f55589,f2709])).
fof(f56118,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f55589,f2710])).
fof(f56119,plain,(
  ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),'F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f55589,f2712])).
fof(f56125,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f55589,f2714])).
fof(f56126,plain,(
  ~sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f426,f55589,f2715])).
fof(f55589,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),
  inference(subsumption_resolution,[],[f55567,f239])).
fof(f55567,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),
  inference(superposition,[],[f51085,f1404])).
fof(f55594,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X28)))) | s(bool,X29) = s(bool,X30) | s(bool,X30) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X28),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f55578,f239])).
fof(f55578,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X28)))) | s(bool,X29) = s(bool,X30) | s(bool,X30) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X28),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f51085,f2758])).
fof(f55592,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X1)))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X1),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f55569,f239])).
fof(f55569,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X1)))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X1),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f51085,f639])).
fof(f52703,plain,(
  ( ! [X103,X101,X109,X107,X105,X102,X100,X108,X106,X104] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X102),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101),s(fun(X100,'type.extreal.extreal'),X103)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X104),s('type.extreal.extreal',X102)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X100,'type.extreal.extreal'),X105),s(X100,sK19(X102,X104,X103,X105,X101,X100)))),s('type.extreal.extreal',happ(s(fun(X100,'type.extreal.extreal'),X103),s(X100,sK19(X102,X104,X103,X105,X101,X100))))))) | s(bool,X106) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X104),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101),s(fun(X100,'type.extreal.extreal'),X105))))) | s(bool,'T_0') = s(bool,X106) | s(bool,X107) = s(bool,X108) | s(bool,X108) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101))) | s(bool,X107) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f52689,f239])).
fof(f52689,plain,(
  ( ! [X103,X101,X109,X107,X105,X102,X100,X108,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X102),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101),s(fun(X100,'type.extreal.extreal'),X103)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X104),s('type.extreal.extreal',X102)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X100,'type.extreal.extreal'),X105),s(X100,sK19(X102,X104,X103,X105,X101,X100)))),s('type.extreal.extreal',happ(s(fun(X100,'type.extreal.extreal'),X103),s(X100,sK19(X102,X104,X103,X105,X101,X100))))))) | s(bool,X106) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X104),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101),s(fun(X100,'type.extreal.extreal'),X105))))) | s(bool,'T_0') = s(bool,X106) | s(bool,X107) = s(bool,X108) | s(bool,X108) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101))) | s(bool,X107) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f3279,f2758])).
fof(f52701,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X21),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20),s(fun(X19,'type.extreal.extreal'),X22)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',X21)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X19,'type.extreal.extreal'),X24),s(X19,sK19(X21,X23,X22,X24,X20,X19)))),s('type.extreal.extreal',happ(s(fun(X19,'type.extreal.extreal'),X22),s(X19,sK19(X21,X23,X22,X24,X20,X19))))))) | s(bool,X25) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X23),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20),s(fun(X19,'type.extreal.extreal'),X24))))) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f52680,f239])).
fof(f52680,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X21),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20),s(fun(X19,'type.extreal.extreal'),X22)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',X21)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X19,'type.extreal.extreal'),X24),s(X19,sK19(X21,X23,X22,X24,X20,X19)))),s('type.extreal.extreal',happ(s(fun(X19,'type.extreal.extreal'),X22),s(X19,sK19(X21,X23,X22,X24,X20,X19))))))) | s(bool,X25) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X23),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20),s(fun(X19,'type.extreal.extreal'),X24))))) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f3279,f639])).
fof(f52700,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X5) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X6),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X7))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X8),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X9)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X7),s(sK14,sK19(X8,X6,X9,X7,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X9),s(sK14,sK19(X8,X6,X9,X7,sK15,sK14))))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(duplicate_literal_removal,[],[f52699])).
fof(f52699,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X5) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X6),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X7))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X8),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X9)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X7),s(sK14,sK19(X8,X6,X9,X7,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X9),s(sK14,sK19(X8,X6,X9,X7,sK15,sK14))))))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(condensation,[],[f52698])).
fof(f52698,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X1),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X2)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X4),s(sK14,sK19(X1,X3,X2,X4,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X2),s(sK14,sK19(X1,X3,X2,X4,sK15,sK14))))))) | s(bool,X5) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X4))))) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f52677,f733])).
fof(f52677,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X1),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X2)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X4),s(sK14,sK19(X1,X3,X2,X4,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X2),s(sK14,sK19(X1,X3,X2,X4,sK15,sK14))))))) | s(bool,X5) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X4))))) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f3279,f9775])).
fof(f52666,plain,(
  ( ! [X103,X101,X109,X107,X105,X102,X100,X108,X106,X104] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X102),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101),s(fun(X100,'type.extreal.extreal'),X103)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X104),s('type.extreal.extreal',X102)))) | p(s(bool,'const.bool.IN_2'(s(X100,sK19(X102,X104,X103,X105,X101,X100)),s(fun(X100,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101)))))) | s(bool,X106) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X104),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101),s(fun(X100,'type.extreal.extreal'),X105))))) | s(bool,'T_0') = s(bool,X106) | s(bool,X107) = s(bool,X108) | s(bool,X108) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101))) | s(bool,X107) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f52652,f239])).
fof(f52652,plain,(
  ( ! [X103,X101,X109,X107,X105,X102,X100,X108,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X102),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101),s(fun(X100,'type.extreal.extreal'),X103)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X104),s('type.extreal.extreal',X102)))) | p(s(bool,'const.bool.IN_2'(s(X100,sK19(X102,X104,X103,X105,X101,X100)),s(fun(X100,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101)))))) | s(bool,X106) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X104),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101),s(fun(X100,'type.extreal.extreal'),X105))))) | s(bool,'T_0') = s(bool,X106) | s(bool,X107) = s(bool,X108) | s(bool,X108) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X100,bool),'type.pair.prod'(fun(fun(X100,bool),bool),fun(fun(X100,bool),'type.realax.real'))),X101))) | s(bool,X107) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f3280,f2758])).
fof(f52664,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X21),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20),s(fun(X19,'type.extreal.extreal'),X22)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',X21)))) | p(s(bool,'const.bool.IN_2'(s(X19,sK19(X21,X23,X22,X24,X20,X19)),s(fun(X19,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20)))))) | s(bool,X25) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X23),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20),s(fun(X19,'type.extreal.extreal'),X24))))) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f52643,f239])).
fof(f52643,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X21),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20),s(fun(X19,'type.extreal.extreal'),X22)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',X21)))) | p(s(bool,'const.bool.IN_2'(s(X19,sK19(X21,X23,X22,X24,X20,X19)),s(fun(X19,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20)))))) | s(bool,X25) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X23),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20),s(fun(X19,'type.extreal.extreal'),X24))))) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X19,bool),'type.pair.prod'(fun(fun(X19,bool),bool),fun(fun(X19,bool),'type.realax.real'))),X20))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f3280,f639])).
fof(f52663,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X5) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X6),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X7))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X8),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X9)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X8,X6,X9,X7,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(duplicate_literal_removal,[],[f52662])).
fof(f52662,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X5) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X6),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X7))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X8),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X9)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X8,X6,X9,X7,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(condensation,[],[f52661])).
fof(f52661,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X1),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X2)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X1,X3,X2,X4,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | s(bool,X5) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X4))))) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f52640,f733])).
fof(f52640,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X1),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X2)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X1,X3,X2,X4,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | s(bool,X5) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X4))))) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f3280,f9775])).
fof(f52528,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f467,f431,f2703])).
fof(f52529,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1193,f431,f2703])).
fof(f52530,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1469,f431,f2703])).
fof(f52531,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f454,f431,f2703])).
fof(f52533,plain,(
  sP7('T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f434,f466,f2703])).
fof(f52534,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f467,f466,f2703])).
fof(f52535,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1193,f466,f2703])).
fof(f52536,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1469,f466,f2703])).
fof(f52537,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f454,f466,f2703])).
fof(f52539,plain,(
  sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f434,f1178,f2703])).
fof(f52540,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f467,f1178,f2703])).
fof(f52541,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1193,f1178,f2703])).
fof(f52542,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1469,f1178,f2703])).
fof(f52543,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f454,f1178,f2703])).
fof(f52545,plain,(
  sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f434,f1454,f2703])).
fof(f52546,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f467,f1454,f2703])).
fof(f52547,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1193,f1454,f2703])).
fof(f52548,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1469,f1454,f2703])).
fof(f52549,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f454,f1454,f2703])).
fof(f52551,plain,(
  ( ! [X0] : (sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f434,f453,f2703])).
fof(f52552,plain,(
  ( ! [X0] : (sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f467,f453,f2703])).
fof(f52553,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f1193,f453,f2703])).
fof(f52554,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f1469,f453,f2703])).
fof(f52555,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f454,f453,f2703])).
fof(f52479,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f467,f431,f2619])).
fof(f52480,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1193,f431,f2619])).
fof(f52481,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1469,f431,f2619])).
fof(f52482,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f454,f431,f2619])).
fof(f52484,plain,(
  sP7('T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f434,f466,f2619])).
fof(f52485,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f467,f466,f2619])).
fof(f52486,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1193,f466,f2619])).
fof(f52487,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1469,f466,f2619])).
fof(f52488,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f454,f466,f2619])).
fof(f52490,plain,(
  sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f434,f1178,f2619])).
fof(f52491,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f467,f1178,f2619])).
fof(f52492,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1193,f1178,f2619])).
fof(f52493,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1469,f1178,f2619])).
fof(f52494,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f454,f1178,f2619])).
fof(f52496,plain,(
  sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f434,f1454,f2619])).
fof(f52497,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f467,f1454,f2619])).
fof(f52498,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1193,f1454,f2619])).
fof(f52499,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1469,f1454,f2619])).
fof(f52500,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f454,f1454,f2619])).
fof(f52502,plain,(
  ( ! [X0] : (sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f434,f453,f2619])).
fof(f52503,plain,(
  ( ! [X0] : (sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f467,f453,f2619])).
fof(f52504,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f1193,f453,f2619])).
fof(f52505,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f1469,f453,f2619])).
fof(f52506,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f454,f453,f2619])).
fof(f52272,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(fun(X91,bool),X92) = s(fun(X91,bool),X93) | s(bool,X94) = s(bool,g(s(X91,sK20(X92,X93,bool,X91)))) | s(bool,'T_0') = s(bool,X94) | s(bool,X95) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X95) = s(bool,f(s(X91,sK20(X92,X93,bool,X91))))) )),
  inference(trivial_inequality_removal,[],[f52268])).
fof(f52268,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X91,bool),X92) = s(fun(X91,bool),X93) | s(bool,X94) = s(bool,g(s(X91,sK20(X92,X93,bool,X91)))) | s(bool,'T_0') = s(bool,X94) | s(bool,X95) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X95) = s(bool,f(s(X91,sK20(X92,X93,bool,X91))))) )),
  inference(superposition,[],[f2957,f9775])).
fof(f52274,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (s(fun(X58,bool),X59) = s(fun(X58,bool),X60) | s(bool,X61) = s(bool,g(s(X58,sK20(X59,X60,bool,X58)))) | s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,X63) | s(bool,X63) = s(bool,f(s(X58,sK20(X59,X60,bool,X58)))) | s(bool,X62) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(trivial_inequality_removal,[],[f52263])).
fof(f52263,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X58,bool),X59) = s(fun(X58,bool),X60) | s(bool,X61) = s(bool,g(s(X58,sK20(X59,X60,bool,X58)))) | s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,X63) | s(bool,X63) = s(bool,f(s(X58,sK20(X59,X60,bool,X58)))) | s(bool,X62) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f2957,f2758])).
fof(f52276,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | s(bool,X7) = s(bool,g(s(X4,sK20(X5,X6,bool,X4)))) | s(bool,'T_0') = s(bool,X7) | s(bool,X8) = s(bool,f(s(X4,sK20(X5,X6,bool,X4)))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(trivial_inequality_removal,[],[f52254])).
fof(f52254,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | s(bool,X7) = s(bool,g(s(X4,sK20(X5,X6,bool,X4)))) | s(bool,'T_0') = s(bool,X7) | s(bool,X8) = s(bool,f(s(X4,sK20(X5,X6,bool,X4)))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2957,f639])).
fof(f51251,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (s(X7,f(s(bool,X6))) != s(X7,g(s(bool,X6))) | s(fun(bool,X7),X8) = s(fun(bool,X7),X9) | s(bool,X10) = s(bool,sK20(X8,X9,X7,bool)) | s(bool,'T_0') = s(bool,X10) | s(bool,X11) = s(bool,X12) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,X6) = s(bool,X11)) )),
  inference(superposition,[],[f2959,f2758])).
fof(f51250,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(bool,X4) = s(bool,sK20(X2,X3,X1,bool)) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f2959,f9775])).
fof(f51121,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X39),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X40)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X39)))) | s(bool,X41) = s(bool,X42) | s(bool,X42) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X39,sK18,X40,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X40),s(sK14,sK19(X39,sK18,X40,sK16,sK15,sK14)))))) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f51109,f239])).
fof(f51109,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X39),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X40)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X39)))) | s(bool,X41) = s(bool,X42) | s(bool,X42) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X39,sK18,X40,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X40),s(sK14,sK19(X39,sK18,X40,sK16,sK15,sK14)))))) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f2689,f2758])).
fof(f51119,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X4)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X3)))) | s(bool,X5) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X3,sK18,X4,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X4),s(sK14,sK19(X3,sK18,X4,sK16,sK15,sK14)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f51100,f239])).
fof(f51100,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X4)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X3)))) | s(bool,X5) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X3,sK18,X4,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X4),s(sK14,sK19(X3,sK18,X4,sK16,sK15,sK14)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2689,f639])).
fof(f51118,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X0))))) )),
  inference(subsumption_resolution,[],[f51098,f239])).
fof(f51098,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X0))))) )),
  inference(superposition,[],[f2689,f480])).
fof(f51088,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X39),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X40)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X39)))) | s(bool,X41) = s(bool,X42) | s(bool,X42) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X39,sK17,X40,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X40),s(sK14,sK19(X39,sK17,X40,sK16,sK15,sK14)))))) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f51076,f239])).
fof(f51076,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X39),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X40)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X39)))) | s(bool,X41) = s(bool,X42) | s(bool,X42) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X39,sK17,X40,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X40),s(sK14,sK19(X39,sK17,X40,sK16,sK15,sK14)))))) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f2605,f2758])).
fof(f51086,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X4)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X3)))) | s(bool,X5) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X3,sK17,X4,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X4),s(sK14,sK19(X3,sK17,X4,sK16,sK15,sK14)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f51067,f239])).
fof(f51067,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X4)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X3)))) | s(bool,X5) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X3,sK17,X4,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X4),s(sK14,sK19(X3,sK17,X4,sK16,sK15,sK14)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2605,f639])).
fof(f51085,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X0))))) )),
  inference(subsumption_resolution,[],[f51065,f239])).
fof(f51065,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X0))))) )),
  inference(superposition,[],[f2605,f480])).
fof(f49921,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X38,sK18,X39,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X38)))) | s(bool,X40) = s(bool,X41) | s(bool,X41) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X38),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X39))))) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f49903,f239])).
fof(f49903,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X38,sK18,X39,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X38)))) | s(bool,X40) = s(bool,X41) | s(bool,X41) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X38),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X39))))) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f2692,f2758])).
fof(f49919,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK18,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X2)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f49894,f239])).
fof(f49894,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK18,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X2)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f2692,f639])).
fof(f49912,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK19(sK17,sK18,sK16,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',sK17))))),
  inference(subsumption_resolution,[],[f49891,f239])).
fof(f49891,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(sK17,sK18,sK16,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',sK17))))),
  inference(superposition,[],[f2692,f1128])).
fof(f49807,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X38,sK17,X39,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X38)))) | s(bool,X40) = s(bool,X41) | s(bool,X41) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X38),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X39))))) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f49789,f239])).
fof(f49789,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X38,sK17,X39,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X38)))) | s(bool,X40) = s(bool,X41) | s(bool,X41) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X38),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X39))))) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f2608,f2758])).
fof(f49805,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK17,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X2)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f49780,f239])).
fof(f49780,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK17,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X2)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f2608,f639])).
fof(f49800,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK19(sK18,sK17,sK16,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),
  inference(subsumption_resolution,[],[f49778,f239])).
fof(f49778,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(sK18,sK17,sK16,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),
  inference(superposition,[],[f2608,f1404])).
fof(f49688,plain,(
  ( ! [X255,X257,X254,X256,X258] : (s(bool,'|dollar|exists_1'(s(fun(X257,bool),X258))) = s(bool,g(s(X254,sK20(X255,X256,bool,X254)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X257,bool),X258))) | s(fun(X254,bool),X255) = s(fun(X254,bool),X256) | s(bool,'T_0') = s(bool,f(s(X254,sK20(X255,X256,bool,X254))))) )),
  inference(subsumption_resolution,[],[f49574,f239])).
fof(f49574,plain,(
  ( ! [X255,X257,X254,X256,X258] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X257,bool),X258))) = s(bool,g(s(X254,sK20(X255,X256,bool,X254)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X257,bool),X258))) | s(fun(X254,bool),X255) = s(fun(X254,bool),X256) | s(bool,'T_0') = s(bool,f(s(X254,sK20(X255,X256,bool,X254))))) )),
  inference(superposition,[],[f5167,f5441])).
fof(f49687,plain,(
  ( ! [X253,X251,X249,X252,X250,X248] : (p(s(bool,X251)) | ~sP13(X252,X251,g(s(X248,sK20(X249,X250,bool,X248)))) | s(bool,X252) = s(bool,X253) | s(bool,'T_0') = s(bool,X253) | s(fun(X248,bool),X249) = s(fun(X248,bool),X250) | s(bool,'T_0') = s(bool,f(s(X248,sK20(X249,X250,bool,X248))))) )),
  inference(subsumption_resolution,[],[f49573,f239])).
fof(f49573,plain,(
  ( ! [X253,X251,X249,X252,X250,X248] : (~p(s(bool,'T_0')) | p(s(bool,X251)) | ~sP13(X252,X251,g(s(X248,sK20(X249,X250,bool,X248)))) | s(bool,X252) = s(bool,X253) | s(bool,'T_0') = s(bool,X253) | s(fun(X248,bool),X249) = s(fun(X248,bool),X250) | s(bool,'T_0') = s(bool,f(s(X248,sK20(X249,X250,bool,X248))))) )),
  inference(superposition,[],[f3276,f5441])).
fof(f49686,plain,(
  ( ! [X243,X246,X244,X242,X247,X245] : (p(s(bool,X245)) | ~sP9(X245,g(s(X242,sK20(X243,X244,bool,X242))),X246) | s(bool,X246) = s(bool,X247) | s(bool,'T_0') = s(bool,X247) | s(fun(X242,bool),X243) = s(fun(X242,bool),X244) | s(bool,'T_0') = s(bool,f(s(X242,sK20(X243,X244,bool,X242))))) )),
  inference(subsumption_resolution,[],[f49572,f239])).
fof(f49572,plain,(
  ( ! [X243,X246,X244,X242,X247,X245] : (~p(s(bool,'T_0')) | p(s(bool,X245)) | ~sP9(X245,g(s(X242,sK20(X243,X244,bool,X242))),X246) | s(bool,X246) = s(bool,X247) | s(bool,'T_0') = s(bool,X247) | s(fun(X242,bool),X243) = s(fun(X242,bool),X244) | s(bool,'T_0') = s(bool,f(s(X242,sK20(X243,X244,bool,X242))))) )),
  inference(superposition,[],[f3268,f5441])).
fof(f49684,plain,(
  ( ! [X146,X144,X142,X145,X143] : (~sP10(g(s(X142,sK20(X143,X144,bool,X142))),X145,X146) | sP11(g(s(X142,sK20(X143,X144,bool,X142))),X145,X146) | s(fun(X142,bool),X143) = s(fun(X142,bool),X144) | s(bool,'T_0') = s(bool,f(s(X142,sK20(X143,X144,bool,X142))))) )),
  inference(subsumption_resolution,[],[f49554,f239])).
fof(f49554,plain,(
  ( ! [X146,X144,X142,X145,X143] : (~p(s(bool,'T_0')) | ~sP10(g(s(X142,sK20(X143,X144,bool,X142))),X145,X146) | sP11(g(s(X142,sK20(X143,X144,bool,X142))),X145,X146) | s(fun(X142,bool),X143) = s(fun(X142,bool),X144) | s(bool,'T_0') = s(bool,f(s(X142,sK20(X143,X144,bool,X142))))) )),
  inference(superposition,[],[f417,f5441])).
fof(f49683,plain,(
  ( ! [X140,X138,X141,X139,X137] : (~sP12(g(s(X137,sK20(X138,X139,bool,X137))),X140,X141) | sP13(g(s(X137,sK20(X138,X139,bool,X137))),X140,X141) | s(fun(X137,bool),X138) = s(fun(X137,bool),X139) | s(bool,'T_0') = s(bool,f(s(X137,sK20(X138,X139,bool,X137))))) )),
  inference(subsumption_resolution,[],[f49553,f239])).
fof(f49553,plain,(
  ( ! [X140,X138,X141,X139,X137] : (~p(s(bool,'T_0')) | ~sP12(g(s(X137,sK20(X138,X139,bool,X137))),X140,X141) | sP13(g(s(X137,sK20(X138,X139,bool,X137))),X140,X141) | s(fun(X137,bool),X138) = s(fun(X137,bool),X139) | s(bool,'T_0') = s(bool,f(s(X137,sK20(X138,X139,bool,X137))))) )),
  inference(superposition,[],[f416,f5441])).
fof(f49682,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,X135)) | p(s(bool,X136)) | ~sP13(g(s(X132,sK20(X133,X134,bool,X132))),X136,X135) | s(fun(X132,bool),X133) = s(fun(X132,bool),X134) | s(bool,'T_0') = s(bool,f(s(X132,sK20(X133,X134,bool,X132))))) )),
  inference(subsumption_resolution,[],[f49552,f239])).
fof(f49552,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,'T_0')) | ~p(s(bool,X135)) | p(s(bool,X136)) | ~sP13(g(s(X132,sK20(X133,X134,bool,X132))),X136,X135) | s(fun(X132,bool),X133) = s(fun(X132,bool),X134) | s(bool,'T_0') = s(bool,f(s(X132,sK20(X133,X134,bool,X132))))) )),
  inference(superposition,[],[f413,f5441])).
fof(f49681,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,X130)) | sP13(g(s(X127,sK20(X128,X129,bool,X127))),X130,X131) | s(fun(X127,bool),X128) = s(fun(X127,bool),X129) | s(bool,'T_0') = s(bool,f(s(X127,sK20(X128,X129,bool,X127))))) )),
  inference(subsumption_resolution,[],[f49551,f239])).
fof(f49551,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | ~p(s(bool,X130)) | sP13(g(s(X127,sK20(X128,X129,bool,X127))),X130,X131) | s(fun(X127,bool),X128) = s(fun(X127,bool),X129) | s(bool,'T_0') = s(bool,f(s(X127,sK20(X128,X129,bool,X127))))) )),
  inference(superposition,[],[f412,f5441])).
fof(f49680,plain,(
  ( ! [X125,X123,X126,X124,X122] : (p(s(bool,X125)) | sP13(g(s(X122,sK20(X123,X124,bool,X122))),X126,X125) | s(fun(X122,bool),X123) = s(fun(X122,bool),X124) | s(bool,'T_0') = s(bool,f(s(X122,sK20(X123,X124,bool,X122))))) )),
  inference(subsumption_resolution,[],[f49550,f239])).
fof(f49550,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,'T_0')) | p(s(bool,X125)) | sP13(g(s(X122,sK20(X123,X124,bool,X122))),X126,X125) | s(fun(X122,bool),X123) = s(fun(X122,bool),X124) | s(bool,'T_0') = s(bool,f(s(X122,sK20(X123,X124,bool,X122))))) )),
  inference(superposition,[],[f411,f5441])).
fof(f49679,plain,(
  ( ! [X121,X118,X120,X119,X117] : (p(s(bool,X120)) | ~p(s(bool,X121)) | ~sP12(X121,X120,g(s(X117,sK20(X118,X119,bool,X117)))) | s(fun(X117,bool),X118) = s(fun(X117,bool),X119) | s(bool,'T_0') = s(bool,f(s(X117,sK20(X118,X119,bool,X117))))) )),
  inference(subsumption_resolution,[],[f49549,f239])).
fof(f49549,plain,(
  ( ! [X121,X118,X120,X119,X117] : (~p(s(bool,'T_0')) | p(s(bool,X120)) | ~p(s(bool,X121)) | ~sP12(X121,X120,g(s(X117,sK20(X118,X119,bool,X117)))) | s(fun(X117,bool),X118) = s(fun(X117,bool),X119) | s(bool,'T_0') = s(bool,f(s(X117,sK20(X118,X119,bool,X117))))) )),
  inference(superposition,[],[f406,f5441])).
fof(f49678,plain,(
  ( ! [X116,X114,X112,X115,X113] : (p(s(bool,X115)) | p(s(bool,X116)) | ~sP11(g(s(X112,sK20(X113,X114,bool,X112))),X116,X115) | s(fun(X112,bool),X113) = s(fun(X112,bool),X114) | s(bool,'T_0') = s(bool,f(s(X112,sK20(X113,X114,bool,X112))))) )),
  inference(subsumption_resolution,[],[f49548,f239])).
fof(f49548,plain,(
  ( ! [X116,X114,X112,X115,X113] : (~p(s(bool,'T_0')) | p(s(bool,X115)) | p(s(bool,X116)) | ~sP11(g(s(X112,sK20(X113,X114,bool,X112))),X116,X115) | s(fun(X112,bool),X113) = s(fun(X112,bool),X114) | s(bool,'T_0') = s(bool,f(s(X112,sK20(X113,X114,bool,X112))))) )),
  inference(superposition,[],[f396,f5441])).
fof(f49677,plain,(
  ( ! [X111,X109,X107,X110,X108] : (~p(s(bool,X110)) | sP11(g(s(X107,sK20(X108,X109,bool,X107))),X110,X111) | s(fun(X107,bool),X108) = s(fun(X107,bool),X109) | s(bool,'T_0') = s(bool,f(s(X107,sK20(X108,X109,bool,X107))))) )),
  inference(subsumption_resolution,[],[f49547,f239])).
fof(f49547,plain,(
  ( ! [X111,X109,X107,X110,X108] : (~p(s(bool,'T_0')) | ~p(s(bool,X110)) | sP11(g(s(X107,sK20(X108,X109,bool,X107))),X110,X111) | s(fun(X107,bool),X108) = s(fun(X107,bool),X109) | s(bool,'T_0') = s(bool,f(s(X107,sK20(X108,X109,bool,X107))))) )),
  inference(superposition,[],[f395,f5441])).
fof(f49676,plain,(
  ( ! [X103,X105,X102,X106,X104] : (~p(s(bool,X105)) | sP11(g(s(X102,sK20(X103,X104,bool,X102))),X106,X105) | s(fun(X102,bool),X103) = s(fun(X102,bool),X104) | s(bool,'T_0') = s(bool,f(s(X102,sK20(X103,X104,bool,X102))))) )),
  inference(subsumption_resolution,[],[f49546,f239])).
fof(f49546,plain,(
  ( ! [X103,X105,X102,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X105)) | sP11(g(s(X102,sK20(X103,X104,bool,X102))),X106,X105) | s(fun(X102,bool),X103) = s(fun(X102,bool),X104) | s(bool,'T_0') = s(bool,f(s(X102,sK20(X103,X104,bool,X102))))) )),
  inference(superposition,[],[f394,f5441])).
fof(f49675,plain,(
  ( ! [X101,X99,X97,X100,X98] : (p(s(bool,X100)) | p(s(bool,X101)) | ~sP10(g(s(X97,sK20(X98,X99,bool,X97))),X100,X101) | s(fun(X97,bool),X98) = s(fun(X97,bool),X99) | s(bool,'T_0') = s(bool,f(s(X97,sK20(X98,X99,bool,X97))))) )),
  inference(subsumption_resolution,[],[f49545,f239])).
fof(f49545,plain,(
  ( ! [X101,X99,X97,X100,X98] : (~p(s(bool,'T_0')) | p(s(bool,X100)) | p(s(bool,X101)) | ~sP10(g(s(X97,sK20(X98,X99,bool,X97))),X100,X101) | s(fun(X97,bool),X98) = s(fun(X97,bool),X99) | s(bool,'T_0') = s(bool,f(s(X97,sK20(X98,X99,bool,X97))))) )),
  inference(superposition,[],[f389,f5441])).
fof(f49674,plain,(
  ( ! [X94,X92,X95,X93,X96] : (~p(s(bool,X95)) | p(s(bool,X96)) | ~sP9(X96,X95,g(s(X92,sK20(X93,X94,bool,X92)))) | s(fun(X92,bool),X93) = s(fun(X92,bool),X94) | s(bool,'T_0') = s(bool,f(s(X92,sK20(X93,X94,bool,X92))))) )),
  inference(subsumption_resolution,[],[f49544,f239])).
fof(f49544,plain,(
  ( ! [X94,X92,X95,X93,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X95)) | p(s(bool,X96)) | ~sP9(X96,X95,g(s(X92,sK20(X93,X94,bool,X92)))) | s(fun(X92,bool),X93) = s(fun(X92,bool),X94) | s(bool,'T_0') = s(bool,f(s(X92,sK20(X93,X94,bool,X92))))) )),
  inference(superposition,[],[f381,f5441])).
fof(f49673,plain,(
  ( ! [X90,X88,X87,X91,X89] : (~p(s(bool,X90)) | p(s(bool,X91)) | ~sP8(X90,g(s(X87,sK20(X88,X89,bool,X87))),X91) | s(fun(X87,bool),X88) = s(fun(X87,bool),X89) | s(bool,'T_0') = s(bool,f(s(X87,sK20(X88,X89,bool,X87))))) )),
  inference(subsumption_resolution,[],[f49543,f239])).
fof(f49543,plain,(
  ( ! [X90,X88,X87,X91,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X90)) | p(s(bool,X91)) | ~sP8(X90,g(s(X87,sK20(X88,X89,bool,X87))),X91) | s(fun(X87,bool),X88) = s(fun(X87,bool),X89) | s(bool,'T_0') = s(bool,f(s(X87,sK20(X88,X89,bool,X87))))) )),
  inference(superposition,[],[f372,f5441])).
fof(f49672,plain,(
  ( ! [X85,X83,X86,X84,X82] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP5(X85,g(s(X82,sK20(X83,X84,bool,X82))),X86) | s(fun(X82,bool),X83) = s(fun(X82,bool),X84) | s(bool,'T_0') = s(bool,f(s(X82,sK20(X83,X84,bool,X82))))) )),
  inference(subsumption_resolution,[],[f49542,f239])).
fof(f49542,plain,(
  ( ! [X85,X83,X86,X84,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP5(X85,g(s(X82,sK20(X83,X84,bool,X82))),X86) | s(fun(X82,bool),X83) = s(fun(X82,bool),X84) | s(bool,'T_0') = s(bool,f(s(X82,sK20(X83,X84,bool,X82))))) )),
  inference(superposition,[],[f353,f5441])).
fof(f49671,plain,(
  ( ! [X80,X78,X81,X79,X77] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP6(X80,g(s(X77,sK20(X78,X79,bool,X77))),X81) | s(fun(X77,bool),X78) = s(fun(X77,bool),X79) | s(bool,'T_0') = s(bool,f(s(X77,sK20(X78,X79,bool,X77))))) )),
  inference(subsumption_resolution,[],[f49541,f239])).
fof(f49541,plain,(
  ( ! [X80,X78,X81,X79,X77] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP6(X80,g(s(X77,sK20(X78,X79,bool,X77))),X81) | s(fun(X77,bool),X78) = s(fun(X77,bool),X79) | s(bool,'T_0') = s(bool,f(s(X77,sK20(X78,X79,bool,X77))))) )),
  inference(superposition,[],[f349,f5441])).
fof(f49670,plain,(
  ( ! [X76,X74,X72,X75,X73] : (sP7(X75,X76,g(s(X72,sK20(X73,X74,bool,X72)))) | ~sP6(X75,g(s(X72,sK20(X73,X74,bool,X72))),X76) | ~sP5(X76,g(s(X72,sK20(X73,X74,bool,X72))),X75) | ~sP4(g(s(X72,sK20(X73,X74,bool,X72))),X76,X75) | s(fun(X72,bool),X73) = s(fun(X72,bool),X74) | s(bool,'T_0') = s(bool,f(s(X72,sK20(X73,X74,bool,X72))))) )),
  inference(subsumption_resolution,[],[f49540,f239])).
fof(f49540,plain,(
  ( ! [X76,X74,X72,X75,X73] : (~p(s(bool,'T_0')) | sP7(X75,X76,g(s(X72,sK20(X73,X74,bool,X72)))) | ~sP6(X75,g(s(X72,sK20(X73,X74,bool,X72))),X76) | ~sP5(X76,g(s(X72,sK20(X73,X74,bool,X72))),X75) | ~sP4(g(s(X72,sK20(X73,X74,bool,X72))),X76,X75) | s(fun(X72,bool),X73) = s(fun(X72,bool),X74) | s(bool,'T_0') = s(bool,f(s(X72,sK20(X73,X74,bool,X72))))) )),
  inference(superposition,[],[f346,f5441])).
fof(f49669,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,X70)) | sP3(X71,X70,g(s(X67,sK20(X68,X69,bool,X67)))) | s(fun(X67,bool),X68) = s(fun(X67,bool),X69) | s(bool,'T_0') = s(bool,f(s(X67,sK20(X68,X69,bool,X67))))) )),
  inference(subsumption_resolution,[],[f49539,f239])).
fof(f49539,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X70)) | sP3(X71,X70,g(s(X67,sK20(X68,X69,bool,X67)))) | s(fun(X67,bool),X68) = s(fun(X67,bool),X69) | s(bool,'T_0') = s(bool,f(s(X67,sK20(X68,X69,bool,X67))))) )),
  inference(superposition,[],[f334,f5441])).
fof(f49668,plain,(
  ( ! [X66,X64,X62,X65,X63] : (~p(s(bool,X65)) | sP2(X65,X66,g(s(X62,sK20(X63,X64,bool,X62)))) | s(fun(X62,bool),X63) = s(fun(X62,bool),X64) | s(bool,'T_0') = s(bool,f(s(X62,sK20(X63,X64,bool,X62))))) )),
  inference(subsumption_resolution,[],[f49538,f239])).
fof(f49538,plain,(
  ( ! [X66,X64,X62,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X65)) | sP2(X65,X66,g(s(X62,sK20(X63,X64,bool,X62)))) | s(fun(X62,bool),X63) = s(fun(X62,bool),X64) | s(bool,'T_0') = s(bool,f(s(X62,sK20(X63,X64,bool,X62))))) )),
  inference(superposition,[],[f327,f5441])).
fof(f49667,plain,(
  ( ! [X61,X59,X57,X60,X58] : (~p(s(bool,X60)) | sP2(X60,g(s(X57,sK20(X58,X59,bool,X57))),X61) | s(fun(X57,bool),X58) = s(fun(X57,bool),X59) | s(bool,'T_0') = s(bool,f(s(X57,sK20(X58,X59,bool,X57))))) )),
  inference(subsumption_resolution,[],[f49537,f239])).
fof(f49537,plain,(
  ( ! [X61,X59,X57,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP2(X60,g(s(X57,sK20(X58,X59,bool,X57))),X61) | s(fun(X57,bool),X58) = s(fun(X57,bool),X59) | s(bool,'T_0') = s(bool,f(s(X57,sK20(X58,X59,bool,X57))))) )),
  inference(superposition,[],[f326,f5441])).
fof(f49666,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,X51)) | sP1(X51,g(s(X48,sK20(X49,X50,bool,X48)))) | s(fun(X48,bool),X49) = s(fun(X48,bool),X50) | s(bool,'T_0') = s(bool,f(s(X48,sK20(X49,X50,bool,X48))))) )),
  inference(subsumption_resolution,[],[f49535,f239])).
fof(f49535,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,'T_0')) | p(s(bool,X51)) | sP1(X51,g(s(X48,sK20(X49,X50,bool,X48)))) | s(fun(X48,bool),X49) = s(fun(X48,bool),X50) | s(bool,'T_0') = s(bool,f(s(X48,sK20(X49,X50,bool,X48))))) )),
  inference(superposition,[],[f316,f5441])).
fof(f49665,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X47)) | sP1(g(s(X44,sK20(X45,X46,bool,X44))),X47) | s(fun(X44,bool),X45) = s(fun(X44,bool),X46) | s(bool,'T_0') = s(bool,f(s(X44,sK20(X45,X46,bool,X44))))) )),
  inference(subsumption_resolution,[],[f49534,f239])).
fof(f49534,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | sP1(g(s(X44,sK20(X45,X46,bool,X44))),X47) | s(fun(X44,bool),X45) = s(fun(X44,bool),X46) | s(bool,'T_0') = s(bool,f(s(X44,sK20(X45,X46,bool,X44))))) )),
  inference(superposition,[],[f313,f5441])).
fof(f49664,plain,(
  ( ! [X43,X41,X42,X40] : (sP0(g(s(X40,sK20(X41,X42,bool,X40))),X43) | s(fun(X40,bool),X41) = s(fun(X40,bool),X42) | s(bool,'T_0') = s(bool,f(s(X40,sK20(X41,X42,bool,X40))))) )),
  inference(subsumption_resolution,[],[f49533,f239])).
fof(f49533,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'T_0')) | sP0(g(s(X40,sK20(X41,X42,bool,X40))),X43) | s(fun(X40,bool),X41) = s(fun(X40,bool),X42) | s(bool,'T_0') = s(bool,f(s(X40,sK20(X41,X42,bool,X40))))) )),
  inference(superposition,[],[f298,f5441])).
fof(f49663,plain,(
  ( ! [X39,X37,X38,X36] : (sP0(X39,g(s(X36,sK20(X37,X38,bool,X36)))) | s(fun(X36,bool),X37) = s(fun(X36,bool),X38) | s(bool,'T_0') = s(bool,f(s(X36,sK20(X37,X38,bool,X36))))) )),
  inference(subsumption_resolution,[],[f49532,f239])).
fof(f49532,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'T_0')) | sP0(X39,g(s(X36,sK20(X37,X38,bool,X36)))) | s(fun(X36,bool),X37) = s(fun(X36,bool),X38) | s(bool,'T_0') = s(bool,f(s(X36,sK20(X37,X38,bool,X36))))) )),
  inference(superposition,[],[f297,f5441])).
fof(f49662,plain,(
  ( ! [X35,X33,X34,X32] : (~p(s(bool,X35)) | s(bool,X35) = s(bool,g(s(X32,sK20(X33,X34,bool,X32)))) | s(fun(X32,bool),X33) = s(fun(X32,bool),X34) | s(bool,'T_0') = s(bool,f(s(X32,sK20(X33,X34,bool,X32))))) )),
  inference(subsumption_resolution,[],[f49531,f239])).
fof(f49531,plain,(
  ( ! [X35,X33,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | s(bool,X35) = s(bool,g(s(X32,sK20(X33,X34,bool,X32)))) | s(fun(X32,bool),X33) = s(fun(X32,bool),X34) | s(bool,'T_0') = s(bool,f(s(X32,sK20(X33,X34,bool,X32))))) )),
  inference(superposition,[],[f288,f5441])).
fof(f49661,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') != s(bool,g(s(X18,sK20(X19,X20,bool,X18)))) | s(fun(X18,bool),X19) = s(fun(X18,bool),X20) | s(bool,'T_0') = s(bool,f(s(X18,sK20(X19,X20,bool,X18))))) )),
  inference(subsumption_resolution,[],[f49527,f239])).
fof(f49527,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,g(s(X18,sK20(X19,X20,bool,X18)))) | s(fun(X18,bool),X19) = s(fun(X18,bool),X20) | s(bool,'T_0') = s(bool,f(s(X18,sK20(X19,X20,bool,X18))))) )),
  inference(superposition,[],[f258,f5441])).
fof(f49633,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,g(s(X12,sK20(X13,X14,bool,X12))))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,f(s(X12,sK20(X13,X14,bool,X12))))) )),
  inference(trivial_inequality_removal,[],[f49525])).
fof(f49525,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,g(s(X12,sK20(X13,X14,bool,X12))))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,f(s(X12,sK20(X13,X14,bool,X12))))) )),
  inference(superposition,[],[f254,f5441])).
fof(f49648,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,f(s(X0,sK20(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK20(X1,X2,bool,X0)))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f49647])).
fof(f49647,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,X42) | s(fun(X39,bool),X40) = s(fun(X39,bool),X41) | s(bool,'T_0') = s(bool,f(s(X39,sK20(X40,X41,bool,X39)))) | s(bool,X43) = s(bool,g(s(X39,sK20(X40,X41,bool,X39)))) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(duplicate_literal_removal,[],[f49512])).
fof(f49512,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,X42) | s(fun(X39,bool),X40) = s(fun(X39,bool),X41) | s(bool,'T_0') = s(bool,f(s(X39,sK20(X40,X41,bool,X39)))) | s(bool,X43) = s(bool,g(s(X39,sK20(X40,X41,bool,X39)))) | s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f5441,f2758])).
fof(f49261,plain,(
  ( ! [X251,X249,X250,X248,X247] : (s(bool,'|dollar|exists_1'(s(fun(X249,bool),X250))) = s(bool,happ(s(fun(X247,bool),X248),s(X247,'const.min.@_1'(s(fun(X247,bool),X248))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X249,bool),X250))) | s(bool,X251) = s(bool,'|dollar|exists_1'(s(fun(X247,bool),X248))) | s(bool,'T_0') = s(bool,X251)) )),
  inference(subsumption_resolution,[],[f48977,f239])).
fof(f48977,plain,(
  ( ! [X251,X249,X250,X248,X247] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X249,bool),X250))) = s(bool,happ(s(fun(X247,bool),X248),s(X247,'const.min.@_1'(s(fun(X247,bool),X248))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X249,bool),X250))) | s(bool,X251) = s(bool,'|dollar|exists_1'(s(fun(X247,bool),X248))) | s(bool,'T_0') = s(bool,X251)) )),
  inference(superposition,[],[f5167,f2743])).
fof(f49260,plain,(
  ( ! [X243,X241,X246,X244,X242,X245] : (p(s(bool,X243)) | ~sP13(X244,X243,happ(s(fun(X241,bool),X242),s(X241,'const.min.@_1'(s(fun(X241,bool),X242))))) | s(bool,X244) = s(bool,X245) | s(bool,'T_0') = s(bool,X245) | s(bool,X246) = s(bool,'|dollar|exists_1'(s(fun(X241,bool),X242))) | s(bool,'T_0') = s(bool,X246)) )),
  inference(subsumption_resolution,[],[f48976,f239])).
fof(f48976,plain,(
  ( ! [X243,X241,X246,X244,X242,X245] : (~p(s(bool,'T_0')) | p(s(bool,X243)) | ~sP13(X244,X243,happ(s(fun(X241,bool),X242),s(X241,'const.min.@_1'(s(fun(X241,bool),X242))))) | s(bool,X244) = s(bool,X245) | s(bool,'T_0') = s(bool,X245) | s(bool,X246) = s(bool,'|dollar|exists_1'(s(fun(X241,bool),X242))) | s(bool,'T_0') = s(bool,X246)) )),
  inference(superposition,[],[f3276,f2743])).
fof(f49259,plain,(
  ( ! [X239,X237,X235,X240,X238,X236] : (p(s(bool,X237)) | ~sP9(X237,happ(s(fun(X235,bool),X236),s(X235,'const.min.@_1'(s(fun(X235,bool),X236)))),X238) | s(bool,X238) = s(bool,X239) | s(bool,'T_0') = s(bool,X239) | s(bool,X240) = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236))) | s(bool,'T_0') = s(bool,X240)) )),
  inference(subsumption_resolution,[],[f48975,f239])).
fof(f48975,plain,(
  ( ! [X239,X237,X235,X240,X238,X236] : (~p(s(bool,'T_0')) | p(s(bool,X237)) | ~sP9(X237,happ(s(fun(X235,bool),X236),s(X235,'const.min.@_1'(s(fun(X235,bool),X236)))),X238) | s(bool,X238) = s(bool,X239) | s(bool,'T_0') = s(bool,X239) | s(bool,X240) = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236))) | s(bool,'T_0') = s(bool,X240)) )),
  inference(superposition,[],[f3268,f2743])).
fof(f49258,plain,(
  ( ! [X146,X144,X147,X145,X148] : (~sP10(happ(s(fun(X144,bool),X145),s(X144,'const.min.@_1'(s(fun(X144,bool),X145)))),X146,X147) | sP11(happ(s(fun(X144,bool),X145),s(X144,'const.min.@_1'(s(fun(X144,bool),X145)))),X146,X147) | s(bool,X148) = s(bool,'|dollar|exists_1'(s(fun(X144,bool),X145))) | s(bool,'T_0') = s(bool,X148)) )),
  inference(subsumption_resolution,[],[f48958,f239])).
fof(f48958,plain,(
  ( ! [X146,X144,X147,X145,X148] : (~p(s(bool,'T_0')) | ~sP10(happ(s(fun(X144,bool),X145),s(X144,'const.min.@_1'(s(fun(X144,bool),X145)))),X146,X147) | sP11(happ(s(fun(X144,bool),X145),s(X144,'const.min.@_1'(s(fun(X144,bool),X145)))),X146,X147) | s(bool,X148) = s(bool,'|dollar|exists_1'(s(fun(X144,bool),X145))) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f417,f2743])).
fof(f49257,plain,(
  ( ! [X142,X140,X143,X141,X139] : (~sP12(happ(s(fun(X139,bool),X140),s(X139,'const.min.@_1'(s(fun(X139,bool),X140)))),X141,X142) | sP13(happ(s(fun(X139,bool),X140),s(X139,'const.min.@_1'(s(fun(X139,bool),X140)))),X141,X142) | s(bool,X143) = s(bool,'|dollar|exists_1'(s(fun(X139,bool),X140))) | s(bool,'T_0') = s(bool,X143)) )),
  inference(subsumption_resolution,[],[f48957,f239])).
fof(f48957,plain,(
  ( ! [X142,X140,X143,X141,X139] : (~p(s(bool,'T_0')) | ~sP12(happ(s(fun(X139,bool),X140),s(X139,'const.min.@_1'(s(fun(X139,bool),X140)))),X141,X142) | sP13(happ(s(fun(X139,bool),X140),s(X139,'const.min.@_1'(s(fun(X139,bool),X140)))),X141,X142) | s(bool,X143) = s(bool,'|dollar|exists_1'(s(fun(X139,bool),X140))) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f416,f2743])).
fof(f49256,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,X131)) | sP13(happ(s(fun(X129,bool),X130),s(X129,'const.min.@_1'(s(fun(X129,bool),X130)))),X131,X132) | s(bool,X133) = s(bool,'|dollar|exists_1'(s(fun(X129,bool),X130))) | s(bool,'T_0') = s(bool,X133)) )),
  inference(subsumption_resolution,[],[f48955,f239])).
fof(f48955,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,'T_0')) | ~p(s(bool,X131)) | sP13(happ(s(fun(X129,bool),X130),s(X129,'const.min.@_1'(s(fun(X129,bool),X130)))),X131,X132) | s(bool,X133) = s(bool,'|dollar|exists_1'(s(fun(X129,bool),X130))) | s(bool,'T_0') = s(bool,X133)) )),
  inference(superposition,[],[f412,f2743])).
fof(f49255,plain,(
  ( ! [X127,X125,X128,X126,X124] : (p(s(bool,X126)) | sP13(happ(s(fun(X124,bool),X125),s(X124,'const.min.@_1'(s(fun(X124,bool),X125)))),X127,X126) | s(bool,X128) = s(bool,'|dollar|exists_1'(s(fun(X124,bool),X125))) | s(bool,'T_0') = s(bool,X128)) )),
  inference(subsumption_resolution,[],[f48954,f239])).
fof(f48954,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,'T_0')) | p(s(bool,X126)) | sP13(happ(s(fun(X124,bool),X125),s(X124,'const.min.@_1'(s(fun(X124,bool),X125)))),X127,X126) | s(bool,X128) = s(bool,'|dollar|exists_1'(s(fun(X124,bool),X125))) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f411,f2743])).
fof(f49254,plain,(
  ( ! [X111,X109,X112,X110,X113] : (~p(s(bool,X111)) | sP11(happ(s(fun(X109,bool),X110),s(X109,'const.min.@_1'(s(fun(X109,bool),X110)))),X111,X112) | s(bool,X113) = s(bool,'|dollar|exists_1'(s(fun(X109,bool),X110))) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f48951,f239])).
fof(f48951,plain,(
  ( ! [X111,X109,X112,X110,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X111)) | sP11(happ(s(fun(X109,bool),X110),s(X109,'const.min.@_1'(s(fun(X109,bool),X110)))),X111,X112) | s(bool,X113) = s(bool,'|dollar|exists_1'(s(fun(X109,bool),X110))) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f395,f2743])).
fof(f49253,plain,(
  ( ! [X107,X105,X108,X106,X104] : (~p(s(bool,X106)) | sP11(happ(s(fun(X104,bool),X105),s(X104,'const.min.@_1'(s(fun(X104,bool),X105)))),X107,X106) | s(bool,X108) = s(bool,'|dollar|exists_1'(s(fun(X104,bool),X105))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f48950,f239])).
fof(f48950,plain,(
  ( ! [X107,X105,X108,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X106)) | sP11(happ(s(fun(X104,bool),X105),s(X104,'const.min.@_1'(s(fun(X104,bool),X105)))),X107,X106) | s(bool,X108) = s(bool,'|dollar|exists_1'(s(fun(X104,bool),X105))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f394,f2743])).
fof(f49252,plain,(
  ( ! [X70,X72,X71,X69,X73] : (~p(s(bool,X71)) | sP3(X72,X71,happ(s(fun(X69,bool),X70),s(X69,'const.min.@_1'(s(fun(X69,bool),X70))))) | s(bool,X73) = s(bool,'|dollar|exists_1'(s(fun(X69,bool),X70))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f48943,f239])).
fof(f48943,plain,(
  ( ! [X70,X72,X71,X69,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X71)) | sP3(X72,X71,happ(s(fun(X69,bool),X70),s(X69,'const.min.@_1'(s(fun(X69,bool),X70))))) | s(bool,X73) = s(bool,'|dollar|exists_1'(s(fun(X69,bool),X70))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f334,f2743])).
fof(f49251,plain,(
  ( ! [X68,X66,X64,X67,X65] : (~p(s(bool,X66)) | sP2(X66,X67,happ(s(fun(X64,bool),X65),s(X64,'const.min.@_1'(s(fun(X64,bool),X65))))) | s(bool,X68) = s(bool,'|dollar|exists_1'(s(fun(X64,bool),X65))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f48942,f239])).
fof(f48942,plain,(
  ( ! [X68,X66,X64,X67,X65] : (~p(s(bool,'T_0')) | ~p(s(bool,X66)) | sP2(X66,X67,happ(s(fun(X64,bool),X65),s(X64,'const.min.@_1'(s(fun(X64,bool),X65))))) | s(bool,X68) = s(bool,'|dollar|exists_1'(s(fun(X64,bool),X65))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f327,f2743])).
fof(f49250,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,X61)) | sP2(X61,happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))),X62) | s(bool,X63) = s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f48941,f239])).
fof(f48941,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | sP2(X61,happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))),X62) | s(bool,X63) = s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f326,f2743])).
fof(f49249,plain,(
  ( ! [X52,X50,X53,X51] : (p(s(bool,X52)) | sP1(X52,happ(s(fun(X50,bool),X51),s(X50,'const.min.@_1'(s(fun(X50,bool),X51))))) | s(bool,X53) = s(bool,'|dollar|exists_1'(s(fun(X50,bool),X51))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f48939,f239])).
fof(f48939,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | sP1(X52,happ(s(fun(X50,bool),X51),s(X50,'const.min.@_1'(s(fun(X50,bool),X51))))) | s(bool,X53) = s(bool,'|dollar|exists_1'(s(fun(X50,bool),X51))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f316,f2743])).
fof(f49248,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X48)) | sP1(happ(s(fun(X46,bool),X47),s(X46,'const.min.@_1'(s(fun(X46,bool),X47)))),X48) | s(bool,X49) = s(bool,'|dollar|exists_1'(s(fun(X46,bool),X47))) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f48938,f239])).
fof(f48938,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | p(s(bool,X48)) | sP1(happ(s(fun(X46,bool),X47),s(X46,'const.min.@_1'(s(fun(X46,bool),X47)))),X48) | s(bool,X49) = s(bool,'|dollar|exists_1'(s(fun(X46,bool),X47))) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f313,f2743])).
fof(f49247,plain,(
  ( ! [X45,X43,X44,X42] : (sP0(happ(s(fun(X42,bool),X43),s(X42,'const.min.@_1'(s(fun(X42,bool),X43)))),X44) | s(bool,X45) = s(bool,'|dollar|exists_1'(s(fun(X42,bool),X43))) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f48937,f239])).
fof(f48937,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X42,bool),X43),s(X42,'const.min.@_1'(s(fun(X42,bool),X43)))),X44) | s(bool,X45) = s(bool,'|dollar|exists_1'(s(fun(X42,bool),X43))) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f298,f2743])).
fof(f49246,plain,(
  ( ! [X39,X41,X38,X40] : (sP0(X40,happ(s(fun(X38,bool),X39),s(X38,'const.min.@_1'(s(fun(X38,bool),X39))))) | s(bool,X41) = s(bool,'|dollar|exists_1'(s(fun(X38,bool),X39))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f48936,f239])).
fof(f48936,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | sP0(X40,happ(s(fun(X38,bool),X39),s(X38,'const.min.@_1'(s(fun(X38,bool),X39))))) | s(bool,X41) = s(bool,'|dollar|exists_1'(s(fun(X38,bool),X39))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f297,f2743])).
fof(f49243,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,happ(s(fun(X20,bool),X21),s(X20,'const.min.@_1'(s(fun(X20,bool),X21))))) | s(bool,X22) = s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f48931,f239])).
fof(f48931,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(X20,bool),X21),s(X20,'const.min.@_1'(s(fun(X20,bool),X21))))) | s(bool,X22) = s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f258,f2743])).
fof(f48819,plain,(
  ( ! [X883,X887,X885,X886,X884] : (s(X885,f(s(bool,'|dollar|exists_1'(s(fun(X886,bool),X887))))) != s(X885,g(s(bool,'|dollar|exists_1'(s(fun(X886,bool),X887))))) | s(fun(bool,X885),X883) = s(fun(bool,X885),X884) | s(bool,'T_0') = s(bool,happ(s(fun(X886,bool),X887),s(X886,'const.min.@_1'(s(fun(X886,bool),X887))))) | s(bool,'T_0') = s(bool,sK20(X883,X884,X885,bool))) )),
  inference(superposition,[],[f321,f2743])).
fof(f48777,plain,(
  ( ! [X697,X699,X700,X698,X701] : (s(bool,'|dollar|exists_1'(s(fun(X700,bool),X701))) != s(bool,f(s(X697,sK20(X698,X699,bool,X697)))) | s(fun(X697,bool),X698) = s(fun(X697,bool),X699) | s(bool,'T_0') = s(bool,happ(s(fun(X700,bool),X701),s(X700,'const.min.@_1'(s(fun(X700,bool),X701))))) | s(bool,'T_0') = s(bool,g(s(X697,sK20(X698,X699,bool,X697))))) )),
  inference(superposition,[],[f321,f2743])).
fof(f49141,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) | s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))))) )),
  inference(condensation,[],[f49140])).
fof(f49140,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21))) = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23))) | s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23))))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,'const.min.@_1'(s(fun(X20,bool),X21)))))) )),
  inference(duplicate_literal_removal,[],[f48564])).
fof(f48564,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21))) = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23))) | s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23))))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,'const.min.@_1'(s(fun(X20,bool),X21))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f2743,f2743])).
fof(f47734,plain,(
  ( ! [X74,X72,X75,X73] : (s('type.extreal.extreal',X72) = s('type.extreal.extreal',X73) | s(bool,X74) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X73),s('type.extreal.extreal',X72))) | s(bool,'T_0') = s(bool,X74) | s(bool,X75) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X75) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X72),s('type.extreal.extreal',X73)))) )),
  inference(subsumption_resolution,[],[f47724,f239])).
fof(f47724,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,'T_0')) | s('type.extreal.extreal',X72) = s('type.extreal.extreal',X73) | s(bool,X74) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X73),s('type.extreal.extreal',X72))) | s(bool,'T_0') = s(bool,X74) | s(bool,X75) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X75) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X72),s('type.extreal.extreal',X73)))) )),
  inference(superposition,[],[f3282,f9775])).
fof(f2828,plain,(
  ( ! [X158,X159] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X158))) = s(bool,happ(s(fun(bool,bool),X158),s(bool,X159))) | s(bool,'T_0') = s(bool,X159) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X158)))) )),
  inference(superposition,[],[f281,f639])).
fof(f2958,plain,(
  ( ! [X152,X153] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X152))) = s(bool,happ(s(fun(bool,bool),X152),s(bool,'T_0'))) | s(bool,X153) = s(bool,'const.min.@_1'(s(fun(bool,bool),X152))) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f281,f639])).
fof(f4372,plain,(
  ( ! [X0,X1] : (~sP7('F_0',happ(s(fun(bool,bool),X0),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1811,f639])).
fof(f4374,plain,(
  ( ! [X4,X5] : (~sP7('F_0',happ(s(fun(bool,bool),X4),s(bool,X5)),'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1811,f639])).
fof(f4420,plain,(
  ( ! [X0,X1] : (sP7('T_0',happ(s(fun(bool,bool),X0),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1812,f639])).
fof(f4422,plain,(
  ( ! [X4,X5] : (sP7('T_0',happ(s(fun(bool,bool),X4),s(bool,X5)),'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1812,f639])).
fof(f4449,plain,(
  ( ! [X0,X1] : (~sP7('F_0','|dollar|exists_1'(s(fun(bool,bool),X0)),happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1813,f639])).
fof(f4451,plain,(
  ( ! [X4,X5] : (~sP7('F_0','|dollar|exists_1'(s(fun(bool,bool),X4)),happ(s(fun(bool,bool),X4),s(bool,X5))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1813,f639])).
fof(f4469,plain,(
  ( ! [X0,X1] : (sP7('T_0','|dollar|exists_1'(s(fun(bool,bool),X0)),happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1814,f639])).
fof(f4471,plain,(
  ( ! [X4,X5] : (sP7('T_0','|dollar|exists_1'(s(fun(bool,bool),X4)),happ(s(fun(bool,bool),X4),s(bool,X5))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1814,f639])).
fof(f4484,plain,(
  ( ! [X2,X0,X1] : (sP0(X1,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1826,f639])).
fof(f4497,plain,(
  ( ! [X2,X0,X1] : (sP0(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1827,f639])).
fof(f4512,plain,(
  ( ! [X2,X0,X1] : (sP1(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1) | p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1828,f639])).
fof(f4523,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1829,f639])).
fof(f4536,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X1,happ(s(fun(bool,bool),X0),s(bool,'T_0')),X2) | ~p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1831,f639])).
fof(f4549,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X1,X2,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | ~p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1832,f639])).
fof(f4564,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X1,X2,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | ~p(s(bool,X2)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1833,f639])).
fof(f4575,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | ~p(s(bool,X2)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1840,f639])).
fof(f4588,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1841,f639])).
fof(f5054,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1815,f639])).
fof(f9225,plain,(
  ( ! [X0,X1] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X0)),happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f4417,f639])).
fof(f9227,plain,(
  ( ! [X4,X5] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X4)),happ(s(fun(bool,bool),X4),s(bool,X5))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f4417,f639])).
fof(f9229,plain,(
  ( ! [X8] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X8)),happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f4417,f247])).
fof(f9255,plain,(
  ( ! [X0,X1] : (sP6('T_0',happ(s(fun(bool,bool),X0),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f4466,f639])).
fof(f9257,plain,(
  ( ! [X4,X5] : (sP6('T_0',happ(s(fun(bool,bool),X4),s(bool,X5)),'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f4466,f639])).
fof(f9259,plain,(
  ( ! [X8] : (sP6('T_0',happ(s(fun(bool,bool),X8),s(bool,'F_0')),'|dollar|exists_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f4466,f247])).
fof(f12070,plain,(
  ( ! [X4,X2,X3] : (sP0(X3,happ(s(fun(bool,bool),X2),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,'const.min.@_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f5168,f639])).
fof(f12193,plain,(
  ( ! [X4,X2,X3] : (sP0(happ(s(fun(bool,bool),X2),s(bool,'T_0')),X3) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,'const.min.@_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f5169,f639])).
fof(f14204,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f5165,f639])).
fof(f14293,plain,(
  ( ! [X4,X2,X3] : (sP1(happ(s(fun(bool,bool),X2),s(bool,'T_0')),X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,'const.min.@_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f5170,f639])).
fof(f14366,plain,(
  ( ! [X4,X2,X3] : (sP1(X3,happ(s(fun(bool,bool),X2),s(bool,'T_0'))) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,'const.min.@_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f5171,f639])).
fof(f14441,plain,(
  ( ! [X6,X4,X5,X3] : (sP2(X4,happ(s(fun(bool,bool),X3),s(bool,'T_0')),X5) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f5172,f639])).
fof(f14516,plain,(
  ( ! [X6,X4,X5,X3] : (sP2(X4,X5,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f5173,f639])).
fof(f14591,plain,(
  ( ! [X6,X4,X5,X3] : (sP3(X4,X5,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | ~p(s(bool,X5)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f5174,f639])).
fof(f14617,plain,(
  ( ! [X6,X4,X5,X3] : (sP11(happ(s(fun(bool,bool),X3),s(bool,'T_0')),X4,X5) | ~p(s(bool,X5)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f5181,f639])).
fof(f14643,plain,(
  ( ! [X6,X4,X5,X3] : (sP11(happ(s(fun(bool,bool),X3),s(bool,'T_0')),X4,X5) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f5182,f639])).
fof(f14668,plain,(
  ( ! [X6,X4,X5,X3] : (sP13(happ(s(fun(bool,bool),X3),s(bool,'T_0')),X4,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f5185,f639])).
fof(f14695,plain,(
  ( ! [X6,X4,X5,X3] : (sP13(happ(s(fun(bool,bool),X3),s(bool,'T_0')),X4,X5) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f5186,f639])).
fof(f15504,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,X134) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X133))) | s(bool,X135) = s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X134) | s(bool,X135) = s(bool,X136) | s(bool,'T_0') = s(bool,X136) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X133)))) )),
  inference(superposition,[],[f2758,f1815])).
fof(f15532,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,X122) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X121))) | s(bool,X122) = s(bool,X123) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X121),s(bool,'F_0'))) | s(bool,X123) = s(bool,X124) | s(bool,'T_0') = s(bool,X124) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X121)))) )),
  inference(superposition,[],[f2758,f1815])).
fof(f15563,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X78))) | s(bool,X79) = s(bool,X80) | s(bool,X80) = s(bool,happ(s(fun(bool,bool),X78),s(bool,'F_0'))) | s(bool,X79) = s(bool,X81) | s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X78)))) )),
  inference(superposition,[],[f2758,f1815])).
fof(f15751,plain,(
  ( ! [X532,X530,X529,X531] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X529))) = s(bool,happ(s(fun(bool,bool),X529),s(bool,X530))) | s(bool,X530) = s(bool,X531) | s(bool,'T_0') = s(bool,X531) | s(bool,X532) = s(bool,'const.min.@_1'(s(fun(bool,bool),X529))) | s(bool,'T_0') = s(bool,X532)) )),
  inference(superposition,[],[f281,f2758])).
fof(f15752,plain,(
  ( ! [X536,X534,X533,X535] : (~sP7('F_0',happ(s(fun(bool,bool),X533),s(bool,X534)),'|dollar|exists_1'(s(fun(bool,bool),X533))) | s(bool,X534) = s(bool,X535) | s(bool,'T_0') = s(bool,X535) | s(bool,X536) = s(bool,'const.min.@_1'(s(fun(bool,bool),X533))) | s(bool,'T_0') = s(bool,X536)) )),
  inference(superposition,[],[f1811,f2758])).
fof(f15753,plain,(
  ( ! [X540,X538,X537,X539] : (sP7('T_0',happ(s(fun(bool,bool),X537),s(bool,X538)),'|dollar|exists_1'(s(fun(bool,bool),X537))) | s(bool,X538) = s(bool,X539) | s(bool,'T_0') = s(bool,X539) | s(bool,X540) = s(bool,'const.min.@_1'(s(fun(bool,bool),X537))) | s(bool,'T_0') = s(bool,X540)) )),
  inference(superposition,[],[f1812,f2758])).
fof(f15754,plain,(
  ( ! [X542,X541,X543,X544] : (~sP7('F_0','|dollar|exists_1'(s(fun(bool,bool),X541)),happ(s(fun(bool,bool),X541),s(bool,X542))) | s(bool,X542) = s(bool,X543) | s(bool,'T_0') = s(bool,X543) | s(bool,X544) = s(bool,'const.min.@_1'(s(fun(bool,bool),X541))) | s(bool,'T_0') = s(bool,X544)) )),
  inference(superposition,[],[f1813,f2758])).
fof(f15755,plain,(
  ( ! [X546,X545,X547,X548] : (sP7('T_0','|dollar|exists_1'(s(fun(bool,bool),X545)),happ(s(fun(bool,bool),X545),s(bool,X546))) | s(bool,X546) = s(bool,X547) | s(bool,'T_0') = s(bool,X547) | s(bool,X548) = s(bool,'const.min.@_1'(s(fun(bool,bool),X545))) | s(bool,'T_0') = s(bool,X548)) )),
  inference(superposition,[],[f1814,f2758])).
fof(f15757,plain,(
  ( ! [X557,X553,X555,X556,X554] : (sP0(X555,happ(s(fun(bool,bool),X553),s(bool,X554))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X553)))) | s(bool,X554) = s(bool,X556) | s(bool,'T_0') = s(bool,X556) | s(bool,X557) = s(bool,'const.min.@_1'(s(fun(bool,bool),X553))) | s(bool,'T_0') = s(bool,X557)) )),
  inference(superposition,[],[f1826,f2758])).
fof(f15758,plain,(
  ( ! [X561,X559,X560,X562,X558] : (sP0(happ(s(fun(bool,bool),X558),s(bool,X559)),X560) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X558)))) | s(bool,X559) = s(bool,X561) | s(bool,'T_0') = s(bool,X561) | s(bool,X562) = s(bool,'const.min.@_1'(s(fun(bool,bool),X558))) | s(bool,'T_0') = s(bool,X562)) )),
  inference(superposition,[],[f1827,f2758])).
fof(f15759,plain,(
  ( ! [X565,X567,X563,X564,X566] : (sP1(happ(s(fun(bool,bool),X563),s(bool,X564)),X565) | p(s(bool,X565)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X563)))) | s(bool,X564) = s(bool,X566) | s(bool,'T_0') = s(bool,X566) | s(bool,X567) = s(bool,'const.min.@_1'(s(fun(bool,bool),X563))) | s(bool,'T_0') = s(bool,X567)) )),
  inference(superposition,[],[f1828,f2758])).
fof(f15770,plain,(
  ( ! [X629,X627,X628,X630] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X627)),happ(s(fun(bool,bool),X627),s(bool,X628))) | s(bool,X628) = s(bool,X629) | s(bool,'T_0') = s(bool,X629) | s(bool,X630) = s(bool,'const.min.@_1'(s(fun(bool,bool),X627))) | s(bool,'T_0') = s(bool,X630)) )),
  inference(superposition,[],[f4417,f2758])).
fof(f15771,plain,(
  ( ! [X633,X631,X632,X634] : (sP6('T_0',happ(s(fun(bool,bool),X631),s(bool,X632)),'|dollar|exists_1'(s(fun(bool,bool),X631))) | s(bool,X632) = s(bool,X633) | s(bool,'T_0') = s(bool,X633) | s(bool,X634) = s(bool,'const.min.@_1'(s(fun(bool,bool),X631))) | s(bool,'T_0') = s(bool,X634)) )),
  inference(superposition,[],[f4466,f2758])).
fof(f15772,plain,(
  ( ! [X635,X636,X638,X637] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X635),s(bool,X636))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X635))) | s(bool,X636) = s(bool,X637) | s(bool,'T_0') = s(bool,X637) | s(bool,X638) = s(bool,'const.min.@_1'(s(fun(bool,bool),X635))) | s(bool,'T_0') = s(bool,X638)) )),
  inference(superposition,[],[f5165,f2758])).
fof(f15773,plain,(
  ( ! [X639,X640,X642,X641,X643] : (sP0(X641,happ(s(fun(bool,bool),X639),s(bool,X640))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X639))) | s(bool,X640) = s(bool,X642) | s(bool,'T_0') = s(bool,X642) | s(bool,X643) = s(bool,'const.min.@_1'(s(fun(bool,bool),X639))) | s(bool,'T_0') = s(bool,X643)) )),
  inference(superposition,[],[f5168,f2758])).
fof(f15774,plain,(
  ( ! [X648,X644,X646,X645,X647] : (sP0(happ(s(fun(bool,bool),X644),s(bool,X645)),X646) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X644))) | s(bool,X645) = s(bool,X647) | s(bool,'T_0') = s(bool,X647) | s(bool,X648) = s(bool,'const.min.@_1'(s(fun(bool,bool),X644))) | s(bool,'T_0') = s(bool,X648)) )),
  inference(superposition,[],[f5169,f2758])).
fof(f15775,plain,(
  ( ! [X652,X650,X653,X649,X651] : (sP1(happ(s(fun(bool,bool),X649),s(bool,X650)),X651) | p(s(bool,X651)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X649))) | s(bool,X650) = s(bool,X652) | s(bool,'T_0') = s(bool,X652) | s(bool,X653) = s(bool,'const.min.@_1'(s(fun(bool,bool),X649))) | s(bool,'T_0') = s(bool,X653)) )),
  inference(superposition,[],[f5170,f2758])).
fof(f15776,plain,(
  ( ! [X656,X658,X654,X657,X655] : (sP1(X656,happ(s(fun(bool,bool),X654),s(bool,X655))) | p(s(bool,X656)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X654))) | s(bool,X655) = s(bool,X657) | s(bool,'T_0') = s(bool,X657) | s(bool,X658) = s(bool,'const.min.@_1'(s(fun(bool,bool),X654))) | s(bool,'T_0') = s(bool,X658)) )),
  inference(superposition,[],[f5171,f2758])).
fof(f15777,plain,(
  ( ! [X660,X662,X664,X661,X663,X659] : (sP2(X661,happ(s(fun(bool,bool),X659),s(bool,X660)),X662) | ~p(s(bool,X661)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X659))) | s(bool,X660) = s(bool,X663) | s(bool,'T_0') = s(bool,X663) | s(bool,X664) = s(bool,'const.min.@_1'(s(fun(bool,bool),X659))) | s(bool,'T_0') = s(bool,X664)) )),
  inference(superposition,[],[f5172,f2758])).
fof(f15778,plain,(
  ( ! [X666,X668,X670,X667,X669,X665] : (sP2(X667,X668,happ(s(fun(bool,bool),X665),s(bool,X666))) | ~p(s(bool,X667)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X665))) | s(bool,X666) = s(bool,X669) | s(bool,'T_0') = s(bool,X669) | s(bool,X670) = s(bool,'const.min.@_1'(s(fun(bool,bool),X665))) | s(bool,'T_0') = s(bool,X670)) )),
  inference(superposition,[],[f5173,f2758])).
fof(f15779,plain,(
  ( ! [X673,X675,X671,X676,X672,X674] : (sP3(X673,X674,happ(s(fun(bool,bool),X671),s(bool,X672))) | ~p(s(bool,X674)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X671))) | s(bool,X672) = s(bool,X675) | s(bool,'T_0') = s(bool,X675) | s(bool,X676) = s(bool,'const.min.@_1'(s(fun(bool,bool),X671))) | s(bool,'T_0') = s(bool,X676)) )),
  inference(superposition,[],[f5174,f2758])).
fof(f15780,plain,(
  ( ! [X681,X677,X679,X680,X682,X678] : (sP11(happ(s(fun(bool,bool),X677),s(bool,X678)),X679,X680) | ~p(s(bool,X680)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X677))) | s(bool,X678) = s(bool,X681) | s(bool,'T_0') = s(bool,X681) | s(bool,X682) = s(bool,'const.min.@_1'(s(fun(bool,bool),X677))) | s(bool,'T_0') = s(bool,X682)) )),
  inference(superposition,[],[f5181,f2758])).
fof(f15781,plain,(
  ( ! [X685,X687,X683,X688,X684,X686] : (sP11(happ(s(fun(bool,bool),X683),s(bool,X684)),X685,X686) | ~p(s(bool,X685)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X683))) | s(bool,X684) = s(bool,X687) | s(bool,'T_0') = s(bool,X687) | s(bool,X688) = s(bool,'const.min.@_1'(s(fun(bool,bool),X683))) | s(bool,'T_0') = s(bool,X688)) )),
  inference(superposition,[],[f5182,f2758])).
fof(f15782,plain,(
  ( ! [X693,X689,X691,X692,X694,X690] : (sP13(happ(s(fun(bool,bool),X689),s(bool,X690)),X691,X692) | p(s(bool,X692)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X689))) | s(bool,X690) = s(bool,X693) | s(bool,'T_0') = s(bool,X693) | s(bool,X694) = s(bool,'const.min.@_1'(s(fun(bool,bool),X689))) | s(bool,'T_0') = s(bool,X694)) )),
  inference(superposition,[],[f5185,f2758])).
fof(f15783,plain,(
  ( ! [X697,X699,X695,X700,X696,X698] : (sP13(happ(s(fun(bool,bool),X695),s(bool,X696)),X697,X698) | ~p(s(bool,X697)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X695))) | s(bool,X696) = s(bool,X699) | s(bool,'T_0') = s(bool,X699) | s(bool,X700) = s(bool,'const.min.@_1'(s(fun(bool,bool),X695))) | s(bool,'T_0') = s(bool,X700)) )),
  inference(superposition,[],[f5186,f2758])).
fof(f16025,plain,(
  ( ! [X532,X530,X529,X531] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X529))) = s(bool,happ(s(fun(bool,bool),X529),s(bool,X530))) | s(bool,X531) = s(bool,'const.min.@_1'(s(fun(bool,bool),X529))) | s(bool,'T_0') = s(bool,X530) | s(bool,X531) = s(bool,X532) | s(bool,'T_0') = s(bool,X532)) )),
  inference(superposition,[],[f281,f2758])).
fof(f16026,plain,(
  ( ! [X536,X534,X533,X535] : (~sP7('F_0',happ(s(fun(bool,bool),X533),s(bool,X534)),'|dollar|exists_1'(s(fun(bool,bool),X533))) | s(bool,X535) = s(bool,'const.min.@_1'(s(fun(bool,bool),X533))) | s(bool,'T_0') = s(bool,X534) | s(bool,X535) = s(bool,X536) | s(bool,'T_0') = s(bool,X536)) )),
  inference(superposition,[],[f1811,f2758])).
fof(f16027,plain,(
  ( ! [X540,X538,X537,X539] : (sP7('T_0',happ(s(fun(bool,bool),X537),s(bool,X538)),'|dollar|exists_1'(s(fun(bool,bool),X537))) | s(bool,X539) = s(bool,'const.min.@_1'(s(fun(bool,bool),X537))) | s(bool,'T_0') = s(bool,X538) | s(bool,X539) = s(bool,X540) | s(bool,'T_0') = s(bool,X540)) )),
  inference(superposition,[],[f1812,f2758])).
fof(f16028,plain,(
  ( ! [X542,X541,X543,X544] : (~sP7('F_0','|dollar|exists_1'(s(fun(bool,bool),X541)),happ(s(fun(bool,bool),X541),s(bool,X542))) | s(bool,X543) = s(bool,'const.min.@_1'(s(fun(bool,bool),X541))) | s(bool,'T_0') = s(bool,X542) | s(bool,X543) = s(bool,X544) | s(bool,'T_0') = s(bool,X544)) )),
  inference(superposition,[],[f1813,f2758])).
fof(f16029,plain,(
  ( ! [X546,X545,X547,X548] : (sP7('T_0','|dollar|exists_1'(s(fun(bool,bool),X545)),happ(s(fun(bool,bool),X545),s(bool,X546))) | s(bool,X547) = s(bool,'const.min.@_1'(s(fun(bool,bool),X545))) | s(bool,'T_0') = s(bool,X546) | s(bool,X547) = s(bool,X548) | s(bool,'T_0') = s(bool,X548)) )),
  inference(superposition,[],[f1814,f2758])).
fof(f16031,plain,(
  ( ! [X557,X553,X555,X556,X554] : (sP0(X555,happ(s(fun(bool,bool),X553),s(bool,X554))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X553)))) | s(bool,X556) = s(bool,'const.min.@_1'(s(fun(bool,bool),X553))) | s(bool,'T_0') = s(bool,X554) | s(bool,X556) = s(bool,X557) | s(bool,'T_0') = s(bool,X557)) )),
  inference(superposition,[],[f1826,f2758])).
fof(f16032,plain,(
  ( ! [X561,X559,X560,X562,X558] : (sP0(happ(s(fun(bool,bool),X558),s(bool,X559)),X560) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X558)))) | s(bool,X561) = s(bool,'const.min.@_1'(s(fun(bool,bool),X558))) | s(bool,'T_0') = s(bool,X559) | s(bool,X561) = s(bool,X562) | s(bool,'T_0') = s(bool,X562)) )),
  inference(superposition,[],[f1827,f2758])).
fof(f16033,plain,(
  ( ! [X565,X567,X563,X564,X566] : (sP1(happ(s(fun(bool,bool),X563),s(bool,X564)),X565) | p(s(bool,X565)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X563)))) | s(bool,X566) = s(bool,'const.min.@_1'(s(fun(bool,bool),X563))) | s(bool,'T_0') = s(bool,X564) | s(bool,X566) = s(bool,X567) | s(bool,'T_0') = s(bool,X567)) )),
  inference(superposition,[],[f1828,f2758])).
fof(f16044,plain,(
  ( ! [X629,X627,X628,X630] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X627)),happ(s(fun(bool,bool),X627),s(bool,X628))) | s(bool,X629) = s(bool,'const.min.@_1'(s(fun(bool,bool),X627))) | s(bool,'T_0') = s(bool,X628) | s(bool,X629) = s(bool,X630) | s(bool,'T_0') = s(bool,X630)) )),
  inference(superposition,[],[f4417,f2758])).
fof(f16045,plain,(
  ( ! [X633,X631,X632,X634] : (sP6('T_0',happ(s(fun(bool,bool),X631),s(bool,X632)),'|dollar|exists_1'(s(fun(bool,bool),X631))) | s(bool,X633) = s(bool,'const.min.@_1'(s(fun(bool,bool),X631))) | s(bool,'T_0') = s(bool,X632) | s(bool,X633) = s(bool,X634) | s(bool,'T_0') = s(bool,X634)) )),
  inference(superposition,[],[f4466,f2758])).
fof(f16046,plain,(
  ( ! [X635,X636,X638,X637] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X635),s(bool,X636))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X635))) | s(bool,X637) = s(bool,'const.min.@_1'(s(fun(bool,bool),X635))) | s(bool,'T_0') = s(bool,X636) | s(bool,X637) = s(bool,X638) | s(bool,'T_0') = s(bool,X638)) )),
  inference(superposition,[],[f5165,f2758])).
fof(f16047,plain,(
  ( ! [X639,X640,X642,X641,X643] : (sP0(X641,happ(s(fun(bool,bool),X639),s(bool,X640))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X639))) | s(bool,X642) = s(bool,'const.min.@_1'(s(fun(bool,bool),X639))) | s(bool,'T_0') = s(bool,X640) | s(bool,X642) = s(bool,X643) | s(bool,'T_0') = s(bool,X643)) )),
  inference(superposition,[],[f5168,f2758])).
fof(f16048,plain,(
  ( ! [X648,X644,X646,X645,X647] : (sP0(happ(s(fun(bool,bool),X644),s(bool,X645)),X646) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X644))) | s(bool,X647) = s(bool,'const.min.@_1'(s(fun(bool,bool),X644))) | s(bool,'T_0') = s(bool,X645) | s(bool,X647) = s(bool,X648) | s(bool,'T_0') = s(bool,X648)) )),
  inference(superposition,[],[f5169,f2758])).
fof(f16049,plain,(
  ( ! [X652,X650,X653,X649,X651] : (sP1(happ(s(fun(bool,bool),X649),s(bool,X650)),X651) | p(s(bool,X651)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X649))) | s(bool,X652) = s(bool,'const.min.@_1'(s(fun(bool,bool),X649))) | s(bool,'T_0') = s(bool,X650) | s(bool,X652) = s(bool,X653) | s(bool,'T_0') = s(bool,X653)) )),
  inference(superposition,[],[f5170,f2758])).
fof(f16050,plain,(
  ( ! [X656,X658,X654,X657,X655] : (sP1(X656,happ(s(fun(bool,bool),X654),s(bool,X655))) | p(s(bool,X656)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X654))) | s(bool,X657) = s(bool,'const.min.@_1'(s(fun(bool,bool),X654))) | s(bool,'T_0') = s(bool,X655) | s(bool,X657) = s(bool,X658) | s(bool,'T_0') = s(bool,X658)) )),
  inference(superposition,[],[f5171,f2758])).
fof(f16051,plain,(
  ( ! [X660,X662,X664,X661,X663,X659] : (sP2(X661,happ(s(fun(bool,bool),X659),s(bool,X660)),X662) | ~p(s(bool,X661)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X659))) | s(bool,X663) = s(bool,'const.min.@_1'(s(fun(bool,bool),X659))) | s(bool,'T_0') = s(bool,X660) | s(bool,X663) = s(bool,X664) | s(bool,'T_0') = s(bool,X664)) )),
  inference(superposition,[],[f5172,f2758])).
fof(f16052,plain,(
  ( ! [X666,X668,X670,X667,X669,X665] : (sP2(X667,X668,happ(s(fun(bool,bool),X665),s(bool,X666))) | ~p(s(bool,X667)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X665))) | s(bool,X669) = s(bool,'const.min.@_1'(s(fun(bool,bool),X665))) | s(bool,'T_0') = s(bool,X666) | s(bool,X669) = s(bool,X670) | s(bool,'T_0') = s(bool,X670)) )),
  inference(superposition,[],[f5173,f2758])).
fof(f16053,plain,(
  ( ! [X673,X675,X671,X676,X672,X674] : (sP3(X673,X674,happ(s(fun(bool,bool),X671),s(bool,X672))) | ~p(s(bool,X674)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X671))) | s(bool,X675) = s(bool,'const.min.@_1'(s(fun(bool,bool),X671))) | s(bool,'T_0') = s(bool,X672) | s(bool,X675) = s(bool,X676) | s(bool,'T_0') = s(bool,X676)) )),
  inference(superposition,[],[f5174,f2758])).
fof(f16054,plain,(
  ( ! [X681,X677,X679,X680,X682,X678] : (sP11(happ(s(fun(bool,bool),X677),s(bool,X678)),X679,X680) | ~p(s(bool,X680)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X677))) | s(bool,X681) = s(bool,'const.min.@_1'(s(fun(bool,bool),X677))) | s(bool,'T_0') = s(bool,X678) | s(bool,X681) = s(bool,X682) | s(bool,'T_0') = s(bool,X682)) )),
  inference(superposition,[],[f5181,f2758])).
fof(f16055,plain,(
  ( ! [X685,X687,X683,X688,X684,X686] : (sP11(happ(s(fun(bool,bool),X683),s(bool,X684)),X685,X686) | ~p(s(bool,X685)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X683))) | s(bool,X687) = s(bool,'const.min.@_1'(s(fun(bool,bool),X683))) | s(bool,'T_0') = s(bool,X684) | s(bool,X687) = s(bool,X688) | s(bool,'T_0') = s(bool,X688)) )),
  inference(superposition,[],[f5182,f2758])).
fof(f16056,plain,(
  ( ! [X693,X689,X691,X692,X694,X690] : (sP13(happ(s(fun(bool,bool),X689),s(bool,X690)),X691,X692) | p(s(bool,X692)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X689))) | s(bool,X693) = s(bool,'const.min.@_1'(s(fun(bool,bool),X689))) | s(bool,'T_0') = s(bool,X690) | s(bool,X693) = s(bool,X694) | s(bool,'T_0') = s(bool,X694)) )),
  inference(superposition,[],[f5185,f2758])).
fof(f16057,plain,(
  ( ! [X697,X699,X695,X700,X696,X698] : (sP13(happ(s(fun(bool,bool),X695),s(bool,X696)),X697,X698) | ~p(s(bool,X697)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X695))) | s(bool,X699) = s(bool,'const.min.@_1'(s(fun(bool,bool),X695))) | s(bool,'T_0') = s(bool,X696) | s(bool,X699) = s(bool,X700) | s(bool,'T_0') = s(bool,X700)) )),
  inference(superposition,[],[f5186,f2758])).
fof(f16181,plain,(
  ( ! [X613,X615,X612,X614] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X612)),happ(s(fun(bool,bool),X612),s(bool,X613))) | s(bool,X613) = s(bool,X614) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X612))) | s(bool,X614) = s(bool,X615) | s(bool,'T_0') = s(bool,X615)) )),
  inference(superposition,[],[f4417,f2758])).
fof(f16182,plain,(
  ( ! [X617,X619,X616,X618] : (sP6('T_0',happ(s(fun(bool,bool),X616),s(bool,X617)),'|dollar|exists_1'(s(fun(bool,bool),X616))) | s(bool,X617) = s(bool,X618) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X616))) | s(bool,X618) = s(bool,X619) | s(bool,'T_0') = s(bool,X619)) )),
  inference(superposition,[],[f4466,f2758])).
fof(f16348,plain,(
  ( ! [X496,X498,X495,X497] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X495))) = s(bool,happ(s(fun(bool,bool),X495),s(bool,'T_0'))) | s(bool,X496) = s(bool,X497) | s(bool,X497) = s(bool,'const.min.@_1'(s(fun(bool,bool),X495))) | s(bool,X496) = s(bool,X498) | s(bool,'T_0') = s(bool,X498)) )),
  inference(superposition,[],[f281,f2758])).
fof(f16349,plain,(
  ( ! [X499,X500,X502,X501] : (~sP7('F_0',happ(s(fun(bool,bool),X499),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X499))) | s(bool,X500) = s(bool,X501) | s(bool,X501) = s(bool,'const.min.@_1'(s(fun(bool,bool),X499))) | s(bool,X500) = s(bool,X502) | s(bool,'T_0') = s(bool,X502)) )),
  inference(superposition,[],[f1811,f2758])).
fof(f16350,plain,(
  ( ! [X505,X504,X506,X503] : (sP7('T_0',happ(s(fun(bool,bool),X503),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X503))) | s(bool,X504) = s(bool,X505) | s(bool,X505) = s(bool,'const.min.@_1'(s(fun(bool,bool),X503))) | s(bool,X504) = s(bool,X506) | s(bool,'T_0') = s(bool,X506)) )),
  inference(superposition,[],[f1812,f2758])).
fof(f16351,plain,(
  ( ! [X509,X507,X508,X510] : (~sP7('F_0','|dollar|exists_1'(s(fun(bool,bool),X507)),happ(s(fun(bool,bool),X507),s(bool,'T_0'))) | s(bool,X508) = s(bool,X509) | s(bool,X509) = s(bool,'const.min.@_1'(s(fun(bool,bool),X507))) | s(bool,X508) = s(bool,X510) | s(bool,'T_0') = s(bool,X510)) )),
  inference(superposition,[],[f1813,f2758])).
fof(f16352,plain,(
  ( ! [X511,X512,X514,X513] : (sP7('T_0','|dollar|exists_1'(s(fun(bool,bool),X511)),happ(s(fun(bool,bool),X511),s(bool,'T_0'))) | s(bool,X512) = s(bool,X513) | s(bool,X513) = s(bool,'const.min.@_1'(s(fun(bool,bool),X511))) | s(bool,X512) = s(bool,X514) | s(bool,'T_0') = s(bool,X514)) )),
  inference(superposition,[],[f1814,f2758])).
fof(f16354,plain,(
  ( ! [X520,X522,X521,X523,X519] : (sP0(X520,happ(s(fun(bool,bool),X519),s(bool,'T_0'))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X519)))) | s(bool,X521) = s(bool,X522) | s(bool,X522) = s(bool,'const.min.@_1'(s(fun(bool,bool),X519))) | s(bool,X521) = s(bool,X523) | s(bool,'T_0') = s(bool,X523)) )),
  inference(superposition,[],[f1826,f2758])).
fof(f16355,plain,(
  ( ! [X528,X524,X526,X525,X527] : (sP0(happ(s(fun(bool,bool),X524),s(bool,'T_0')),X525) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X524)))) | s(bool,X526) = s(bool,X527) | s(bool,X527) = s(bool,'const.min.@_1'(s(fun(bool,bool),X524))) | s(bool,X526) = s(bool,X528) | s(bool,'T_0') = s(bool,X528)) )),
  inference(superposition,[],[f1827,f2758])).
fof(f16356,plain,(
  ( ! [X532,X530,X533,X529,X531] : (sP1(happ(s(fun(bool,bool),X529),s(bool,'T_0')),X530) | p(s(bool,X530)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X529)))) | s(bool,X531) = s(bool,X532) | s(bool,X532) = s(bool,'const.min.@_1'(s(fun(bool,bool),X529))) | s(bool,X531) = s(bool,X533) | s(bool,'T_0') = s(bool,X533)) )),
  inference(superposition,[],[f1828,f2758])).
fof(f16367,plain,(
  ( ! [X596,X594,X593,X595] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X593)),happ(s(fun(bool,bool),X593),s(bool,'T_0'))) | s(bool,X594) = s(bool,X595) | s(bool,X595) = s(bool,'const.min.@_1'(s(fun(bool,bool),X593))) | s(bool,X594) = s(bool,X596) | s(bool,'T_0') = s(bool,X596)) )),
  inference(superposition,[],[f4417,f2758])).
fof(f16368,plain,(
  ( ! [X600,X598,X597,X599] : (sP6('T_0',happ(s(fun(bool,bool),X597),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X597))) | s(bool,X598) = s(bool,X599) | s(bool,X599) = s(bool,'const.min.@_1'(s(fun(bool,bool),X597))) | s(bool,X598) = s(bool,X600) | s(bool,'T_0') = s(bool,X600)) )),
  inference(superposition,[],[f4466,f2758])).
fof(f16369,plain,(
  ( ! [X602,X601,X603,X604] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X601),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X601))) | s(bool,X602) = s(bool,X603) | s(bool,X603) = s(bool,'const.min.@_1'(s(fun(bool,bool),X601))) | s(bool,X602) = s(bool,X604) | s(bool,'T_0') = s(bool,X604)) )),
  inference(superposition,[],[f5165,f2758])).
fof(f16370,plain,(
  ( ! [X606,X609,X605,X607,X608] : (sP0(X606,happ(s(fun(bool,bool),X605),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X605))) | s(bool,X607) = s(bool,X608) | s(bool,X608) = s(bool,'const.min.@_1'(s(fun(bool,bool),X605))) | s(bool,X607) = s(bool,X609) | s(bool,'T_0') = s(bool,X609)) )),
  inference(superposition,[],[f5168,f2758])).
fof(f16371,plain,(
  ( ! [X613,X611,X612,X614,X610] : (sP0(happ(s(fun(bool,bool),X610),s(bool,'T_0')),X611) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X610))) | s(bool,X612) = s(bool,X613) | s(bool,X613) = s(bool,'const.min.@_1'(s(fun(bool,bool),X610))) | s(bool,X612) = s(bool,X614) | s(bool,'T_0') = s(bool,X614)) )),
  inference(superposition,[],[f5169,f2758])).
fof(f16372,plain,(
  ( ! [X617,X619,X615,X616,X618] : (sP1(happ(s(fun(bool,bool),X615),s(bool,'T_0')),X616) | p(s(bool,X616)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X615))) | s(bool,X617) = s(bool,X618) | s(bool,X618) = s(bool,'const.min.@_1'(s(fun(bool,bool),X615))) | s(bool,X617) = s(bool,X619) | s(bool,'T_0') = s(bool,X619)) )),
  inference(superposition,[],[f5170,f2758])).
fof(f16373,plain,(
  ( ! [X621,X623,X624,X620,X622] : (sP1(X621,happ(s(fun(bool,bool),X620),s(bool,'T_0'))) | p(s(bool,X621)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X620))) | s(bool,X622) = s(bool,X623) | s(bool,X623) = s(bool,'const.min.@_1'(s(fun(bool,bool),X620))) | s(bool,X622) = s(bool,X624) | s(bool,'T_0') = s(bool,X624)) )),
  inference(superposition,[],[f5171,f2758])).
fof(f16374,plain,(
  ( ! [X629,X625,X627,X628,X630,X626] : (sP2(X626,happ(s(fun(bool,bool),X625),s(bool,'T_0')),X627) | ~p(s(bool,X626)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X625))) | s(bool,X628) = s(bool,X629) | s(bool,X629) = s(bool,'const.min.@_1'(s(fun(bool,bool),X625))) | s(bool,X628) = s(bool,X630) | s(bool,'T_0') = s(bool,X630)) )),
  inference(superposition,[],[f5172,f2758])).
fof(f16375,plain,(
  ( ! [X633,X635,X631,X636,X632,X634] : (sP2(X632,X633,happ(s(fun(bool,bool),X631),s(bool,'T_0'))) | ~p(s(bool,X632)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X631))) | s(bool,X634) = s(bool,X635) | s(bool,X635) = s(bool,'const.min.@_1'(s(fun(bool,bool),X631))) | s(bool,X634) = s(bool,X636) | s(bool,'T_0') = s(bool,X636)) )),
  inference(superposition,[],[f5173,f2758])).
fof(f16376,plain,(
  ( ! [X639,X640,X642,X638,X641,X637] : (sP3(X638,X639,happ(s(fun(bool,bool),X637),s(bool,'T_0'))) | ~p(s(bool,X639)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X637))) | s(bool,X640) = s(bool,X641) | s(bool,X641) = s(bool,'const.min.@_1'(s(fun(bool,bool),X637))) | s(bool,X640) = s(bool,X642) | s(bool,'T_0') = s(bool,X642)) )),
  inference(superposition,[],[f5174,f2758])).
fof(f16377,plain,(
  ( ! [X648,X644,X646,X645,X647,X643] : (sP11(happ(s(fun(bool,bool),X643),s(bool,'T_0')),X644,X645) | ~p(s(bool,X645)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X643))) | s(bool,X646) = s(bool,X647) | s(bool,X647) = s(bool,'const.min.@_1'(s(fun(bool,bool),X643))) | s(bool,X646) = s(bool,X648) | s(bool,'T_0') = s(bool,X648)) )),
  inference(superposition,[],[f5181,f2758])).
fof(f16378,plain,(
  ( ! [X652,X654,X650,X653,X649,X651] : (sP11(happ(s(fun(bool,bool),X649),s(bool,'T_0')),X650,X651) | ~p(s(bool,X650)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X649))) | s(bool,X652) = s(bool,X653) | s(bool,X653) = s(bool,'const.min.@_1'(s(fun(bool,bool),X649))) | s(bool,X652) = s(bool,X654) | s(bool,'T_0') = s(bool,X654)) )),
  inference(superposition,[],[f5182,f2758])).
fof(f16379,plain,(
  ( ! [X660,X656,X658,X657,X659,X655] : (sP13(happ(s(fun(bool,bool),X655),s(bool,'T_0')),X656,X657) | p(s(bool,X657)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X655))) | s(bool,X658) = s(bool,X659) | s(bool,X659) = s(bool,'const.min.@_1'(s(fun(bool,bool),X655))) | s(bool,X658) = s(bool,X660) | s(bool,'T_0') = s(bool,X660)) )),
  inference(superposition,[],[f5185,f2758])).
fof(f16380,plain,(
  ( ! [X666,X662,X664,X661,X663,X665] : (sP13(happ(s(fun(bool,bool),X661),s(bool,'T_0')),X662,X663) | ~p(s(bool,X662)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X661))) | s(bool,X664) = s(bool,X665) | s(bool,X665) = s(bool,'const.min.@_1'(s(fun(bool,bool),X661))) | s(bool,X664) = s(bool,X666) | s(bool,'T_0') = s(bool,X666)) )),
  inference(superposition,[],[f5186,f2758])).
fof(f18288,plain,(
  ( ! [X300,X298,X297,X299] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X298))) = s(bool,happ(s(fun(bool,bool),X298),s(bool,X297))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X298))) | s(bool,X297) = s(bool,X299) | s(bool,X299) = s(bool,X300) | s(bool,'T_0') = s(bool,X300)) )),
  inference(subsumption_resolution,[],[f16124,f477])).
fof(f16124,plain,(
  ( ! [X300,X298,X297,X299] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X298))) = s(bool,happ(s(fun(bool,bool),X298),s(bool,X297))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X298))) | s(bool,X297) = s(bool,X299) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X299) = s(bool,X300) | s(bool,'T_0') = s(bool,X300)) )),
  inference(superposition,[],[f1815,f2758])).
fof(f18291,plain,(
  ( ! [X316,X318,X317,X319] : (~sP7('F_0',happ(s(fun(bool,bool),X317),s(bool,X316)),'|dollar|exists_1'(s(fun(bool,bool),X317))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X317))) | s(bool,X316) = s(bool,X318) | s(bool,X318) = s(bool,X319) | s(bool,'T_0') = s(bool,X319)) )),
  inference(subsumption_resolution,[],[f16127,f477])).
fof(f16127,plain,(
  ( ! [X316,X318,X317,X319] : (~sP7('F_0',happ(s(fun(bool,bool),X317),s(bool,X316)),'|dollar|exists_1'(s(fun(bool,bool),X317))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X317))) | s(bool,X316) = s(bool,X318) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X318) = s(bool,X319) | s(bool,'T_0') = s(bool,X319)) )),
  inference(superposition,[],[f4376,f2758])).
fof(f18292,plain,(
  ( ! [X321,X323,X320,X322] : (sP7('T_0',happ(s(fun(bool,bool),X321),s(bool,X320)),'|dollar|exists_1'(s(fun(bool,bool),X321))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X321))) | s(bool,X320) = s(bool,X322) | s(bool,X322) = s(bool,X323) | s(bool,'T_0') = s(bool,X323)) )),
  inference(subsumption_resolution,[],[f16128,f477])).
fof(f16128,plain,(
  ( ! [X321,X323,X320,X322] : (sP7('T_0',happ(s(fun(bool,bool),X321),s(bool,X320)),'|dollar|exists_1'(s(fun(bool,bool),X321))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X321))) | s(bool,X320) = s(bool,X322) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X322) = s(bool,X323) | s(bool,'T_0') = s(bool,X323)) )),
  inference(superposition,[],[f4424,f2758])).
fof(f18293,plain,(
  ( ! [X325,X327,X324,X326] : (~sP7('F_0','|dollar|exists_1'(s(fun(bool,bool),X325)),happ(s(fun(bool,bool),X325),s(bool,X324))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X325))) | s(bool,X324) = s(bool,X326) | s(bool,X326) = s(bool,X327) | s(bool,'T_0') = s(bool,X327)) )),
  inference(subsumption_resolution,[],[f16129,f477])).
fof(f16129,plain,(
  ( ! [X325,X327,X324,X326] : (~sP7('F_0','|dollar|exists_1'(s(fun(bool,bool),X325)),happ(s(fun(bool,bool),X325),s(bool,X324))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X325))) | s(bool,X324) = s(bool,X326) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X326) = s(bool,X327) | s(bool,'T_0') = s(bool,X327)) )),
  inference(superposition,[],[f4453,f2758])).
fof(f18294,plain,(
  ( ! [X329,X331,X328,X330] : (sP7('T_0','|dollar|exists_1'(s(fun(bool,bool),X329)),happ(s(fun(bool,bool),X329),s(bool,X328))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X329))) | s(bool,X328) = s(bool,X330) | s(bool,X330) = s(bool,X331) | s(bool,'T_0') = s(bool,X331)) )),
  inference(subsumption_resolution,[],[f16130,f477])).
fof(f16130,plain,(
  ( ! [X329,X331,X328,X330] : (sP7('T_0','|dollar|exists_1'(s(fun(bool,bool),X329)),happ(s(fun(bool,bool),X329),s(bool,X328))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X329))) | s(bool,X328) = s(bool,X330) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X330) = s(bool,X331) | s(bool,'T_0') = s(bool,X331)) )),
  inference(superposition,[],[f4473,f2758])).
fof(f18514,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP7('F_0',happ(s(fun(bool,bool),X1),s(bool,X0)),'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f4372,f2758])).
fof(f18726,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('T_0',happ(s(fun(bool,bool),X1),s(bool,X0)),'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f4420,f2758])).
fof(f19649,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP7('F_0','|dollar|exists_1'(s(fun(bool,bool),X1)),happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f4449,f2758])).
fof(f19869,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('T_0','|dollar|exists_1'(s(fun(bool,bool),X1)),happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f4469,f2758])).
fof(f20419,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f20407,f239])).
fof(f20407,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f5061,f639])).
fof(f20426,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X27) | s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f20415,f239])).
fof(f20415,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X27) | s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f5061,f2758])).
fof(f20500,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) )),
  inference(subsumption_resolution,[],[f20487,f426])).
fof(f20487,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5062,f247])).
fof(f29776,plain,(
  ( ! [X412,X413] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X412))) = s(bool,happ(s(fun(bool,bool),X412),s(bool,X413))) | s(bool,X413) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X412)))) )),
  inference(superposition,[],[f281,f9775])).
fof(f30070,plain,(
  ( ! [X393,X392] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X392))) = s(bool,happ(s(fun(bool,bool),X392),s(bool,'T_0'))) | s(bool,X393) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X393) = s(bool,'const.min.@_1'(s(fun(bool,bool),X392)))) )),
  inference(superposition,[],[f281,f9775])).
fof(f34188,plain,(
  ( ! [X23,X21,X19,X22,X20] : (sP1(X21,happ(s(fun(bool,bool),X19),s(bool,X20))) | p(s(bool,X21)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X19)))) | s(bool,X22) = s(bool,'const.min.@_1'(s(fun(bool,bool),X19))) | s(bool,'T_0') = s(bool,X22) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f1829,f2758])).
fof(f34192,plain,(
  ( ! [X39,X43,X41,X42,X40] : (sP1(X40,happ(s(fun(bool,bool),X39),s(bool,'T_0'))) | p(s(bool,X40)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X39)))) | s(bool,X41) = s(bool,X42) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,X41) = s(bool,'const.min.@_1'(s(fun(bool,bool),X39)))) )),
  inference(superposition,[],[f1829,f2758])).
fof(f34193,plain,(
  ( ! [X47,X45,X48,X46,X44] : (sP1(X46,happ(s(fun(bool,bool),X44),s(bool,X45))) | p(s(bool,X46)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X44)))) | s(bool,X47) = s(bool,'const.min.@_1'(s(fun(bool,bool),X44))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f1829,f2758])).
fof(f34300,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (sP2(X27,happ(s(fun(bool,bool),X25),s(bool,X26)),X28) | ~p(s(bool,X27)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25)))) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1831,f2758])).
fof(f34304,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (sP2(X50,happ(s(fun(bool,bool),X49),s(bool,'T_0')),X51) | ~p(s(bool,X50)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) | s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X52) = s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f1831,f2758])).
fof(f34305,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (sP2(X57,happ(s(fun(bool,bool),X55),s(bool,X56)),X58) | ~p(s(bool,X57)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X55)))) | s(bool,X59) = s(bool,'const.min.@_1'(s(fun(bool,bool),X55))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f1831,f2758])).
fof(f34461,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (sP2(X27,X28,happ(s(fun(bool,bool),X25),s(bool,X26))) | ~p(s(bool,X27)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25)))) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1832,f2758])).
fof(f34465,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (sP2(X50,X51,happ(s(fun(bool,bool),X49),s(bool,'T_0'))) | ~p(s(bool,X50)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) | s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X52) = s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f1832,f2758])).
fof(f34466,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (sP2(X57,X58,happ(s(fun(bool,bool),X55),s(bool,X56))) | ~p(s(bool,X57)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X55)))) | s(bool,X59) = s(bool,'const.min.@_1'(s(fun(bool,bool),X55))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f1832,f2758])).
fof(f34671,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (sP3(X27,X28,happ(s(fun(bool,bool),X25),s(bool,X26))) | ~p(s(bool,X28)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25)))) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1833,f2758])).
fof(f34675,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (sP3(X50,X51,happ(s(fun(bool,bool),X49),s(bool,'T_0'))) | ~p(s(bool,X51)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) | s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X52) = s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f1833,f2758])).
fof(f34676,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (sP3(X57,X58,happ(s(fun(bool,bool),X55),s(bool,X56))) | ~p(s(bool,X58)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X55)))) | s(bool,X59) = s(bool,'const.min.@_1'(s(fun(bool,bool),X55))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f1833,f2758])).
fof(f34778,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (sP11(happ(s(fun(bool,bool),X25),s(bool,X26)),X27,X28) | ~p(s(bool,X28)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25)))) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1840,f2758])).
fof(f34782,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (sP11(happ(s(fun(bool,bool),X49),s(bool,'T_0')),X50,X51) | ~p(s(bool,X51)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) | s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X52) = s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f1840,f2758])).
fof(f34783,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (sP11(happ(s(fun(bool,bool),X55),s(bool,X56)),X57,X58) | ~p(s(bool,X58)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X55)))) | s(bool,X59) = s(bool,'const.min.@_1'(s(fun(bool,bool),X55))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f1840,f2758])).
fof(f34954,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (sP11(happ(s(fun(bool,bool),X25),s(bool,X26)),X27,X28) | ~p(s(bool,X27)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25)))) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1841,f2758])).
fof(f34958,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (sP11(happ(s(fun(bool,bool),X49),s(bool,'T_0')),X50,X51) | ~p(s(bool,X50)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) | s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X52) = s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f1841,f2758])).
fof(f34959,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (sP11(happ(s(fun(bool,bool),X55),s(bool,X56)),X57,X58) | ~p(s(bool,X57)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X55)))) | s(bool,X59) = s(bool,'const.min.@_1'(s(fun(bool,bool),X55))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f1841,f2758])).
fof(f36158,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (sP13(happ(s(fun(bool,bool),X25),s(bool,X26)),X27,X28) | p(s(bool,X28)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25)))) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1844,f2758])).
fof(f36162,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (sP13(happ(s(fun(bool,bool),X49),s(bool,'T_0')),X50,X51) | p(s(bool,X51)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) | s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X52) = s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f1844,f2758])).
fof(f36163,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (sP13(happ(s(fun(bool,bool),X55),s(bool,X56)),X57,X58) | p(s(bool,X58)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X55)))) | s(bool,X59) = s(bool,'const.min.@_1'(s(fun(bool,bool),X55))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f1844,f2758])).
fof(f36193,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (sP13(happ(s(fun(bool,bool),X25),s(bool,X26)),X27,X28) | ~p(s(bool,X27)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25)))) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1845,f2758])).
fof(f36197,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (sP13(happ(s(fun(bool,bool),X49),s(bool,'T_0')),X50,X51) | ~p(s(bool,X50)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) | s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X52) = s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f1845,f2758])).
fof(f36198,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (sP13(happ(s(fun(bool,bool),X55),s(bool,X56)),X57,X58) | ~p(s(bool,X57)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X55)))) | s(bool,X59) = s(bool,'const.min.@_1'(s(fun(bool,bool),X55))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f1845,f2758])).
fof(f36232,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X13),s(bool,X14))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13))) | s(bool,X15) = s(bool,'const.min.@_1'(s(fun(bool,bool),X13))) | s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f1816,f2758])).
fof(f36236,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X29),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X29))) | s(bool,X30) = s(bool,X31) | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) = s(bool,'const.min.@_1'(s(fun(bool,bool),X29)))) )),
  inference(superposition,[],[f1816,f2758])).
fof(f36237,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X33),s(bool,X34))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X33))) | s(bool,X35) = s(bool,'const.min.@_1'(s(fun(bool,bool),X33))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1816,f2758])).
fof(f36240,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X43),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43))) | s(bool,X44) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X44) = s(bool,'const.min.@_1'(s(fun(bool,bool),X43)))) )),
  inference(superposition,[],[f1816,f9775])).
fof(f36357,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X41),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X41))) | s(bool,X42) = s(bool,'const.min.@_1'(s(fun(bool,bool),X41))) | s(bool,'T_0') = s(bool,X42)) )),
  inference(forward_demodulation,[],[f36239,f481])).
fof(f36239,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X41),s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X41))) | s(bool,X42) = s(bool,'const.min.@_1'(s(fun(bool,bool),X41))) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f1816,f9775])).
fof(f37219,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (sP13(happ(s(fun(bool,bool),X25),s(bool,X26)),X27,X28) | ~sP12(happ(s(fun(bool,bool),X25),s(bool,'const.min.@_1'(s(fun(bool,bool),X25)))),X27,X28) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25)))) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1847,f2758])).
fof(f37223,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (sP13(happ(s(fun(bool,bool),X49),s(bool,'T_0')),X50,X51) | ~sP12(happ(s(fun(bool,bool),X49),s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))),X50,X51) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) | s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X52) = s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f1847,f2758])).
fof(f37224,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (sP13(happ(s(fun(bool,bool),X55),s(bool,X56)),X57,X58) | ~sP12(happ(s(fun(bool,bool),X55),s(bool,'const.min.@_1'(s(fun(bool,bool),X55)))),X57,X58) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X55)))) | s(bool,X59) = s(bool,'const.min.@_1'(s(fun(bool,bool),X55))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f1847,f2758])).
fof(f37292,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (sP11(happ(s(fun(bool,bool),X25),s(bool,X26)),X27,X28) | ~sP10(happ(s(fun(bool,bool),X25),s(bool,'const.min.@_1'(s(fun(bool,bool),X25)))),X27,X28) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25)))) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1848,f2758])).
fof(f37296,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (sP11(happ(s(fun(bool,bool),X49),s(bool,'T_0')),X50,X51) | ~sP10(happ(s(fun(bool,bool),X49),s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))),X50,X51) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) | s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X52) = s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f1848,f2758])).
fof(f37297,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (sP11(happ(s(fun(bool,bool),X55),s(bool,X56)),X57,X58) | ~sP10(happ(s(fun(bool,bool),X55),s(bool,'const.min.@_1'(s(fun(bool,bool),X55)))),X57,X58) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X55)))) | s(bool,X59) = s(bool,'const.min.@_1'(s(fun(bool,bool),X55))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f1848,f2758])).
fof(f39457,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X79))) | sP11(happ(s(fun(bool,bool),X79),s(bool,X80)),X81,X82) | ~sP10(happ(s(fun(bool,bool),X79),s(bool,'const.min.@_1'(s(fun(bool,bool),X79)))),X81,X82) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X79)))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(forward_demodulation,[],[f37302,f481])).
fof(f37302,plain,(
  ( ! [X80,X81,X79,X82] : (sP11(happ(s(fun(bool,bool),X79),s(bool,X80)),X81,X82) | ~sP10(happ(s(fun(bool,bool),X79),s(bool,'const.min.@_1'(s(fun(bool,bool),X79)))),X81,X82) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X79)))) | s(bool,'const.min.@_1'(s(fun(bool,bool),X79))) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f1848,f9775])).
fof(f39459,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,X74) | sP11(happ(s(fun(bool,bool),X71),s(bool,'T_0')),X72,X73) | ~sP10(happ(s(fun(bool,bool),X71),s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))),X72,X73) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) | s(bool,X74) = s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))) )),
  inference(forward_demodulation,[],[f37300,f481])).
fof(f37300,plain,(
  ( ! [X74,X72,X71,X73] : (sP11(happ(s(fun(bool,bool),X71),s(bool,'T_0')),X72,X73) | ~sP10(happ(s(fun(bool,bool),X71),s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))),X72,X73) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) | s(bool,X74) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X74) = s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))) )),
  inference(superposition,[],[f1848,f9775])).
fof(f39460,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X79))) | sP13(happ(s(fun(bool,bool),X79),s(bool,X80)),X81,X82) | ~sP12(happ(s(fun(bool,bool),X79),s(bool,'const.min.@_1'(s(fun(bool,bool),X79)))),X81,X82) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X79)))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(forward_demodulation,[],[f37229,f481])).
fof(f37229,plain,(
  ( ! [X80,X81,X79,X82] : (sP13(happ(s(fun(bool,bool),X79),s(bool,X80)),X81,X82) | ~sP12(happ(s(fun(bool,bool),X79),s(bool,'const.min.@_1'(s(fun(bool,bool),X79)))),X81,X82) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X79)))) | s(bool,'const.min.@_1'(s(fun(bool,bool),X79))) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f1847,f9775])).
fof(f39462,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,X74) | sP13(happ(s(fun(bool,bool),X71),s(bool,'T_0')),X72,X73) | ~sP12(happ(s(fun(bool,bool),X71),s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))),X72,X73) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) | s(bool,X74) = s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))) )),
  inference(forward_demodulation,[],[f37227,f481])).
fof(f37227,plain,(
  ( ! [X74,X72,X71,X73] : (sP13(happ(s(fun(bool,bool),X71),s(bool,'T_0')),X72,X73) | ~sP12(happ(s(fun(bool,bool),X71),s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))),X72,X73) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) | s(bool,X74) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X74) = s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))) )),
  inference(superposition,[],[f1847,f9775])).
fof(f39464,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X79))) | sP13(happ(s(fun(bool,bool),X79),s(bool,X80)),X81,X82) | ~p(s(bool,X81)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X79)))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(forward_demodulation,[],[f36203,f481])).
fof(f36203,plain,(
  ( ! [X80,X81,X79,X82] : (sP13(happ(s(fun(bool,bool),X79),s(bool,X80)),X81,X82) | ~p(s(bool,X81)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X79)))) | s(bool,'const.min.@_1'(s(fun(bool,bool),X79))) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f1845,f9775])).
fof(f39466,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,X74) | sP13(happ(s(fun(bool,bool),X71),s(bool,'T_0')),X72,X73) | ~p(s(bool,X72)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) | s(bool,X74) = s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))) )),
  inference(forward_demodulation,[],[f36201,f481])).
fof(f36201,plain,(
  ( ! [X74,X72,X71,X73] : (sP13(happ(s(fun(bool,bool),X71),s(bool,'T_0')),X72,X73) | ~p(s(bool,X72)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) | s(bool,X74) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X74) = s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))) )),
  inference(superposition,[],[f1845,f9775])).
fof(f39467,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X79))) | sP13(happ(s(fun(bool,bool),X79),s(bool,X80)),X81,X82) | p(s(bool,X82)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X79)))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(forward_demodulation,[],[f36168,f481])).
fof(f36168,plain,(
  ( ! [X80,X81,X79,X82] : (sP13(happ(s(fun(bool,bool),X79),s(bool,X80)),X81,X82) | p(s(bool,X82)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X79)))) | s(bool,'const.min.@_1'(s(fun(bool,bool),X79))) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f1844,f9775])).
fof(f39469,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,X74) | sP13(happ(s(fun(bool,bool),X71),s(bool,'T_0')),X72,X73) | p(s(bool,X73)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) | s(bool,X74) = s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))) )),
  inference(forward_demodulation,[],[f36166,f481])).
fof(f36166,plain,(
  ( ! [X74,X72,X71,X73] : (sP13(happ(s(fun(bool,bool),X71),s(bool,'T_0')),X72,X73) | p(s(bool,X73)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) | s(bool,X74) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X74) = s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))) )),
  inference(superposition,[],[f1844,f9775])).
fof(f41846,plain,(
  ( ! [X43,X44] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X43),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43))) | s(bool,X44) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X44) = s(bool,'const.min.@_1'(s(fun(bool,bool),X43)))) )),
  inference(superposition,[],[f5165,f9775])).
fof(f41847,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X45),s(bool,X46))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X45))) | s(bool,X46) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X45)))) )),
  inference(superposition,[],[f5165,f9775])).
fof(f40436,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,X71) = s(bool,'|dollar|exists_1'(s(fun(X72,bool),X73))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X72,bool),X73))) | s(bool,X74) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X71) = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f40406,f239])).
fof(f40406,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | s(bool,X71) = s(bool,'|dollar|exists_1'(s(fun(X72,bool),X73))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X72,bool),X73))) | s(bool,X74) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X71) = s(bool,X74)) )),
  inference(superposition,[],[f5167,f9775])).
fof(f4488,plain,(
  ( ! [X12,X13] : (sP0(X13,happ(s(fun(bool,bool),X12),s(bool,'F_0'))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X12)))) )),
  inference(superposition,[],[f1826,f247])).
fof(f4501,plain,(
  ( ! [X12,X13] : (sP0(happ(s(fun(bool,bool),X12),s(bool,'F_0')),X13) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X12)))) )),
  inference(superposition,[],[f1827,f247])).
fof(f5074,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X32)))) | sP7(X33,X34,happ(s(fun(bool,bool),X32),s(bool,'F_0'))) | ~sP6(X33,happ(s(fun(bool,bool),X32),s(bool,'F_0')),X34) | ~sP5(X34,happ(s(fun(bool,bool),X32),s(bool,'F_0')),X33) | ~sP4(happ(s(fun(bool,bool),X32),s(bool,'F_0')),X34,X33) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X32)))) )),
  inference(superposition,[],[f346,f1815])).
fof(f5087,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) | ~sP12(happ(s(fun(bool,bool),X71),s(bool,'F_0')),X72,X73) | sP13(happ(s(fun(bool,bool),X71),s(bool,'F_0')),X72,X73) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X71)))) )),
  inference(superposition,[],[f416,f1815])).
fof(f5088,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X74)))) | ~sP10(happ(s(fun(bool,bool),X74),s(bool,'F_0')),X75,X76) | sP11(happ(s(fun(bool,bool),X74),s(bool,'F_0')),X75,X76) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X74)))) )),
  inference(superposition,[],[f417,f1815])).
fof(f5993,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X46)))) | sP1(X47,happ(s(fun(bool,bool),X46),s(bool,'F_0'))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X46)))) )),
  inference(superposition,[],[f3259,f1815])).
fof(f8870,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X46)))) | sP1(happ(s(fun(bool,bool),X46),s(bool,'F_0')),X47) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X46)))) )),
  inference(superposition,[],[f3260,f1815])).
fof(f11674,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X57)))) | p(s(bool,X58)) | ~sP9(X58,happ(s(fun(bool,bool),X57),s(bool,'F_0')),X59) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X57)))) )),
  inference(superposition,[],[f3268,f1815])).
fof(f11955,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) | p(s(bool,X61)) | ~sP13(X62,X61,happ(s(fun(bool,bool),X60),s(bool,'F_0'))) | s(bool,X62) = s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f3276,f1815])).
fof(f12069,plain,(
  ( ! [X0,X1] : (sP0(X1,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5168,f247])).
fof(f12192,plain,(
  ( ! [X0,X1] : (sP0(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5169,f247])).
fof(f14203,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5165,f247])).
fof(f16170,plain,(
  ( ! [X549,X551,X552,X548,X550] : (sP1(happ(s(fun(bool,bool),X548),s(bool,X549)),X550) | p(s(bool,X550)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X548)))) | s(bool,X549) = s(bool,X551) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X548))) | s(bool,X551) = s(bool,X552) | s(bool,'T_0') = s(bool,X552)) )),
  inference(superposition,[],[f1828,f2758])).
fof(f16186,plain,(
  ( ! [X635,X636,X638,X634,X637] : (sP1(happ(s(fun(bool,bool),X634),s(bool,X635)),X636) | p(s(bool,X636)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X634))) | s(bool,X635) = s(bool,X637) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X634))) | s(bool,X637) = s(bool,X638) | s(bool,'T_0') = s(bool,X638)) )),
  inference(superposition,[],[f5170,f2758])).
fof(f16187,plain,(
  ( ! [X639,X640,X642,X641,X643] : (sP1(X641,happ(s(fun(bool,bool),X639),s(bool,X640))) | p(s(bool,X641)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X639))) | s(bool,X640) = s(bool,X642) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X639))) | s(bool,X642) = s(bool,X643) | s(bool,'T_0') = s(bool,X643)) )),
  inference(superposition,[],[f5171,f2758])).
fof(f16188,plain,(
  ( ! [X648,X644,X646,X649,X645,X647] : (sP2(X646,happ(s(fun(bool,bool),X644),s(bool,X645)),X647) | ~p(s(bool,X646)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X644))) | s(bool,X645) = s(bool,X648) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X644))) | s(bool,X648) = s(bool,X649) | s(bool,'T_0') = s(bool,X649)) )),
  inference(superposition,[],[f5172,f2758])).
fof(f16189,plain,(
  ( ! [X652,X654,X650,X653,X655,X651] : (sP2(X652,X653,happ(s(fun(bool,bool),X650),s(bool,X651))) | ~p(s(bool,X652)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X650))) | s(bool,X651) = s(bool,X654) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X650))) | s(bool,X654) = s(bool,X655) | s(bool,'T_0') = s(bool,X655)) )),
  inference(superposition,[],[f5173,f2758])).
fof(f16190,plain,(
  ( ! [X660,X656,X658,X661,X657,X659] : (sP3(X658,X659,happ(s(fun(bool,bool),X656),s(bool,X657))) | ~p(s(bool,X659)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X656))) | s(bool,X657) = s(bool,X660) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X656))) | s(bool,X660) = s(bool,X661) | s(bool,'T_0') = s(bool,X661)) )),
  inference(superposition,[],[f5174,f2758])).
fof(f16191,plain,(
  ( ! [X666,X662,X664,X667,X663,X665] : (sP11(happ(s(fun(bool,bool),X662),s(bool,X663)),X664,X665) | ~p(s(bool,X665)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X662))) | s(bool,X663) = s(bool,X666) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X662))) | s(bool,X666) = s(bool,X667) | s(bool,'T_0') = s(bool,X667)) )),
  inference(superposition,[],[f5181,f2758])).
fof(f16192,plain,(
  ( ! [X668,X670,X673,X669,X671,X672] : (sP11(happ(s(fun(bool,bool),X668),s(bool,X669)),X670,X671) | ~p(s(bool,X670)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X668))) | s(bool,X669) = s(bool,X672) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X668))) | s(bool,X672) = s(bool,X673) | s(bool,'T_0') = s(bool,X673)) )),
  inference(superposition,[],[f5182,f2758])).
fof(f16193,plain,(
  ( ! [X677,X679,X675,X676,X678,X674] : (sP13(happ(s(fun(bool,bool),X674),s(bool,X675)),X676,X677) | p(s(bool,X677)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X674))) | s(bool,X675) = s(bool,X678) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X674))) | s(bool,X678) = s(bool,X679) | s(bool,'T_0') = s(bool,X679)) )),
  inference(superposition,[],[f5185,f2758])).
fof(f16194,plain,(
  ( ! [X685,X681,X683,X684,X680,X682] : (sP13(happ(s(fun(bool,bool),X680),s(bool,X681)),X682,X683) | ~p(s(bool,X682)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X680))) | s(bool,X681) = s(bool,X684) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X680))) | s(bool,X684) = s(bool,X685) | s(bool,'T_0') = s(bool,X685)) )),
  inference(superposition,[],[f5186,f2758])).
fof(f18295,plain,(
  ( ! [X333,X335,X332,X334,X336] : (sP0(X333,happ(s(fun(bool,bool),X334),s(bool,X332))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X334)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X334))) | s(bool,X332) = s(bool,X335) | s(bool,X335) = s(bool,X336) | s(bool,'T_0') = s(bool,X336)) )),
  inference(subsumption_resolution,[],[f16131,f477])).
fof(f16131,plain,(
  ( ! [X333,X335,X332,X334,X336] : (sP0(X333,happ(s(fun(bool,bool),X334),s(bool,X332))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X334)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X334))) | s(bool,X332) = s(bool,X335) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X335) = s(bool,X336) | s(bool,'T_0') = s(bool,X336)) )),
  inference(superposition,[],[f4488,f2758])).
fof(f18296,plain,(
  ( ! [X341,X337,X339,X340,X338] : (sP0(happ(s(fun(bool,bool),X338),s(bool,X337)),X339) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X338)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X338))) | s(bool,X337) = s(bool,X340) | s(bool,X340) = s(bool,X341) | s(bool,'T_0') = s(bool,X341)) )),
  inference(subsumption_resolution,[],[f16132,f477])).
fof(f16132,plain,(
  ( ! [X341,X337,X339,X340,X338] : (sP0(happ(s(fun(bool,bool),X338),s(bool,X337)),X339) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X338)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X338))) | s(bool,X337) = s(bool,X340) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X340) = s(bool,X341) | s(bool,'T_0') = s(bool,X341)) )),
  inference(superposition,[],[f4501,f2758])).
fof(f34194,plain,(
  ( ! [X52,X50,X53,X51,X49] : (sP1(X51,happ(s(fun(bool,bool),X49),s(bool,X50))) | p(s(bool,X51)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) | s(bool,X50) = s(bool,X52) | s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f1829,f2758])).
fof(f34306,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (sP2(X63,happ(s(fun(bool,bool),X61),s(bool,X62)),X64) | ~p(s(bool,X63)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61)))) | s(bool,X62) = s(bool,X65) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))) )),
  inference(superposition,[],[f1831,f2758])).
fof(f34467,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (sP2(X63,X64,happ(s(fun(bool,bool),X61),s(bool,X62))) | ~p(s(bool,X63)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61)))) | s(bool,X62) = s(bool,X65) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))) )),
  inference(superposition,[],[f1832,f2758])).
fof(f34677,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (sP3(X63,X64,happ(s(fun(bool,bool),X61),s(bool,X62))) | ~p(s(bool,X64)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61)))) | s(bool,X62) = s(bool,X65) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))) )),
  inference(superposition,[],[f1833,f2758])).
fof(f34784,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (sP11(happ(s(fun(bool,bool),X61),s(bool,X62)),X63,X64) | ~p(s(bool,X64)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61)))) | s(bool,X62) = s(bool,X65) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))) )),
  inference(superposition,[],[f1840,f2758])).
fof(f34960,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (sP11(happ(s(fun(bool,bool),X61),s(bool,X62)),X63,X64) | ~p(s(bool,X63)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61)))) | s(bool,X62) = s(bool,X65) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))) )),
  inference(superposition,[],[f1841,f2758])).
fof(f4486,plain,(
  ( ! [X6,X8,X7] : (sP0(X8,happ(s(fun(bool,bool),X6),s(bool,X7))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1826,f639])).
fof(f4499,plain,(
  ( ! [X6,X8,X7] : (sP0(happ(s(fun(bool,bool),X6),s(bool,X7)),X8) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1827,f639])).
fof(f4514,plain,(
  ( ! [X6,X8,X7] : (sP1(happ(s(fun(bool,bool),X6),s(bool,X7)),X8) | p(s(bool,X8)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1828,f639])).
fof(f4516,plain,(
  ( ! [X12,X13] : (sP1(happ(s(fun(bool,bool),X12),s(bool,'F_0')),X13) | p(s(bool,X13)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X12)))) )),
  inference(superposition,[],[f1828,f247])).
fof(f4525,plain,(
  ( ! [X6,X8,X7] : (sP1(X8,happ(s(fun(bool,bool),X6),s(bool,X7))) | p(s(bool,X8)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1829,f639])).
fof(f4527,plain,(
  ( ! [X12,X13] : (sP1(X13,happ(s(fun(bool,bool),X12),s(bool,'F_0'))) | p(s(bool,X13)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X12)))) )),
  inference(superposition,[],[f1829,f247])).
fof(f4538,plain,(
  ( ! [X10,X8,X11,X9] : (sP2(X10,happ(s(fun(bool,bool),X8),s(bool,X9)),X11) | ~p(s(bool,X10)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1831,f639])).
fof(f4540,plain,(
  ( ! [X17,X18,X16] : (sP2(X17,happ(s(fun(bool,bool),X16),s(bool,'F_0')),X18) | ~p(s(bool,X17)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1831,f247])).
fof(f4551,plain,(
  ( ! [X10,X8,X11,X9] : (sP2(X10,X11,happ(s(fun(bool,bool),X8),s(bool,X9))) | ~p(s(bool,X10)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1832,f639])).
fof(f4553,plain,(
  ( ! [X17,X18,X16] : (sP2(X17,X18,happ(s(fun(bool,bool),X16),s(bool,'F_0'))) | ~p(s(bool,X17)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1832,f247])).
fof(f4566,plain,(
  ( ! [X10,X8,X11,X9] : (sP3(X10,X11,happ(s(fun(bool,bool),X8),s(bool,X9))) | ~p(s(bool,X11)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1833,f639])).
fof(f4568,plain,(
  ( ! [X17,X18,X16] : (sP3(X17,X18,happ(s(fun(bool,bool),X16),s(bool,'F_0'))) | ~p(s(bool,X18)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1833,f247])).
fof(f4577,plain,(
  ( ! [X10,X8,X11,X9] : (sP11(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | ~p(s(bool,X11)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1840,f639])).
fof(f4579,plain,(
  ( ! [X17,X18,X16] : (sP11(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | ~p(s(bool,X18)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1840,f247])).
fof(f4590,plain,(
  ( ! [X10,X8,X11,X9] : (sP11(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | ~p(s(bool,X10)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1841,f639])).
fof(f4592,plain,(
  ( ! [X17,X18,X16] : (sP11(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | ~p(s(bool,X17)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1841,f247])).
fof(f5075,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X35)))) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP6(X36,happ(s(fun(bool,bool),X35),s(bool,'F_0')),X37) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X35)))) )),
  inference(superposition,[],[f349,f1815])).
fof(f5076,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X38)))) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,happ(s(fun(bool,bool),X38),s(bool,'F_0')),X40) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X38)))) )),
  inference(superposition,[],[f353,f1815])).
fof(f5077,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X41)))) | ~p(s(bool,X42)) | p(s(bool,X43)) | ~sP8(X42,happ(s(fun(bool,bool),X41),s(bool,'F_0')),X43) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X41)))) )),
  inference(superposition,[],[f372,f1815])).
fof(f5078,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X44)))) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP9(X46,X45,happ(s(fun(bool,bool),X44),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X44)))) )),
  inference(superposition,[],[f381,f1815])).
fof(f5079,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X47)))) | p(s(bool,X48)) | p(s(bool,X49)) | ~sP10(happ(s(fun(bool,bool),X47),s(bool,'F_0')),X48,X49) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X47)))) )),
  inference(superposition,[],[f389,f1815])).
fof(f5082,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP11(happ(s(fun(bool,bool),X56),s(bool,'F_0')),X58,X57) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f396,f1815])).
fof(f5083,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X59)))) | p(s(bool,X60)) | ~p(s(bool,X61)) | ~sP12(X61,X60,happ(s(fun(bool,bool),X59),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X59)))) )),
  inference(superposition,[],[f406,f1815])).
fof(f5086,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X68)))) | ~p(s(bool,X69)) | p(s(bool,X70)) | ~sP13(happ(s(fun(bool,bool),X68),s(bool,'F_0')),X70,X69) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X68)))) )),
  inference(superposition,[],[f413,f1815])).
fof(f5934,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X28))) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X28),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X28)))) )),
  inference(superposition,[],[f3255,f1815])).
fof(f12071,plain,(
  ( ! [X6,X7,X5] : (sP0(X7,happ(s(fun(bool,bool),X5),s(bool,X6))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f5168,f639])).
fof(f12194,plain,(
  ( ! [X6,X7,X5] : (sP0(happ(s(fun(bool,bool),X5),s(bool,X6)),X7) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f5169,f639])).
fof(f14205,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X3),s(bool,X4))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3)))) )),
  inference(superposition,[],[f5165,f639])).
fof(f14292,plain,(
  ( ! [X0,X1] : (sP1(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5170,f247])).
fof(f14294,plain,(
  ( ! [X6,X7,X5] : (sP1(happ(s(fun(bool,bool),X5),s(bool,X6)),X7) | p(s(bool,X7)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f5170,f639])).
fof(f14365,plain,(
  ( ! [X0,X1] : (sP1(X1,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5171,f247])).
fof(f14367,plain,(
  ( ! [X6,X7,X5] : (sP1(X7,happ(s(fun(bool,bool),X5),s(bool,X6))) | p(s(bool,X7)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f5171,f639])).
fof(f14440,plain,(
  ( ! [X2,X0,X1] : (sP2(X1,happ(s(fun(bool,bool),X0),s(bool,'F_0')),X2) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5172,f247])).
fof(f14442,plain,(
  ( ! [X10,X8,X7,X9] : (sP2(X9,happ(s(fun(bool,bool),X7),s(bool,X8)),X10) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f5172,f639])).
fof(f14515,plain,(
  ( ! [X2,X0,X1] : (sP2(X1,X2,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5173,f247])).
fof(f14517,plain,(
  ( ! [X10,X8,X7,X9] : (sP2(X9,X10,happ(s(fun(bool,bool),X7),s(bool,X8))) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f5173,f639])).
fof(f14590,plain,(
  ( ! [X2,X0,X1] : (sP3(X1,X2,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | ~p(s(bool,X2)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5174,f247])).
fof(f14592,plain,(
  ( ! [X10,X8,X7,X9] : (sP3(X9,X10,happ(s(fun(bool,bool),X7),s(bool,X8))) | ~p(s(bool,X10)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f5174,f639])).
fof(f14616,plain,(
  ( ! [X2,X0,X1] : (sP11(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1,X2) | ~p(s(bool,X2)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5181,f247])).
fof(f14618,plain,(
  ( ! [X10,X8,X7,X9] : (sP11(happ(s(fun(bool,bool),X7),s(bool,X8)),X9,X10) | ~p(s(bool,X10)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f5181,f639])).
fof(f14642,plain,(
  ( ! [X2,X0,X1] : (sP11(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1,X2) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5182,f247])).
fof(f14644,plain,(
  ( ! [X10,X8,X7,X9] : (sP11(happ(s(fun(bool,bool),X7),s(bool,X8)),X9,X10) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f5182,f639])).
fof(f14667,plain,(
  ( ! [X2,X0,X1] : (sP13(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1,X2) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5185,f247])).
fof(f14669,plain,(
  ( ! [X10,X8,X7,X9] : (sP13(happ(s(fun(bool,bool),X7),s(bool,X8)),X9,X10) | p(s(bool,X10)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f5185,f639])).
fof(f14694,plain,(
  ( ! [X2,X0,X1] : (sP13(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1,X2) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f5186,f247])).
fof(f14696,plain,(
  ( ! [X10,X8,X7,X9] : (sP13(happ(s(fun(bool,bool),X7),s(bool,X8)),X9,X10) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f5186,f639])).
fof(f16183,plain,(
  ( ! [X621,X623,X620,X622] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X620),s(bool,X621))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X620))) | s(bool,X621) = s(bool,X622) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X620))) | s(bool,X622) = s(bool,X623) | s(bool,'T_0') = s(bool,X623)) )),
  inference(superposition,[],[f5165,f2758])).
fof(f16184,plain,(
  ( ! [X625,X627,X628,X624,X626] : (sP0(X626,happ(s(fun(bool,bool),X624),s(bool,X625))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X624))) | s(bool,X625) = s(bool,X627) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X624))) | s(bool,X627) = s(bool,X628) | s(bool,'T_0') = s(bool,X628)) )),
  inference(superposition,[],[f5168,f2758])).
fof(f16185,plain,(
  ( ! [X633,X629,X631,X632,X630] : (sP0(happ(s(fun(bool,bool),X629),s(bool,X630)),X631) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X629))) | s(bool,X630) = s(bool,X632) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X629))) | s(bool,X632) = s(bool,X633) | s(bool,'T_0') = s(bool,X633)) )),
  inference(superposition,[],[f5169,f2758])).
fof(f36164,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (sP13(happ(s(fun(bool,bool),X61),s(bool,X62)),X63,X64) | p(s(bool,X64)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61)))) | s(bool,X62) = s(bool,X65) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))) )),
  inference(superposition,[],[f1844,f2758])).
fof(f36199,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (sP13(happ(s(fun(bool,bool),X61),s(bool,X62)),X63,X64) | ~p(s(bool,X63)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61)))) | s(bool,X62) = s(bool,X65) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))) )),
  inference(superposition,[],[f1845,f2758])).
fof(f36238,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X37),s(bool,X38))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37))) | s(bool,X38) = s(bool,X39) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X37)))) )),
  inference(superposition,[],[f1816,f2758])).
fof(f36241,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X45),s(bool,X46))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X45))) | s(bool,X46) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X45)))) )),
  inference(superposition,[],[f1816,f9775])).
fof(f39463,plain,(
  ( ! [X47,X48] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X47))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X47),s(bool,X48))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X47))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(forward_demodulation,[],[f36242,f481])).
fof(f36242,plain,(
  ( ! [X47,X48] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X47),s(bool,X48))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X47))) | s(bool,'const.min.@_1'(s(fun(bool,bool),X47))) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f1816,f9775])).
fof(f37225,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (sP13(happ(s(fun(bool,bool),X61),s(bool,X62)),X63,X64) | ~sP12(happ(s(fun(bool,bool),X61),s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))),X63,X64) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61)))) | s(bool,X62) = s(bool,X65) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))) )),
  inference(superposition,[],[f1847,f2758])).
fof(f37298,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (sP11(happ(s(fun(bool,bool),X61),s(bool,X62)),X63,X64) | ~sP10(happ(s(fun(bool,bool),X61),s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))),X63,X64) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61)))) | s(bool,X62) = s(bool,X65) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))) )),
  inference(superposition,[],[f1848,f2758])).
fof(f16353,plain,(
  ( ! [X516,X518,X517,X515] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X515),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X515))) | s(bool,X516) = s(bool,X517) | s(bool,X517) = s(bool,'const.min.@_1'(s(fun(bool,bool),X515))) | s(bool,X516) = s(bool,X518) | s(bool,'T_0') = s(bool,X518)) )),
  inference(superposition,[],[f1816,f2758])).
fof(f16167,plain,(
  ( ! [X536,X534,X537,X535] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X534),s(bool,X535))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X534))) | s(bool,X535) = s(bool,X536) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X534))) | s(bool,X536) = s(bool,X537) | s(bool,'T_0') = s(bool,X537)) )),
  inference(superposition,[],[f1816,f2758])).
fof(f16030,plain,(
  ( ! [X549,X551,X552,X550] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X549),s(bool,X550))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X549))) | s(bool,X551) = s(bool,'const.min.@_1'(s(fun(bool,bool),X549))) | s(bool,'T_0') = s(bool,X550) | s(bool,X551) = s(bool,X552) | s(bool,'T_0') = s(bool,X552)) )),
  inference(superposition,[],[f1816,f2758])).
fof(f15756,plain,(
  ( ! [X549,X551,X552,X550] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X549),s(bool,X550))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X549))) | s(bool,X550) = s(bool,X551) | s(bool,'T_0') = s(bool,X551) | s(bool,X552) = s(bool,'const.min.@_1'(s(fun(bool,bool),X549))) | s(bool,'T_0') = s(bool,X552)) )),
  inference(superposition,[],[f1816,f2758])).
fof(f5109,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X4),s(bool,X5))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1816,f639])).
fof(f5107,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1816,f639])).
fof(f5056,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X9))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X9),s(bool,'F_0')))) )),
  inference(superposition,[],[f1815,f247])).
fof(f4624,plain,(
  ( ! [X17,X18,X16] : (sP13(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | ~p(s(bool,X17)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1845,f247])).
fof(f4622,plain,(
  ( ! [X10,X8,X11,X9] : (sP13(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | ~p(s(bool,X10)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1845,f639])).
fof(f4620,plain,(
  ( ! [X2,X0,X3,X1] : (sP13(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1845,f639])).
fof(f4609,plain,(
  ( ! [X17,X18,X16] : (sP13(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | p(s(bool,X18)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1844,f247])).
fof(f4607,plain,(
  ( ! [X10,X8,X11,X9] : (sP13(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | p(s(bool,X11)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1844,f639])).
fof(f4605,plain,(
  ( ! [X2,X0,X3,X1] : (sP13(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | p(s(bool,X2)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1844,f639])).
fof(f1815,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f281,f247])).
fof(f5059,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X2)))) )),
  inference(superposition,[],[f254,f1815])).
fof(f5060,plain,(
  ( ! [X3] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3)))) )),
  inference(superposition,[],[f255,f1815])).
fof(f5061,plain,(
  ( ! [X4] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4)))) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4)))) )),
  inference(superposition,[],[f258,f1815])).
fof(f5062,plain,(
  ( ! [X5] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5)))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f259,f1815])).
fof(f5063,plain,(
  ( ! [X6,X7] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6)))) | p(s(bool,X7)) | s(bool,X7) = s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X6)))) )),
  inference(superposition,[],[f285,f1815])).
fof(f5065,plain,(
  ( ! [X10,X11] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X10)))) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),X10),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X10)))) )),
  inference(superposition,[],[f288,f1815])).
fof(f5495,plain,(
  ( ! [X2,X0,X3,X1] : (sP13(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | ~sP12(happ(s(fun(bool,bool),X0),s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))),X1,X2) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1847,f639])).
fof(f5497,plain,(
  ( ! [X10,X8,X11,X9] : (sP13(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | ~sP12(happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))),X10,X11) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1847,f639])).
fof(f5499,plain,(
  ( ! [X17,X18,X16] : (sP13(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | ~sP12(happ(s(fun(bool,bool),X16),s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))),X17,X18) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1847,f247])).
fof(f5523,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | ~sP10(happ(s(fun(bool,bool),X0),s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))),X1,X2) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1848,f639])).
fof(f5525,plain,(
  ( ! [X10,X8,X11,X9] : (sP11(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | ~sP10(happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))),X10,X11) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1848,f639])).
fof(f5527,plain,(
  ( ! [X17,X18,X16] : (sP11(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | ~sP10(happ(s(fun(bool,bool),X16),s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))),X17,X18) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1848,f247])).
fof(f9429,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X46)))) | s(bool,X47) != s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0'))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X46)))) )),
  inference(superposition,[],[f2880,f1815])).
fof(f9758,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X48)))) | s(bool,X49) = s(bool,happ(s(fun(bool,bool),X48),s(bool,'F_0'))) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X48)))) )),
  inference(superposition,[],[f3256,f1815])).
fof(f12017,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X46)))) | s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48))) = s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X46)))) )),
  inference(superposition,[],[f5167,f1815])).
fof(f15760,plain,(
  ( ! [X569,X571,X572,X568,X570] : (sP1(X570,happ(s(fun(bool,bool),X568),s(bool,X569))) | p(s(bool,X570)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X568)))) | s(bool,X569) = s(bool,X571) | s(bool,'T_0') = s(bool,X571) | s(bool,X572) = s(bool,'const.min.@_1'(s(fun(bool,bool),X568))) | s(bool,'T_0') = s(bool,X572)) )),
  inference(superposition,[],[f1829,f2758])).
fof(f15761,plain,(
  ( ! [X573,X575,X576,X578,X574,X577] : (sP2(X575,happ(s(fun(bool,bool),X573),s(bool,X574)),X576) | ~p(s(bool,X575)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X573)))) | s(bool,X574) = s(bool,X577) | s(bool,'T_0') = s(bool,X577) | s(bool,X578) = s(bool,'const.min.@_1'(s(fun(bool,bool),X573))) | s(bool,'T_0') = s(bool,X578)) )),
  inference(superposition,[],[f1831,f2758])).
fof(f15762,plain,(
  ( ! [X584,X580,X582,X581,X583,X579] : (sP2(X581,X582,happ(s(fun(bool,bool),X579),s(bool,X580))) | ~p(s(bool,X581)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X579)))) | s(bool,X580) = s(bool,X583) | s(bool,'T_0') = s(bool,X583) | s(bool,X584) = s(bool,'const.min.@_1'(s(fun(bool,bool),X579))) | s(bool,'T_0') = s(bool,X584)) )),
  inference(superposition,[],[f1832,f2758])).
fof(f15763,plain,(
  ( ! [X588,X590,X586,X589,X585,X587] : (sP3(X587,X588,happ(s(fun(bool,bool),X585),s(bool,X586))) | ~p(s(bool,X588)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X585)))) | s(bool,X586) = s(bool,X589) | s(bool,'T_0') = s(bool,X589) | s(bool,X590) = s(bool,'const.min.@_1'(s(fun(bool,bool),X585))) | s(bool,'T_0') = s(bool,X590)) )),
  inference(superposition,[],[f1833,f2758])).
fof(f15764,plain,(
  ( ! [X596,X592,X594,X593,X595,X591] : (sP11(happ(s(fun(bool,bool),X591),s(bool,X592)),X593,X594) | ~p(s(bool,X594)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X591)))) | s(bool,X592) = s(bool,X595) | s(bool,'T_0') = s(bool,X595) | s(bool,X596) = s(bool,'const.min.@_1'(s(fun(bool,bool),X591))) | s(bool,'T_0') = s(bool,X596)) )),
  inference(superposition,[],[f1840,f2758])).
fof(f15765,plain,(
  ( ! [X600,X602,X598,X601,X597,X599] : (sP11(happ(s(fun(bool,bool),X597),s(bool,X598)),X599,X600) | ~p(s(bool,X599)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X597)))) | s(bool,X598) = s(bool,X601) | s(bool,'T_0') = s(bool,X601) | s(bool,X602) = s(bool,'const.min.@_1'(s(fun(bool,bool),X597))) | s(bool,'T_0') = s(bool,X602)) )),
  inference(superposition,[],[f1841,f2758])).
fof(f15766,plain,(
  ( ! [X606,X605,X607,X603,X608,X604] : (sP13(happ(s(fun(bool,bool),X603),s(bool,X604)),X605,X606) | p(s(bool,X606)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X603)))) | s(bool,X604) = s(bool,X607) | s(bool,'T_0') = s(bool,X607) | s(bool,X608) = s(bool,'const.min.@_1'(s(fun(bool,bool),X603))) | s(bool,'T_0') = s(bool,X608)) )),
  inference(superposition,[],[f1844,f2758])).
fof(f15767,plain,(
  ( ! [X613,X609,X611,X612,X614,X610] : (sP13(happ(s(fun(bool,bool),X609),s(bool,X610)),X611,X612) | ~p(s(bool,X611)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X609)))) | s(bool,X610) = s(bool,X613) | s(bool,'T_0') = s(bool,X613) | s(bool,X614) = s(bool,'const.min.@_1'(s(fun(bool,bool),X609))) | s(bool,'T_0') = s(bool,X614)) )),
  inference(superposition,[],[f1845,f2758])).
fof(f15768,plain,(
  ( ! [X617,X619,X615,X620,X616,X618] : (sP13(happ(s(fun(bool,bool),X615),s(bool,X616)),X617,X618) | ~sP12(happ(s(fun(bool,bool),X615),s(bool,'const.min.@_1'(s(fun(bool,bool),X615)))),X617,X618) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X615)))) | s(bool,X616) = s(bool,X619) | s(bool,'T_0') = s(bool,X619) | s(bool,X620) = s(bool,'const.min.@_1'(s(fun(bool,bool),X615))) | s(bool,'T_0') = s(bool,X620)) )),
  inference(superposition,[],[f1847,f2758])).
fof(f15769,plain,(
  ( ! [X625,X621,X623,X624,X626,X622] : (sP11(happ(s(fun(bool,bool),X621),s(bool,X622)),X623,X624) | ~sP10(happ(s(fun(bool,bool),X621),s(bool,'const.min.@_1'(s(fun(bool,bool),X621)))),X623,X624) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X621)))) | s(bool,X622) = s(bool,X625) | s(bool,'T_0') = s(bool,X625) | s(bool,X626) = s(bool,'const.min.@_1'(s(fun(bool,bool),X621))) | s(bool,'T_0') = s(bool,X626)) )),
  inference(superposition,[],[f1848,f2758])).
fof(f16034,plain,(
  ( ! [X569,X571,X572,X568,X570] : (sP1(X570,happ(s(fun(bool,bool),X568),s(bool,X569))) | p(s(bool,X570)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X568)))) | s(bool,X571) = s(bool,'const.min.@_1'(s(fun(bool,bool),X568))) | s(bool,'T_0') = s(bool,X569) | s(bool,X571) = s(bool,X572) | s(bool,'T_0') = s(bool,X572)) )),
  inference(superposition,[],[f1829,f2758])).
fof(f16035,plain,(
  ( ! [X573,X575,X576,X578,X574,X577] : (sP2(X575,happ(s(fun(bool,bool),X573),s(bool,X574)),X576) | ~p(s(bool,X575)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X573)))) | s(bool,X577) = s(bool,'const.min.@_1'(s(fun(bool,bool),X573))) | s(bool,'T_0') = s(bool,X574) | s(bool,X577) = s(bool,X578) | s(bool,'T_0') = s(bool,X578)) )),
  inference(superposition,[],[f1831,f2758])).
fof(f16036,plain,(
  ( ! [X584,X580,X582,X581,X583,X579] : (sP2(X581,X582,happ(s(fun(bool,bool),X579),s(bool,X580))) | ~p(s(bool,X581)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X579)))) | s(bool,X583) = s(bool,'const.min.@_1'(s(fun(bool,bool),X579))) | s(bool,'T_0') = s(bool,X580) | s(bool,X583) = s(bool,X584) | s(bool,'T_0') = s(bool,X584)) )),
  inference(superposition,[],[f1832,f2758])).
fof(f16037,plain,(
  ( ! [X588,X590,X586,X589,X585,X587] : (sP3(X587,X588,happ(s(fun(bool,bool),X585),s(bool,X586))) | ~p(s(bool,X588)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X585)))) | s(bool,X589) = s(bool,'const.min.@_1'(s(fun(bool,bool),X585))) | s(bool,'T_0') = s(bool,X586) | s(bool,X589) = s(bool,X590) | s(bool,'T_0') = s(bool,X590)) )),
  inference(superposition,[],[f1833,f2758])).
fof(f16038,plain,(
  ( ! [X596,X592,X594,X593,X595,X591] : (sP11(happ(s(fun(bool,bool),X591),s(bool,X592)),X593,X594) | ~p(s(bool,X594)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X591)))) | s(bool,X595) = s(bool,'const.min.@_1'(s(fun(bool,bool),X591))) | s(bool,'T_0') = s(bool,X592) | s(bool,X595) = s(bool,X596) | s(bool,'T_0') = s(bool,X596)) )),
  inference(superposition,[],[f1840,f2758])).
fof(f16039,plain,(
  ( ! [X600,X602,X598,X601,X597,X599] : (sP11(happ(s(fun(bool,bool),X597),s(bool,X598)),X599,X600) | ~p(s(bool,X599)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X597)))) | s(bool,X601) = s(bool,'const.min.@_1'(s(fun(bool,bool),X597))) | s(bool,'T_0') = s(bool,X598) | s(bool,X601) = s(bool,X602) | s(bool,'T_0') = s(bool,X602)) )),
  inference(superposition,[],[f1841,f2758])).
fof(f16040,plain,(
  ( ! [X606,X605,X607,X603,X608,X604] : (sP13(happ(s(fun(bool,bool),X603),s(bool,X604)),X605,X606) | p(s(bool,X606)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X603)))) | s(bool,X607) = s(bool,'const.min.@_1'(s(fun(bool,bool),X603))) | s(bool,'T_0') = s(bool,X604) | s(bool,X607) = s(bool,X608) | s(bool,'T_0') = s(bool,X608)) )),
  inference(superposition,[],[f1844,f2758])).
fof(f16041,plain,(
  ( ! [X613,X609,X611,X612,X614,X610] : (sP13(happ(s(fun(bool,bool),X609),s(bool,X610)),X611,X612) | ~p(s(bool,X611)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X609)))) | s(bool,X613) = s(bool,'const.min.@_1'(s(fun(bool,bool),X609))) | s(bool,'T_0') = s(bool,X610) | s(bool,X613) = s(bool,X614) | s(bool,'T_0') = s(bool,X614)) )),
  inference(superposition,[],[f1845,f2758])).
fof(f16042,plain,(
  ( ! [X617,X619,X615,X620,X616,X618] : (sP13(happ(s(fun(bool,bool),X615),s(bool,X616)),X617,X618) | ~sP12(happ(s(fun(bool,bool),X615),s(bool,'const.min.@_1'(s(fun(bool,bool),X615)))),X617,X618) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X615)))) | s(bool,X619) = s(bool,'const.min.@_1'(s(fun(bool,bool),X615))) | s(bool,'T_0') = s(bool,X616) | s(bool,X619) = s(bool,X620) | s(bool,'T_0') = s(bool,X620)) )),
  inference(superposition,[],[f1847,f2758])).
fof(f16043,plain,(
  ( ! [X625,X621,X623,X624,X626,X622] : (sP11(happ(s(fun(bool,bool),X621),s(bool,X622)),X623,X624) | ~sP10(happ(s(fun(bool,bool),X621),s(bool,'const.min.@_1'(s(fun(bool,bool),X621)))),X623,X624) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X621)))) | s(bool,X625) = s(bool,'const.min.@_1'(s(fun(bool,bool),X621))) | s(bool,'T_0') = s(bool,X622) | s(bool,X625) = s(bool,X626) | s(bool,'T_0') = s(bool,X626)) )),
  inference(superposition,[],[f1848,f2758])).
fof(f16171,plain,(
  ( ! [X557,X553,X555,X556,X554] : (sP1(X555,happ(s(fun(bool,bool),X553),s(bool,X554))) | p(s(bool,X555)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X553)))) | s(bool,X554) = s(bool,X556) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X553))) | s(bool,X556) = s(bool,X557) | s(bool,'T_0') = s(bool,X557)) )),
  inference(superposition,[],[f1829,f2758])).
fof(f16172,plain,(
  ( ! [X561,X563,X559,X560,X562,X558] : (sP2(X560,happ(s(fun(bool,bool),X558),s(bool,X559)),X561) | ~p(s(bool,X560)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X558)))) | s(bool,X559) = s(bool,X562) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X558))) | s(bool,X562) = s(bool,X563) | s(bool,'T_0') = s(bool,X563)) )),
  inference(superposition,[],[f1831,f2758])).
fof(f16173,plain,(
  ( ! [X569,X565,X567,X568,X564,X566] : (sP2(X566,X567,happ(s(fun(bool,bool),X564),s(bool,X565))) | ~p(s(bool,X566)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X564)))) | s(bool,X565) = s(bool,X568) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X564))) | s(bool,X568) = s(bool,X569) | s(bool,'T_0') = s(bool,X569)) )),
  inference(superposition,[],[f1832,f2758])).
fof(f16174,plain,(
  ( ! [X573,X575,X571,X572,X574,X570] : (sP3(X572,X573,happ(s(fun(bool,bool),X570),s(bool,X571))) | ~p(s(bool,X573)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X570)))) | s(bool,X571) = s(bool,X574) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X570))) | s(bool,X574) = s(bool,X575) | s(bool,'T_0') = s(bool,X575)) )),
  inference(superposition,[],[f1833,f2758])).
fof(f16175,plain,(
  ( ! [X580,X576,X578,X581,X577,X579] : (sP11(happ(s(fun(bool,bool),X576),s(bool,X577)),X578,X579) | ~p(s(bool,X579)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X576)))) | s(bool,X577) = s(bool,X580) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X576))) | s(bool,X580) = s(bool,X581) | s(bool,'T_0') = s(bool,X581)) )),
  inference(superposition,[],[f1840,f2758])).
fof(f16176,plain,(
  ( ! [X584,X586,X582,X585,X587,X583] : (sP11(happ(s(fun(bool,bool),X582),s(bool,X583)),X584,X585) | ~p(s(bool,X584)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X582)))) | s(bool,X583) = s(bool,X586) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X582))) | s(bool,X586) = s(bool,X587) | s(bool,'T_0') = s(bool,X587)) )),
  inference(superposition,[],[f1841,f2758])).
fof(f16177,plain,(
  ( ! [X592,X588,X590,X593,X589,X591] : (sP13(happ(s(fun(bool,bool),X588),s(bool,X589)),X590,X591) | p(s(bool,X591)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X588)))) | s(bool,X589) = s(bool,X592) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X588))) | s(bool,X592) = s(bool,X593) | s(bool,'T_0') = s(bool,X593)) )),
  inference(superposition,[],[f1844,f2758])).
fof(f16178,plain,(
  ( ! [X596,X598,X594,X597,X599,X595] : (sP13(happ(s(fun(bool,bool),X594),s(bool,X595)),X596,X597) | ~p(s(bool,X596)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X594)))) | s(bool,X595) = s(bool,X598) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X594))) | s(bool,X598) = s(bool,X599) | s(bool,'T_0') = s(bool,X599)) )),
  inference(superposition,[],[f1845,f2758])).
fof(f16179,plain,(
  ( ! [X600,X602,X605,X601,X603,X604] : (sP13(happ(s(fun(bool,bool),X600),s(bool,X601)),X602,X603) | ~sP12(happ(s(fun(bool,bool),X600),s(bool,'const.min.@_1'(s(fun(bool,bool),X600)))),X602,X603) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X600)))) | s(bool,X601) = s(bool,X604) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X600))) | s(bool,X604) = s(bool,X605) | s(bool,'T_0') = s(bool,X605)) )),
  inference(superposition,[],[f1847,f2758])).
fof(f16180,plain,(
  ( ! [X606,X609,X611,X607,X608,X610] : (sP11(happ(s(fun(bool,bool),X606),s(bool,X607)),X608,X609) | ~sP10(happ(s(fun(bool,bool),X606),s(bool,'const.min.@_1'(s(fun(bool,bool),X606)))),X608,X609) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X606)))) | s(bool,X607) = s(bool,X610) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X606))) | s(bool,X610) = s(bool,X611) | s(bool,'T_0') = s(bool,X611)) )),
  inference(superposition,[],[f1848,f2758])).
fof(f16357,plain,(
  ( ! [X536,X538,X534,X537,X535] : (sP1(X535,happ(s(fun(bool,bool),X534),s(bool,'T_0'))) | p(s(bool,X535)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X534)))) | s(bool,X536) = s(bool,X537) | s(bool,X537) = s(bool,'const.min.@_1'(s(fun(bool,bool),X534))) | s(bool,X536) = s(bool,X538) | s(bool,'T_0') = s(bool,X538)) )),
  inference(superposition,[],[f1829,f2758])).
fof(f16358,plain,(
  ( ! [X540,X542,X541,X543,X539,X544] : (sP2(X540,happ(s(fun(bool,bool),X539),s(bool,'T_0')),X541) | ~p(s(bool,X540)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X539)))) | s(bool,X542) = s(bool,X543) | s(bool,X543) = s(bool,'const.min.@_1'(s(fun(bool,bool),X539))) | s(bool,X542) = s(bool,X544) | s(bool,'T_0') = s(bool,X544)) )),
  inference(superposition,[],[f1831,f2758])).
fof(f16359,plain,(
  ( ! [X546,X549,X545,X547,X548,X550] : (sP2(X546,X547,happ(s(fun(bool,bool),X545),s(bool,'T_0'))) | ~p(s(bool,X546)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X545)))) | s(bool,X548) = s(bool,X549) | s(bool,X549) = s(bool,'const.min.@_1'(s(fun(bool,bool),X545))) | s(bool,X548) = s(bool,X550) | s(bool,'T_0') = s(bool,X550)) )),
  inference(superposition,[],[f1832,f2758])).
fof(f16360,plain,(
  ( ! [X553,X555,X551,X556,X552,X554] : (sP3(X552,X553,happ(s(fun(bool,bool),X551),s(bool,'T_0'))) | ~p(s(bool,X553)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X551)))) | s(bool,X554) = s(bool,X555) | s(bool,X555) = s(bool,'const.min.@_1'(s(fun(bool,bool),X551))) | s(bool,X554) = s(bool,X556) | s(bool,'T_0') = s(bool,X556)) )),
  inference(superposition,[],[f1833,f2758])).
fof(f16361,plain,(
  ( ! [X561,X557,X559,X560,X562,X558] : (sP11(happ(s(fun(bool,bool),X557),s(bool,'T_0')),X558,X559) | ~p(s(bool,X559)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X557)))) | s(bool,X560) = s(bool,X561) | s(bool,X561) = s(bool,'const.min.@_1'(s(fun(bool,bool),X557))) | s(bool,X560) = s(bool,X562) | s(bool,'T_0') = s(bool,X562)) )),
  inference(superposition,[],[f1840,f2758])).
fof(f16362,plain,(
  ( ! [X565,X567,X563,X568,X564,X566] : (sP11(happ(s(fun(bool,bool),X563),s(bool,'T_0')),X564,X565) | ~p(s(bool,X564)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X563)))) | s(bool,X566) = s(bool,X567) | s(bool,X567) = s(bool,'const.min.@_1'(s(fun(bool,bool),X563))) | s(bool,X566) = s(bool,X568) | s(bool,'T_0') = s(bool,X568)) )),
  inference(superposition,[],[f1841,f2758])).
fof(f16363,plain,(
  ( ! [X573,X569,X571,X572,X574,X570] : (sP13(happ(s(fun(bool,bool),X569),s(bool,'T_0')),X570,X571) | p(s(bool,X571)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X569)))) | s(bool,X572) = s(bool,X573) | s(bool,X573) = s(bool,'const.min.@_1'(s(fun(bool,bool),X569))) | s(bool,X572) = s(bool,X574) | s(bool,'T_0') = s(bool,X574)) )),
  inference(superposition,[],[f1844,f2758])).
fof(f16364,plain,(
  ( ! [X575,X580,X576,X578,X577,X579] : (sP13(happ(s(fun(bool,bool),X575),s(bool,'T_0')),X576,X577) | ~p(s(bool,X576)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X575)))) | s(bool,X578) = s(bool,X579) | s(bool,X579) = s(bool,'const.min.@_1'(s(fun(bool,bool),X575))) | s(bool,X578) = s(bool,X580) | s(bool,'T_0') = s(bool,X580)) )),
  inference(superposition,[],[f1845,f2758])).
fof(f16365,plain,(
  ( ! [X584,X586,X582,X585,X581,X583] : (sP13(happ(s(fun(bool,bool),X581),s(bool,'T_0')),X582,X583) | ~sP12(happ(s(fun(bool,bool),X581),s(bool,'const.min.@_1'(s(fun(bool,bool),X581)))),X582,X583) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X581)))) | s(bool,X584) = s(bool,X585) | s(bool,X585) = s(bool,'const.min.@_1'(s(fun(bool,bool),X581))) | s(bool,X584) = s(bool,X586) | s(bool,'T_0') = s(bool,X586)) )),
  inference(superposition,[],[f1847,f2758])).
fof(f16366,plain,(
  ( ! [X592,X588,X590,X589,X591,X587] : (sP11(happ(s(fun(bool,bool),X587),s(bool,'T_0')),X588,X589) | ~sP10(happ(s(fun(bool,bool),X587),s(bool,'const.min.@_1'(s(fun(bool,bool),X587)))),X588,X589) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X587)))) | s(bool,X590) = s(bool,X591) | s(bool,X591) = s(bool,'const.min.@_1'(s(fun(bool,bool),X587))) | s(bool,X590) = s(bool,X592) | s(bool,'T_0') = s(bool,X592)) )),
  inference(superposition,[],[f1848,f2758])).
fof(f18638,plain,(
  ( ! [X76,X75] : (~sP7('F_0',happ(s(fun(bool,bool),X76),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X75)))),'|dollar|exists_1'(s(fun(bool,bool),X76))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X76))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X75),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X75)))) )),
  inference(superposition,[],[f4374,f1815])).
fof(f18854,plain,(
  ( ! [X76,X75] : (sP7('T_0',happ(s(fun(bool,bool),X76),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X75)))),'|dollar|exists_1'(s(fun(bool,bool),X76))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X76))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X75),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X75)))) )),
  inference(superposition,[],[f4422,f1815])).
fof(f19775,plain,(
  ( ! [X76,X75] : (~sP7('F_0','|dollar|exists_1'(s(fun(bool,bool),X76)),happ(s(fun(bool,bool),X76),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X75))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X76))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X75),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X75)))) )),
  inference(superposition,[],[f4451,f1815])).
fof(f20000,plain,(
  ( ! [X76,X75] : (sP7('T_0','|dollar|exists_1'(s(fun(bool,bool),X76)),happ(s(fun(bool,bool),X76),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X75))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X76))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X75),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X75)))) )),
  inference(superposition,[],[f4471,f1815])).
fof(f20257,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X27) | s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(trivial_inequality_removal,[],[f20252])).
fof(f20252,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X27) | s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f5059,f2758])).
fof(f20259,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(trivial_inequality_removal,[],[f20244])).
fof(f20244,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f5059,f639])).
fof(f4376,plain,(
  ( ! [X8] : (~sP7('F_0',happ(s(fun(bool,bool),X8),s(bool,'F_0')),'|dollar|exists_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f1811,f247])).
fof(f4424,plain,(
  ( ! [X8] : (sP7('T_0',happ(s(fun(bool,bool),X8),s(bool,'F_0')),'|dollar|exists_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f1812,f247])).
fof(f4453,plain,(
  ( ! [X8] : (~sP7('F_0','|dollar|exists_1'(s(fun(bool,bool),X8)),happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f1813,f247])).
fof(f4473,plain,(
  ( ! [X8] : (sP7('T_0','|dollar|exists_1'(s(fun(bool,bool),X8)),happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f1814,f247])).
fof(f32120,plain,(
  ( ! [X59,X60] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60))) | s(bool,'T_0') = s(bool,happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))))) )),
  inference(forward_demodulation,[],[f32065,f481])).
fof(f32065,plain,(
  ( ! [X59,X60] : (s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60))) != s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))))) )),
  inference(superposition,[],[f30109,f281])).
fof(f32104,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f32103])).
fof(f32103,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X22) != s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,X23) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(duplicate_literal_removal,[],[f32051])).
fof(f32051,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X22) != s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X21) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f30109,f2758])).
fof(f2692,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3)))))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK18,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X2))))) )),
  inference(subsumption_resolution,[],[f2691,f239])).
fof(f2691,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK18,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X2))))) )),
  inference(forward_demodulation,[],[f2690,f481])).
fof(f2690,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK18,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X2)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))))) )),
  inference(subsumption_resolution,[],[f2651,f239])).
fof(f2651,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK18,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X2)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))))) )),
  inference(superposition,[],[f283,f1404])).
fof(f2689,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X0,sK18,X1,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X1),s(sK14,sK19(X0,sK18,X1,sK16,sK15,sK14))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X1)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X0))))) )),
  inference(subsumption_resolution,[],[f2688,f239])).
fof(f2688,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X0,sK18,X1,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X1),s(sK14,sK19(X0,sK18,X1,sK16,sK15,sK14))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X1)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X0))))) )),
  inference(forward_demodulation,[],[f2687,f481])).
fof(f2687,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X0,sK18,X1,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X1),s(sK14,sK19(X0,sK18,X1,sK16,sK15,sK14))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X1)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))))) )),
  inference(subsumption_resolution,[],[f2650,f239])).
fof(f2650,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X0,sK18,X1,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X1),s(sK14,sK19(X0,sK18,X1,sK16,sK15,sK14))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X1)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))))) )),
  inference(superposition,[],[f284,f1404])).
fof(f2608,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3)))))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK17,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X2))))) )),
  inference(subsumption_resolution,[],[f2607,f239])).
fof(f2607,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK17,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X2))))) )),
  inference(forward_demodulation,[],[f2606,f481])).
fof(f2606,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK17,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X2)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))))) )),
  inference(subsumption_resolution,[],[f2567,f239])).
fof(f2567,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK14,sK19(X2,sK17,X3,sK16,sK15,sK14)),s(fun(sK14,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X2),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X3)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X2)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))))) )),
  inference(superposition,[],[f283,f1128])).
fof(f2605,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X0,sK17,X1,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X1),s(sK14,sK19(X0,sK17,X1,sK16,sK15,sK14))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X1)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X0))))) )),
  inference(subsumption_resolution,[],[f2604,f239])).
fof(f2604,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X0,sK17,X1,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X1),s(sK14,sK19(X0,sK17,X1,sK16,sK15,sK14))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X1)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X0))))) )),
  inference(forward_demodulation,[],[f2603,f481])).
fof(f2603,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X0,sK17,X1,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X1),s(sK14,sK19(X0,sK17,X1,sK16,sK15,sK14))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X1)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))))) )),
  inference(subsumption_resolution,[],[f2566,f239])).
fof(f2566,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),sK16),s(sK14,sK19(X0,sK17,X1,sK16,sK15,sK14)))),s('type.extreal.extreal',happ(s(fun(sK14,'type.extreal.extreal'),X1),s(sK14,sK19(X0,sK17,X1,sK16,sK15,sK14))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X0),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),X1)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))))) )),
  inference(superposition,[],[f284,f1128])).
fof(f283,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X4),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK19(X5,X4,X3,X2,X1,X0)),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X5),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X3)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X4),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X5),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK19(X5,X4,X3,X2,X1,X0)),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,sK19(X5,X4,X3,X2,X1,X0)))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X3),s(X0,sK19(X5,X4,X3,X2,X1,X0)))))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X4),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X5),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X3)))))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X6))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X3),s(X0,X6))))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X4),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X5),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X3)))))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X6))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X3),s(X0,X6)))))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X4),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X5),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X3)))))) & ! [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X6))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X3),s(X0,X6)))))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X4),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X5),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X3)))))) & ! [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,X6))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X3),s(X0,X6)))))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X7,X8,X19,X9,X10] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X9),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),X8)))))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X10),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),X19)))))) & ! [X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X6,'type.extreal.extreal'),X8),s(X6,X12))),s('type.extreal.extreal',happ(s(fun(X6,'type.extreal.extreal'),X19),s(X6,X12)))))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X10)))))),
  file('thm.lebesgue.psfis_unique','thm.lebesgue.psfis_mono')).
fof(f284,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X4),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X2),s(X0,sK19(X5,X4,X3,X2,X1,X0)))),s('type.extreal.extreal',happ(s(fun(X0,'type.extreal.extreal'),X3),s(X0,sK19(X5,X4,X3,X2,X1,X0))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X5),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X3)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) )),
  inference(cnf_transformation,[],[f176])).
fof(f3279,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X118,bool),'type.pair.prod'(fun(fun(X118,bool),bool),fun(fun(X118,bool),'type.realax.real'))),X119)))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X121),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X118,bool),'type.pair.prod'(fun(fun(X118,bool),bool),fun(fun(X118,bool),'type.realax.real'))),X119),s(fun(X118,'type.extreal.extreal'),X122)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X117),s('type.extreal.extreal',X121)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X118,'type.extreal.extreal'),X120),s(X118,sK19(X121,X117,X122,X120,X119,X118)))),s('type.extreal.extreal',happ(s(fun(X118,'type.extreal.extreal'),X122),s(X118,sK19(X121,X117,X122,X120,X119,X118))))))) | s(bool,X123) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X117),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X118,bool),'type.pair.prod'(fun(fun(X118,bool),bool),fun(fun(X118,bool),'type.realax.real'))),X119),s(fun(X118,'type.extreal.extreal'),X120))))) | s(bool,'T_0') = s(bool,X123)) )),
  inference(subsumption_resolution,[],[f2951,f239])).
fof(f2951,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X118,'type.extreal.extreal'),X120),s(X118,sK19(X121,X117,X122,X120,X119,X118)))),s('type.extreal.extreal',happ(s(fun(X118,'type.extreal.extreal'),X122),s(X118,sK19(X121,X117,X122,X120,X119,X118))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X121),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X118,bool),'type.pair.prod'(fun(fun(X118,bool),bool),fun(fun(X118,bool),'type.realax.real'))),X119),s(fun(X118,'type.extreal.extreal'),X122)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X117),s('type.extreal.extreal',X121)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X118,bool),'type.pair.prod'(fun(fun(X118,bool),bool),fun(fun(X118,bool),'type.realax.real'))),X119)))) | s(bool,X123) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X117),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X118,bool),'type.pair.prod'(fun(fun(X118,bool),bool),fun(fun(X118,bool),'type.realax.real'))),X119),s(fun(X118,'type.extreal.extreal'),X120))))) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f284,f639])).
fof(f3280,plain,(
  ( ! [X127,X125,X130,X128,X126,X124,X129] : (~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X125,bool),'type.pair.prod'(fun(fun(X125,bool),bool),fun(fun(X125,bool),'type.realax.real'))),X126)))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X128),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X125,bool),'type.pair.prod'(fun(fun(X125,bool),bool),fun(fun(X125,bool),'type.realax.real'))),X126),s(fun(X125,'type.extreal.extreal'),X129)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X124),s('type.extreal.extreal',X128)))) | p(s(bool,'const.bool.IN_2'(s(X125,sK19(X128,X124,X129,X127,X126,X125)),s(fun(X125,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X125,bool),'type.pair.prod'(fun(fun(X125,bool),bool),fun(fun(X125,bool),'type.realax.real'))),X126)))))) | s(bool,X130) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X124),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X125,bool),'type.pair.prod'(fun(fun(X125,bool),bool),fun(fun(X125,bool),'type.realax.real'))),X126),s(fun(X125,'type.extreal.extreal'),X127))))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(subsumption_resolution,[],[f2952,f239])).
fof(f2952,plain,(
  ( ! [X127,X125,X130,X128,X126,X124,X129] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X125,sK19(X128,X124,X129,X127,X126,X125)),s(fun(X125,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X125,bool),'type.pair.prod'(fun(fun(X125,bool),bool),fun(fun(X125,bool),'type.realax.real'))),X126)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X128),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X125,bool),'type.pair.prod'(fun(fun(X125,bool),bool),fun(fun(X125,bool),'type.realax.real'))),X126),s(fun(X125,'type.extreal.extreal'),X129)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X124),s('type.extreal.extreal',X128)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X125,bool),'type.pair.prod'(fun(fun(X125,bool),bool),fun(fun(X125,bool),'type.realax.real'))),X126)))) | s(bool,X130) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X124),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X125,bool),'type.pair.prod'(fun(fun(X125,bool),bool),fun(fun(X125,bool),'type.realax.real'))),X126),s(fun(X125,'type.extreal.extreal'),X127))))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f283,f639])).
fof(f18350,plain,(
  ( ! [X358,X364,X366,X360,X362,X359,X365,X361,X363] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X359,'type.extreal.extreal'),X361),s(X359,sK19(X362,X358,X363,X361,X360,X359)))),s('type.extreal.extreal',happ(s(fun(X359,'type.extreal.extreal'),X363),s(X359,sK19(X362,X358,X363,X361,X360,X359))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X362),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X359,bool),'type.pair.prod'(fun(fun(X359,bool),bool),fun(fun(X359,bool),'type.realax.real'))),X360),s(fun(X359,'type.extreal.extreal'),X363)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X358),s('type.extreal.extreal',X362)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X359,bool),'type.pair.prod'(fun(fun(X359,bool),bool),fun(fun(X359,bool),'type.realax.real'))),X360)))) | s(bool,X364) = s(bool,X365) | s(bool,X365) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X358),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X359,bool),'type.pair.prod'(fun(fun(X359,bool),bool),fun(fun(X359,bool),'type.realax.real'))),X360),s(fun(X359,'type.extreal.extreal'),X361))))) | s(bool,X364) = s(bool,X366) | s(bool,'T_0') = s(bool,X366)) )),
  inference(subsumption_resolution,[],[f16328,f239])).
fof(f16328,plain,(
  ( ! [X358,X364,X366,X360,X362,X359,X365,X361,X363] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',happ(s(fun(X359,'type.extreal.extreal'),X361),s(X359,sK19(X362,X358,X363,X361,X360,X359)))),s('type.extreal.extreal',happ(s(fun(X359,'type.extreal.extreal'),X363),s(X359,sK19(X362,X358,X363,X361,X360,X359))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X362),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X359,bool),'type.pair.prod'(fun(fun(X359,bool),bool),fun(fun(X359,bool),'type.realax.real'))),X360),s(fun(X359,'type.extreal.extreal'),X363)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X358),s('type.extreal.extreal',X362)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X359,bool),'type.pair.prod'(fun(fun(X359,bool),bool),fun(fun(X359,bool),'type.realax.real'))),X360)))) | s(bool,X364) = s(bool,X365) | s(bool,X365) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X358),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X359,bool),'type.pair.prod'(fun(fun(X359,bool),bool),fun(fun(X359,bool),'type.realax.real'))),X360),s(fun(X359,'type.extreal.extreal'),X361))))) | s(bool,X364) = s(bool,X366) | s(bool,'T_0') = s(bool,X366)) )),
  inference(superposition,[],[f284,f2758])).
fof(f18351,plain,(
  ( ! [X372,X374,X368,X370,X367,X373,X375,X369,X371] : (p(s(bool,'const.bool.IN_2'(s(X368,sK19(X371,X367,X372,X370,X369,X368)),s(fun(X368,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X368,bool),'type.pair.prod'(fun(fun(X368,bool),bool),fun(fun(X368,bool),'type.realax.real'))),X369)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X371),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X368,bool),'type.pair.prod'(fun(fun(X368,bool),bool),fun(fun(X368,bool),'type.realax.real'))),X369),s(fun(X368,'type.extreal.extreal'),X372)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X367),s('type.extreal.extreal',X371)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X368,bool),'type.pair.prod'(fun(fun(X368,bool),bool),fun(fun(X368,bool),'type.realax.real'))),X369)))) | s(bool,X373) = s(bool,X374) | s(bool,X374) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X367),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X368,bool),'type.pair.prod'(fun(fun(X368,bool),bool),fun(fun(X368,bool),'type.realax.real'))),X369),s(fun(X368,'type.extreal.extreal'),X370))))) | s(bool,X373) = s(bool,X375) | s(bool,'T_0') = s(bool,X375)) )),
  inference(subsumption_resolution,[],[f16329,f239])).
fof(f16329,plain,(
  ( ! [X372,X374,X368,X370,X367,X373,X375,X369,X371] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X368,sK19(X371,X367,X372,X370,X369,X368)),s(fun(X368,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X368,bool),'type.pair.prod'(fun(fun(X368,bool),bool),fun(fun(X368,bool),'type.realax.real'))),X369)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X371),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X368,bool),'type.pair.prod'(fun(fun(X368,bool),bool),fun(fun(X368,bool),'type.realax.real'))),X369),s(fun(X368,'type.extreal.extreal'),X372)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X367),s('type.extreal.extreal',X371)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X368,bool),'type.pair.prod'(fun(fun(X368,bool),bool),fun(fun(X368,bool),'type.realax.real'))),X369)))) | s(bool,X373) = s(bool,X374) | s(bool,X374) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X367),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X368,bool),'type.pair.prod'(fun(fun(X368,bool),bool),fun(fun(X368,bool),'type.realax.real'))),X369),s(fun(X368,'type.extreal.extreal'),X370))))) | s(bool,X373) = s(bool,X375) | s(bool,'T_0') = s(bool,X375)) )),
  inference(superposition,[],[f283,f2758])).
fof(f293,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1)) & (s('type.extreal.extreal',X0) != s('type.extreal.extreal',X1) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : (((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1)) & (s('type.extreal.extreal',X0) != s('type.extreal.extreal',X1) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) <=> s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X12,X18] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X18)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X12))))) <=> s('type.extreal.extreal',X12) = s('type.extreal.extreal',X18))),
  file('thm.lebesgue.psfis_unique','thm.extreal.le_antisym')).
fof(f295,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal',X0) != s('type.extreal.extreal',X1) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f178])).
fof(f3282,plain,(
  ( ! [X140,X141,X139] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X140),s('type.extreal.extreal',X139)))) | s('type.extreal.extreal',X139) = s('type.extreal.extreal',X140) | s(bool,X141) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X139),s('type.extreal.extreal',X140))) | s(bool,'T_0') = s(bool,X141)) )),
  inference(subsumption_resolution,[],[f2955,f239])).
fof(f2955,plain,(
  ( ! [X140,X141,X139] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X140),s('type.extreal.extreal',X139)))) | s('type.extreal.extreal',X139) = s('type.extreal.extreal',X140) | s(bool,X141) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X139),s('type.extreal.extreal',X140))) | s(bool,'T_0') = s(bool,X141)) )),
  inference(superposition,[],[f293,f639])).
fof(f18362,plain,(
  ( ! [X472,X469,X471,X473,X470] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X470),s('type.extreal.extreal',X469)))) | s('type.extreal.extreal',X469) = s('type.extreal.extreal',X470) | s(bool,X471) = s(bool,X472) | s(bool,X472) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X469),s('type.extreal.extreal',X470))) | s(bool,X471) = s(bool,X473) | s(bool,'T_0') = s(bool,X473)) )),
  inference(subsumption_resolution,[],[f16344,f239])).
fof(f16344,plain,(
  ( ! [X472,X469,X471,X473,X470] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X470),s('type.extreal.extreal',X469)))) | s('type.extreal.extreal',X469) = s('type.extreal.extreal',X470) | s(bool,X471) = s(bool,X472) | s(bool,X472) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X469),s('type.extreal.extreal',X470))) | s(bool,X471) = s(bool,X473) | s(bool,'T_0') = s(bool,X473)) )),
  inference(superposition,[],[f293,f2758])).
fof(f30109,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f30087])).
fof(f30087,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(equality_factoring,[],[f9775])).
fof(f30081,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X0) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(equality_factoring,[],[f9775])).
fof(f30080,plain,(
  ( ! [X416,X418,X415,X417] : (s(X417,f(s(bool,'T_0'))) != s(X417,g(s(bool,'T_0'))) | s(fun(bool,X417),X415) = s(fun(bool,X417),X416) | s(bool,X418) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X418) = s(bool,sK20(X415,X416,X417,bool))) )),
  inference(superposition,[],[f321,f9775])).
fof(f30069,plain,(
  ( ! [X389,X391,X388,X390] : (s(bool,'T_0') != s(bool,f(s(X388,sK20(X389,X390,bool,X388)))) | s(fun(X388,bool),X389) = s(fun(X388,bool),X390) | s(bool,X391) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X391) = s(bool,g(s(X388,sK20(X389,X390,bool,X388))))) )),
  inference(superposition,[],[f321,f9775])).
fof(f30699,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X6) = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f29924,f239])).
fof(f29924,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X6) = s(bool,X7)) )),
  inference(superposition,[],[f258,f9775])).
fof(f29786,plain,(
  ( ! [X436,X438,X437,X435] : (s(X437,f(s(bool,X438))) != s(X437,g(s(bool,X438))) | s(fun(bool,X437),X435) = s(fun(bool,X437),X436) | s(bool,X438) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,sK20(X435,X436,X437,bool))) )),
  inference(superposition,[],[f321,f9775])).
fof(f29775,plain,(
  ( ! [X408,X410,X409,X411] : (s(bool,X411) != s(bool,f(s(X408,sK20(X409,X410,bool,X408)))) | s(fun(X408,bool),X409) = s(fun(X408,bool),X410) | s(bool,X411) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,g(s(X408,sK20(X409,X410,bool,X408))))) )),
  inference(superposition,[],[f321,f9775])).
fof(f30585,plain,(
  ( ! [X362,X361,X363] : (s(bool,X361) != s(bool,X362) | s(bool,X362) = s(bool,X363) | s(bool,'T_0') = s(bool,X363) | s(bool,X361) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(subsumption_resolution,[],[f29759,f477])).
fof(f29759,plain,(
  ( ! [X362,X361,X363] : (s(bool,X361) != s(bool,X362) | s(bool,X362) = s(bool,X363) | s(bool,'T_0') = s(bool,X363) | s(bool,X361) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f3255,f9775])).
fof(f30584,plain,(
  ( ! [X360] : (s(bool,'T_0') != s(bool,X360) | s(bool,X360) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(subsumption_resolution,[],[f29758,f477])).
fof(f29758,plain,(
  ( ! [X360] : (s(bool,'T_0') != s(bool,X360) | s(bool,X360) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f477,f9775])).
fof(f29313,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68))) | s(bool,X69) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X69) = s(bool,happ(s(fun(X67,bool),X68),s(X67,'const.min.@_1'(s(fun(X67,bool),X68)))))) )),
  inference(superposition,[],[f9775,f281])).
fof(f30267,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X49) = s(bool,X51) | s(bool,X49) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(duplicate_literal_removal,[],[f29307])).
fof(f29307,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X49) = s(bool,X51) | s(bool,X49) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f9775,f9775])).
fof(f29277,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X5) = s(bool,X6) | s(bool,X7) = s(bool,X8) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,X4) = s(bool,X7)) )),
  inference(superposition,[],[f9775,f2758])).
fof(f29276,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f9775,f9775])).
fof(f29273,plain,(
  ( ! [X92,X90,X91] : (s(bool,X92) = s(bool,'|dollar|exists_1'(s(fun(X90,bool),X91))) | s(bool,X92) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,happ(s(fun(X90,bool),X91),s(X90,'const.min.@_1'(s(fun(X90,bool),X91)))))) )),
  inference(superposition,[],[f9775,f281])).
fof(f30299,plain,(
  ( ! [X61,X62,X60] : (s(bool,X61) = s(bool,X62) | s(bool,X62) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X60) | s(bool,X61) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(duplicate_literal_removal,[],[f29257])).
fof(f29257,plain,(
  ( ! [X61,X62,X60] : (s(bool,X61) = s(bool,X62) | s(bool,X62) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X60) | s(bool,X61) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f9775,f9775])).
fof(f30317,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f29242])).
fof(f29242,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f9775,f247])).
fof(f30340,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') = s(bool,X58) | s(bool,X57) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X59) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X57) = s(bool,X59)) )),
  inference(duplicate_literal_removal,[],[f29222])).
fof(f29222,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') = s(bool,X58) | s(bool,X57) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X58) | s(bool,X59) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X57) = s(bool,X59)) )),
  inference(superposition,[],[f9775,f9775])).
fof(f30344,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X2) = s(bool,X3) | s(bool,X1) = s(bool,X3) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f30343])).
fof(f30343,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X36) = s(bool,X37) | s(bool,X34) = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(duplicate_literal_removal,[],[f29217])).
fof(f29217,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,X37) | s(bool,X34) = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f9775,f2758])).
fof(f258,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.lebesgue.psfis_unique','thm.bool.EQ_CLAUSES|split|2')).
fof(f259,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f162])).
fof(f426,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f242,f239])).
fof(f242,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.lebesgue.psfis_unique','thm.bool.NOT_CLAUSES|split|2')).
fof(f1816,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3)))) )),
  inference(superposition,[],[f281,f247])).
fof(f1875,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK20(X0,X1,X2,bool))) )),
  inference(superposition,[],[f321,f247])).
fof(f1876,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK20(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK20(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f321,f247])).
fof(f5165,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,happ(s(fun(X10,bool),X11),s(X10,'const.min.@_1'(s(fun(X10,bool),X11))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11)))) )),
  inference(subsumption_resolution,[],[f5121,f239])).
fof(f5121,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(X10,bool),X11),s(X10,'const.min.@_1'(s(fun(X10,bool),X11))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11)))) )),
  inference(superposition,[],[f258,f1816])).
fof(f5167,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X22)) | s(bool,X22) = s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) )),
  inference(forward_demodulation,[],[f5166,f281])).
fof(f5166,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X22)) | s(bool,X22) = s(bool,happ(s(fun(X20,bool),X21),s(X20,'const.min.@_1'(s(fun(X20,bool),X21))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) )),
  inference(subsumption_resolution,[],[f5125,f239])).
fof(f5125,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X22)) | s(bool,X22) = s(bool,happ(s(fun(X20,bool),X21),s(X20,'const.min.@_1'(s(fun(X20,bool),X21))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) )),
  inference(superposition,[],[f288,f1816])).
fof(f5168,plain,(
  ( ! [X24,X23,X25] : (sP0(X25,happ(s(fun(X23,bool),X24),s(X23,'const.min.@_1'(s(fun(X23,bool),X24))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X23,bool),X24)))) )),
  inference(subsumption_resolution,[],[f5126,f239])).
fof(f5126,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | sP0(X25,happ(s(fun(X23,bool),X24),s(X23,'const.min.@_1'(s(fun(X23,bool),X24))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X23,bool),X24)))) )),
  inference(superposition,[],[f297,f1816])).
fof(f5169,plain,(
  ( ! [X28,X26,X27] : (sP0(happ(s(fun(X26,bool),X27),s(X26,'const.min.@_1'(s(fun(X26,bool),X27)))),X28) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27)))) )),
  inference(subsumption_resolution,[],[f5127,f239])).
fof(f5127,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X26,bool),X27),s(X26,'const.min.@_1'(s(fun(X26,bool),X27)))),X28) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27)))) )),
  inference(superposition,[],[f298,f1816])).
fof(f5170,plain,(
  ( ! [X30,X31,X29] : (sP1(happ(s(fun(X29,bool),X30),s(X29,'const.min.@_1'(s(fun(X29,bool),X30)))),X31) | p(s(bool,X31)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X29,bool),X30)))) )),
  inference(subsumption_resolution,[],[f5128,f239])).
fof(f5128,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | sP1(happ(s(fun(X29,bool),X30),s(X29,'const.min.@_1'(s(fun(X29,bool),X30)))),X31) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X29,bool),X30)))) )),
  inference(superposition,[],[f313,f1816])).
fof(f5171,plain,(
  ( ! [X33,X34,X32] : (sP1(X34,happ(s(fun(X32,bool),X33),s(X32,'const.min.@_1'(s(fun(X32,bool),X33))))) | p(s(bool,X34)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X32,bool),X33)))) )),
  inference(subsumption_resolution,[],[f5129,f239])).
fof(f5129,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X34)) | sP1(X34,happ(s(fun(X32,bool),X33),s(X32,'const.min.@_1'(s(fun(X32,bool),X33))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X32,bool),X33)))) )),
  inference(superposition,[],[f316,f1816])).
fof(f5172,plain,(
  ( ! [X39,X41,X42,X40] : (sP2(X41,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))),X42) | ~p(s(bool,X41)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40)))) )),
  inference(subsumption_resolution,[],[f5131,f239])).
fof(f5131,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | sP2(X41,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))),X42) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40)))) )),
  inference(superposition,[],[f326,f1816])).
fof(f5173,plain,(
  ( ! [X45,X43,X46,X44] : (sP2(X45,X46,happ(s(fun(X43,bool),X44),s(X43,'const.min.@_1'(s(fun(X43,bool),X44))))) | ~p(s(bool,X45)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44)))) )),
  inference(subsumption_resolution,[],[f5132,f239])).
fof(f5132,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | sP2(X45,X46,happ(s(fun(X43,bool),X44),s(X43,'const.min.@_1'(s(fun(X43,bool),X44))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44)))) )),
  inference(superposition,[],[f327,f1816])).
fof(f5174,plain,(
  ( ! [X47,X50,X48,X49] : (sP3(X50,X49,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48))))) | ~p(s(bool,X49)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48)))) )),
  inference(subsumption_resolution,[],[f5133,f239])).
fof(f5133,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | sP3(X50,X49,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48)))) )),
  inference(superposition,[],[f334,f1816])).
fof(f5175,plain,(
  ( ! [X54,X52,X53,X51] : (sP7(X53,X54,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52))))) | ~sP6(X53,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))),X54) | ~sP5(X54,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))),X53) | ~sP4(happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))),X54,X53) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52)))) )),
  inference(subsumption_resolution,[],[f5134,f239])).
fof(f5134,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | sP7(X53,X54,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52))))) | ~sP6(X53,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))),X54) | ~sP5(X54,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))),X53) | ~sP4(happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))),X54,X53) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52)))) )),
  inference(superposition,[],[f346,f1816])).
fof(f5176,plain,(
  ( ! [X57,X58,X56,X55] : (~sP6(X57,happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56)))),X58) | p(s(bool,X58)) | ~p(s(bool,X57)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) )),
  inference(subsumption_resolution,[],[f5135,f239])).
fof(f5135,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP6(X57,happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56)))),X58) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) )),
  inference(superposition,[],[f349,f1816])).
fof(f5177,plain,(
  ( ! [X61,X59,X62,X60] : (~sP5(X61,happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))),X62) | p(s(bool,X62)) | ~p(s(bool,X61)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60)))) )),
  inference(subsumption_resolution,[],[f5136,f239])).
fof(f5136,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP5(X61,happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))),X62) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60)))) )),
  inference(superposition,[],[f353,f1816])).
fof(f5178,plain,(
  ( ! [X66,X64,X65,X63] : (~sP8(X65,happ(s(fun(X63,bool),X64),s(X63,'const.min.@_1'(s(fun(X63,bool),X64)))),X66) | p(s(bool,X66)) | ~p(s(bool,X65)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X63,bool),X64)))) )),
  inference(subsumption_resolution,[],[f5137,f239])).
fof(f5137,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X65,happ(s(fun(X63,bool),X64),s(X63,'const.min.@_1'(s(fun(X63,bool),X64)))),X66) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X63,bool),X64)))) )),
  inference(superposition,[],[f372,f1816])).
fof(f5179,plain,(
  ( ! [X70,X68,X69,X67] : (~sP9(X70,X69,happ(s(fun(X67,bool),X68),s(X67,'const.min.@_1'(s(fun(X67,bool),X68))))) | p(s(bool,X70)) | ~p(s(bool,X69)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68)))) )),
  inference(subsumption_resolution,[],[f5138,f239])).
fof(f5138,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X69)) | p(s(bool,X70)) | ~sP9(X70,X69,happ(s(fun(X67,bool),X68),s(X67,'const.min.@_1'(s(fun(X67,bool),X68))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68)))) )),
  inference(superposition,[],[f381,f1816])).
fof(f5180,plain,(
  ( ! [X74,X72,X71,X73] : (~sP10(happ(s(fun(X71,bool),X72),s(X71,'const.min.@_1'(s(fun(X71,bool),X72)))),X73,X74) | p(s(bool,X74)) | p(s(bool,X73)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X71,bool),X72)))) )),
  inference(subsumption_resolution,[],[f5139,f239])).
fof(f5139,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X73)) | p(s(bool,X74)) | ~sP10(happ(s(fun(X71,bool),X72),s(X71,'const.min.@_1'(s(fun(X71,bool),X72)))),X73,X74) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X71,bool),X72)))) )),
  inference(superposition,[],[f389,f1816])).
fof(f5181,plain,(
  ( ! [X78,X76,X77,X75] : (sP11(happ(s(fun(X75,bool),X76),s(X75,'const.min.@_1'(s(fun(X75,bool),X76)))),X78,X77) | ~p(s(bool,X77)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X75,bool),X76)))) )),
  inference(subsumption_resolution,[],[f5140,f239])).
fof(f5140,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | sP11(happ(s(fun(X75,bool),X76),s(X75,'const.min.@_1'(s(fun(X75,bool),X76)))),X78,X77) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X75,bool),X76)))) )),
  inference(superposition,[],[f394,f1816])).
fof(f5182,plain,(
  ( ! [X80,X81,X79,X82] : (sP11(happ(s(fun(X79,bool),X80),s(X79,'const.min.@_1'(s(fun(X79,bool),X80)))),X81,X82) | ~p(s(bool,X81)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X79,bool),X80)))) )),
  inference(subsumption_resolution,[],[f5141,f239])).
fof(f5141,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | sP11(happ(s(fun(X79,bool),X80),s(X79,'const.min.@_1'(s(fun(X79,bool),X80)))),X81,X82) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X79,bool),X80)))) )),
  inference(superposition,[],[f395,f1816])).
fof(f5183,plain,(
  ( ! [X85,X83,X86,X84] : (~sP11(happ(s(fun(X83,bool),X84),s(X83,'const.min.@_1'(s(fun(X83,bool),X84)))),X86,X85) | p(s(bool,X86)) | p(s(bool,X85)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X83,bool),X84)))) )),
  inference(subsumption_resolution,[],[f5142,f239])).
fof(f5142,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X85)) | p(s(bool,X86)) | ~sP11(happ(s(fun(X83,bool),X84),s(X83,'const.min.@_1'(s(fun(X83,bool),X84)))),X86,X85) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X83,bool),X84)))) )),
  inference(superposition,[],[f396,f1816])).
fof(f5184,plain,(
  ( ! [X90,X88,X87,X89] : (~sP12(X90,X89,happ(s(fun(X87,bool),X88),s(X87,'const.min.@_1'(s(fun(X87,bool),X88))))) | ~p(s(bool,X90)) | p(s(bool,X89)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X87,bool),X88)))) )),
  inference(subsumption_resolution,[],[f5143,f239])).
fof(f5143,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | p(s(bool,X89)) | ~p(s(bool,X90)) | ~sP12(X90,X89,happ(s(fun(X87,bool),X88),s(X87,'const.min.@_1'(s(fun(X87,bool),X88))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X87,bool),X88)))) )),
  inference(superposition,[],[f406,f1816])).
fof(f5185,plain,(
  ( ! [X94,X92,X93,X91] : (sP13(happ(s(fun(X91,bool),X92),s(X91,'const.min.@_1'(s(fun(X91,bool),X92)))),X94,X93) | p(s(bool,X93)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X91,bool),X92)))) )),
  inference(subsumption_resolution,[],[f5144,f239])).
fof(f5144,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | p(s(bool,X93)) | sP13(happ(s(fun(X91,bool),X92),s(X91,'const.min.@_1'(s(fun(X91,bool),X92)))),X94,X93) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X91,bool),X92)))) )),
  inference(superposition,[],[f411,f1816])).
fof(f5186,plain,(
  ( ! [X97,X95,X98,X96] : (sP13(happ(s(fun(X95,bool),X96),s(X95,'const.min.@_1'(s(fun(X95,bool),X96)))),X97,X98) | ~p(s(bool,X97)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X95,bool),X96)))) )),
  inference(subsumption_resolution,[],[f5145,f239])).
fof(f5145,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X97)) | sP13(happ(s(fun(X95,bool),X96),s(X95,'const.min.@_1'(s(fun(X95,bool),X96)))),X97,X98) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X95,bool),X96)))) )),
  inference(superposition,[],[f412,f1816])).
fof(f5187,plain,(
  ( ! [X101,X99,X102,X100] : (~sP13(happ(s(fun(X99,bool),X100),s(X99,'const.min.@_1'(s(fun(X99,bool),X100)))),X102,X101) | p(s(bool,X102)) | ~p(s(bool,X101)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X99,bool),X100)))) )),
  inference(subsumption_resolution,[],[f5146,f239])).
fof(f5146,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,'T_0')) | ~p(s(bool,X101)) | p(s(bool,X102)) | ~sP13(happ(s(fun(X99,bool),X100),s(X99,'const.min.@_1'(s(fun(X99,bool),X100)))),X102,X101) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X99,bool),X100)))) )),
  inference(superposition,[],[f413,f1816])).
fof(f5188,plain,(
  ( ! [X103,X105,X106,X104] : (sP13(happ(s(fun(X103,bool),X104),s(X103,'const.min.@_1'(s(fun(X103,bool),X104)))),X105,X106) | ~sP12(happ(s(fun(X103,bool),X104),s(X103,'const.min.@_1'(s(fun(X103,bool),X104)))),X105,X106) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X103,bool),X104)))) )),
  inference(subsumption_resolution,[],[f5147,f239])).
fof(f5147,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,'T_0')) | ~sP12(happ(s(fun(X103,bool),X104),s(X103,'const.min.@_1'(s(fun(X103,bool),X104)))),X105,X106) | sP13(happ(s(fun(X103,bool),X104),s(X103,'const.min.@_1'(s(fun(X103,bool),X104)))),X105,X106) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X103,bool),X104)))) )),
  inference(superposition,[],[f416,f1816])).
fof(f5189,plain,(
  ( ! [X109,X107,X110,X108] : (sP11(happ(s(fun(X107,bool),X108),s(X107,'const.min.@_1'(s(fun(X107,bool),X108)))),X109,X110) | ~sP10(happ(s(fun(X107,bool),X108),s(X107,'const.min.@_1'(s(fun(X107,bool),X108)))),X109,X110) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X107,bool),X108)))) )),
  inference(subsumption_resolution,[],[f5148,f239])).
fof(f5148,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,'T_0')) | ~sP10(happ(s(fun(X107,bool),X108),s(X107,'const.min.@_1'(s(fun(X107,bool),X108)))),X109,X110) | sP11(happ(s(fun(X107,bool),X108),s(X107,'const.min.@_1'(s(fun(X107,bool),X108)))),X109,X110) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X107,bool),X108)))) )),
  inference(superposition,[],[f417,f1816])).
fof(f5933,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X25,bool),X26))) | s(bool,X27) = s(bool,happ(s(fun(X25,bool),X26),s(X25,'const.min.@_1'(s(fun(X25,bool),X26))))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f3255,f281])).
fof(f9778,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun(X40,bool),X41),s(X40,'const.min.@_1'(s(fun(X40,bool),X41))))) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X40,bool),X41)))) )),
  inference(subsumption_resolution,[],[f9756,f239])).
fof(f9756,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'T_0')) | s(bool,X42) = s(bool,happ(s(fun(X40,bool),X41),s(X40,'const.min.@_1'(s(fun(X40,bool),X41))))) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X40,bool),X41)))) )),
  inference(superposition,[],[f3256,f1816])).
fof(f11691,plain,(
  ( ! [X47,X50,X48,X51,X49] : (p(s(bool,X49)) | ~sP9(X49,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X50) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48)))) )),
  inference(subsumption_resolution,[],[f11672,f239])).
fof(f11672,plain,(
  ( ! [X47,X50,X48,X51,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | ~sP9(X49,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X50) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48)))) )),
  inference(superposition,[],[f3268,f1816])).
fof(f11973,plain,(
  ( ! [X54,X52,X50,X53,X51] : (p(s(bool,X52)) | ~sP13(X53,X52,happ(s(fun(X50,bool),X51),s(X50,'const.min.@_1'(s(fun(X50,bool),X51))))) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X50,bool),X51)))) )),
  inference(subsumption_resolution,[],[f11953,f239])).
fof(f11953,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | ~sP13(X53,X52,happ(s(fun(X50,bool),X51),s(X50,'const.min.@_1'(s(fun(X50,bool),X51))))) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X50,bool),X51)))) )),
  inference(superposition,[],[f3276,f1816])).
fof(f12036,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'|dollar|exists_1'(s(fun(X40,bool),X41))) = s(bool,happ(s(fun(X38,bool),X39),s(X38,'const.min.@_1'(s(fun(X38,bool),X39))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X40,bool),X41))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X38,bool),X39)))) )),
  inference(subsumption_resolution,[],[f12015,f239])).
fof(f12015,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X40,bool),X41))) = s(bool,happ(s(fun(X38,bool),X39),s(X38,'const.min.@_1'(s(fun(X38,bool),X39))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X40,bool),X41))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X38,bool),X39)))) )),
  inference(superposition,[],[f5167,f1816])).
fof(f18346,plain,(
  ( ! [X241,X239,X237,X242,X240,X238] : (s(bool,X237) = s(bool,'|dollar|exists_1'(s(fun(X238,bool),X239))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X238,bool),X239))) | s(bool,X240) = s(bool,X241) | s(bool,X237) = s(bool,X241) | s(bool,X240) = s(bool,X242) | s(bool,'T_0') = s(bool,X242)) )),
  inference(subsumption_resolution,[],[f16298,f239])).
fof(f16298,plain,(
  ( ! [X241,X239,X237,X242,X240,X238] : (~p(s(bool,'T_0')) | s(bool,X237) = s(bool,'|dollar|exists_1'(s(fun(X238,bool),X239))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X238,bool),X239))) | s(bool,X240) = s(bool,X241) | s(bool,X237) = s(bool,X241) | s(bool,X240) = s(bool,X242) | s(bool,'T_0') = s(bool,X242)) )),
  inference(superposition,[],[f5167,f2758])).
fof(f2827,plain,(
  ( ! [X156,X154,X157,X155] : (s(bool,X157) != s(bool,f(s(X154,sK20(X155,X156,bool,X154)))) | s(fun(X154,bool),X155) = s(fun(X154,bool),X156) | s(bool,'T_0') = s(bool,X157) | s(bool,'T_0') = s(bool,g(s(X154,sK20(X155,X156,bool,X154))))) )),
  inference(superposition,[],[f321,f639])).
fof(f2957,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'T_0') != s(bool,f(s(X148,sK20(X149,X150,bool,X148)))) | s(fun(X148,bool),X149) = s(fun(X148,bool),X150) | s(bool,X151) = s(bool,g(s(X148,sK20(X149,X150,bool,X148)))) | s(bool,'T_0') = s(bool,X151)) )),
  inference(superposition,[],[f321,f639])).
fof(f5441,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK20(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK20(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f5440])).
fof(f5440,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK20(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK20(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f5433,f247])).
fof(f5433,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK20(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK20(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1876,f639])).
fof(f15749,plain,(
  ( ! [X520,X522,X518,X521,X523,X519] : (s(bool,X521) != s(bool,f(s(X518,sK20(X519,X520,bool,X518)))) | s(fun(X518,bool),X519) = s(fun(X518,bool),X520) | s(bool,X521) = s(bool,X522) | s(bool,'T_0') = s(bool,X522) | s(bool,X523) = s(bool,g(s(X518,sK20(X519,X520,bool,X518)))) | s(bool,'T_0') = s(bool,X523)) )),
  inference(superposition,[],[f321,f2758])).
fof(f16023,plain,(
  ( ! [X520,X522,X518,X521,X523,X519] : (s(bool,X521) != s(bool,f(s(X518,sK20(X519,X520,bool,X518)))) | s(fun(X518,bool),X519) = s(fun(X518,bool),X520) | s(bool,X522) = s(bool,g(s(X518,sK20(X519,X520,bool,X518)))) | s(bool,'T_0') = s(bool,X521) | s(bool,X522) = s(bool,X523) | s(bool,'T_0') = s(bool,X523)) )),
  inference(superposition,[],[f321,f2758])).
fof(f16160,plain,(
  ( ! [X505,X507,X508,X504,X506,X503] : (s(bool,X506) != s(bool,f(s(X503,sK20(X504,X505,bool,X503)))) | s(fun(X503,bool),X504) = s(fun(X503,bool),X505) | s(bool,X506) = s(bool,X507) | s(bool,'T_0') = s(bool,g(s(X503,sK20(X504,X505,bool,X503)))) | s(bool,X507) = s(bool,X508) | s(bool,'T_0') = s(bool,X508)) )),
  inference(superposition,[],[f321,f2758])).
fof(f16346,plain,(
  ( ! [X484,X486,X488,X485,X487,X489] : (s(bool,'T_0') != s(bool,f(s(X484,sK20(X485,X486,bool,X484)))) | s(fun(X484,bool),X485) = s(fun(X484,bool),X486) | s(bool,X487) = s(bool,X488) | s(bool,X488) = s(bool,g(s(X484,sK20(X485,X486,bool,X484)))) | s(bool,X487) = s(bool,X489) | s(bool,'T_0') = s(bool,X489)) )),
  inference(superposition,[],[f321,f2758])).
fof(f18269,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK20(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,f(s(X0,sK20(X1,X2,bool,X0)))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f18268])).
fof(f18268,plain,(
  ( ! [X364,X360,X362,X365,X361,X363] : (s(fun(X360,bool),X361) = s(fun(X360,bool),X362) | s(bool,'T_0') = s(bool,g(s(X360,sK20(X361,X362,bool,X360)))) | s(bool,X364) = s(bool,f(s(X360,sK20(X361,X362,bool,X360)))) | s(bool,'T_0') = s(bool,X363) | s(bool,X364) = s(bool,X365) | s(bool,'T_0') = s(bool,X365)) )),
  inference(subsumption_resolution,[],[f15996,f247])).
fof(f15996,plain,(
  ( ! [X364,X360,X362,X365,X361,X363] : (s(bool,'F_0') != s(bool,X363) | s(fun(X360,bool),X361) = s(fun(X360,bool),X362) | s(bool,'T_0') = s(bool,g(s(X360,sK20(X361,X362,bool,X360)))) | s(bool,X364) = s(bool,f(s(X360,sK20(X361,X362,bool,X360)))) | s(bool,'T_0') = s(bool,X363) | s(bool,X364) = s(bool,X365) | s(bool,'T_0') = s(bool,X365)) )),
  inference(superposition,[],[f1876,f2758])).
fof(f281,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X6,bool),X5))) = s(bool,happ(s(fun(X6,bool),X5),s(X6,'const.min.@_1'(s(fun(X6,bool),X5)))))),
  file('thm.lebesgue.psfis_unique','thm.bool.EXISTS_DEF')).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f282,f281])).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X11,X12] : (p(s(bool,happ(s(fun(X6,bool),X11),s(X6,X12)))) => p(s(bool,happ(s(fun(X6,bool),X11),s(X6,'const.min.@_1'(s(fun(X6,bool),X11)))))))),
  file('thm.lebesgue.psfis_unique','thm.bool.SELECT_AX')).
fof(f1826,plain,(
  ( ! [X23,X21,X22] : (sP0(X23,happ(s(fun(X21,bool),X22),s(X21,'const.min.@_1'(s(fun(X21,bool),X22))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X22))))) )),
  inference(superposition,[],[f297,f281])).
fof(f1827,plain,(
  ( ! [X26,X24,X25] : (sP0(happ(s(fun(X24,bool),X25),s(X24,'const.min.@_1'(s(fun(X24,bool),X25)))),X26) | ~p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25))))) )),
  inference(superposition,[],[f298,f281])).
fof(f1828,plain,(
  ( ! [X28,X29,X27] : (sP1(happ(s(fun(X27,bool),X28),s(X27,'const.min.@_1'(s(fun(X27,bool),X28)))),X29) | p(s(bool,X29)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X27,bool),X28))))) )),
  inference(superposition,[],[f313,f281])).
fof(f1829,plain,(
  ( ! [X30,X31,X32] : (sP1(X32,happ(s(fun(X30,bool),X31),s(X30,'const.min.@_1'(s(fun(X30,bool),X31))))) | p(s(bool,X32)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X30,bool),X31))))) )),
  inference(superposition,[],[f316,f281])).
fof(f1831,plain,(
  ( ! [X39,X37,X38,X40] : (sP2(X39,happ(s(fun(X37,bool),X38),s(X37,'const.min.@_1'(s(fun(X37,bool),X38)))),X40) | ~p(s(bool,X39)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X37,bool),X38))))) )),
  inference(superposition,[],[f326,f281])).
fof(f1832,plain,(
  ( ! [X43,X41,X44,X42] : (sP2(X43,X44,happ(s(fun(X41,bool),X42),s(X41,'const.min.@_1'(s(fun(X41,bool),X42))))) | ~p(s(bool,X43)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X41,bool),X42))))) )),
  inference(superposition,[],[f327,f281])).
fof(f1833,plain,(
  ( ! [X47,X45,X48,X46] : (sP3(X48,X47,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46))))) | ~p(s(bool,X47)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X45,bool),X46))))) )),
  inference(superposition,[],[f334,f281])).
fof(f1834,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'|dollar|exists_1'(s(fun(X49,bool),X50)))) | sP7(X51,X52,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50))))) | ~sP6(X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))),X52) | ~sP5(X52,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))),X51) | ~sP4(happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))),X52,X51)) )),
  inference(superposition,[],[f346,f281])).
fof(f1835,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'|dollar|exists_1'(s(fun(X53,bool),X54)))) | ~p(s(bool,X55)) | p(s(bool,X56)) | ~sP6(X55,happ(s(fun(X53,bool),X54),s(X53,'const.min.@_1'(s(fun(X53,bool),X54)))),X56)) )),
  inference(superposition,[],[f349,f281])).
fof(f1836,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'|dollar|exists_1'(s(fun(X57,bool),X58)))) | ~p(s(bool,X59)) | p(s(bool,X60)) | ~sP5(X59,happ(s(fun(X57,bool),X58),s(X57,'const.min.@_1'(s(fun(X57,bool),X58)))),X60)) )),
  inference(superposition,[],[f353,f281])).
fof(f1837,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'|dollar|exists_1'(s(fun(X61,bool),X62)))) | ~p(s(bool,X63)) | p(s(bool,X64)) | ~sP8(X63,happ(s(fun(X61,bool),X62),s(X61,'const.min.@_1'(s(fun(X61,bool),X62)))),X64)) )),
  inference(superposition,[],[f372,f281])).
fof(f1838,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'|dollar|exists_1'(s(fun(X65,bool),X66)))) | ~p(s(bool,X67)) | p(s(bool,X68)) | ~sP9(X68,X67,happ(s(fun(X65,bool),X66),s(X65,'const.min.@_1'(s(fun(X65,bool),X66)))))) )),
  inference(superposition,[],[f381,f281])).
fof(f1839,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'|dollar|exists_1'(s(fun(X69,bool),X70)))) | p(s(bool,X71)) | p(s(bool,X72)) | ~sP10(happ(s(fun(X69,bool),X70),s(X69,'const.min.@_1'(s(fun(X69,bool),X70)))),X71,X72)) )),
  inference(superposition,[],[f389,f281])).
fof(f1840,plain,(
  ( ! [X76,X74,X75,X73] : (sP11(happ(s(fun(X73,bool),X74),s(X73,'const.min.@_1'(s(fun(X73,bool),X74)))),X76,X75) | ~p(s(bool,X75)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X73,bool),X74))))) )),
  inference(superposition,[],[f394,f281])).
fof(f1841,plain,(
  ( ! [X80,X78,X79,X77] : (sP11(happ(s(fun(X77,bool),X78),s(X77,'const.min.@_1'(s(fun(X77,bool),X78)))),X79,X80) | ~p(s(bool,X79)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X77,bool),X78))))) )),
  inference(superposition,[],[f395,f281])).
fof(f1842,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),X82)))) | p(s(bool,X83)) | p(s(bool,X84)) | ~sP11(happ(s(fun(X81,bool),X82),s(X81,'const.min.@_1'(s(fun(X81,bool),X82)))),X84,X83)) )),
  inference(superposition,[],[f396,f281])).
fof(f1843,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),X86)))) | p(s(bool,X87)) | ~p(s(bool,X88)) | ~sP12(X88,X87,happ(s(fun(X85,bool),X86),s(X85,'const.min.@_1'(s(fun(X85,bool),X86)))))) )),
  inference(superposition,[],[f406,f281])).
fof(f1844,plain,(
  ( ! [X92,X90,X91,X89] : (sP13(happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90)))),X92,X91) | p(s(bool,X91)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X89,bool),X90))))) )),
  inference(superposition,[],[f411,f281])).
fof(f1845,plain,(
  ( ! [X94,X95,X93,X96] : (sP13(happ(s(fun(X93,bool),X94),s(X93,'const.min.@_1'(s(fun(X93,bool),X94)))),X95,X96) | ~p(s(bool,X95)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X93,bool),X94))))) )),
  inference(superposition,[],[f412,f281])).
fof(f1846,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98)))) | ~p(s(bool,X99)) | p(s(bool,X100)) | ~sP13(happ(s(fun(X97,bool),X98),s(X97,'const.min.@_1'(s(fun(X97,bool),X98)))),X100,X99)) )),
  inference(superposition,[],[f413,f281])).
fof(f1847,plain,(
  ( ! [X103,X101,X102,X104] : (sP13(happ(s(fun(X101,bool),X102),s(X101,'const.min.@_1'(s(fun(X101,bool),X102)))),X103,X104) | ~sP12(happ(s(fun(X101,bool),X102),s(X101,'const.min.@_1'(s(fun(X101,bool),X102)))),X103,X104) | ~p(s(bool,'|dollar|exists_1'(s(fun(X101,bool),X102))))) )),
  inference(superposition,[],[f416,f281])).
fof(f1848,plain,(
  ( ! [X107,X105,X108,X106] : (sP11(happ(s(fun(X105,bool),X106),s(X105,'const.min.@_1'(s(fun(X105,bool),X106)))),X107,X108) | ~sP10(happ(s(fun(X105,bool),X106),s(X105,'const.min.@_1'(s(fun(X105,bool),X106)))),X107,X108) | ~p(s(bool,'|dollar|exists_1'(s(fun(X105,bool),X106))))) )),
  inference(superposition,[],[f417,f281])).
fof(f2743,plain,(
  ( ! [X24,X23,X25] : (s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun(X23,bool),X24))) | s(bool,'T_0') = s(bool,happ(s(fun(X23,bool),X24),s(X23,'const.min.@_1'(s(fun(X23,bool),X24))))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f639,f281])).
fof(f3281,plain,(
  ( ! [X134,X132,X133,X131] : (p(s(bool,'|dollar|exists_1'(s(fun(X131,bool),X132)))) | s(bool,X134) = s(bool,happ(s(fun(X131,bool),X132),s(X131,X133))) | s(bool,'T_0') = s(bool,X134)) )),
  inference(subsumption_resolution,[],[f2953,f239])).
fof(f2953,plain,(
  ( ! [X134,X132,X133,X131] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X131,bool),X132)))) | s(bool,X134) = s(bool,happ(s(fun(X131,bool),X132),s(X131,X133))) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f427,f639])).
fof(f5205,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP6(X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3) | p(s(bool,X3)) | ~p(s(bool,X2)) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5200,f239])).
fof(f5200,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP6(X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1835,f639])).
fof(f5223,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5(X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3) | p(s(bool,X3)) | ~p(s(bool,X2)) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5218,f239])).
fof(f5218,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP5(X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1836,f639])).
fof(f5245,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8(X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3) | p(s(bool,X3)) | ~p(s(bool,X2)) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5240,f239])).
fof(f5240,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP8(X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1837,f639])).
fof(f5261,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP9(X3,X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) | p(s(bool,X3)) | ~p(s(bool,X2)) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5256,f239])).
fof(f5256,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP9(X3,X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1838,f639])).
fof(f5277,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP10(happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X2,X3) | p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5272,f239])).
fof(f5272,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,X3)) | ~sP10(happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X2,X3) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1839,f639])).
fof(f5305,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP11(happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3,X2) | p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5300,f239])).
fof(f5300,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,X3)) | ~sP11(happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3,X2) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1842,f639])).
fof(f5329,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP12(X3,X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5324,f239])).
fof(f5324,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | ~p(s(bool,X3)) | ~sP12(X3,X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1843,f639])).
fof(f5365,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP13(happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3,X2) | p(s(bool,X3)) | ~p(s(bool,X2)) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5360,f239])).
fof(f5360,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP13(happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3,X2) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1846,f639])).
fof(f5609,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP7(X2,X3,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) | ~sP6(X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3) | ~sP5(X3,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X2) | ~sP4(happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3,X2) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5604,f239])).
fof(f5604,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | sP7(X2,X3,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) | ~sP6(X2,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3) | ~sP5(X3,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X2) | ~sP4(happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),X3,X2) | s(bool,X4) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1834,f639])).
fof(f5992,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,'|dollar|exists_1'(s(fun(X42,bool),X43)))) | sP1(X44,happ(s(fun(X42,bool),X43),s(X42,'const.min.@_1'(s(fun(X42,bool),X43))))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f3259,f281])).
fof(f8869,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,'|dollar|exists_1'(s(fun(X42,bool),X43)))) | sP1(happ(s(fun(X42,bool),X43),s(X42,'const.min.@_1'(s(fun(X42,bool),X43)))),X44) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f3260,f281])).
fof(f11673,plain,(
  ( ! [X54,X52,X56,X55,X53] : (~p(s(bool,'|dollar|exists_1'(s(fun(X52,bool),X53)))) | p(s(bool,X54)) | ~sP9(X54,happ(s(fun(X52,bool),X53),s(X52,'const.min.@_1'(s(fun(X52,bool),X53)))),X55) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f3268,f281])).
fof(f11954,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) | p(s(bool,X57)) | ~sP13(X58,X57,happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56))))) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f3276,f281])).
fof(f12759,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f12754,f426])).
fof(f12754,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | s(bool,X2) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f3281,f247])).
fof(f15503,plain,(
  ( ! [X132,X130,X128,X131,X129] : (s(bool,X130) = s(bool,'|dollar|exists_1'(s(fun(X128,bool),X129))) | s(bool,X131) = s(bool,happ(s(fun(X128,bool),X129),s(X128,'const.min.@_1'(s(fun(X128,bool),X129))))) | s(bool,'T_0') = s(bool,X130) | s(bool,X131) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f2758,f281])).
fof(f15531,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,X118) = s(bool,'|dollar|exists_1'(s(fun(X116,bool),X117))) | s(bool,X118) = s(bool,X119) | s(bool,'T_0') = s(bool,happ(s(fun(X116,bool),X117),s(X116,'const.min.@_1'(s(fun(X116,bool),X117))))) | s(bool,X119) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f2758,f281])).
fof(f15562,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X73,bool),X74))) | s(bool,X75) = s(bool,X76) | s(bool,X76) = s(bool,happ(s(fun(X73,bool),X74),s(X73,'const.min.@_1'(s(fun(X73,bool),X74))))) | s(bool,X75) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f2758,f281])).
fof(f18352,plain,(
  ( ! [X377,X379,X380,X382,X376,X378,X381] : (sP7(X378,X379,happ(s(fun(X376,bool),X377),s(X376,'const.min.@_1'(s(fun(X376,bool),X377))))) | ~sP6(X378,happ(s(fun(X376,bool),X377),s(X376,'const.min.@_1'(s(fun(X376,bool),X377)))),X379) | ~sP5(X379,happ(s(fun(X376,bool),X377),s(X376,'const.min.@_1'(s(fun(X376,bool),X377)))),X378) | ~sP4(happ(s(fun(X376,bool),X377),s(X376,'const.min.@_1'(s(fun(X376,bool),X377)))),X379,X378) | s(bool,X380) = s(bool,X381) | s(bool,X381) = s(bool,'|dollar|exists_1'(s(fun(X376,bool),X377))) | s(bool,X380) = s(bool,X382) | s(bool,'T_0') = s(bool,X382)) )),
  inference(subsumption_resolution,[],[f16330,f239])).
fof(f16330,plain,(
  ( ! [X377,X379,X380,X382,X376,X378,X381] : (~p(s(bool,'T_0')) | sP7(X378,X379,happ(s(fun(X376,bool),X377),s(X376,'const.min.@_1'(s(fun(X376,bool),X377))))) | ~sP6(X378,happ(s(fun(X376,bool),X377),s(X376,'const.min.@_1'(s(fun(X376,bool),X377)))),X379) | ~sP5(X379,happ(s(fun(X376,bool),X377),s(X376,'const.min.@_1'(s(fun(X376,bool),X377)))),X378) | ~sP4(happ(s(fun(X376,bool),X377),s(X376,'const.min.@_1'(s(fun(X376,bool),X377)))),X379,X378) | s(bool,X380) = s(bool,X381) | s(bool,X381) = s(bool,'|dollar|exists_1'(s(fun(X376,bool),X377))) | s(bool,X380) = s(bool,X382) | s(bool,'T_0') = s(bool,X382)) )),
  inference(superposition,[],[f1834,f2758])).
fof(f18353,plain,(
  ( ! [X383,X389,X385,X387,X388,X384,X386] : (~p(s(bool,X385)) | p(s(bool,X386)) | ~sP6(X385,happ(s(fun(X383,bool),X384),s(X383,'const.min.@_1'(s(fun(X383,bool),X384)))),X386) | s(bool,X387) = s(bool,X388) | s(bool,X388) = s(bool,'|dollar|exists_1'(s(fun(X383,bool),X384))) | s(bool,X387) = s(bool,X389) | s(bool,'T_0') = s(bool,X389)) )),
  inference(subsumption_resolution,[],[f16331,f239])).
fof(f16331,plain,(
  ( ! [X383,X389,X385,X387,X388,X384,X386] : (~p(s(bool,'T_0')) | ~p(s(bool,X385)) | p(s(bool,X386)) | ~sP6(X385,happ(s(fun(X383,bool),X384),s(X383,'const.min.@_1'(s(fun(X383,bool),X384)))),X386) | s(bool,X387) = s(bool,X388) | s(bool,X388) = s(bool,'|dollar|exists_1'(s(fun(X383,bool),X384))) | s(bool,X387) = s(bool,X389) | s(bool,'T_0') = s(bool,X389)) )),
  inference(superposition,[],[f1835,f2758])).
fof(f18354,plain,(
  ( ! [X391,X393,X395,X390,X396,X392,X394] : (~p(s(bool,X392)) | p(s(bool,X393)) | ~sP5(X392,happ(s(fun(X390,bool),X391),s(X390,'const.min.@_1'(s(fun(X390,bool),X391)))),X393) | s(bool,X394) = s(bool,X395) | s(bool,X395) = s(bool,'|dollar|exists_1'(s(fun(X390,bool),X391))) | s(bool,X394) = s(bool,X396) | s(bool,'T_0') = s(bool,X396)) )),
  inference(subsumption_resolution,[],[f16332,f239])).
fof(f16332,plain,(
  ( ! [X391,X393,X395,X390,X396,X392,X394] : (~p(s(bool,'T_0')) | ~p(s(bool,X392)) | p(s(bool,X393)) | ~sP5(X392,happ(s(fun(X390,bool),X391),s(X390,'const.min.@_1'(s(fun(X390,bool),X391)))),X393) | s(bool,X394) = s(bool,X395) | s(bool,X395) = s(bool,'|dollar|exists_1'(s(fun(X390,bool),X391))) | s(bool,X394) = s(bool,X396) | s(bool,'T_0') = s(bool,X396)) )),
  inference(superposition,[],[f1836,f2758])).
fof(f18355,plain,(
  ( ! [X397,X399,X401,X403,X398,X400,X402] : (~p(s(bool,X399)) | p(s(bool,X400)) | ~sP8(X399,happ(s(fun(X397,bool),X398),s(X397,'const.min.@_1'(s(fun(X397,bool),X398)))),X400) | s(bool,X401) = s(bool,X402) | s(bool,X402) = s(bool,'|dollar|exists_1'(s(fun(X397,bool),X398))) | s(bool,X401) = s(bool,X403) | s(bool,'T_0') = s(bool,X403)) )),
  inference(subsumption_resolution,[],[f16333,f239])).
fof(f16333,plain,(
  ( ! [X397,X399,X401,X403,X398,X400,X402] : (~p(s(bool,'T_0')) | ~p(s(bool,X399)) | p(s(bool,X400)) | ~sP8(X399,happ(s(fun(X397,bool),X398),s(X397,'const.min.@_1'(s(fun(X397,bool),X398)))),X400) | s(bool,X401) = s(bool,X402) | s(bool,X402) = s(bool,'|dollar|exists_1'(s(fun(X397,bool),X398))) | s(bool,X401) = s(bool,X403) | s(bool,'T_0') = s(bool,X403)) )),
  inference(superposition,[],[f1837,f2758])).
fof(f18356,plain,(
  ( ! [X408,X410,X405,X407,X409,X404,X406] : (~p(s(bool,X406)) | p(s(bool,X407)) | ~sP9(X407,X406,happ(s(fun(X404,bool),X405),s(X404,'const.min.@_1'(s(fun(X404,bool),X405))))) | s(bool,X408) = s(bool,X409) | s(bool,X409) = s(bool,'|dollar|exists_1'(s(fun(X404,bool),X405))) | s(bool,X408) = s(bool,X410) | s(bool,'T_0') = s(bool,X410)) )),
  inference(subsumption_resolution,[],[f16334,f239])).
fof(f16334,plain,(
  ( ! [X408,X410,X405,X407,X409,X404,X406] : (~p(s(bool,'T_0')) | ~p(s(bool,X406)) | p(s(bool,X407)) | ~sP9(X407,X406,happ(s(fun(X404,bool),X405),s(X404,'const.min.@_1'(s(fun(X404,bool),X405))))) | s(bool,X408) = s(bool,X409) | s(bool,X409) = s(bool,'|dollar|exists_1'(s(fun(X404,bool),X405))) | s(bool,X408) = s(bool,X410) | s(bool,'T_0') = s(bool,X410)) )),
  inference(superposition,[],[f1838,f2758])).
fof(f18357,plain,(
  ( ! [X412,X414,X416,X413,X415,X411,X417] : (p(s(bool,X413)) | p(s(bool,X414)) | ~sP10(happ(s(fun(X411,bool),X412),s(X411,'const.min.@_1'(s(fun(X411,bool),X412)))),X413,X414) | s(bool,X415) = s(bool,X416) | s(bool,X416) = s(bool,'|dollar|exists_1'(s(fun(X411,bool),X412))) | s(bool,X415) = s(bool,X417) | s(bool,'T_0') = s(bool,X417)) )),
  inference(subsumption_resolution,[],[f16335,f239])).
fof(f16335,plain,(
  ( ! [X412,X414,X416,X413,X415,X411,X417] : (~p(s(bool,'T_0')) | p(s(bool,X413)) | p(s(bool,X414)) | ~sP10(happ(s(fun(X411,bool),X412),s(X411,'const.min.@_1'(s(fun(X411,bool),X412)))),X413,X414) | s(bool,X415) = s(bool,X416) | s(bool,X416) = s(bool,'|dollar|exists_1'(s(fun(X411,bool),X412))) | s(bool,X415) = s(bool,X417) | s(bool,'T_0') = s(bool,X417)) )),
  inference(superposition,[],[f1839,f2758])).
fof(f18358,plain,(
  ( ! [X420,X422,X418,X424,X421,X423,X419] : (p(s(bool,X420)) | p(s(bool,X421)) | ~sP11(happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419)))),X421,X420) | s(bool,X422) = s(bool,X423) | s(bool,X423) = s(bool,'|dollar|exists_1'(s(fun(X418,bool),X419))) | s(bool,X422) = s(bool,X424) | s(bool,'T_0') = s(bool,X424)) )),
  inference(subsumption_resolution,[],[f16336,f239])).
fof(f16336,plain,(
  ( ! [X420,X422,X418,X424,X421,X423,X419] : (~p(s(bool,'T_0')) | p(s(bool,X420)) | p(s(bool,X421)) | ~sP11(happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419)))),X421,X420) | s(bool,X422) = s(bool,X423) | s(bool,X423) = s(bool,'|dollar|exists_1'(s(fun(X418,bool),X419))) | s(bool,X422) = s(bool,X424) | s(bool,'T_0') = s(bool,X424)) )),
  inference(superposition,[],[f1842,f2758])).
fof(f18359,plain,(
  ( ! [X428,X430,X426,X429,X431,X425,X427] : (p(s(bool,X427)) | ~p(s(bool,X428)) | ~sP12(X428,X427,happ(s(fun(X425,bool),X426),s(X425,'const.min.@_1'(s(fun(X425,bool),X426))))) | s(bool,X429) = s(bool,X430) | s(bool,X430) = s(bool,'|dollar|exists_1'(s(fun(X425,bool),X426))) | s(bool,X429) = s(bool,X431) | s(bool,'T_0') = s(bool,X431)) )),
  inference(subsumption_resolution,[],[f16337,f239])).
fof(f16337,plain,(
  ( ! [X428,X430,X426,X429,X431,X425,X427] : (~p(s(bool,'T_0')) | p(s(bool,X427)) | ~p(s(bool,X428)) | ~sP12(X428,X427,happ(s(fun(X425,bool),X426),s(X425,'const.min.@_1'(s(fun(X425,bool),X426))))) | s(bool,X429) = s(bool,X430) | s(bool,X430) = s(bool,'|dollar|exists_1'(s(fun(X425,bool),X426))) | s(bool,X429) = s(bool,X431) | s(bool,'T_0') = s(bool,X431)) )),
  inference(superposition,[],[f1843,f2758])).
fof(f18360,plain,(
  ( ! [X436,X438,X432,X434,X437,X433,X435] : (~p(s(bool,X434)) | p(s(bool,X435)) | ~sP13(happ(s(fun(X432,bool),X433),s(X432,'const.min.@_1'(s(fun(X432,bool),X433)))),X435,X434) | s(bool,X436) = s(bool,X437) | s(bool,X437) = s(bool,'|dollar|exists_1'(s(fun(X432,bool),X433))) | s(bool,X436) = s(bool,X438) | s(bool,'T_0') = s(bool,X438)) )),
  inference(subsumption_resolution,[],[f16338,f239])).
fof(f16338,plain,(
  ( ! [X436,X438,X432,X434,X437,X433,X435] : (~p(s(bool,'T_0')) | ~p(s(bool,X434)) | p(s(bool,X435)) | ~sP13(happ(s(fun(X432,bool),X433),s(X432,'const.min.@_1'(s(fun(X432,bool),X433)))),X435,X434) | s(bool,X436) = s(bool,X437) | s(bool,X437) = s(bool,'|dollar|exists_1'(s(fun(X432,bool),X433))) | s(bool,X436) = s(bool,X438) | s(bool,'T_0') = s(bool,X438)) )),
  inference(superposition,[],[f1846,f2758])).
fof(f18361,plain,(
  ( ! [X447,X449,X451,X446,X448,X450] : (p(s(bool,'|dollar|exists_1'(s(fun(X446,bool),X447)))) | s(bool,X449) = s(bool,X450) | s(bool,X450) = s(bool,happ(s(fun(X446,bool),X447),s(X446,X448))) | s(bool,X449) = s(bool,X451) | s(bool,'T_0') = s(bool,X451)) )),
  inference(subsumption_resolution,[],[f16340,f239])).
fof(f16340,plain,(
  ( ! [X447,X449,X451,X446,X448,X450] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X446,bool),X447)))) | s(bool,X449) = s(bool,X450) | s(bool,X450) = s(bool,happ(s(fun(X446,bool),X447),s(X446,X448))) | s(bool,X449) = s(bool,X451) | s(bool,'T_0') = s(bool,X451)) )),
  inference(superposition,[],[f427,f2758])).
fof(f20595,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s('type.extreal.extreal',X44) = s('type.extreal.extreal',X45) | s(bool,X46) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X45),s('type.extreal.extreal',X44))) | s(bool,'T_0') = s(bool,X46) | s(bool,X47) = s(bool,X48) | s(bool,X48) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X44),s('type.extreal.extreal',X45))) | s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f20584,f239])).
fof(f20584,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | s('type.extreal.extreal',X44) = s('type.extreal.extreal',X45) | s(bool,X46) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X45),s('type.extreal.extreal',X44))) | s(bool,'T_0') = s(bool,X46) | s(bool,X47) = s(bool,X48) | s(bool,X48) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X44),s('type.extreal.extreal',X45))) | s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f3282,f2758])).
fof(f20588,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.extreal.extreal',X4) = s('type.extreal.extreal',X5) | s(bool,X6) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X5),s('type.extreal.extreal',X4))) | s(bool,'T_0') = s(bool,X6) | s(bool,X7) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f20576,f239])).
fof(f20576,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | s('type.extreal.extreal',X4) = s('type.extreal.extreal',X5) | s(bool,X6) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X5),s('type.extreal.extreal',X4))) | s(bool,'T_0') = s(bool,X6) | s(bool,X7) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X5))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f3282,f639])).
fof(f20508,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X34)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X33),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X33))) | s(bool,X35) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X33))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f20498,f254])).
fof(f20498,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X34)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X33),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X33))) | s(bool,X35) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X33))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f5062,f2758])).
fof(f20503,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X17))) | s(bool,X19) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X17))) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f20494,f254])).
fof(f20494,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X17))) | s(bool,X19) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f5062,f2758])).
fof(f16881,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2)) )),
  inference(equality_factoring,[],[f2758])).
fof(f18449,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f16874,f639])).
fof(f16874,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(equality_factoring,[],[f2758])).
fof(f18406,plain,(
  ( ! [X466,X463,X465,X467,X462,X468,X464] : (p(s(bool,X464)) | s(bool,X465) = s(bool,happ(s(fun(X462,bool),X463),s(X462,X466))) | s(bool,'T_0') = s(bool,X465) | s(bool,X467) = s(bool,'|dollar|exists_1'(s(fun(X462,bool),X463))) | s(bool,X467) = s(bool,X468) | s(bool,'T_0') = s(bool,X468)) )),
  inference(subsumption_resolution,[],[f16473,f254])).
fof(f16473,plain,(
  ( ! [X466,X463,X465,X467,X462,X468,X464] : (p(s(bool,X464)) | s(bool,X465) = s(bool,happ(s(fun(X462,bool),X463),s(X462,X466))) | s(bool,'T_0') = s(bool,X465) | s(bool,X467) = s(bool,'|dollar|exists_1'(s(fun(X462,bool),X463))) | s(bool,X467) = s(bool,X468) | s(bool,'T_0') = s(bool,X468) | s(bool,'T_0') = s(bool,X464)) )),
  inference(superposition,[],[f3281,f2758])).
fof(f18373,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X225,X227) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230)) )),
  inference(subsumption_resolution,[],[f16425,f254])).
fof(f16425,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X225,X227) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230) | s(bool,'T_0') = s(bool,X226)) )),
  inference(superposition,[],[f3260,f2758])).
fof(f18372,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (p(s(bool,X220)) | sP1(X221,X219) | s(bool,X221) = s(bool,X222) | s(bool,'T_0') = s(bool,X222) | s(bool,X219) = s(bool,X223) | s(bool,X223) = s(bool,X224) | s(bool,'T_0') = s(bool,X224)) )),
  inference(subsumption_resolution,[],[f16424,f254])).
fof(f16424,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (p(s(bool,X220)) | sP1(X221,X219) | s(bool,X221) = s(bool,X222) | s(bool,'T_0') = s(bool,X222) | s(bool,X219) = s(bool,X223) | s(bool,X223) = s(bool,X224) | s(bool,'T_0') = s(bool,X224) | s(bool,'T_0') = s(bool,X220)) )),
  inference(superposition,[],[f3259,f2758])).
fof(f16381,plain,(
  ( ! [X668,X670,X667,X669,X671,X672] : (s(X669,f(s(bool,'T_0'))) != s(X669,g(s(bool,'T_0'))) | s(fun(bool,X669),X667) = s(fun(bool,X669),X668) | s(bool,X670) = s(bool,X671) | s(bool,X671) = s(bool,sK20(X667,X668,X669,bool)) | s(bool,X670) = s(bool,X672) | s(bool,'T_0') = s(bool,X672)) )),
  inference(superposition,[],[f321,f2758])).
fof(f18345,plain,(
  ( ! [X231,X235,X233,X230,X236,X234,X232] : (p(s(bool,X231)) | ~sP13(X232,X231,X230) | s(bool,X232) = s(bool,X233) | s(bool,'T_0') = s(bool,X233) | s(bool,X234) = s(bool,X235) | s(bool,X230) = s(bool,X235) | s(bool,X234) = s(bool,X236) | s(bool,'T_0') = s(bool,X236)) )),
  inference(subsumption_resolution,[],[f16297,f239])).
fof(f16297,plain,(
  ( ! [X231,X235,X233,X230,X236,X234,X232] : (~p(s(bool,'T_0')) | p(s(bool,X231)) | ~sP13(X232,X231,X230) | s(bool,X232) = s(bool,X233) | s(bool,'T_0') = s(bool,X233) | s(bool,X234) = s(bool,X235) | s(bool,X230) = s(bool,X235) | s(bool,X234) = s(bool,X236) | s(bool,'T_0') = s(bool,X236)) )),
  inference(superposition,[],[f3276,f2758])).
fof(f18344,plain,(
  ( ! [X229,X227,X225,X223,X228,X226,X224] : (p(s(bool,X224)) | ~sP9(X224,X223,X225) | s(bool,X225) = s(bool,X226) | s(bool,'T_0') = s(bool,X226) | s(bool,X227) = s(bool,X228) | s(bool,X223) = s(bool,X228) | s(bool,X227) = s(bool,X229) | s(bool,'T_0') = s(bool,X229)) )),
  inference(subsumption_resolution,[],[f16296,f239])).
fof(f16296,plain,(
  ( ! [X229,X227,X225,X223,X228,X226,X224] : (~p(s(bool,'T_0')) | p(s(bool,X224)) | ~sP9(X224,X223,X225) | s(bool,X225) = s(bool,X226) | s(bool,'T_0') = s(bool,X226) | s(bool,X227) = s(bool,X228) | s(bool,X223) = s(bool,X228) | s(bool,X227) = s(bool,X229) | s(bool,'T_0') = s(bool,X229)) )),
  inference(superposition,[],[f3268,f2758])).
fof(f18343,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~sP10(X163,X164,X165) | sP11(X163,X164,X165) | s(bool,X166) = s(bool,X167) | s(bool,X163) = s(bool,X167) | s(bool,X166) = s(bool,X168) | s(bool,'T_0') = s(bool,X168)) )),
  inference(subsumption_resolution,[],[f16285,f239])).
fof(f16285,plain,(
  ( ! [X167,X165,X163,X166,X164,X168] : (~p(s(bool,'T_0')) | ~sP10(X163,X164,X165) | sP11(X163,X164,X165) | s(bool,X166) = s(bool,X167) | s(bool,X163) = s(bool,X167) | s(bool,X166) = s(bool,X168) | s(bool,'T_0') = s(bool,X168)) )),
  inference(superposition,[],[f417,f2758])).
fof(f18342,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (~sP12(X157,X158,X159) | sP13(X157,X158,X159) | s(bool,X160) = s(bool,X161) | s(bool,X157) = s(bool,X161) | s(bool,X160) = s(bool,X162) | s(bool,'T_0') = s(bool,X162)) )),
  inference(subsumption_resolution,[],[f16284,f239])).
fof(f16284,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (~p(s(bool,'T_0')) | ~sP12(X157,X158,X159) | sP13(X157,X158,X159) | s(bool,X160) = s(bool,X161) | s(bool,X157) = s(bool,X161) | s(bool,X160) = s(bool,X162) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f416,f2758])).
fof(f18341,plain,(
  ( ! [X156,X154,X152,X151,X155,X153] : (~p(s(bool,X152)) | p(s(bool,X153)) | ~sP13(X151,X153,X152) | s(bool,X154) = s(bool,X155) | s(bool,X151) = s(bool,X155) | s(bool,X154) = s(bool,X156) | s(bool,'T_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f16283,f239])).
fof(f16283,plain,(
  ( ! [X156,X154,X152,X151,X155,X153] : (~p(s(bool,'T_0')) | ~p(s(bool,X152)) | p(s(bool,X153)) | ~sP13(X151,X153,X152) | s(bool,X154) = s(bool,X155) | s(bool,X151) = s(bool,X155) | s(bool,X154) = s(bool,X156) | s(bool,'T_0') = s(bool,X156)) )),
  inference(superposition,[],[f413,f2758])).
fof(f18340,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (~p(s(bool,X146)) | sP13(X145,X146,X147) | s(bool,X148) = s(bool,X149) | s(bool,X145) = s(bool,X149) | s(bool,X148) = s(bool,X150) | s(bool,'T_0') = s(bool,X150)) )),
  inference(subsumption_resolution,[],[f16282,f239])).
fof(f16282,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (~p(s(bool,'T_0')) | ~p(s(bool,X146)) | sP13(X145,X146,X147) | s(bool,X148) = s(bool,X149) | s(bool,X145) = s(bool,X149) | s(bool,X148) = s(bool,X150) | s(bool,'T_0') = s(bool,X150)) )),
  inference(superposition,[],[f412,f2758])).
fof(f18339,plain,(
  ( ! [X144,X142,X140,X143,X141,X139] : (p(s(bool,X140)) | sP13(X139,X141,X140) | s(bool,X142) = s(bool,X143) | s(bool,X139) = s(bool,X143) | s(bool,X142) = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) )),
  inference(subsumption_resolution,[],[f16281,f239])).
fof(f16281,plain,(
  ( ! [X144,X142,X140,X143,X141,X139] : (~p(s(bool,'T_0')) | p(s(bool,X140)) | sP13(X139,X141,X140) | s(bool,X142) = s(bool,X143) | s(bool,X139) = s(bool,X143) | s(bool,X142) = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f411,f2758])).
fof(f18338,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (p(s(bool,X134)) | ~p(s(bool,X135)) | ~sP12(X135,X134,X133) | s(bool,X136) = s(bool,X137) | s(bool,X133) = s(bool,X137) | s(bool,X136) = s(bool,X138) | s(bool,'T_0') = s(bool,X138)) )),
  inference(subsumption_resolution,[],[f16280,f239])).
fof(f16280,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (~p(s(bool,'T_0')) | p(s(bool,X134)) | ~p(s(bool,X135)) | ~sP12(X135,X134,X133) | s(bool,X136) = s(bool,X137) | s(bool,X133) = s(bool,X137) | s(bool,X136) = s(bool,X138) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f406,f2758])).
fof(f18337,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (p(s(bool,X128)) | p(s(bool,X129)) | ~sP11(X127,X129,X128) | s(bool,X130) = s(bool,X131) | s(bool,X127) = s(bool,X131) | s(bool,X130) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(subsumption_resolution,[],[f16279,f239])).
fof(f16279,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X128)) | p(s(bool,X129)) | ~sP11(X127,X129,X128) | s(bool,X130) = s(bool,X131) | s(bool,X127) = s(bool,X131) | s(bool,X130) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f396,f2758])).
fof(f18336,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (~p(s(bool,X122)) | sP11(X121,X122,X123) | s(bool,X124) = s(bool,X125) | s(bool,X121) = s(bool,X125) | s(bool,X124) = s(bool,X126) | s(bool,'T_0') = s(bool,X126)) )),
  inference(subsumption_resolution,[],[f16278,f239])).
fof(f16278,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (~p(s(bool,'T_0')) | ~p(s(bool,X122)) | sP11(X121,X122,X123) | s(bool,X124) = s(bool,X125) | s(bool,X121) = s(bool,X125) | s(bool,X124) = s(bool,X126) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f395,f2758])).
fof(f18335,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (~p(s(bool,X116)) | sP11(X115,X117,X116) | s(bool,X118) = s(bool,X119) | s(bool,X115) = s(bool,X119) | s(bool,X118) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f16277,f239])).
fof(f16277,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (~p(s(bool,'T_0')) | ~p(s(bool,X116)) | sP11(X115,X117,X116) | s(bool,X118) = s(bool,X119) | s(bool,X115) = s(bool,X119) | s(bool,X118) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f394,f2758])).
fof(f18334,plain,(
  ( ! [X111,X109,X114,X112,X110,X113] : (p(s(bool,X110)) | p(s(bool,X111)) | ~sP10(X109,X110,X111) | s(bool,X112) = s(bool,X113) | s(bool,X109) = s(bool,X113) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f16276,f239])).
fof(f16276,plain,(
  ( ! [X111,X109,X114,X112,X110,X113] : (~p(s(bool,'T_0')) | p(s(bool,X110)) | p(s(bool,X111)) | ~sP10(X109,X110,X111) | s(bool,X112) = s(bool,X113) | s(bool,X109) = s(bool,X113) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f389,f2758])).
fof(f18333,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (~p(s(bool,X104)) | p(s(bool,X105)) | ~sP9(X105,X104,X103) | s(bool,X106) = s(bool,X107) | s(bool,X103) = s(bool,X107) | s(bool,X106) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f16275,f239])).
fof(f16275,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X104)) | p(s(bool,X105)) | ~sP9(X105,X104,X103) | s(bool,X106) = s(bool,X107) | s(bool,X103) = s(bool,X107) | s(bool,X106) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f381,f2758])).
fof(f18332,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (~p(s(bool,X98)) | p(s(bool,X99)) | ~sP8(X98,X97,X99) | s(bool,X100) = s(bool,X101) | s(bool,X97) = s(bool,X101) | s(bool,X100) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f16274,f239])).
fof(f16274,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (~p(s(bool,'T_0')) | ~p(s(bool,X98)) | p(s(bool,X99)) | ~sP8(X98,X97,X99) | s(bool,X100) = s(bool,X101) | s(bool,X97) = s(bool,X101) | s(bool,X100) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f372,f2758])).
fof(f18331,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (~p(s(bool,X92)) | p(s(bool,X93)) | ~sP5(X92,X91,X93) | s(bool,X94) = s(bool,X95) | s(bool,X91) = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f16273,f239])).
fof(f16273,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X92)) | p(s(bool,X93)) | ~sP5(X92,X91,X93) | s(bool,X94) = s(bool,X95) | s(bool,X91) = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f353,f2758])).
fof(f18330,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~sP6(X86,X85,X87) | s(bool,X88) = s(bool,X89) | s(bool,X85) = s(bool,X89) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f16272,f239])).
fof(f16272,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,'T_0')) | ~p(s(bool,X86)) | p(s(bool,X87)) | ~sP6(X86,X85,X87) | s(bool,X88) = s(bool,X89) | s(bool,X85) = s(bool,X89) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f349,f2758])).
fof(f18329,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (sP7(X80,X81,X79) | ~sP6(X80,X79,X81) | ~sP5(X81,X79,X80) | ~sP4(X79,X81,X80) | s(bool,X82) = s(bool,X83) | s(bool,X79) = s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f16271,f239])).
fof(f16271,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (~p(s(bool,'T_0')) | sP7(X80,X81,X79) | ~sP6(X80,X79,X81) | ~sP5(X81,X79,X80) | ~sP4(X79,X81,X80) | s(bool,X82) = s(bool,X83) | s(bool,X79) = s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f346,f2758])).
fof(f18328,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (~p(s(bool,X74)) | sP3(X75,X74,X73) | s(bool,X76) = s(bool,X77) | s(bool,X73) = s(bool,X77) | s(bool,X76) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f16270,f239])).
fof(f16270,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X74)) | sP3(X75,X74,X73) | s(bool,X76) = s(bool,X77) | s(bool,X73) = s(bool,X77) | s(bool,X76) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f334,f2758])).
fof(f18327,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,X68)) | sP2(X68,X69,X67) | s(bool,X70) = s(bool,X71) | s(bool,X67) = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f16269,f239])).
fof(f16269,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X68)) | sP2(X68,X69,X67) | s(bool,X70) = s(bool,X71) | s(bool,X67) = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f327,f2758])).
fof(f18326,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,X62)) | sP2(X62,X61,X63) | s(bool,X64) = s(bool,X65) | s(bool,X61) = s(bool,X65) | s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f16268,f239])).
fof(f16268,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X62)) | sP2(X62,X61,X63) | s(bool,X64) = s(bool,X65) | s(bool,X61) = s(bool,X65) | s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f326,f2758])).
fof(f18325,plain,(
  ( ! [X54,X52,X50,X53,X51] : (p(s(bool,X51)) | sP1(X51,X50) | s(bool,X52) = s(bool,X53) | s(bool,X50) = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f16266,f239])).
fof(f16266,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X51)) | sP1(X51,X50) | s(bool,X52) = s(bool,X53) | s(bool,X50) = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f316,f2758])).
fof(f18324,plain,(
  ( ! [X47,X45,X48,X46,X49] : (p(s(bool,X46)) | sP1(X45,X46) | s(bool,X47) = s(bool,X48) | s(bool,X45) = s(bool,X48) | s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f16265,f239])).
fof(f16265,plain,(
  ( ! [X47,X45,X48,X46,X49] : (~p(s(bool,'T_0')) | p(s(bool,X46)) | sP1(X45,X46) | s(bool,X47) = s(bool,X48) | s(bool,X45) = s(bool,X48) | s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f313,f2758])).
fof(f18323,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP0(X40,X41) | s(bool,X42) = s(bool,X43) | s(bool,X40) = s(bool,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f16264,f239])).
fof(f16264,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'T_0')) | sP0(X40,X41) | s(bool,X42) = s(bool,X43) | s(bool,X40) = s(bool,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f298,f2758])).
fof(f18322,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP0(X36,X35) | s(bool,X37) = s(bool,X38) | s(bool,X35) = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f16263,f239])).
fof(f16263,plain,(
  ( ! [X39,X37,X35,X38,X36] : (~p(s(bool,'T_0')) | sP0(X36,X35) | s(bool,X37) = s(bool,X38) | s(bool,X35) = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f297,f2758])).
fof(f18321,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,X31)) | s(bool,X30) = s(bool,X31) | s(bool,X32) = s(bool,X33) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f16262,f239])).
fof(f16262,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | s(bool,X30) = s(bool,X31) | s(bool,X32) = s(bool,X33) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f288,f2758])).
fof(f18320,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,X12) | s(bool,X13) = s(bool,X14) | s(bool,X12) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f16258,f239])).
fof(f16258,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X12) | s(bool,X13) = s(bool,X14) | s(bool,X12) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f258,f2758])).
fof(f16198,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | s(bool,X6) = s(bool,X7) | s(bool,X4) = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f254,f2758])).
fof(f18290,plain,(
  ( ! [X313,X315,X312,X314,X311] : (s(bool,X311) != s(bool,X312) | s(bool,X312) = s(bool,X313) | s(bool,'T_0') = s(bool,X313) | s(bool,X311) = s(bool,X314) | s(bool,X314) = s(bool,X315) | s(bool,'T_0') = s(bool,X315)) )),
  inference(subsumption_resolution,[],[f16126,f477])).
fof(f16126,plain,(
  ( ! [X313,X315,X312,X314,X311] : (s(bool,X311) != s(bool,X312) | s(bool,X312) = s(bool,X313) | s(bool,'T_0') = s(bool,X313) | s(bool,X311) = s(bool,X314) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X314) = s(bool,X315) | s(bool,'T_0') = s(bool,X315)) )),
  inference(superposition,[],[f3255,f2758])).
fof(f18289,plain,(
  ( ! [X302,X304,X306,X301,X303,X305] : (s(X302,f(s(bool,X301))) != s(X302,g(s(bool,X301))) | s(fun(bool,X302),X303) = s(fun(bool,X302),X304) | s(bool,'T_0') = s(bool,sK20(X303,X304,X302,bool)) | s(bool,X301) = s(bool,X305) | s(bool,X305) = s(bool,X306) | s(bool,'T_0') = s(bool,X306)) )),
  inference(subsumption_resolution,[],[f16125,f477])).
fof(f16125,plain,(
  ( ! [X302,X304,X306,X301,X303,X305] : (s(X302,f(s(bool,X301))) != s(X302,g(s(bool,X301))) | s(fun(bool,X302),X303) = s(fun(bool,X302),X304) | s(bool,'T_0') = s(bool,sK20(X303,X304,X302,bool)) | s(bool,X301) = s(bool,X305) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X305) = s(bool,X306) | s(bool,'T_0') = s(bool,X306)) )),
  inference(superposition,[],[f1875,f2758])).
fof(f18287,plain,(
  ( ! [X294,X296,X295] : (s(bool,'T_0') != s(bool,X294) | s(bool,X294) = s(bool,X295) | s(bool,X295) = s(bool,X296) | s(bool,'T_0') = s(bool,X296)) )),
  inference(subsumption_resolution,[],[f16123,f477])).
fof(f16123,plain,(
  ( ! [X294,X296,X295] : (s(bool,'T_0') != s(bool,X294) | s(bool,X294) = s(bool,X295) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X295) = s(bool,X296) | s(bool,'T_0') = s(bool,X296)) )),
  inference(superposition,[],[f477,f2758])).
fof(f18272,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f16061,f254])).
fof(f16061,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f254,f2758])).
fof(f16058,plain,(
  ( ! [X705,X704,X702,X706,X701,X703] : (s(X703,f(s(bool,X704))) != s(X703,g(s(bool,X704))) | s(fun(bool,X703),X701) = s(fun(bool,X703),X702) | s(bool,X705) = s(bool,sK20(X701,X702,X703,bool)) | s(bool,'T_0') = s(bool,X704) | s(bool,X705) = s(bool,X706) | s(bool,'T_0') = s(bool,X706)) )),
  inference(superposition,[],[f321,f2758])).
fof(f18270,plain,(
  ( ! [X466,X463,X465,X467,X462,X468,X464] : (p(s(bool,X464)) | s(bool,X465) = s(bool,happ(s(fun(X462,bool),X463),s(X462,X466))) | s(bool,'T_0') = s(bool,X465) | s(bool,X467) = s(bool,'|dollar|exists_1'(s(fun(X462,bool),X463))) | s(bool,X467) = s(bool,X468) | s(bool,'T_0') = s(bool,X468)) )),
  inference(subsumption_resolution,[],[f16013,f254])).
fof(f16013,plain,(
  ( ! [X466,X463,X465,X467,X462,X468,X464] : (p(s(bool,X464)) | s(bool,X465) = s(bool,happ(s(fun(X462,bool),X463),s(X462,X466))) | s(bool,'T_0') = s(bool,X465) | s(bool,X467) = s(bool,'|dollar|exists_1'(s(fun(X462,bool),X463))) | s(bool,'T_0') = s(bool,X464) | s(bool,X467) = s(bool,X468) | s(bool,'T_0') = s(bool,X468)) )),
  inference(superposition,[],[f3281,f2758])).
fof(f18236,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X225,X227) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230)) )),
  inference(subsumption_resolution,[],[f15965,f254])).
fof(f15965,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (p(s(bool,X226)) | sP1(X225,X227) | s(bool,X227) = s(bool,X228) | s(bool,'T_0') = s(bool,X228) | s(bool,X225) = s(bool,X229) | s(bool,'T_0') = s(bool,X226) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230)) )),
  inference(superposition,[],[f3260,f2758])).
fof(f18234,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (p(s(bool,X220)) | sP1(X221,X219) | s(bool,X221) = s(bool,X222) | s(bool,'T_0') = s(bool,X222) | s(bool,X219) = s(bool,X223) | s(bool,X223) = s(bool,X224) | s(bool,'T_0') = s(bool,X224)) )),
  inference(subsumption_resolution,[],[f15964,f254])).
fof(f15964,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (p(s(bool,X220)) | sP1(X221,X219) | s(bool,X221) = s(bool,X222) | s(bool,'T_0') = s(bool,X222) | s(bool,X219) = s(bool,X223) | s(bool,'T_0') = s(bool,X220) | s(bool,X223) = s(bool,X224) | s(bool,'T_0') = s(bool,X224)) )),
  inference(superposition,[],[f3259,f2758])).
fof(f18228,plain,(
  ( ! [X198,X196,X194,X199,X197,X195] : (p(s(bool,X195)) | s(bool,X194) != s(bool,X196) | s(bool,X196) = s(bool,X197) | s(bool,'T_0') = s(bool,X197) | s(bool,X194) = s(bool,X198) | s(bool,X198) = s(bool,X199) | s(bool,'T_0') = s(bool,X199)) )),
  inference(subsumption_resolution,[],[f15959,f254])).
fof(f15959,plain,(
  ( ! [X198,X196,X194,X199,X197,X195] : (p(s(bool,X195)) | s(bool,X194) != s(bool,X196) | s(bool,X196) = s(bool,X197) | s(bool,'T_0') = s(bool,X197) | s(bool,X194) = s(bool,X198) | s(bool,'T_0') = s(bool,X195) | s(bool,X198) = s(bool,X199) | s(bool,'T_0') = s(bool,X199)) )),
  inference(superposition,[],[f2880,f2758])).
fof(f18224,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f18223])).
fof(f18223,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X20) = s(bool,X22) | s(bool,X20) = s(bool,X23) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f15928,f254])).
fof(f15928,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X20) = s(bool,X22) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X21) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f285,f2758])).
fof(f15784,plain,(
  ( ! [X705,X704,X702,X706,X701,X703] : (s(X703,f(s(bool,X704))) != s(X703,g(s(bool,X704))) | s(fun(bool,X703),X701) = s(fun(bool,X703),X702) | s(bool,X704) = s(bool,X705) | s(bool,'T_0') = s(bool,X705) | s(bool,X706) = s(bool,sK20(X701,X702,X703,bool)) | s(bool,'T_0') = s(bool,X706)) )),
  inference(superposition,[],[f321,f2758])).
fof(f15535,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X1) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f2758,f2758])).
fof(f17758,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f17757])).
fof(f17757,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f15507])).
fof(f15507,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2758,f247])).
fof(f15448,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,X134) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X133))) | s(bool,X134) = s(bool,X135) | s(bool,'T_0') = s(bool,X135) | s(bool,X136) = s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X136) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X133)))) )),
  inference(superposition,[],[f2758,f1815])).
fof(f15447,plain,(
  ( ! [X132,X130,X128,X131,X129] : (s(bool,X130) = s(bool,'|dollar|exists_1'(s(fun(X128,bool),X129))) | s(bool,X130) = s(bool,X131) | s(bool,'T_0') = s(bool,X131) | s(bool,X132) = s(bool,happ(s(fun(X128,bool),X129),s(X128,'const.min.@_1'(s(fun(X128,bool),X129))))) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f2758,f281])).
fof(f17912,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X1) = s(bool,X3)) )),
  inference(condensation,[],[f17911])).
fof(f17911,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,X80) = s(bool,X81) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X83) | s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,X80) = s(bool,X84)) )),
  inference(duplicate_literal_removal,[],[f15436])).
fof(f15436,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,X80) = s(bool,X81) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X83) | s(bool,'T_0') = s(bool,X83) | s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,X80) = s(bool,X84)) )),
  inference(superposition,[],[f2758,f2758])).
fof(f15237,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f1277,f2717])).
fof(f15242,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f1278,f2717])).
fof(f15204,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f1325,f2716])).
fof(f15173,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f1553,f2633])).
fof(f15178,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f1554,f2633])).
fof(f14727,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f1601,f2632])).
fof(f12851,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f467,f431,f1720])).
fof(f12854,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f454,f431,f1720])).
fof(f12855,plain,(
  sP7('T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f434,f466,f1720])).
fof(f12856,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f467,f466,f1720])).
fof(f12857,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f1193,f466,f1720])).
fof(f12858,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f1469,f466,f1720])).
fof(f12859,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f454,f466,f1720])).
fof(f12860,plain,(
  sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f434,f1178,f1720])).
fof(f12861,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f467,f1178,f1720])).
fof(f12862,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f1193,f1178,f1720])).
fof(f12863,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f1469,f1178,f1720])).
fof(f12864,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f454,f1178,f1720])).
fof(f12865,plain,(
  sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f434,f1454,f1720])).
fof(f12866,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f467,f1454,f1720])).
fof(f12867,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f1193,f1454,f1720])).
fof(f12868,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f1469,f1454,f1720])).
fof(f12869,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f454,f1454,f1720])).
fof(f12870,plain,(
  ( ! [X0] : (sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f434,f453,f1720])).
fof(f12871,plain,(
  ( ! [X0] : (sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f467,f453,f1720])).
fof(f12872,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f1193,f453,f1720])).
fof(f12873,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f1469,f453,f1720])).
fof(f12874,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f454,f453,f1720])).
fof(f12341,plain,(
  ( ! [X0] : (sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f467,f431,f1722])).
fof(f12344,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f454,f431,f1722])).
fof(f12345,plain,(
  ( ! [X0] : (sP7('T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f434,f466,f1722])).
fof(f12346,plain,(
  ( ! [X0] : (sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f467,f466,f1722])).
fof(f12347,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f1193,f466,f1722])).
fof(f12348,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f1469,f466,f1722])).
fof(f12349,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f454,f466,f1722])).
fof(f12350,plain,(
  ( ! [X0] : (sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f434,f1178,f1722])).
fof(f12351,plain,(
  ( ! [X0] : (sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f467,f1178,f1722])).
fof(f12352,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f1193,f1178,f1722])).
fof(f12353,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f1469,f1178,f1722])).
fof(f12354,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f454,f1178,f1722])).
fof(f12355,plain,(
  ( ! [X0] : (sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f434,f1454,f1722])).
fof(f12356,plain,(
  ( ! [X0] : (sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f467,f1454,f1722])).
fof(f12357,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f1193,f1454,f1722])).
fof(f12358,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f1469,f1454,f1722])).
fof(f12359,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f454,f1454,f1722])).
fof(f12360,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f434,f453,f1722])).
fof(f12361,plain,(
  ( ! [X0,X1] : (sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f467,f453,f1722])).
fof(f12362,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f1193,f453,f1722])).
fof(f12363,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f1469,f453,f1722])).
fof(f12364,plain,(
  ( ! [X2,X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f454,f453,f1722])).
fof(f12021,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) = s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f12004,f239])).
fof(f12004,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f5167,f639])).
fof(f11974,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X65)) | ~sP13(X66,X65,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X64),s('type.extreal.extreal',X64))) | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f11956,f239])).
fof(f11956,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | ~sP13(X66,X65,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X64),s('type.extreal.extreal',X64))) | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f3276,f480])).
fof(f11972,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X47)) | ~sP13(X48,X47,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f11952,f239])).
fof(f11952,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | ~sP13(X48,X47,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f3276,f1404])).
fof(f11971,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | ~sP13(X45,X44,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f11951,f239])).
fof(f11951,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'T_0')) | p(s(bool,X44)) | ~sP13(X45,X44,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f3276,f1128])).
fof(f11970,plain,(
  ( ! [X43,X41,X42] : (~sP13(X42,X41,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | p(s(bool,X41)) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f11950,f239])).
fof(f11950,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'T_0')) | p(s(bool,X41)) | ~sP13(X42,X41,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f3276,f481])).
fof(f11965,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | ~sP13(X0,X2,'T_0') | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f11964])).
fof(f11964,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | ~sP13(X0,X2,'T_0') | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f11963])).
fof(f11963,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | ~sP13(X26,X25,'T_0') | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f11946,f733])).
fof(f11946,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | ~sP13(X26,X25,'T_0') | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f3276,f639])).
fof(f11962,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,X5)) | ~sP13(X6,X5,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f11942,f239])).
fof(f11942,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | ~sP13(X6,X5,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f3276,f639])).
fof(f11692,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,X62)) | ~sP9(X62,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal',X61)),X63) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f11675,f239])).
fof(f11675,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'T_0')) | p(s(bool,X62)) | ~sP9(X62,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal',X61)),X63) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f3268,f480])).
fof(f11690,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | ~sP9(X44,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X45) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f11671,f239])).
fof(f11671,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'T_0')) | p(s(bool,X44)) | ~sP9(X44,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X45) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f3268,f1404])).
fof(f11689,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~sP9(X41,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X42) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f11670,f239])).
fof(f11670,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'T_0')) | p(s(bool,X41)) | ~sP9(X41,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X42) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f3268,f1128])).
fof(f11688,plain,(
  ( ! [X39,X38,X40] : (~sP9(X38,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X39) | p(s(bool,X38)) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f11669,f239])).
fof(f11669,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,'T_0')) | p(s(bool,X38)) | ~sP9(X38,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X39) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f3268,f481])).
fof(f11684,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | ~sP9(X2,'T_0',X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f11683])).
fof(f11683,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | ~sP9(X2,'T_0',X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f11682])).
fof(f11682,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | ~sP9(X25,'T_0',X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f11666,f733])).
fof(f11666,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | ~sP9(X25,'T_0',X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f3268,f639])).
fof(f11681,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,X5)) | ~sP9(X5,X4,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f11662,f239])).
fof(f11662,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | ~sP9(X5,X4,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f3268,f639])).
fof(f9780,plain,(
  ( ! [X52,X53,X51] : (s(bool,X52) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X51),s('type.extreal.extreal',X51))) | s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f9759,f239])).
fof(f9759,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X52) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X51),s('type.extreal.extreal',X51))) | s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f3256,f480])).
fof(f9777,plain,(
  ( ! [X39,X38] : (s(bool,X38) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f9755,f239])).
fof(f9755,plain,(
  ( ! [X39,X38] : (~p(s(bool,'T_0')) | s(bool,X38) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f3256,f1404])).
fof(f9776,plain,(
  ( ! [X37,X36] : (s(bool,X36) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f9754,f239])).
fof(f9754,plain,(
  ( ! [X37,X36] : (~p(s(bool,'T_0')) | s(bool,X36) = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f3256,f1128])).
fof(f9775,plain,(
  ( ! [X35,X34] : (s(bool,X34) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f9753,f239])).
fof(f9753,plain,(
  ( ! [X35,X34] : (~p(s(bool,'T_0')) | s(bool,X34) = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f3256,f481])).
fof(f8878,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f8863,f254])).
fof(f8863,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f3260,f247])).
fof(f8877,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,X17) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f8860,f426])).
fof(f8860,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | sP1(X16,X17) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f3260,f247])).
fof(f8876,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f8859,f254])).
fof(f8859,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f3260,f639])).
fof(f8874,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f8858,f254])).
fof(f8858,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f3260,f639])).
fof(f6002,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f5986,f254])).
fof(f5986,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f3259,f247])).
fof(f6001,plain,(
  ( ! [X17,X18,X16] : (sP1(X17,X16) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f5983,f426])).
fof(f5983,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | sP1(X17,X16) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f3259,f247])).
fof(f6000,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f5982,f254])).
fof(f5982,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f3259,f639])).
fof(f5997,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f5981,f254])).
fof(f5981,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f3259,f639])).
fof(f5916,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f3255,f247])).
fof(f5427,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK20(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK20(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1876,f639])).
fof(f5426,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK20(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK20(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1876,f639])).
fof(f5394,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK20(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1875,f247])).
fof(f5393,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK20(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1875,f639])).
fof(f5392,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK20(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1875,f639])).
fof(f5391,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK20(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1875,f639])).
fof(f5390,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK20(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1875,f639])).
fof(f5230,plain,(
  ( ! [X0] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f472,f722])).
fof(f5233,plain,(
  ( ! [X0,X1] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f459,f722])).
fof(f5235,plain,(
  ( ! [X0] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f473,f722])).
fof(f5238,plain,(
  ( ! [X0,X1] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f460,f722])).
fof(f5208,plain,(
  ( ! [X0] : (sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f474,f721])).
fof(f5211,plain,(
  ( ! [X0,X1] : (sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f461,f721])).
fof(f5052,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1815,f639])).
fof(f4972,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f472,f1022])).
fof(f4973,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f1277,f1022])).
fof(f4975,plain,(
  ( ! [X2,X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f459,f1022])).
fof(f4977,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f473,f1022])).
fof(f4978,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f1278,f1022])).
fof(f4979,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f1554,f1022])).
fof(f4980,plain,(
  ( ! [X2,X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f460,f1022])).
fof(f4951,plain,(
  ( ! [X0,X1] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f474,f1021])).
fof(f4952,plain,(
  ( ! [X0,X1] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f1325,f1021])).
fof(f4953,plain,(
  ( ! [X0,X1] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f1601,f1021])).
fof(f4954,plain,(
  ( ! [X2,X0,X1] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f461,f1021])).
fof(f4907,plain,(
  ~sP7('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1463,f343])).
fof(f4893,plain,(
  ~sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1448,f344])).
fof(f4776,plain,(
  ( ! [X0] : (~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f426,f4771,f405])).
fof(f4771,plain,(
  ( ! [X0] : (~sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f426,f1118])).
fof(f4754,plain,(
  ( ! [X0] : (~sP13('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f4734,f401])).
fof(f4748,plain,(
  ~sP7('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1187,f343])).
fof(f4734,plain,(
  ( ! [X0] : (~sP12('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f426,f239,f1086])).
fof(f4728,plain,(
  ~sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1172,f344])).
fof(f4718,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f4691,f384])).
fof(f4691,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f426,f1016])).
fof(f4679,plain,(
  ( ! [X0] : (~sP9('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f4661,f365])).
fof(f4661,plain,(
  ( ! [X0] : (~sP8('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f239,f426,f1014])).
fof(f4642,plain,(
  ( ! [X0] : (~sP7('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f4628,f344])).
fof(f4628,plain,(
  ( ! [X0] : (~sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f239,f426,f1012])).
fof(f4466,plain,(
  ( ! [X0,X1] : (sP6('T_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f1814,f344])).
fof(f4417,plain,(
  ( ! [X0,X1] : (sP6('T_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1812,f344])).
fof(f2768,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),X15))) | s(bool,X16) = s(bool,happ(s(fun(X14,bool),X15),s(X14,'const.min.@_1'(s(fun(X14,bool),X15))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f639,f281])).
fof(f3325,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3324])).
fof(f3324,plain,(
  ( ! [X109,X107,X110,X108] : (~sP10('T_0',X108,X109) | sP11('T_0',X108,X109) | s(bool,X107) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f2991,f733])).
fof(f2991,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,X107)) | ~sP10('T_0',X108,X109) | sP11('T_0',X108,X109) | s(bool,X107) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f417,f639])).
fof(f3323,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP12('T_0',X1,X2) | sP13('T_0',X1,X2)) )),
  inference(condensation,[],[f3322])).
fof(f3322,plain,(
  ( ! [X103,X105,X106,X104] : (~sP12('T_0',X104,X105) | sP13('T_0',X104,X105) | s(bool,X103) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(subsumption_resolution,[],[f2990,f733])).
fof(f2990,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,X103)) | ~sP12('T_0',X104,X105) | sP13('T_0',X104,X105) | s(bool,X103) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f416,f639])).
fof(f3321,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP13('T_0',X2,X1)) )),
  inference(condensation,[],[f3320])).
fof(f3320,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,X100)) | p(s(bool,X101)) | ~sP13('T_0',X101,X100) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f2989,f733])).
fof(f2989,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,X99)) | ~p(s(bool,X100)) | p(s(bool,X101)) | ~sP13('T_0',X101,X100) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f413,f639])).
fof(f3319,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP13('T_0',X1,X2)) )),
  inference(condensation,[],[f3318])).
fof(f3318,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X96)) | sP13('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f2988,f733])).
fof(f2988,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X95)) | ~p(s(bool,X96)) | sP13('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f412,f639])).
fof(f3317,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP13('T_0',X2,X1)) )),
  inference(condensation,[],[f3316])).
fof(f3316,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X92)) | sP13('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f2987,f733])).
fof(f2987,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | p(s(bool,X92)) | sP13('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f411,f639])).
fof(f3315,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP12(X2,X1,'T_0')) )),
  inference(condensation,[],[f3314])).
fof(f3314,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X88)) | ~p(s(bool,X89)) | ~sP12(X89,X88,'T_0') | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f2986,f733])).
fof(f2986,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | p(s(bool,X88)) | ~p(s(bool,X89)) | ~sP12(X89,X88,'T_0') | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f406,f639])).
fof(f3313,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3312])).
fof(f3312,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f2985,f733])).
fof(f2985,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f396,f639])).
fof(f3311,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3310])).
fof(f3310,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f2984,f733])).
fof(f2984,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f395,f639])).
fof(f3309,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3308])).
fof(f3308,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f2983,f733])).
fof(f2983,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | ~p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f394,f639])).
fof(f3307,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f3306])).
fof(f3306,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | p(s(bool,X73)) | ~sP10('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2982,f733])).
fof(f2982,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | p(s(bool,X73)) | ~sP10('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f389,f639])).
fof(f3305,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X2,X1,'T_0')) )),
  inference(condensation,[],[f3304])).
fof(f3304,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X69,X68,'T_0') | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f2981,f733])).
fof(f2981,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | ~p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X69,X68,'T_0') | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f381,f639])).
fof(f3303,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X1,'T_0',X2)) )),
  inference(condensation,[],[f3302])).
fof(f3302,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | p(s(bool,X65)) | ~sP8(X64,'T_0',X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f2980,f733])).
fof(f2980,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | p(s(bool,X65)) | ~sP8(X64,'T_0',X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f372,f639])).
fof(f3301,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f3300])).
fof(f3300,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | ~sP5(X60,'T_0',X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f2979,f733])).
fof(f2979,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | p(s(bool,X61)) | ~sP5(X60,'T_0',X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f353,f639])).
fof(f3299,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f3298])).
fof(f3298,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP6(X56,'T_0',X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f2978,f733])).
fof(f2978,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP6(X56,'T_0',X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f349,f639])).
fof(f3297,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP7(X1,X2,'T_0') | ~sP6(X1,'T_0',X2) | ~sP5(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f3296,f439])).
fof(f3296,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP7(X1,X2,'T_0') | ~sP6(X1,'T_0',X2) | ~sP5(X2,'T_0',X1) | ~sP4('T_0',X2,X1)) )),
  inference(condensation,[],[f3295])).
fof(f3295,plain,(
  ( ! [X54,X52,X53,X51] : (sP7(X52,X53,'T_0') | ~sP6(X52,'T_0',X53) | ~sP5(X53,'T_0',X52) | ~sP4('T_0',X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f2977,f733])).
fof(f2977,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | sP7(X52,X53,'T_0') | ~sP6(X52,'T_0',X53) | ~sP5(X53,'T_0',X52) | ~sP4('T_0',X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f346,f639])).
fof(f3294,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP3(X2,X1,'T_0')) )),
  inference(condensation,[],[f3293])).
fof(f3293,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | sP3(X49,X48,'T_0') | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f2976,f733])).
fof(f2976,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | sP3(X49,X48,'T_0') | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f334,f639])).
fof(f3292,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f3291])).
fof(f3291,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | sP2(X44,X45,'T_0') | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f2975,f733])).
fof(f2975,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | sP2(X44,X45,'T_0') | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f327,f639])).
fof(f3290,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f3289])).
fof(f3289,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f2972,f733])).
fof(f2972,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f316,f639])).
fof(f3288,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f3287])).
fof(f3287,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f2971,f733])).
fof(f2971,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f313,f639])).
fof(f2959,plain,(
  ( ! [X156,X154,X157,X155] : (s(X156,f(s(bool,'T_0'))) != s(X156,g(s(bool,'T_0'))) | s(fun(bool,X156),X154) = s(fun(bool,X156),X155) | s(bool,X157) = s(bool,sK20(X154,X155,X156,bool)) | s(bool,'T_0') = s(bool,X157)) )),
  inference(superposition,[],[f321,f639])).
fof(f3278,plain,(
  ( ! [X103,X105,X106,X104] : (~sP10(X103,X104,X105) | sP11(X103,X104,X105) | s(bool,X103) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(subsumption_resolution,[],[f2944,f239])).
fof(f2944,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,'T_0')) | ~sP10(X103,X104,X105) | sP11(X103,X104,X105) | s(bool,X103) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f417,f639])).
fof(f3277,plain,(
  ( ! [X101,X99,X102,X100] : (~sP12(X99,X100,X101) | sP13(X99,X100,X101) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f2943,f239])).
fof(f2943,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,'T_0')) | ~sP12(X99,X100,X101) | sP13(X99,X100,X101) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f416,f639])).
fof(f3276,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X96)) | p(s(bool,X97)) | ~sP13(X95,X97,X96) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f2942,f239])).
fof(f2942,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X96)) | p(s(bool,X97)) | ~sP13(X95,X97,X96) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f413,f639])).
fof(f3275,plain,(
  ( ! [X94,X92,X93,X91] : (sP13(X91,X92,X93) | ~p(s(bool,X92)) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f2941,f239])).
fof(f2941,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~p(s(bool,X92)) | sP13(X91,X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f412,f639])).
fof(f3274,plain,(
  ( ! [X90,X88,X87,X89] : (sP13(X87,X89,X88) | p(s(bool,X88)) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f2940,f239])).
fof(f2940,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | p(s(bool,X88)) | sP13(X87,X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f411,f639])).
fof(f3273,plain,(
  ( ! [X85,X83,X86,X84] : (~sP12(X85,X84,X83) | ~p(s(bool,X85)) | p(s(bool,X84)) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f2939,f239])).
fof(f2939,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X84)) | ~p(s(bool,X85)) | ~sP12(X85,X84,X83) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f406,f639])).
fof(f3272,plain,(
  ( ! [X80,X81,X79,X82] : (~sP11(X79,X81,X80) | p(s(bool,X81)) | p(s(bool,X80)) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f2938,f239])).
fof(f2938,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f396,f639])).
fof(f3271,plain,(
  ( ! [X78,X76,X77,X75] : (sP11(X75,X76,X77) | ~p(s(bool,X76)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f2937,f239])).
fof(f2937,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f395,f639])).
fof(f3270,plain,(
  ( ! [X74,X72,X71,X73] : (sP11(X71,X73,X72) | ~p(s(bool,X72)) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2936,f239])).
fof(f2936,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f394,f639])).
fof(f3269,plain,(
  ( ! [X70,X68,X69,X67] : (~sP10(X67,X68,X69) | p(s(bool,X69)) | p(s(bool,X68)) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f2935,f239])).
fof(f2935,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP10(X67,X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f389,f639])).
fof(f3268,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X65,X64,X63) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f2934,f239])).
fof(f2934,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X65,X64,X63) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f381,f639])).
fof(f3267,plain,(
  ( ! [X61,X59,X62,X60] : (~sP8(X60,X59,X61) | p(s(bool,X61)) | ~p(s(bool,X60)) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f2933,f239])).
fof(f2933,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | p(s(bool,X61)) | ~sP8(X60,X59,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f372,f639])).
fof(f3266,plain,(
  ( ! [X57,X58,X56,X55] : (~sP5(X56,X55,X57) | p(s(bool,X57)) | ~p(s(bool,X56)) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f2932,f239])).
fof(f2932,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP5(X56,X55,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f353,f639])).
fof(f3265,plain,(
  ( ! [X54,X52,X53,X51] : (~sP6(X52,X51,X53) | p(s(bool,X53)) | ~p(s(bool,X52)) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f2931,f239])).
fof(f2931,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP6(X52,X51,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f349,f639])).
fof(f3264,plain,(
  ( ! [X47,X50,X48,X49] : (~sP5(X49,X47,X48) | ~sP6(X48,X47,X49) | sP7(X48,X49,X47) | ~sP4(X47,X49,X48) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f2930,f239])).
fof(f2930,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | sP7(X48,X49,X47) | ~sP6(X48,X47,X49) | ~sP5(X49,X47,X48) | ~sP4(X47,X49,X48) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f346,f639])).
fof(f3263,plain,(
  ( ! [X45,X43,X46,X44] : (sP3(X45,X44,X43) | ~p(s(bool,X44)) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f2929,f239])).
fof(f2929,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | sP3(X45,X44,X43) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f334,f639])).
fof(f3262,plain,(
  ( ! [X39,X41,X42,X40] : (sP2(X40,X41,X39) | ~p(s(bool,X40)) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f2928,f239])).
fof(f2928,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | sP2(X40,X41,X39) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f327,f639])).
fof(f3261,plain,(
  ( ! [X37,X35,X38,X36] : (sP2(X36,X35,X37) | ~p(s(bool,X36)) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2927,f239])).
fof(f2927,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | sP2(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f326,f639])).
fof(f3260,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f2925,f239])).
fof(f2925,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f316,f639])).
fof(f3259,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f2924,f239])).
fof(f2924,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f313,f639])).
fof(f3258,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f2923,f239])).
fof(f2923,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f298,f639])).
fof(f3257,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f2922,f239])).
fof(f2922,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f297,f639])).
fof(f3256,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f2921,f239])).
fof(f2921,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f288,f639])).
fof(f3255,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2917,f239])).
fof(f2917,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f258,f639])).
fof(f3253,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3252])).
fof(f3252,plain,(
  ( ! [X109,X107,X110,X108] : (~sP10('T_0',X108,X109) | sP11('T_0',X108,X109) | s(bool,X107) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f2911,f733])).
fof(f2911,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,X107)) | ~sP10('T_0',X108,X109) | sP11('T_0',X108,X109) | s(bool,X107) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f417,f639])).
fof(f3250,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP12('T_0',X1,X2) | sP13('T_0',X1,X2)) )),
  inference(condensation,[],[f3249])).
fof(f3249,plain,(
  ( ! [X103,X105,X106,X104] : (~sP12('T_0',X104,X105) | sP13('T_0',X104,X105) | s(bool,X103) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(subsumption_resolution,[],[f2910,f733])).
fof(f2910,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,X103)) | ~sP12('T_0',X104,X105) | sP13('T_0',X104,X105) | s(bool,X103) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f416,f639])).
fof(f3247,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP13('T_0',X2,X1)) )),
  inference(condensation,[],[f3246])).
fof(f3246,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,X100)) | p(s(bool,X101)) | ~sP13('T_0',X101,X100) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f2909,f733])).
fof(f2909,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,X99)) | ~p(s(bool,X100)) | p(s(bool,X101)) | ~sP13('T_0',X101,X100) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f413,f639])).
fof(f3244,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP13('T_0',X1,X2)) )),
  inference(condensation,[],[f3243])).
fof(f3243,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X96)) | sP13('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f2908,f733])).
fof(f2908,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X95)) | ~p(s(bool,X96)) | sP13('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f412,f639])).
fof(f3241,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP13('T_0',X2,X1)) )),
  inference(condensation,[],[f3240])).
fof(f3240,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X92)) | sP13('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f2907,f733])).
fof(f2907,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | p(s(bool,X92)) | sP13('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f411,f639])).
fof(f3238,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP12(X2,X1,'T_0')) )),
  inference(condensation,[],[f3237])).
fof(f3237,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X88)) | ~p(s(bool,X89)) | ~sP12(X89,X88,'T_0') | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f2906,f733])).
fof(f2906,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | p(s(bool,X88)) | ~p(s(bool,X89)) | ~sP12(X89,X88,'T_0') | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f406,f639])).
fof(f3235,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3234])).
fof(f3234,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f2905,f733])).
fof(f2905,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f396,f639])).
fof(f3232,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3231])).
fof(f3231,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f2904,f733])).
fof(f2904,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f395,f639])).
fof(f3229,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3228])).
fof(f3228,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f2903,f733])).
fof(f2903,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | ~p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f394,f639])).
fof(f3226,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f3225])).
fof(f3225,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | p(s(bool,X73)) | ~sP10('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2902,f733])).
fof(f2902,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | p(s(bool,X73)) | ~sP10('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f389,f639])).
fof(f3223,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X2,X1,'T_0')) )),
  inference(condensation,[],[f3222])).
fof(f3222,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X69,X68,'T_0') | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f2901,f733])).
fof(f2901,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | ~p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X69,X68,'T_0') | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f381,f639])).
fof(f3220,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X1,'T_0',X2)) )),
  inference(condensation,[],[f3219])).
fof(f3219,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | p(s(bool,X65)) | ~sP8(X64,'T_0',X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f2900,f733])).
fof(f2900,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | p(s(bool,X65)) | ~sP8(X64,'T_0',X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f372,f639])).
fof(f3217,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f3216])).
fof(f3216,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | ~sP5(X60,'T_0',X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f2899,f733])).
fof(f2899,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | p(s(bool,X61)) | ~sP5(X60,'T_0',X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f353,f639])).
fof(f3214,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f3213])).
fof(f3213,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP6(X56,'T_0',X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f2898,f733])).
fof(f2898,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP6(X56,'T_0',X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f349,f639])).
fof(f3211,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP7(X1,X2,'T_0') | ~sP6(X1,'T_0',X2) | ~sP5(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f3210,f439])).
fof(f3210,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP7(X1,X2,'T_0') | ~sP6(X1,'T_0',X2) | ~sP5(X2,'T_0',X1) | ~sP4('T_0',X2,X1)) )),
  inference(condensation,[],[f3209])).
fof(f3209,plain,(
  ( ! [X54,X52,X53,X51] : (sP7(X52,X53,'T_0') | ~sP6(X52,'T_0',X53) | ~sP5(X53,'T_0',X52) | ~sP4('T_0',X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f2897,f733])).
fof(f2897,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | sP7(X52,X53,'T_0') | ~sP6(X52,'T_0',X53) | ~sP5(X53,'T_0',X52) | ~sP4('T_0',X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f346,f639])).
fof(f3207,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP3(X2,X1,'T_0')) )),
  inference(condensation,[],[f3206])).
fof(f3206,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | sP3(X49,X48,'T_0') | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f2896,f733])).
fof(f2896,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | sP3(X49,X48,'T_0') | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f334,f639])).
fof(f3204,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f3203])).
fof(f3203,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | sP2(X44,X45,'T_0') | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f2895,f733])).
fof(f2895,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | sP2(X44,X45,'T_0') | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f327,f639])).
fof(f3201,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f3200])).
fof(f3200,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f2892,f733])).
fof(f2892,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f316,f639])).
fof(f3197,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f3196])).
fof(f3196,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f2891,f733])).
fof(f2891,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f313,f639])).
fof(f2880,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X2)) | s(bool,X1) != s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f254,f639])).
fof(f2829,plain,(
  ( ! [X163,X161,X162,X160] : (s(X162,f(s(bool,X163))) != s(X162,g(s(bool,X163))) | s(fun(bool,X162),X160) = s(fun(bool,X162),X161) | s(bool,'T_0') = s(bool,X163) | s(bool,'T_0') = s(bool,sK20(X160,X161,X162,bool))) )),
  inference(superposition,[],[f321,f639])).
fof(f2758,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f639,f639])).
fof(f2717,plain,(
  ( ! [X47,X48] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X47,X48) | ~sP10('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X47,X48)) )),
  inference(subsumption_resolution,[],[f2682,f239])).
fof(f2682,plain,(
  ( ! [X47,X48] : (~p(s(bool,'T_0')) | ~sP10('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X47,X48) | sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X47,X48)) )),
  inference(superposition,[],[f417,f1404])).
fof(f2716,plain,(
  ( ! [X45,X46] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X45,X46) | ~sP12('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X45,X46)) )),
  inference(subsumption_resolution,[],[f2681,f239])).
fof(f2681,plain,(
  ( ! [X45,X46] : (~p(s(bool,'T_0')) | ~sP12('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X45,X46) | sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X45,X46)) )),
  inference(superposition,[],[f416,f1404])).
fof(f2715,plain,(
  ( ! [X43,X44] : (~sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X44,X43) | p(s(bool,X44)) | ~p(s(bool,X43))) )),
  inference(subsumption_resolution,[],[f2680,f239])).
fof(f2680,plain,(
  ( ! [X43,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X44,X43)) )),
  inference(superposition,[],[f413,f1404])).
fof(f2714,plain,(
  ( ! [X41,X42] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X41,X42) | ~p(s(bool,X41))) )),
  inference(subsumption_resolution,[],[f2679,f239])).
fof(f2679,plain,(
  ( ! [X41,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X41,X42)) )),
  inference(superposition,[],[f412,f1404])).
fof(f2713,plain,(
  ( ! [X39,X40] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X40,X39) | p(s(bool,X39))) )),
  inference(subsumption_resolution,[],[f2678,f239])).
fof(f2678,plain,(
  ( ! [X39,X40] : (~p(s(bool,'T_0')) | p(s(bool,X39)) | sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X40,X39)) )),
  inference(superposition,[],[f411,f1404])).
fof(f2712,plain,(
  ( ! [X37,X38] : (~sP12(X38,X37,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~p(s(bool,X38)) | p(s(bool,X37))) )),
  inference(subsumption_resolution,[],[f2677,f239])).
fof(f2677,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~sP12(X38,X37,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(superposition,[],[f406,f1404])).
fof(f2711,plain,(
  ( ! [X35,X36] : (~sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X36,X35) | p(s(bool,X36)) | p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f2676,f239])).
fof(f2676,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | p(s(bool,X35)) | p(s(bool,X36)) | ~sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X36,X35)) )),
  inference(superposition,[],[f396,f1404])).
fof(f2710,plain,(
  ( ! [X33,X34] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X33,X34) | ~p(s(bool,X33))) )),
  inference(subsumption_resolution,[],[f2675,f239])).
fof(f2675,plain,(
  ( ! [X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X33,X34)) )),
  inference(superposition,[],[f395,f1404])).
fof(f2709,plain,(
  ( ! [X31,X32] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X32,X31) | ~p(s(bool,X31))) )),
  inference(subsumption_resolution,[],[f2674,f239])).
fof(f2674,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X32,X31)) )),
  inference(superposition,[],[f394,f1404])).
fof(f2708,plain,(
  ( ! [X30,X29] : (~sP10('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X29,X30) | p(s(bool,X30)) | p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f2673,f239])).
fof(f2673,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | p(s(bool,X30)) | ~sP10('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X29,X30)) )),
  inference(superposition,[],[f389,f1404])).
fof(f2707,plain,(
  ( ! [X28,X27] : (~sP9(X28,X27,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | p(s(bool,X28)) | ~p(s(bool,X27))) )),
  inference(subsumption_resolution,[],[f2672,f239])).
fof(f2672,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | p(s(bool,X28)) | ~sP9(X28,X27,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(superposition,[],[f381,f1404])).
fof(f2706,plain,(
  ( ! [X26,X25] : (~sP8(X25,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X26) | p(s(bool,X26)) | ~p(s(bool,X25))) )),
  inference(subsumption_resolution,[],[f2671,f239])).
fof(f2671,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP8(X25,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X26)) )),
  inference(superposition,[],[f372,f1404])).
fof(f2705,plain,(
  ( ! [X24,X23] : (~sP5(X23,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X24) | p(s(bool,X24)) | ~p(s(bool,X23))) )),
  inference(subsumption_resolution,[],[f2670,f239])).
fof(f2670,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP5(X23,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X24)) )),
  inference(superposition,[],[f353,f1404])).
fof(f2704,plain,(
  ( ! [X21,X22] : (~sP6(X21,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X22) | p(s(bool,X22)) | ~p(s(bool,X21))) )),
  inference(subsumption_resolution,[],[f2669,f239])).
fof(f2669,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | p(s(bool,X22)) | ~sP6(X21,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X22)) )),
  inference(superposition,[],[f349,f1404])).
fof(f2703,plain,(
  ( ! [X19,X20] : (sP7(X19,X20,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~sP6(X19,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X20) | ~sP5(X20,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X19)) )),
  inference(subsumption_resolution,[],[f2702,f1472])).
fof(f2702,plain,(
  ( ! [X19,X20] : (sP7(X19,X20,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~sP6(X19,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X20) | ~sP5(X20,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X19) | ~sP4('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X20,X19)) )),
  inference(subsumption_resolution,[],[f2668,f239])).
fof(f2668,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | sP7(X19,X20,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~sP6(X19,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X20) | ~sP5(X20,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X19) | ~sP4('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X20,X19)) )),
  inference(superposition,[],[f346,f1404])).
fof(f2701,plain,(
  ( ! [X17,X18] : (sP3(X18,X17,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~p(s(bool,X17))) )),
  inference(subsumption_resolution,[],[f2667,f239])).
fof(f2667,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | sP3(X18,X17,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(superposition,[],[f334,f1404])).
fof(f2700,plain,(
  ( ! [X15,X16] : (sP2(X15,X16,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~p(s(bool,X15))) )),
  inference(subsumption_resolution,[],[f2666,f239])).
fof(f2666,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | sP2(X15,X16,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(superposition,[],[f327,f1404])).
fof(f2699,plain,(
  ( ! [X10] : (sP1(X10,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | p(s(bool,X10))) )),
  inference(subsumption_resolution,[],[f2663,f239])).
fof(f2663,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | p(s(bool,X10)) | sP1(X10,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(superposition,[],[f316,f1404])).
fof(f2698,plain,(
  ( ! [X9] : (sP1('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X9) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f2662,f239])).
fof(f2662,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | sP1('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X9)) )),
  inference(superposition,[],[f313,f1404])).
fof(f2648,plain,(
  sP7('T_0','T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f1404,f362])).
fof(f2633,plain,(
  ( ! [X47,X48] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X47,X48) | ~sP10('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X47,X48)) )),
  inference(subsumption_resolution,[],[f2598,f239])).
fof(f2598,plain,(
  ( ! [X47,X48] : (~p(s(bool,'T_0')) | ~sP10('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X47,X48) | sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X47,X48)) )),
  inference(superposition,[],[f417,f1128])).
fof(f2632,plain,(
  ( ! [X45,X46] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X45,X46) | ~sP12('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X45,X46)) )),
  inference(subsumption_resolution,[],[f2597,f239])).
fof(f2597,plain,(
  ( ! [X45,X46] : (~p(s(bool,'T_0')) | ~sP12('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X45,X46) | sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X45,X46)) )),
  inference(superposition,[],[f416,f1128])).
fof(f2631,plain,(
  ( ! [X43,X44] : (~sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X44,X43) | p(s(bool,X44)) | ~p(s(bool,X43))) )),
  inference(subsumption_resolution,[],[f2596,f239])).
fof(f2596,plain,(
  ( ! [X43,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X44,X43)) )),
  inference(superposition,[],[f413,f1128])).
fof(f2630,plain,(
  ( ! [X41,X42] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X41,X42) | ~p(s(bool,X41))) )),
  inference(subsumption_resolution,[],[f2595,f239])).
fof(f2595,plain,(
  ( ! [X41,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X41,X42)) )),
  inference(superposition,[],[f412,f1128])).
fof(f2629,plain,(
  ( ! [X39,X40] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X40,X39) | p(s(bool,X39))) )),
  inference(subsumption_resolution,[],[f2594,f239])).
fof(f2594,plain,(
  ( ! [X39,X40] : (~p(s(bool,'T_0')) | p(s(bool,X39)) | sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X40,X39)) )),
  inference(superposition,[],[f411,f1128])).
fof(f2628,plain,(
  ( ! [X37,X38] : (~sP12(X38,X37,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~p(s(bool,X38)) | p(s(bool,X37))) )),
  inference(subsumption_resolution,[],[f2593,f239])).
fof(f2593,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~sP12(X38,X37,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(superposition,[],[f406,f1128])).
fof(f2627,plain,(
  ( ! [X35,X36] : (~sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X36,X35) | p(s(bool,X36)) | p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f2592,f239])).
fof(f2592,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | p(s(bool,X35)) | p(s(bool,X36)) | ~sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X36,X35)) )),
  inference(superposition,[],[f396,f1128])).
fof(f2626,plain,(
  ( ! [X33,X34] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X33,X34) | ~p(s(bool,X33))) )),
  inference(subsumption_resolution,[],[f2591,f239])).
fof(f2591,plain,(
  ( ! [X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X33,X34)) )),
  inference(superposition,[],[f395,f1128])).
fof(f2625,plain,(
  ( ! [X31,X32] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X32,X31) | ~p(s(bool,X31))) )),
  inference(subsumption_resolution,[],[f2590,f239])).
fof(f2590,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X32,X31)) )),
  inference(superposition,[],[f394,f1128])).
fof(f2624,plain,(
  ( ! [X30,X29] : (~sP10('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X29,X30) | p(s(bool,X30)) | p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f2589,f239])).
fof(f2589,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | p(s(bool,X30)) | ~sP10('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X29,X30)) )),
  inference(superposition,[],[f389,f1128])).
fof(f2623,plain,(
  ( ! [X28,X27] : (~sP9(X28,X27,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | p(s(bool,X28)) | ~p(s(bool,X27))) )),
  inference(subsumption_resolution,[],[f2588,f239])).
fof(f2588,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | p(s(bool,X28)) | ~sP9(X28,X27,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(superposition,[],[f381,f1128])).
fof(f2622,plain,(
  ( ! [X26,X25] : (~sP8(X25,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X26) | p(s(bool,X26)) | ~p(s(bool,X25))) )),
  inference(subsumption_resolution,[],[f2587,f239])).
fof(f2587,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP8(X25,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X26)) )),
  inference(superposition,[],[f372,f1128])).
fof(f2621,plain,(
  ( ! [X24,X23] : (~sP5(X23,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X24) | p(s(bool,X24)) | ~p(s(bool,X23))) )),
  inference(subsumption_resolution,[],[f2586,f239])).
fof(f2586,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP5(X23,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X24)) )),
  inference(superposition,[],[f353,f1128])).
fof(f2620,plain,(
  ( ! [X21,X22] : (~sP6(X21,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X22) | p(s(bool,X22)) | ~p(s(bool,X21))) )),
  inference(subsumption_resolution,[],[f2585,f239])).
fof(f2585,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | p(s(bool,X22)) | ~sP6(X21,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X22)) )),
  inference(superposition,[],[f349,f1128])).
fof(f2619,plain,(
  ( ! [X19,X20] : (sP7(X19,X20,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~sP6(X19,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X20) | ~sP5(X20,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X19)) )),
  inference(subsumption_resolution,[],[f2618,f1196])).
fof(f2618,plain,(
  ( ! [X19,X20] : (sP7(X19,X20,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~sP6(X19,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X20) | ~sP5(X20,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X19) | ~sP4('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X20,X19)) )),
  inference(subsumption_resolution,[],[f2584,f239])).
fof(f2584,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | sP7(X19,X20,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~sP6(X19,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X20) | ~sP5(X20,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X19) | ~sP4('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X20,X19)) )),
  inference(superposition,[],[f346,f1128])).
fof(f2617,plain,(
  ( ! [X17,X18] : (sP3(X18,X17,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~p(s(bool,X17))) )),
  inference(subsumption_resolution,[],[f2583,f239])).
fof(f2583,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | sP3(X18,X17,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(superposition,[],[f334,f1128])).
fof(f2616,plain,(
  ( ! [X15,X16] : (sP2(X15,X16,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | ~p(s(bool,X15))) )),
  inference(subsumption_resolution,[],[f2582,f239])).
fof(f2582,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | sP2(X15,X16,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(superposition,[],[f327,f1128])).
fof(f2615,plain,(
  ( ! [X10] : (sP1(X10,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))) | p(s(bool,X10))) )),
  inference(subsumption_resolution,[],[f2579,f239])).
fof(f2579,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | p(s(bool,X10)) | sP1(X10,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(superposition,[],[f316,f1128])).
fof(f2614,plain,(
  ( ! [X9] : (sP1('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X9) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f2578,f239])).
fof(f2578,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | sP1('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X9)) )),
  inference(superposition,[],[f313,f1128])).
fof(f2564,plain,(
  sP7('T_0','T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f1128,f362])).
fof(f2372,plain,(
  sP7('F_0','F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f433,f436,f1472,f347])).
fof(f2359,plain,(
  sP7('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f433,f435,f1471,f348])).
fof(f2349,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f432,f436,f1470,f347])).
fof(f2335,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f439,f431,f1469,f346])).
fof(f2339,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f439,f1178,f1469,f346])).
fof(f2340,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f439,f1454,f1469,f346])).
fof(f2332,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1462,f343])).
fof(f2327,plain,(
  ~sP7('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1455,f343])).
fof(f2311,plain,(
  sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f439,f434,f1454,f346])).
fof(f2313,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f439,f467,f1454,f346])).
fof(f2315,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f439,f1193,f1454,f346])).
fof(f2317,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f439,f454,f1454,f346])).
fof(f2306,plain,(
  ~sP7('T_0','F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1447,f344])).
fof(f2119,plain,(
  sP7('F_0','F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f433,f436,f1196,f347])).
fof(f2107,plain,(
  sP7('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f433,f435,f1195,f348])).
fof(f2098,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f432,f436,f1194,f347])).
fof(f2086,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f439,f431,f1193,f346])).
fof(f2089,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f439,f1178,f1193,f346])).
fof(f2083,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1186,f343])).
fof(f2078,plain,(
  ~sP7('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1179,f343])).
fof(f2064,plain,(
  sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f439,f434,f1178,f346])).
fof(f2066,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f439,f454,f1178,f346])).
fof(f2068,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f439,f467,f1178,f346])).
fof(f2059,plain,(
  ~sP7('T_0','F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f1171,f344])).
fof(f2014,plain,(
  ~sP12('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f2009,f405])).
fof(f2009,plain,(
  ~sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f1117])).
fof(f2000,plain,(
  ~sP13('T_0','F_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f1988,f401])).
fof(f1988,plain,(
  ~sP12('T_0','F_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f426,f239,f1085])).
fof(f1976,plain,(
  ~sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1960,f384])).
fof(f1960,plain,(
  ~sP10('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f426,f948])).
fof(f1950,plain,(
  ~sP9('F_0','T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f1938,f365])).
fof(f1938,plain,(
  ~sP8('T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f871])).
fof(f1924,plain,(
  ~sP7('T_0','F_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f1912,f344])).
fof(f1912,plain,(
  ~sP6('T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f806])).
fof(f1904,plain,(
  ( ! [X0,X1] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f239,f1020])).
fof(f1811,plain,(
  ( ! [X0,X1] : (~sP7('F_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f426,f281,f364])).
fof(f1812,plain,(
  ( ! [X0,X1] : (sP7('T_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f239,f281,f362])).
fof(f1813,plain,(
  ( ! [X0,X1] : (~sP7('F_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f426,f281,f364])).
fof(f1814,plain,(
  ( ! [X0,X1] : (sP7('T_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f239,f281,f362])).
fof(f1773,plain,(
  ~sP5('T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f470,f780,f431,f426,f348])).
fof(f1776,plain,(
  ( ! [X0] : (~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f438,f971,f431,f426,f348])).
fof(f1783,plain,(
  sP7('F_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f469,f435,f432,f426,f348])).
fof(f1784,plain,(
  ( ! [X0] : (sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f456,f435,f432,f426,f348])).
fof(f1740,plain,(
  sP7('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f437,f434,f432,f426,f347])).
fof(f1741,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f455,f454,f432,f426,f347])).
fof(f1742,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f468,f467,f432,f426,f347])).
fof(f1745,plain,(
  sP7('F_0','F_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f470,f436,f433,f426,f347])).
fof(f1746,plain,(
  ( ! [X0] : (sP7('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f457,f436,f433,f426,f347])).
fof(f1722,plain,(
  ( ! [X10,X8,X9] : (sP7(X9,X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8))) | ~sP6(X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X10) | ~sP5(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X9)) )),
  inference(subsumption_resolution,[],[f1721,f457])).
fof(f1721,plain,(
  ( ! [X10,X8,X9] : (sP7(X9,X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8))) | ~sP6(X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X10) | ~sP5(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X9) | ~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X10,X9)) )),
  inference(subsumption_resolution,[],[f1718,f239])).
fof(f1718,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | sP7(X9,X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8))) | ~sP6(X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X10) | ~sP5(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X9) | ~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X10,X9)) )),
  inference(superposition,[],[f346,f480])).
fof(f1720,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | ~sP6(X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7) | ~sP5(X7,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X6)) )),
  inference(subsumption_resolution,[],[f1719,f470])).
fof(f1719,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | ~sP6(X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7) | ~sP5(X7,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X6) | ~sP4('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7,X6)) )),
  inference(subsumption_resolution,[],[f1717,f239])).
fof(f1717,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP7(X6,X7,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | ~sP6(X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7) | ~sP5(X7,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X6) | ~sP4('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7,X6)) )),
  inference(superposition,[],[f346,f481])).
fof(f1688,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f439,f434,f431,f239,f346])).
fof(f1691,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f455,f454,f431,f239,f346])).
fof(f1692,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f439,f467,f431,f239,f346])).
fof(f1698,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f455,f454,f453,f239,f346])).
fof(f1701,plain,(
  ( ! [X0] : (sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f468,f467,f453,f239,f346])).
fof(f1705,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f455,f454,f466,f239,f346])).
fof(f1706,plain,(
  sP7('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f439,f467,f466,f239,f346])).
fof(f1710,plain,(
  ~sP5('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f439,f433,f778,f239,f346])).
fof(f1714,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f456,f453,f969,f239,f346])).
fof(f1686,plain,(
  ( ! [X0,X1] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f1019])).
fof(f1678,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f239,f1018])).
fof(f1404,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f237,f288])).
fof(f1406,plain,(
  ( ! [X0] : (sP0(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f297])).
fof(f1407,plain,(
  ( ! [X0] : (sP0('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f237,f298])).
fof(f1408,plain,(
  ~sP1('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f237,f312])).
fof(f1410,plain,(
  ~sP1('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f237,f312])).
fof(f1411,plain,(
  ~sP1('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f237,f237,f312])).
fof(f1412,plain,(
  sP1('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f237,f313])).
fof(f1413,plain,(
  sP1('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f237,f316])).
fof(f1424,plain,(
  ( ! [X0] : (sP2('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f237,f327])).
fof(f1427,plain,(
  ( ! [X0] : (sP2('T_0',X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f239,f237,f327])).
fof(f1428,plain,(
  ( ! [X0] : (sP2('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f237,f327])).
fof(f1430,plain,(
  ( ! [X0] : (sP3(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f237,f334])).
fof(f1433,plain,(
  ( ! [X0] : (sP3(X0,'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f239,f237,f334])).
fof(f1434,plain,(
  ( ! [X0] : (sP3(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f237,f334])).
fof(f1438,plain,(
  ~sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f237,f345])).
fof(f1439,plain,(
  ~sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f237,f426,f237,f345])).
fof(f1440,plain,(
  ~sP6('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f237,f349])).
fof(f1447,plain,(
  ~sP6('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f237,f349])).
fof(f1448,plain,(
  ~sP6('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f237,f426,f237,f349])).
fof(f1454,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f350])).
fof(f1455,plain,(
  ~sP5('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f237,f353])).
fof(f1462,plain,(
  ~sP5('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f237,f353])).
fof(f1463,plain,(
  ~sP5('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f237,f426,f237,f353])).
fof(f1469,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f354])).
fof(f1470,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f358])).
fof(f1471,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f237,f359])).
fof(f1472,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f237,f360])).
fof(f1473,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f481,f237,f362])).
fof(f1474,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f480,f237,f362])).
fof(f1475,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f481,f237,f362])).
fof(f1476,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f480,f237,f362])).
fof(f1477,plain,(
  ~sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f477,f237,f363])).
fof(f1478,plain,(
  ~sP7('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f477,f237,f363])).
fof(f1479,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f237,f366])).
fof(f1480,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f426,f237,f367])).
fof(f1483,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f458,f237,f368])).
fof(f1485,plain,(
  sP9('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f471,f237,f368])).
fof(f1490,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f239,f458,f237,f368])).
fof(f1491,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f458,f237,f368])).
fof(f1492,plain,(
  sP9('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f471,f237,f368])).
fof(f1493,plain,(
  sP9('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f237,f471,f237,f368])).
fof(f1495,plain,(
  ~sP8('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f237,f372])).
fof(f1502,plain,(
  ~sP8('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f237,f372])).
fof(f1503,plain,(
  ~sP8('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f237,f426,f237,f372])).
fof(f1509,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f373])).
fof(f1510,plain,(
  sP9('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f239,f237,f378])).
fof(f1515,plain,(
  sP9('T_0','T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f239,f237,f378])).
fof(f1517,plain,(
  sP9('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f237,f237,f378])).
fof(f1520,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f239,f237,f378])).
fof(f1521,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f237,f239,f237,f378])).
fof(f1522,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f237,f237,f378])).
fof(f1523,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f237,f237,f237,f378])).
fof(f1524,plain,(
  ~sP9('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f237,f381])).
fof(f1527,plain,(
  ~sP9('F_0','T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f426,f237,f381])).
fof(f1528,plain,(
  ~sP9('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f237,f426,f237,f381])).
fof(f1530,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f426,f237,f382])).
fof(f1535,plain,(
  ( ! [X0] : (~sP11('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f426,f237,f383])).
fof(f1546,plain,(
  ~sP10('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f426,f237,f389])).
fof(f1553,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f390])).
fof(f1554,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f237,f391])).
fof(f1555,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f239,f237,f394])).
fof(f1559,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f237,f394])).
fof(f1560,plain,(
  ( ! [X0] : (sP11('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f239,f237,f395])).
fof(f1564,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f237,f237,f395])).
fof(f1565,plain,(
  ~sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f426,f237,f396])).
fof(f1566,plain,(
  ( ! [X0] : (~sP13('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f426,f237,f400])).
fof(f1587,plain,(
  ~sP12('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f426,f239,f237,f406])).
fof(f1594,plain,(
  ~sP12('T_0','F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f239,f237,f406])).
fof(f1595,plain,(
  ~sP12('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f237,f237,f406])).
fof(f1601,plain,(
  ( ! [X0,X1] : (sP12(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f237,f408])).
fof(f1602,plain,(
  sP13('F_0','F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f426,f237,f410])).
fof(f1605,plain,(
  ( ! [X0] : (sP13('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f239,f237,f412])).
fof(f1609,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f237,f237,f412])).
fof(f1610,plain,(
  ~sP13('T_0','F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f426,f237,f413])).
fof(f1615,plain,(
  ~sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f237,f413])).
fof(f1616,plain,(
  ~sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f237,f426,f237,f413])).
fof(f1617,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f446,f237,f416])).
fof(f1618,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f447,f237,f416])).
fof(f1619,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f474,f237,f416])).
fof(f1620,plain,(
  ( ! [X0,X1] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f461,f237,f416])).
fof(f1621,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f443,f237,f417])).
fof(f1622,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f459,f237,f417])).
fof(f1623,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f472,f237,f417])).
fof(f1624,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f444,f237,f417])).
fof(f1625,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f473,f237,f417])).
fof(f1626,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f460,f237,f417])).
fof(f1627,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f237,f419])).
fof(f1628,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f237,f420])).
fof(f1629,plain,(
  ( ! [X0] : (sP2('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f237,f715])).
fof(f1630,plain,(
  ( ! [X0] : (sP3(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f237,f716])).
fof(f1631,plain,(
  ( ! [X0] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f717])).
fof(f1632,plain,(
  ( ! [X0] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f237,f718])).
fof(f1633,plain,(
  ( ! [X0] : (sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f237,f720])).
fof(f1634,plain,(
  ( ! [X0,X1] : (sP2('const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f237,f1010])).
fof(f1635,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f237,f1011])).
fof(f1636,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f237,f1017])).
fof(f1399,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f1017])).
fof(f1128,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f236,f288])).
fof(f1130,plain,(
  ( ! [X0] : (sP0(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f236,f297])).
fof(f1131,plain,(
  ( ! [X0] : (sP0('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f236,f298])).
fof(f1132,plain,(
  ~sP1('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f236,f312])).
fof(f1134,plain,(
  ~sP1('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f236,f312])).
fof(f1135,plain,(
  ~sP1('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f236,f236,f312])).
fof(f1136,plain,(
  sP1('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f236,f313])).
fof(f1137,plain,(
  sP1('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f236,f316])).
fof(f1148,plain,(
  ( ! [X0] : (sP2('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f236,f327])).
fof(f1151,plain,(
  ( ! [X0] : (sP2('T_0',X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f239,f236,f327])).
fof(f1152,plain,(
  ( ! [X0] : (sP2('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f236,f236,f327])).
fof(f1154,plain,(
  ( ! [X0] : (sP3(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f236,f334])).
fof(f1157,plain,(
  ( ! [X0] : (sP3(X0,'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f239,f236,f334])).
fof(f1158,plain,(
  ( ! [X0] : (sP3(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f236,f236,f334])).
fof(f1162,plain,(
  ~sP7('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f236,f345])).
fof(f1163,plain,(
  ~sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f236,f426,f236,f345])).
fof(f1164,plain,(
  ~sP6('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f236,f349])).
fof(f1171,plain,(
  ~sP6('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f236,f349])).
fof(f1172,plain,(
  ~sP6('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f236,f426,f236,f349])).
fof(f1178,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f236,f350])).
fof(f1179,plain,(
  ~sP5('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f236,f353])).
fof(f1186,plain,(
  ~sP5('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f236,f353])).
fof(f1187,plain,(
  ~sP5('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f236,f426,f236,f353])).
fof(f1193,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f236,f354])).
fof(f1194,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f236,f358])).
fof(f1195,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f236,f359])).
fof(f1196,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f236,f360])).
fof(f1197,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f481,f236,f362])).
fof(f1198,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f480,f236,f362])).
fof(f1199,plain,(
  sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f481,f236,f362])).
fof(f1200,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f480,f236,f362])).
fof(f1201,plain,(
  ~sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f477,f236,f363])).
fof(f1202,plain,(
  ~sP7('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f477,f236,f363])).
fof(f1203,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f236,f366])).
fof(f1204,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f426,f236,f367])).
fof(f1207,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f458,f236,f368])).
fof(f1209,plain,(
  sP9('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f471,f236,f368])).
fof(f1214,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f239,f458,f236,f368])).
fof(f1215,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f236,f458,f236,f368])).
fof(f1216,plain,(
  sP9('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f471,f236,f368])).
fof(f1217,plain,(
  sP9('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f236,f471,f236,f368])).
fof(f1219,plain,(
  ~sP8('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f236,f372])).
fof(f1226,plain,(
  ~sP8('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f236,f372])).
fof(f1227,plain,(
  ~sP8('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0')),
  inference(unit_resulting_resolution,[],[f236,f426,f236,f372])).
fof(f1233,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f236,f373])).
fof(f1234,plain,(
  sP9('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f239,f236,f378])).
fof(f1239,plain,(
  sP9('T_0','T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f239,f236,f378])).
fof(f1241,plain,(
  sP9('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f236,f236,f378])).
fof(f1244,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f239,f236,f378])).
fof(f1245,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f236,f239,f236,f378])).
fof(f1246,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f236,f236,f378])).
fof(f1247,plain,(
  sP9('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f236,f236,f236,f378])).
fof(f1248,plain,(
  ~sP9('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f236,f381])).
fof(f1251,plain,(
  ~sP9('F_0','T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f426,f236,f381])).
fof(f1252,plain,(
  ~sP9('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f236,f426,f236,f381])).
fof(f1254,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f426,f236,f382])).
fof(f1259,plain,(
  ( ! [X0] : (~sP11('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f426,f236,f383])).
fof(f1270,plain,(
  ~sP10('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f426,f236,f389])).
fof(f1277,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f236,f390])).
fof(f1278,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f236,f391])).
fof(f1279,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f239,f236,f394])).
fof(f1283,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f236,f236,f394])).
fof(f1284,plain,(
  ( ! [X0] : (sP11('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f239,f236,f395])).
fof(f1288,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f236,f236,f395])).
fof(f1289,plain,(
  ~sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f426,f236,f396])).
fof(f1290,plain,(
  ( ! [X0] : (~sP13('F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f426,f236,f400])).
fof(f1311,plain,(
  ~sP12('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f426,f239,f236,f406])).
fof(f1318,plain,(
  ~sP12('T_0','F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f239,f236,f406])).
fof(f1319,plain,(
  ~sP12('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f236,f236,f406])).
fof(f1325,plain,(
  ( ! [X0,X1] : (sP12(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f236,f408])).
fof(f1326,plain,(
  sP13('F_0','F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f426,f426,f236,f410])).
fof(f1329,plain,(
  ( ! [X0] : (sP13('T_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f239,f236,f412])).
fof(f1333,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f236,f236,f412])).
fof(f1334,plain,(
  ~sP13('T_0','F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f239,f426,f236,f413])).
fof(f1339,plain,(
  ~sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f236,f413])).
fof(f1340,plain,(
  ~sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'F_0','const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))),
  inference(unit_resulting_resolution,[],[f236,f426,f236,f413])).
fof(f1341,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f446,f236,f416])).
fof(f1342,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f447,f236,f416])).
fof(f1343,plain,(
  ( ! [X0] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f474,f236,f416])).
fof(f1344,plain,(
  ( ! [X0,X1] : (sP13('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f461,f236,f416])).
fof(f1345,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f443,f236,f417])).
fof(f1346,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f459,f236,f417])).
fof(f1347,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f472,f236,f417])).
fof(f1348,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f444,f236,f417])).
fof(f1349,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f473,f236,f417])).
fof(f1350,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f460,f236,f417])).
fof(f1351,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f236,f419])).
fof(f1352,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X1)) )),
  inference(unit_resulting_resolution,[],[f236,f420])).
fof(f1353,plain,(
  ( ! [X0] : (sP2('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f236,f715])).
fof(f1354,plain,(
  ( ! [X0] : (sP3(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f236,f716])).
fof(f1355,plain,(
  ( ! [X0] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) )),
  inference(unit_resulting_resolution,[],[f236,f717])).
fof(f1356,plain,(
  ( ! [X0] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f236,f718])).
fof(f1357,plain,(
  ( ! [X0] : (sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0)) )),
  inference(unit_resulting_resolution,[],[f236,f720])).
fof(f1358,plain,(
  ( ! [X0,X1] : (sP2('const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f236,f1010])).
fof(f1359,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f236,f1011])).
fof(f1122,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f239,f1011])).
fof(f1118,plain,(
  ( ! [X10,X8,X9] : (~sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X10,X9) | p(s(bool,X10)) | ~p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f1114,f239])).
fof(f1114,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | p(s(bool,X10)) | ~sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X10,X9)) )),
  inference(superposition,[],[f413,f480])).
fof(f1117,plain,(
  ( ! [X6,X7] : (~sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7,X6) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1113,f239])).
fof(f1113,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7,X6)) )),
  inference(superposition,[],[f413,f481])).
fof(f1103,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f239,f1010])).
fof(f1090,plain,(
  ~sP13('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1078,f401])).
fof(f1086,plain,(
  ( ! [X10,X8,X9] : (~sP12(X10,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8))) | ~p(s(bool,X10)) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f1082,f239])).
fof(f1082,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | ~p(s(bool,X10)) | ~sP12(X10,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)))) )),
  inference(superposition,[],[f406,f480])).
fof(f1085,plain,(
  ( ! [X6,X7] : (~sP12(X7,X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | ~p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1081,f239])).
fof(f1081,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | ~p(s(bool,X7)) | ~sP12(X7,X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(superposition,[],[f406,f481])).
fof(f1078,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f426,f239,f406])).
fof(f1072,plain,(
  ( ! [X0] : (sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f426,f1009])).
fof(f1068,plain,(
  sP13('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f448,f239,f426,f405])).
fof(f1063,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f1008])).
fof(f1040,plain,(
  ( ! [X10,X8,X9] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X10,X9) | p(s(bool,X10)) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f1034,f239])).
fof(f1034,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | p(s(bool,X10)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X8)),X10,X9)) )),
  inference(superposition,[],[f396,f480])).
fof(f1039,plain,(
  ( ! [X6,X7] : (~sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7,X6) | p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1033,f239])).
fof(f1033,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7,X6)) )),
  inference(superposition,[],[f396,f481])).
fof(f1022,plain,(
  ( ! [X61,X62,X63] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal',X61)),X62,X63) | ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal',X61)),X62,X63)) )),
  inference(subsumption_resolution,[],[f1000,f239])).
fof(f1000,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'T_0')) | ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal',X61)),X62,X63) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal',X61)),X62,X63)) )),
  inference(superposition,[],[f417,f480])).
fof(f1021,plain,(
  ( ! [X59,X60,X58] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X58),s('type.extreal.extreal',X58)),X59,X60) | ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',X58),s('type.extreal.extreal',X58)),X59,X60)) )),
  inference(subsumption_resolution,[],[f999,f239])).
fof(f999,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',X58),s('type.extreal.extreal',X58)),X59,X60) | sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X58),s('type.extreal.extreal',X58)),X59,X60)) )),
  inference(superposition,[],[f416,f480])).
fof(f1020,plain,(
  ( ! [X57,X56,X55] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X55),s('type.extreal.extreal',X55)),X56,X57) | ~p(s(bool,X56))) )),
  inference(subsumption_resolution,[],[f998,f239])).
fof(f998,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X55),s('type.extreal.extreal',X55)),X56,X57)) )),
  inference(superposition,[],[f412,f480])).
fof(f1019,plain,(
  ( ! [X54,X52,X53] : (sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X52),s('type.extreal.extreal',X52)),X54,X53) | p(s(bool,X53))) )),
  inference(subsumption_resolution,[],[f997,f239])).
fof(f997,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | sP13('const.extreal.extreal_le_2'(s('type.extreal.extreal',X52),s('type.extreal.extreal',X52)),X54,X53)) )),
  inference(superposition,[],[f411,f480])).
fof(f1018,plain,(
  ( ! [X50,X51,X49] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X49),s('type.extreal.extreal',X49)),X50,X51) | ~p(s(bool,X50))) )),
  inference(subsumption_resolution,[],[f996,f239])).
fof(f996,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X49),s('type.extreal.extreal',X49)),X50,X51)) )),
  inference(superposition,[],[f395,f480])).
fof(f1017,plain,(
  ( ! [X47,X48,X46] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X46),s('type.extreal.extreal',X46)),X48,X47) | ~p(s(bool,X47))) )),
  inference(subsumption_resolution,[],[f995,f239])).
fof(f995,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X47)) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X46),s('type.extreal.extreal',X46)),X48,X47)) )),
  inference(superposition,[],[f394,f480])).
fof(f1016,plain,(
  ( ! [X45,X43,X44] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X43),s('type.extreal.extreal',X43)),X44,X45) | p(s(bool,X45)) | p(s(bool,X44))) )),
  inference(subsumption_resolution,[],[f994,f239])).
fof(f994,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | p(s(bool,X44)) | p(s(bool,X45)) | ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X43),s('type.extreal.extreal',X43)),X44,X45)) )),
  inference(superposition,[],[f389,f480])).
fof(f1015,plain,(
  ( ! [X41,X42,X40] : (~sP9(X42,X41,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X40),s('type.extreal.extreal',X40))) | p(s(bool,X42)) | ~p(s(bool,X41))) )),
  inference(subsumption_resolution,[],[f993,f239])).
fof(f993,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP9(X42,X41,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X40),s('type.extreal.extreal',X40)))) )),
  inference(superposition,[],[f381,f480])).
fof(f1014,plain,(
  ( ! [X39,X37,X38] : (~sP8(X38,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal',X37)),X39) | p(s(bool,X39)) | ~p(s(bool,X38))) )),
  inference(subsumption_resolution,[],[f992,f239])).
fof(f992,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~sP8(X38,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal',X37)),X39)) )),
  inference(superposition,[],[f372,f480])).
fof(f1013,plain,(
  ( ! [X35,X36,X34] : (~sP5(X35,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X34),s('type.extreal.extreal',X34)),X36) | p(s(bool,X36)) | ~p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f991,f239])).
fof(f991,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP5(X35,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X34),s('type.extreal.extreal',X34)),X36)) )),
  inference(superposition,[],[f353,f480])).
fof(f1012,plain,(
  ( ! [X33,X31,X32] : (~sP6(X32,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X31),s('type.extreal.extreal',X31)),X33) | p(s(bool,X33)) | ~p(s(bool,X32))) )),
  inference(subsumption_resolution,[],[f990,f239])).
fof(f990,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP6(X32,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X31),s('type.extreal.extreal',X31)),X33)) )),
  inference(superposition,[],[f349,f480])).
fof(f1011,plain,(
  ( ! [X30,X28,X29] : (sP3(X30,X29,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X28),s('type.extreal.extreal',X28))) | ~p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f989,f239])).
fof(f989,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | sP3(X30,X29,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X28),s('type.extreal.extreal',X28)))) )),
  inference(superposition,[],[f334,f480])).
fof(f1010,plain,(
  ( ! [X26,X27,X25] : (sP2(X26,X27,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X25),s('type.extreal.extreal',X25))) | ~p(s(bool,X26))) )),
  inference(subsumption_resolution,[],[f988,f239])).
fof(f988,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | sP2(X26,X27,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X25),s('type.extreal.extreal',X25)))) )),
  inference(superposition,[],[f327,f480])).
fof(f1009,plain,(
  ( ! [X17,X18] : (sP1(X18,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X17),s('type.extreal.extreal',X17))) | p(s(bool,X18))) )),
  inference(subsumption_resolution,[],[f985,f239])).
fof(f985,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP1(X18,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X17),s('type.extreal.extreal',X17)))) )),
  inference(superposition,[],[f316,f480])).
fof(f1008,plain,(
  ( ! [X15,X16] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal',X15)),X16) | p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f984,f239])).
fof(f984,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal',X15)),X16)) )),
  inference(superposition,[],[f313,f480])).
fof(f969,plain,(
  ( ! [X0] : (~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f426,f480,f364])).
fof(f970,plain,(
  ( ! [X0] : (sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f480,f362])).
fof(f971,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f426,f480,f364])).
fof(f972,plain,(
  ( ! [X0] : (sP7('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f239,f480,f362])).
fof(f952,plain,(
  ~sP11('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f940,f384])).
fof(f948,plain,(
  ( ! [X6,X7] : (~sP10('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X6,X7) | p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f943,f239])).
fof(f943,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP10('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X6,X7)) )),
  inference(superposition,[],[f389,f481])).
fof(f940,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f426,f239,f389])).
fof(f932,plain,(
  ( ! [X0] : (sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f239,f720])).
fof(f928,plain,(
  sP11('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f445,f426,f426,f388])).
fof(f920,plain,(
  ( ! [X0] : (sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f719])).
fof(f910,plain,(
  ( ! [X0] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f239,f718])).
fof(f900,plain,(
  ( ! [X0] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f717])).
fof(f899,plain,(
  ( ! [X6,X7] : (~sP9(X7,X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f896,f239])).
fof(f896,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP9(X7,X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(superposition,[],[f381,f481])).
fof(f875,plain,(
  ~sP9('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f865,f365])).
fof(f871,plain,(
  ( ! [X6,X7] : (~sP8(X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f868,f239])).
fof(f868,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP8(X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7)) )),
  inference(superposition,[],[f372,f481])).
fof(f865,plain,(
  ~sP8('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f239,f239,f372])).
fof(f853,plain,(
  sP9('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f440,f239,f239,f368])).
fof(f854,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f458,f239,f239,f368])).
fof(f855,plain,(
  sP9('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f471,f239,f239,f368])).
fof(f849,plain,(
  ~sP7('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f836,f342])).
fof(f836,plain,(
  ~sP4('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f426,f426,f357])).
fof(f828,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f818,f343])).
fof(f824,plain,(
  ( ! [X6,X7] : (~sP5(X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f821,f239])).
fof(f821,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP5(X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7)) )),
  inference(superposition,[],[f353,f481])).
fof(f818,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f239,f239,f353])).
fof(f811,plain,(
  ( ! [X0] : (sP3(X0,'T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f239,f716])).
fof(f806,plain,(
  ( ! [X6,X7] : (~sP6(X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f803,f239])).
fof(f803,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP6(X6,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X7)) )),
  inference(superposition,[],[f349,f481])).
fof(f800,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f239,f239,f349])).
fof(f790,plain,(
  ( ! [X0] : (sP2('T_0',X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f239,f715])).
fof(f778,plain,(
  ~sP7('F_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f426,f481,f364])).
fof(f780,plain,(
  ~sP7('F_0','T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f426,f481,f364])).
fof(f769,plain,(
  ~sP7('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f477,f363])).
fof(f771,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f239,f477,f363])).
fof(f758,plain,(
  sP7('T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f239,f481,f362])).
fof(f759,plain,(
  sP7('T_0','T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f239,f481,f362])).
fof(f750,plain,(
  sP7('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f426,f477,f361])).
fof(f751,plain,(
  sP7('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f477,f361])).
fof(f747,plain,(
  sP1('F_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f426,f714])).
fof(f734,plain,(
  sP1('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f426,f713])).
fof(f733,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f728,f426])).
fof(f728,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f285,f247])).
fof(f722,plain,(
  ( ! [X23,X22] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X22,X23) | ~sP10('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X22,X23)) )),
  inference(subsumption_resolution,[],[f707,f239])).
fof(f707,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | ~sP10('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X22,X23) | sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X22,X23)) )),
  inference(superposition,[],[f417,f481])).
fof(f721,plain,(
  ( ! [X21,X20] : (sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X20,X21) | ~sP12('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X20,X21)) )),
  inference(subsumption_resolution,[],[f706,f239])).
fof(f706,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | ~sP12('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X20,X21) | sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X20,X21)) )),
  inference(superposition,[],[f416,f481])).
fof(f720,plain,(
  ( ! [X19,X18] : (sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X18,X19) | ~p(s(bool,X18))) )),
  inference(subsumption_resolution,[],[f705,f239])).
fof(f705,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X18,X19)) )),
  inference(superposition,[],[f412,f481])).
fof(f719,plain,(
  ( ! [X17,X16] : (sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X17,X16) | p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f704,f239])).
fof(f704,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | sP13('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X17,X16)) )),
  inference(superposition,[],[f411,f481])).
fof(f718,plain,(
  ( ! [X14,X15] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X14,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f703,f239])).
fof(f703,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X14,X15)) )),
  inference(superposition,[],[f395,f481])).
fof(f717,plain,(
  ( ! [X12,X13] : (sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X13,X12) | ~p(s(bool,X12))) )),
  inference(subsumption_resolution,[],[f702,f239])).
fof(f702,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X12)) | sP11('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X13,X12)) )),
  inference(superposition,[],[f394,f481])).
fof(f716,plain,(
  ( ! [X10,X11] : (sP3(X11,X10,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | ~p(s(bool,X10))) )),
  inference(subsumption_resolution,[],[f701,f239])).
fof(f701,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | sP3(X11,X10,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(superposition,[],[f334,f481])).
fof(f715,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | ~p(s(bool,X8))) )),
  inference(subsumption_resolution,[],[f700,f239])).
fof(f700,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | sP2(X8,X9,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(superposition,[],[f327,f481])).
fof(f714,plain,(
  ( ! [X3] : (sP1(X3,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f697,f239])).
fof(f697,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | sP1(X3,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(superposition,[],[f316,f481])).
fof(f713,plain,(
  ( ! [X2] : (sP1('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f696,f239])).
fof(f696,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X2)) )),
  inference(superposition,[],[f313,f481])).
fof(f639,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f247,f247])).
fof(f620,plain,(
  ( ! [X0,X1] : (sP11('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f459,f239,f417])).
fof(f621,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f472,f239,f417])).
fof(f623,plain,(
  ( ! [X0] : (sP11('T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f473,f239,f417])).
fof(f624,plain,(
  ( ! [X0,X1] : (sP11('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f460,f239,f417])).
fof(f603,plain,(
  ( ! [X0] : (sP13('T_0','const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f474,f239,f416])).
fof(f604,plain,(
  ( ! [X0,X1] : (sP13('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f461,f239,f416])).
fof(f593,plain,(
  ( ! [X0] : (sP13('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f239,f239,f412])).
fof(f587,plain,(
  ( ! [X0] : (sP13('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f239,f411])).
fof(f584,plain,(
  ( ! [X0] : (~sP13('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f239,f426,f400])).
fof(f583,plain,(
  ( ! [X0] : (~sP13('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f426,f399])).
fof(f579,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f239,f239,f395])).
fof(f572,plain,(
  ( ! [X0] : (sP11('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f239,f394])).
fof(f570,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f239,f426,f383])).
fof(f569,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f426,f382])).
fof(f562,plain,(
  ( ! [X0] : (sP9('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f426,f426,f377])).
fof(f559,plain,(
  ( ! [X0] : (sP9('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f426,f376])).
fof(f556,plain,(
  ( ! [X0] : (~sP9('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f239,f426,f367])).
fof(f551,plain,(
  ( ! [X0] : (~sP9('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f239,f426,f366])).
fof(f544,plain,(
  ( ! [X0] : (sP3(X0,'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f239,f334])).
fof(f539,plain,(
  ( ! [X0] : (~sP3('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f426,f426,f333])).
fof(f534,plain,(
  ( ! [X0] : (~sP3('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f426,f332])).
fof(f529,plain,(
  ( ! [X0] : (sP2('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f239,f327])).
fof(f520,plain,(
  ( ! [X0] : (~sP2('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f426,f426,f323])).
fof(f515,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f426,f322])).
fof(f506,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f426,f428])).
fof(f497,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f426,f239,f316])).
fof(f492,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f239,f313])).
fof(f489,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f239,f239,f312])).
fof(f488,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f426,f426,f311])).
fof(f480,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f246,f255])).
fof(f481,plain,(
  s(bool,'T_0') = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))),
  inference(unit_resulting_resolution,[],[f235,f255])).
fof(f477,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f426,f254])).
fof(f464,plain,(
  ( ! [X0] : (sP0(X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f235,f297])).
fof(f465,plain,(
  ( ! [X0] : (sP0('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f235,f298])).
fof(f466,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f235,f350])).
fof(f467,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f235,f354])).
fof(f468,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f235,f358])).
fof(f469,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f235,f359])).
fof(f470,plain,(
  ( ! [X0,X1] : (sP4('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f235,f360])).
fof(f471,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f235,f373])).
fof(f472,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) )),
  inference(unit_resulting_resolution,[],[f235,f390])).
fof(f473,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f235,f391])).
fof(f474,plain,(
  ( ! [X0,X1] : (sP12(X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f235,f408])).
fof(f475,plain,(
  ( ! [X0,X1] : (sP3('const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f235,f419])).
fof(f476,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f235,f420])).
fof(f451,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f246,f297])).
fof(f452,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f246,f298])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f246,f350])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f246,f354])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f246,f358])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f246,f359])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f246,f360])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f246,f373])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f246,f390])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f246,f391])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f246,f408])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f246,f419])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f246,f420])).
fof(f450,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f239,f420])).
fof(f449,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f239,f419])).
fof(f448,plain,(
  ( ! [X0,X1] : (sP12('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f426,f409])).
fof(f447,plain,(
  ( ! [X0,X1] : (sP12(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f239,f408])).
fof(f446,plain,(
  ( ! [X0,X1] : (sP12(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f426,f407])).
fof(f445,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f426,f392])).
fof(f444,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f239,f391])).
fof(f443,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f390])).
fof(f442,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f426,f375])).
fof(f441,plain,(
  ( ! [X0,X1] : (sP8(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f426,f374])).
fof(f440,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f373])).
fof(f439,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f239,f360])).
fof(f438,plain,(
  ( ! [X0,X1] : (sP4(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f239,f359])).
fof(f437,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f358])).
fof(f436,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f426,f356])).
fof(f435,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f426,f355])).
fof(f434,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f354])).
fof(f433,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f426,f352])).
fof(f432,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f426,f351])).
fof(f431,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f350])).
fof(f430,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f239,f298])).
fof(f429,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f239,f297])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP13(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP13(X2,X0,X1)) & (~sP13(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP13(X2,X0,X1)) & (~sP13(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP13(X2,X0,X1))),
  inference(definition_folding,[],[f134,f155,f154])).
fof(f154,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f155,plain,(
  ! [X2,X0,X1] : (sP13(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP12(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f134,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) => p(s(bool,X20)))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (~p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.lebesgue.psfis_unique','thm.sat.sat_27_prove_ltb')).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP13(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP13(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP13(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP12(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f231])).
fof(f231,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f154])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f232])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f232])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f232])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (~sP13(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((~sP13(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP12(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP12(X0,X1,X2) | sP13(X0,X1,X2)))),
  inference(rectify,[],[f228])).
fof(f228,plain,(
  ! [X2,X0,X1] : ((~sP13(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP12(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP12(X2,X0,X1) | sP13(X2,X0,X1)))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : ((~sP13(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP12(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP12(X2,X0,X1)) | sP13(X2,X0,X1)))),
  inference(nnf_transformation,[],[f155])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (~sP13(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f229])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~sP13(X0,X1,X2) | sP12(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f229])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP12(X0,X1,X2) | sP13(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f402])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~sP12(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~sP12(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f229])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~sP12(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~sP12(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f229])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f101,f152,f151])).
fof(f151,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f152,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f101,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) | p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.lebesgue.psfis_unique','thm.sat.sat_25_prove_ltb')).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f226])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f226])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f226])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f223])).
fof(f223,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f151])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f224])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f224])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f224])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f220])).
fof(f220,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f219])).
fof(f219,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f152])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f221])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f221])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f385])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f221])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f221])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f99,f149,f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : (sP8(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f149,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f99,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) & p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.lebesgue.psfis_unique','thm.sat.sat_23_prove_ltb')).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(nnf_transformation,[],[f148])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f216])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f216])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f216])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP8(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP8(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f213])).
fof(f213,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f149])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f214])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f371])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f214])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP7(X2,X1,X0)) & (~sP7(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP7(X2,X1,X0))),
  inference(definition_folding,[],[f97,f146,f145,f144,f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f144,plain,(
  ! [X1,X0,X2] : (sP5(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f145,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f146,plain,(
  ! [X2,X1,X0] : (sP7(X2,X1,X0) <=> (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X22)) | p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.lebesgue.psfis_unique','thm.sat.sat_2|49|_prove_ltb')).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f211])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f211])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f143])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f210])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f210])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f210])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f207])).
fof(f207,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(flattening,[],[f206])).
fof(f206,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(nnf_transformation,[],[f144])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f208])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f208])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f208])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f204])).
fof(f204,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f203])).
fof(f203,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f145])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f205])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f205])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f205])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP4(X2,X1,X0) & sP5(X1,X2,X0) & sP6(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X2,X1,X0) | ~sP5(X1,X2,X0) | ~sP6(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f201])).
fof(f201,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP7(X2,X1,X0)))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP7(X2,X1,X0)))),
  inference(nnf_transformation,[],[f146])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP5(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f202])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f202])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP7(X0,X1,X2) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X2,X1) | p(s(bool,X1)) | sP7(X0,X1,X2) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X2,X1) | p(s(bool,X0)) | sP7(X0,X1,X2) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X2)) | p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X2)) & ~p(s(bool,X0))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f196])).
fof(f196,plain,(
  ! [X0,X2,X1] : ((~sP3(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (~p(s(bool,X2)) & ~p(s(bool,X0))) | sP3(X0,X2,X1)))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X0,X2,X1] : ((~sP3(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (~p(s(bool,X2)) & ~p(s(bool,X0)))) | sP3(X0,X2,X1)))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X2,X1] : (sP3(X0,X2,X1) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f197])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f197])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f337])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f197])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f191])).
fof(f191,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f324])).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f192])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f192])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f192])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f192])).
fof(f321,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK20(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK20(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK20(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK20(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.lebesgue.psfis_unique',aHL_EXT)).
fof(f311,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f186])).
fof(f186,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f312,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f313,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f187])).
fof(f316,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f187])).
fof(f428,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f296,f426])).
fof(f296,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f297,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f180])).
fof(f298,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f180])).
fof(f285,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.lebesgue.psfis_unique','thm.bool.IMP_ANTISYM_AX')).
fof(f288,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f120])).
fof(f254,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.lebesgue.psfis_unique','thm.bool.EQ_CLAUSES|split|1')).
fof(f255,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f160])).
fof(f247,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.lebesgue.psfis_unique',aHL_BOOL_CASES)).
fof(f246,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X12] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X12))))),
  file('thm.lebesgue.psfis_unique','thm.extreal.le_refl')).
fof(f239,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.lebesgue.psfis_unique',aHL_TRUTH)).
fof(f235,plain,(
  p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15))))),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15)))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16)))))) & s('type.extreal.extreal',sK17) != s('type.extreal.extreal',sK18)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15,sK16,sK17,sK18])],[f105])).
fof(f105,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X4),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) & s('type.extreal.extreal',X3) != s('type.extreal.extreal',X4))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X4),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2))))))) & s('type.extreal.extreal',X3) != s('type.extreal.extreal',X4))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ~! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X4),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2))))))) => s('type.extreal.extreal',X3) = s('type.extreal.extreal',X4))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X3),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X4),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2))))))) => s('type.extreal.extreal',X3) = s('type.extreal.extreal',X4))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X9),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),X8)))))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X10),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),X8))))))) => s('type.extreal.extreal',X9) = s('type.extreal.extreal',X10))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X9),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),X8)))))) & p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',X10),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),X8))))))) => s('type.extreal.extreal',X9) = s('type.extreal.extreal',X10))),
  file('thm.lebesgue.psfis_unique','thm.lebesgue.psfis_unique')).
fof(f236,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',sK17),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))))),
  inference(cnf_transformation,[],[f157])).
fof(f237,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.extreal.extreal',sK18),s(fun('type.extreal.extreal',bool),'const.lebesgue.psfis_2'(s('type.pair.prod'(fun(sK14,bool),'type.pair.prod'(fun(fun(sK14,bool),bool),fun(fun(sK14,bool),'type.realax.real'))),sK15),s(fun(sK14,'type.extreal.extreal'),sK16))))))),
  inference(cnf_transformation,[],[f157])).
fof(f238,plain,(
  s('type.extreal.extreal',sK17) != s('type.extreal.extreal',sK18)),
  inference(cnf_transformation,[],[f157])).
% SZS output end Proof for thm.lebesgue
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1394
% Passive clauses: 6321
% Generated clauses: 59029
% Final active clauses: 802
% Final passive clauses: 1003
% Input formulas: 41
% Initial clauses: 181
% 
% Duplicate literals: 3447
% Trivial inequalities: 181
% Fw subsumption resolutions: 2841
% Fw demodulations: 6003
% Bw demodulations: 1448
% Condensations: 4541
% Global subsumptions: 364
% 
% Simple tautologies: 199
% Equational tautologies: 868
% Forward subsumptions: 32199
% Bw demodulations to eq. taut.: 350
% 
% Unit resulting resolution: 17635
% Factoring: 1
% Forward superposition: 10083
% Backward superposition: 8953
% Self superposition: 262
% Equality factoring: 196
% Equality resolution: 12
% 
% Split clauses: 420
% Split components: 845
% SAT solver clauses: 8481
% SAT solver unit clauses: 1258
% SAT solver binary clauses: 1745
% 
% Sat splits: 420
% Sat splitting refutations: 954
% 
% Memory used [KB]: 28400
% Time elapsed: 4.148 s
% ------------------------------
% ------------------------------
% Success in time 5.467 s
