Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (95108)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1168
% Passive clauses: 7085
% Generated clauses: 35881
% Final active clauses: 1143
% Final passive clauses: 5525
% Input formulas: 517
% Initial clauses: 914
% 
% Duplicate literals: 223
% Trivial inequalities: 224
% Fw subsumption resolutions: 1962
% Fw demodulations: 4060
% Bw demodulations: 16
% 
% Simple tautologies: 368
% Equational tautologies: 17
% Forward subsumptions: 21725
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 6881
% Factoring: 809
% Forward superposition: 11132
% Backward superposition: 9441
% Self superposition: 47
% Equality factoring: 2
% Equality resolution: 20
% 
% Split clauses: 358
% Split components: 718
% SAT solver clauses: 593
% SAT solver unit clauses: 26
% SAT solver binary clauses: 97
% 
% Sat splits: 358
% 
% TWLsolver clauses: 593
% TWLsolver calls for satisfiability: 412
% 
% Memory used [KB]: 18293
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (95120)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 767
% Passive clauses: 3911
% Generated clauses: 29077
% Final active clauses: 666
% Final passive clauses: 2514
% Input formulas: 517
% Initial clauses: 914
% 
% Duplicate literals: 94
% Trivial inequalities: 97
% Fw subsumption resolutions: 1146
% Fw demodulations: 4498
% Bw demodulations: 652
% Condensations: 101
% Global subsumptions: 27
% 
% Simple tautologies: 323
% Equational tautologies: 14
% Forward subsumptions: 14113
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 3027
% Factoring: 82
% Forward superposition: 8119
% Backward superposition: 5946
% Self superposition: 86
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 110
% Split components: 220
% SAT solver clauses: 5421
% SAT solver unit clauses: 1152
% SAT solver binary clauses: 306
% SAT solver learnt clauses: 54
% SAT solver learnt literals: 19
% 
% Sat splits: 110
% Sat splitting refutations: 77
% 
% TWLsolver clauses: 5314
% TWLsolver calls for satisfiability: 5361
% 
% Memory used [KB]: 23155
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (95133)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3180
% Passive clauses: 24692
% Generated clauses: 237488
% Final active clauses: 2860
% Final passive clauses: 6542
% Input formulas: 517
% Initial clauses: 914
% Discarded non-redundant clauses: 109052
% 
% Duplicate literals: 13
% Trivial inequalities: 103
% Fw subsumption resolutions: 3473
% Fw demodulations: 51848
% Bw demodulations: 3830
% Condensations: 129
% Global subsumptions: 4
% 
% Simple tautologies: 192
% Equational tautologies: 161
% Forward subsumptions: 115485
% Fw demodulations to eq. taut.: 844
% Bw demodulations to eq. taut.: 60
% 
% Unit resulting resolution: 100780
% Factoring: 10
% Forward superposition: 42958
% Backward superposition: 29852
% Self superposition: 145
% Equality factoring: 2
% Equality resolution: 34
% 
% Split clauses: 44
% Split components: 88
% SAT solver clauses: 27839
% SAT solver unit clauses: 10759
% SAT solver binary clauses: 12453
% 
% Sat splits: 44
% Sat splitting refutations: 3052
% 
% Memory used [KB]: 117055
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f12112,plain,(
  $false),
  inference(subsumption_resolution,[],[f12111,f1746])).
fof(f1746,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK17))))))) )),
  inference(subsumption_resolution,[],[f1745,f1532])).
fof(f1532,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f755])).
fof(f755,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f158])).
fof(f158,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_LT_CALCULATE|split|3','thm.numeral.numeral_distrib|split|20')).
fof(f1745,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK17))))))) )),
  inference(forward_demodulation,[],[f1744,f1181])).
fof(f1181,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f585])).
fof(f585,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f129])).
fof(f129,axiom,(
  ! [X6,X7] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))),
  file('thm.integer.INT_LT_CALCULATE|split|3','thm.integer.INT_LT_CALCULATE|split|0')).
fof(f1744,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK17))))))) )),
  inference(forward_demodulation,[],[f1556,f1192])).
fof(f1192,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f599])).
fof(f599,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X9] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_LT_CALCULATE|split|3','thm.integer.INT_NEG_GT|48|')).
fof(f1556,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK17))))))) )),
  inference(superposition,[],[f1063,f1158])).
fof(f1158,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f915])).
fof(f915,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))))),
  inference(flattening,[],[f914])).
fof(f914,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f563])).
fof(f563,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) <=> (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f562])).
fof(f562,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.integer.INT_LT_CALCULATE|split|3','thm.integer.INT_LT_CALCULATE|split|2')).
fof(f1063,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK17))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18)))))))),
  inference(cnf_transformation,[],[f881])).
fof(f881,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK17))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18])],[f762])).
fof(f762,plain,(
  ? [X0,X1] : s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0,X1] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))))),
  file('thm.integer.INT_LT_CALCULATE|split|3','thm.integer.INT_LT_CALCULATE|split|3')).
fof(f12111,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK17))))))),
  inference(subsumption_resolution,[],[f11547,f1499])).
fof(f1499,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1047])).
fof(f1047,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f1046])).
fof(f1046,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f842])).
fof(f842,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f730])).
fof(f730,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_LT_CALCULATE|split|3','thm.bool.IMP_CLAUSES|split|2')).
fof(f11547,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK17))))))),
  inference(superposition,[],[f1150,f1733])).
fof(f1733,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK17))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18)))))))),
  inference(trivial_inequality_removal,[],[f1732])).
fof(f1732,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK17))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18)))))))),
  inference(superposition,[],[f1063,f1529])).
fof(f1529,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f752])).
fof(f752,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f302])).
fof(f302,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X12))),
  file('thm.integer.INT_LT_CALCULATE|split|3','thm.bool.BOOL_CASES_AX')).
fof(f1150,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f786])).
fof(f786,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f556])).
fof(f556,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f167])).
fof(f167,axiom,(
  ! [X9,X10] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X10),s('type.integer.int',X9)))))),
  file('thm.integer.INT_LT_CALCULATE|split|3','thm.integer.INT_LT_ANTISYM')).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 482
% Passive clauses: 4288
% Generated clauses: 10487
% Final active clauses: 474
% Final passive clauses: 4037
% Input formulas: 517
% Initial clauses: 479
% 
% Selected by SInE selection: 195
% SInE iterations: 1
% 
% Duplicate literals: 133
% Trivial inequalities: 77
% Fw subsumption resolutions: 993
% Fw demodulations: 1379
% Bw demodulations: 4
% 
% Simple tautologies: 251
% Equational tautologies: 3
% Forward subsumptions: 2627
% 
% Binary resolution: 3007
% Forward superposition: 3655
% Backward superposition: 1198
% Equality resolution: 1
% 
% Split clauses: 365
% Split components: 743
% SAT solver clauses: 397
% SAT solver unit clauses: 31
% SAT solver binary clauses: 289
% 
% Sat splits: 365
% Sat splitting refutations: 31
% 
% Memory used [KB]: 6012
% Time elapsed: 0.200 s
% ------------------------------
% ------------------------------
% Success in time 11.922 s
