Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f4090,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4087,f266,f481,f479,f4048,f265,f480,f478])).
fof(f478,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK13))) | $spl0),
  inference(cnf_transformation,[],[f478_D])).
fof(f478_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK13))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f480,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) | $spl2),
  inference(cnf_transformation,[],[f480_D])).
fof(f480_D,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f265,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK13)))),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK13))) | s('type.rat.rat',sK12) = s('type.rat.rat',sK13)) & (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK13))) | s('type.rat.rat',sK12) != s('type.rat.rat',sK13))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f166])).
fof(f166,plain,(
  ? [X0,X1] : ((s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) & (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',X0) != s('type.rat.rat',X1)))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ? [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) <~> s('type.rat.rat',X0) = s('type.rat.rat',X1))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ~! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) <=> s('type.rat.rat',X0) = s('type.rat.rat',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) <=> s('type.rat.rat',X6) = s('type.rat.rat',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) <=> s('type.rat.rat',X6) = s('type.rat.rat',X7))),
  file('thm.rat.RAT_EQ_|48|SUB','thm.rat.RAT_EQ_|48|SUB')).
fof(f4048,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4045,f481])).
fof(f4045,plain,(
  s('type.rat.rat',sK12) = s('type.rat.rat',sK13) | $spl0),
  inference(trivial_inequality_removal,[],[f4044])).
fof(f4044,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.rat.rat',sK12) = s('type.rat.rat',sK13) | $spl0),
  inference(superposition,[],[f351,f478])).
fof(f351,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) & (s('type.rat.rat',X0) != s('type.rat.rat',X1) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) <=> s('type.rat.rat',X0) = s('type.rat.rat',X1))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6,X7] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) <=> s('type.rat.rat',X6) = s('type.rat.rat',X7))),
  file('thm.rat.RAT_EQ_|48|SUB','thm.rat.RAT_EQ_SUB|48|')).
fof(f479,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK13))) | $spl1),
  inference(cnf_transformation,[],[f479_D])).
fof(f479_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK13))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f481,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) | $spl3),
  inference(cnf_transformation,[],[f481_D])).
fof(f481_D,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f266,plain,(
  s('type.rat.rat',sK12) != s('type.rat.rat',sK13) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK13)))),
  inference(cnf_transformation,[],[f167])).
fof(f4087,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f4082])).
fof(f4082,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl2)),
  inference(superposition,[],[f4053,f4072])).
fof(f4072,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(equality_resolution,[],[f352])).
fof(f352,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) != s('type.rat.rat',X1) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) )),
  inference(cnf_transformation,[],[f210])).
fof(f4053,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',sK12),s('type.rat.rat',sK12))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f479,f480])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 311
% Passive clauses: 1165
% Generated clauses: 3818
% Final active clauses: 309
% Final passive clauses: 705
% Input formulas: 51
% Initial clauses: 206
% 
% Duplicate literals: 141
% Trivial inequalities: 16
% Fw subsumption resolutions: 251
% Fw demodulations: 9
% 
% Simple tautologies: 170
% Equational tautologies: 21
% Forward subsumptions: 2036
% 
% Binary resolution: 2269
% Factoring: 706
% Forward superposition: 108
% Backward superposition: 66
% Self superposition: 3
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 88
% Split components: 177
% SAT solver clauses: 166
% SAT solver unit clauses: 8
% SAT solver binary clauses: 42
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 88
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 164
% TWLsolver calls for satisfiability: 133
% 
% Memory used [KB]: 1918
% Time elapsed: 0.085 s
% ------------------------------
% ------------------------------
% Success in time 0.086 s
