Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.measure
% (72096)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 198
% Passive clauses: 2422
% Generated clauses: 20374
% Final active clauses: 175
% Final passive clauses: 1960
% Input formulas: 29
% Initial clauses: 79
% 
% Duplicate literals: 858
% Trivial inequalities: 226
% Fw subsumption resolutions: 3308
% Fw demodulations: 3446
% Bw demodulations: 19
% 
% Simple tautologies: 390
% Equational tautologies: 22
% Forward subsumptions: 7857
% 
% Binary resolution: 1056
% Forward superposition: 6197
% Backward superposition: 4744
% Self superposition: 114
% Equality factoring: 55
% Equality resolution: 43
% 
% Split clauses: 998
% Split components: 2012
% SAT solver clauses: 1802
% SAT solver unit clauses: 18
% SAT solver binary clauses: 157
% 
% Sat splits: 998
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1801
% TWLsolver calls for satisfiability: 289
% 
% Memory used [KB]: 12153
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.measure
% (72116)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 148
% Passive clauses: 1159
% Generated clauses: 14082
% Final active clauses: 133
% Final passive clauses: 743
% Input formulas: 29
% Initial clauses: 79
% 
% Duplicate literals: 786
% Trivial inequalities: 208
% Fw subsumption resolutions: 1472
% Fw demodulations: 1244
% Bw demodulations: 190
% Condensations: 938
% Global subsumptions: 8
% 
% Simple tautologies: 332
% Equational tautologies: 22
% Forward subsumptions: 6121
% Bw demodulations to eq. taut.: 92
% 
% Binary resolution: 606
% Factoring: 1
% Forward superposition: 3888
% Backward superposition: 3303
% Self superposition: 114
% Equality factoring: 68
% Equality resolution: 25
% 
% Split clauses: 227
% Split components: 472
% SAT solver clauses: 1794
% SAT solver unit clauses: 86
% SAT solver binary clauses: 50
% SAT solver learnt clauses: 16
% 
% Sat splits: 227
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1762
% TWLsolver calls for satisfiability: 459
% 
% Memory used [KB]: 13432
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.measure
% (72142)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 703
% Passive clauses: 5172
% Generated clauses: 154039
% Final active clauses: 75
% Final passive clauses: 161
% Input formulas: 29
% Initial clauses: 79
% 
% Duplicate literals: 7334
% Trivial inequalities: 949
% Fw subsumption resolutions: 13664
% Fw demodulations: 18455
% Bw demodulations: 3210
% Condensations: 15611
% Global subsumptions: 786
% 
% Simple tautologies: 1838
% Equational tautologies: 2117
% Forward subsumptions: 82118
% Fw demodulations to eq. taut.: 70
% Bw demodulations to eq. taut.: 428
% 
% Unit resulting resolution: 10112
% Forward superposition: 37753
% Backward superposition: 26832
% Self superposition: 745
% Equality factoring: 374
% Equality resolution: 147
% 
% Split clauses: 410
% Split components: 831
% SAT solver clauses: 7625
% SAT solver unit clauses: 69
% SAT solver binary clauses: 1004
% 
% Sat splits: 410
% Sat splitting refutations: 844
% 
% Memory used [KB]: 67163
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.measure
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.measure
% SZS output start Proof for thm.measure
fof(f4341,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f203,f204,f205,f206,f207,f209,f229,f210,f220,f227,f243,f244,f245,f246,f247,f249,f230,f275,f250,f260,f274,f300,f338,f369,f371,f367,f366,f404,f405,f372,f401,f400,f454,f451,f436,f435,f469,f470,f471,f472,f473,f474,f475,f476,f531,f533,f534,f535,f538,f548,f558,f568,f570,f571,f572,f575,f585,f592,f722,f724,f731,f726,f733,f694,f696,f697,f699,f701,f712,f734,f1068,f1069,f1121,f1123,f1124,f1125,f1128,f1138,f1152,f1512,f1514,f1515,f1516,f1519,f1529,f1537,f1556,f1558,f1559,f1560,f1563,f1573,f1580,f1601,f1603,f1604,f1606,f1608,f1619,f1628,f1635,f1744,f1747,f1749,f1753,f1754,f1756,f1757,f1758,f1760,f1764,f1771,f1772,f1773,f1775,f1745,f1782,f1784,f1785,f1786,f1789,f1799,f1748,f1817,f1819,f1820,f1821,f1824,f1834,f1841,f1882,f1884,f1885,f1887,f1889,f1900,f1910,f1925,f1927,f1928,f1930,f1932,f1943,f1952,f2310,f2312,f2313,f2314,f2317,f2327,f2337,f2371,f2373,f2374,f2375,f2378,f2388,f2398,f2636,f2638,f2639,f2640,f2643,f2653,f2660,f2802,f2805,f2807,f2811,f2812,f2814,f2815,f2816,f2818,f2822,f2829,f2830,f2831,f2833,f2835,f2803,f2840,f2842,f2843,f2844,f2847,f2857,f2806,f2874,f2876,f2877,f2879,f2881,f2892,f2901,f3024,f3027,f3031,f3032,f3033,f3035,f3038,f3039,f3040,f3041,f3050,f3051,f3054,f3055,f3056,f3059,f3025,f3060,f3063,f3065,f3066,f3068,f3070,f3081,f3028,f3124,f3126,f3127,f3128,f3131,f3141,f3149,f3219,f3218,f3217,f3221,f3220,f3222,f3233,f3235,f3234,f3273,f3275,f3276,f3278,f3280,f3291,f3300,f3359,f3361,f3362,f3364,f3366,f3377,f3386,f3407,f3409,f3410,f3412,f3414,f3425,f3433,f3445,f3447,f3448,f3450,f3452,f3463,f3474,f3665,f3667,f3668,f3670,f3672,f3683,f3694,f3902,f3905,f3907,f3911,f3912,f3914,f3915,f3916,f3918,f3922,f3929,f3930,f3931,f370,f3933,f1776,f3935,f3936,f3937,f2836,f3938,f3939,f3940,f3903,f3941,f3944,f3946,f3947,f3949,f3951,f3962,f3973,f3987,f3989,f3990,f3991,f3994,f4004,f3906,f4019,f4021,f4022,f4024,f4026,f4037,f4048,f4068,f4070,f4071,f4073,f4075,f4086,f4097,f4194,f4196,f4197,f4199,f4201,f4212,f4222,f4232,f4234,f4235,f4237,f4239,f4250,f4258,f4275,f4277,f4278,f228,f4279,f232,f4282,f231,f4292,f4300,f233,f4335])).
fof(f4335,plain,(
  $false | ($spl26 | $spl31)),
  inference(subsumption_resolution,[],[f4334,f191])).
fof(f191,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.measure.MEASURABLE_UP_LIFT',aHL_TRUTH)).
fof(f4334,plain,(
  ~p(s(bool,'T_0')) | ($spl26 | $spl31)),
  inference(forward_demodulation,[],[f4306,f1832])).
fof(f1832,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | $spl26),
  inference(resolution,[],[f724,f186])).
fof(f186,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.EQ_CLAUSES|split|1')).
fof(f4306,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | $spl31),
  inference(resolution,[],[f2872,f557])).
fof(f557,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),X0))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),X0),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))) )),
  inference(forward_demodulation,[],[f556,f196])).
fof(f196,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X14,X15] : s(fun(X6,bool),X14) = s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X14),s(fun(fun(X6,bool),bool),X15)))))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.measure.space_def')).
fof(f556,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))))))),s(fun(fun(sK0,bool),bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),X0),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))) )),
  inference(forward_demodulation,[],[f528,f195])).
fof(f195,plain,(
  ( ! [X2,X0,X1] : (s(fun(fun(X0,bool),bool),X2) = s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : s(fun(fun(X0,bool),bool),X2) = s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1,X2] : s(fun(fun(X0,bool),bool),X2) = s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X14,X15] : s(fun(fun(X6,bool),bool),X15) = s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X14),s(fun(fun(X6,bool),bool),X15)))))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.measure.subsets_def')).
fof(f528,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),X0))),s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))))))),s(fun(fun(sK0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),X0),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))) )),
  inference(resolution,[],[f122,f132])).
fof(f132,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X6))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X6))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))))) & (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK8(X4,X3,X2,X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),sK8(X4,X3,X2,X1,X0)))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X6))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))))) & (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))))),
  inference(rectify,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))))) & (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))))))) & ((~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) <=> (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))))))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) <=> (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))))))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) <=> (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),X5))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X7,X6] : ! [X9,X10,X12] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X12),s(fun(fun(X6,X7),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X9),s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X10)))))) <=> (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X9)))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X10)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X12),s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X9))),s(fun(X7,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X10)))))))) & ! [X21] : (p(s(bool,'const.bool.IN_2'(s(fun(X7,bool),X21),s(fun(fun(X7,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X10)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X12),s(fun(X7,bool),X21))),s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X9))))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X9)))))))))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.measure.IN_MEASURABLE')).
fof(f122,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))) & ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f67])).
fof(f67,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X6),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X3))),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X5)))))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X4)))))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X3),s(fun(fun(X0,bool),bool),X4)))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X6),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X4))),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X5)))))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X6),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X3))),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X5)))))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X4)))))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X3),s(fun(fun(X0,bool),bool),X4))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X6),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X4))),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X5)))))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X6),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X3))),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X5)))))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X4)))))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X3),s(fun(fun(X0,bool),bool),X4))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X6),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X4))),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X5)))))))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X6),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X3))),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X5)))))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X4)))))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X3),s(fun(fun(X0,bool),bool),X4))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X6),s(fun(fun(X0,X1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),'const.pair.,_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),X4))),s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X5)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11,X12] : ((p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X12),s(fun(fun(X6,X7),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),X9))),s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X11)))))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),X10)))))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X6,bool),bool),X9),s(fun(fun(X6,bool),bool),X10))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X12),s(fun(fun(X6,X7),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),X10))),s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X11)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11,X12] : ((p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X12),s(fun(fun(X6,X7),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),X9))),s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X11)))))) & p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),X10)))))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X6,bool),bool),X9),s(fun(fun(X6,bool),bool),X10))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X12),s(fun(fun(X6,X7),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),'const.pair.,_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),X10))),s('type.pair.prod'(fun(X7,bool),fun(fun(X7,bool),bool)),X11)))))))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.measure.MEASURABLE_UP_LIFT')).
fof(f2872,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK3)))) | $spl31),
  inference(resolution,[],[f733,f202])).
fof(f202,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X0),s(fun(fun(sK0,bool),bool),sK3))))) )),
  inference(resolution,[],[f124,f126])).
fof(f126,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X21,X13] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X21),s(fun(X6,bool),X13)))) <=> ! [X14] : (p(s(bool,'const.bool.IN_2'(s(X6,X14),s(fun(X6,bool),X21)))) => p(s(bool,'const.bool.IN_2'(s(X6,X14),s(fun(X6,bool),X13))))))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.pred_set.SUBSET_DEF')).
fof(f124,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4))))),
  inference(cnf_transformation,[],[f80])).
fof(f233,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) | $spl12),
  inference(cnf_transformation,[],[f233_D])).
fof(f233_D,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4300,plain,(
  ( ! [X4,X2] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,sK7(X2,sK2,sK0)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2)))) ) | $spl114),
  inference(cnf_transformation,[],[f4300_D])).
fof(f4300_D,plain,(
  ( ! [X4,X2] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,sK7(X2,sK2,sK0)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f4292,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,sK7(X43,sK2,sK0)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X43)))) )),
  inference(resolution,[],[f3403,f200])).
fof(f200,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f139])).
fof(f139,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X14,X19,X15,X20] : ((s(bool,X14) = s(bool,X19) & (p(s(bool,X19)) => s(bool,X15) = s(bool,X20))) => ((p(s(bool,X14)) => p(s(bool,X15))) <=> (p(s(bool,X19)) => p(s(bool,X20)))))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.IMP_CONG')).
fof(f3403,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,sK7(X0,sK2,sK0)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f2651,f127])).
fof(f127,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f2651,plain,(
  ( ! [X46] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X46),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X46))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f1553,f186])).
fof(f1553,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X0))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) )),
  inference(resolution,[],[f555,f192])).
fof(f192,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) | ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))),s(fun(X1,bool),X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK9(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK9(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) | ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))),s(fun(X1,bool),X4)))))) & (? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))))),
  inference(rectify,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) | ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4)))))) & (? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) <=> ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4))))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) <=> ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) <=> ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X7] : ! [X12,X22,X23] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,X7),X12),s(fun(fun(X6,X7),bool),'const.util_prob.FUNSET_2'(s(fun(X6,bool),X22),s(fun(X7,bool),X23)))))) <=> ! [X14] : (p(s(bool,'const.bool.IN_2'(s(X6,X14),s(fun(X6,bool),X22)))) => p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X12),s(X6,X14))),s(fun(X7,bool),X23))))))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.util_prob.IN_FUNSET')).
fof(f555,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))))),
  inference(forward_demodulation,[],[f527,f196])).
fof(f527,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))))),
  inference(resolution,[],[f122,f131])).
fof(f131,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f231,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,X8))) ) | $spl8),
  inference(cnf_transformation,[],[f231_D])).
fof(f231_D,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,X8))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f4282,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,sK7(X30,sK2,sK0)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X30))) | s(bool,X32) != s(bool,X33)) )),
  inference(resolution,[],[f3403,f142])).
fof(f142,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f89])).
fof(f232,plain,(
  ( ! [X14] : (~p(s(bool,X14))) ) | $spl10),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  ( ! [X14] : (~p(s(bool,X14))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f4279,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,sK7(X18,sK2,sK0)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X18)))) )),
  inference(resolution,[],[f3403,f141])).
fof(f141,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f89])).
fof(f228,plain,(
  ( ! [X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X1))) ) | $spl2),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  ( ! [X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X1))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f4278,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,sK7(X14,sK2,sK0)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X14)))) )),
  inference(resolution,[],[f3403,f140])).
fof(f140,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f89])).
fof(f4277,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,sK7(X10,sK2,sK0)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X10))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f3403,f138])).
fof(f138,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f89])).
fof(f4275,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,sK7(X2,sK2,sK0)))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2)))) )),
  inference(resolution,[],[f3403,f136])).
fof(f136,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f89])).
fof(f4258,plain,(
  ( ! [X14,X17,X15] : (p(s(bool,X14)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK2))) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X15))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) | $spl112),
  inference(cnf_transformation,[],[f4258_D])).
fof(f4258_D,plain,(
  ( ! [X14,X17,X15] : (p(s(bool,X14)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK2))) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X15))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f4250,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X74)) | s(bool,X74) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X75))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X77) != s(bool,'const.bool.IN_2'(s(sK0,X75),s(fun(sK0,bool),sK2)))) ) | $spl58),
  inference(resolution,[],[f2660,f200])).
fof(f4239,plain,(
  ( ! [X52,X50,X53,X51,X49] : (p(s(bool,X49)) | s(bool,X49) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X50))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(sK0,X50),s(fun(sK0,bool),sK2))) | s(bool,X51) != s(bool,X53)) ) | $spl58),
  inference(resolution,[],[f2660,f142])).
fof(f4237,plain,(
  ( ! [X39,X43,X41,X42,X40] : (p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X40))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(sK0,X40),s(fun(sK0,bool),sK2)))) ) | $spl58),
  inference(resolution,[],[f2660,f141])).
fof(f4235,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X30))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK0,X30),s(fun(sK0,bool),sK2)))) ) | $spl58),
  inference(resolution,[],[f2660,f140])).
fof(f4234,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X25))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(sK0,X25),s(fun(sK0,bool),sK2))) | s(bool,X26) != s(bool,X28)) ) | $spl58),
  inference(resolution,[],[f2660,f138])).
fof(f4232,plain,(
  ( ! [X14,X17,X15,X18,X16] : (p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X15))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK2)))) ) | $spl58),
  inference(resolution,[],[f2660,f136])).
fof(f4222,plain,(
  ( ! [X6,X8,X5] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),X6)))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),X6))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),sK3))))) ) | $spl110),
  inference(cnf_transformation,[],[f4222_D])).
fof(f4222_D,plain,(
  ( ! [X6,X8,X5] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),X6)))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),X6))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),sK3))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f4212,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X65),s(fun(fun(sK0,bool),bool),X66)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X65),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X68) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),X66)))) )),
  inference(resolution,[],[f1510,f200])).
fof(f1510,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),X6)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),X6)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),sK3))))) )),
  inference(resolution,[],[f202,f126])).
fof(f4201,plain,(
  ( ! [X43,X41,X44,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X40),s(fun(fun(sK0,bool),bool),X41)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X40),s(fun(fun(sK0,bool),bool),sK3)))) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X43) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),X41))) | s(bool,X42) != s(bool,X44)) )),
  inference(resolution,[],[f1510,f142])).
fof(f4199,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X30),s(fun(fun(sK0,bool),bool),X31)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X30),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),X31)))) )),
  inference(resolution,[],[f1510,f141])).
fof(f4197,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X20),s(fun(fun(sK0,bool),bool),X21)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X20),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),X21)))) )),
  inference(resolution,[],[f1510,f140])).
fof(f4196,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),X16)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),sK3)))) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X18) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),X16))) | s(bool,X17) != s(bool,X19)) )),
  inference(resolution,[],[f1510,f138])).
fof(f4194,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),X6)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),X6)))) )),
  inference(resolution,[],[f1510,f136])).
fof(f4097,plain,(
  ( ! [X12,X10,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X8,sK0),X9),s(fun(fun(X8,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(X8,bool),'const.pred_set.INTER_2'(s(fun(X8,bool),'const.util_prob.PREIMAGE_2'(s(fun(X8,sK0),X9),s(fun(sK0,bool),sK8(X9,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X10,X8,sK0)))),s(fun(X8,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10))))),s(fun(fun(X8,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10))))) | ~p(s(bool,X12)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X8,sK0),X9),s(fun(fun(X8,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X8,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10))),s(fun(sK0,bool),sK2))))))) ) | $spl108),
  inference(cnf_transformation,[],[f4097_D])).
fof(f4097_D,plain,(
  ( ! [X12,X10,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X8,sK0),X9),s(fun(fun(X8,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(X8,bool),'const.pred_set.INTER_2'(s(fun(X8,bool),'const.util_prob.PREIMAGE_2'(s(fun(X8,sK0),X9),s(fun(sK0,bool),sK8(X9,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X10,X8,sK0)))),s(fun(X8,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10))))),s(fun(fun(X8,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10))))) | ~p(s(bool,X12)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X8,sK0),X9),s(fun(fun(X8,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X8,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10))),s(fun(sK0,bool),sK2))))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f4086,plain,(
  ( ! [X90,X88,X87,X89,X86] : (p(s(bool,'const.bool.IN_2'(s(fun(X86,sK0),X87),s(fun(fun(X86,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X86,bool),fun(fun(X86,bool),bool)),X88),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X86,sK0),X87),s(fun(fun(X86,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X86,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X86,bool),fun(fun(X86,bool),bool)),X88))),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X86,bool),fun(fun(X86,bool),bool)),X88)))) | p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X90) != s(bool,'const.bool.IN_2'(s(fun(X86,bool),'const.pred_set.INTER_2'(s(fun(X86,bool),'const.util_prob.PREIMAGE_2'(s(fun(X86,sK0),X87),s(fun(sK0,bool),sK8(X87,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X88,X86,sK0)))),s(fun(X86,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X86,bool),fun(fun(X86,bool),bool)),X88))))),s(fun(fun(X86,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X86,bool),fun(fun(X86,bool),bool)),X88)))))) )),
  inference(resolution,[],[f273,f200])).
fof(f273,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(fun(X9,bool),'const.pred_set.INTER_2'(s(fun(X9,bool),'const.util_prob.PREIMAGE_2'(s(fun(X9,sK0),X10),s(fun(sK0,bool),sK8(X10,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X11,X9,sK0)))),s(fun(X9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X9,bool),fun(fun(X9,bool),bool)),X11))))),s(fun(fun(X9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X9,bool),fun(fun(X9,bool),bool)),X11)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X9,sK0),X10),s(fun(fun(X9,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X9,bool),fun(fun(X9,bool),bool)),X11),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X9,sK0),X10),s(fun(fun(X9,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X9,bool),fun(fun(X9,bool),bool)),X11))),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X9,bool),fun(fun(X9,bool),bool)),X11))))) )),
  inference(forward_demodulation,[],[f242,f196])).
fof(f242,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X9,sK0),X10),s(fun(fun(X9,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X9,bool),fun(fun(X9,bool),bool)),X11),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X9,bool),'const.pred_set.INTER_2'(s(fun(X9,bool),'const.util_prob.PREIMAGE_2'(s(fun(X9,sK0),X10),s(fun(sK0,bool),sK8(X10,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X11,X9,sK0)))),s(fun(X9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X9,bool),fun(fun(X9,bool),bool)),X11))))),s(fun(fun(X9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X9,bool),fun(fun(X9,bool),bool)),X11)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X9,sK0),X10),s(fun(fun(X9,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X9,bool),fun(fun(X9,bool),bool)),X11))),s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X9,bool),fun(fun(X9,bool),bool)),X11))))) )),
  inference(resolution,[],[f123,f134])).
fof(f134,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,bool),'const.pred_set.INTER_2'(s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),X4),s(fun(X0,bool),sK8(X4,X3,X2,X1,X0)))),s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))),s(fun(fun(X1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f123,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))),
  inference(cnf_transformation,[],[f80])).
fof(f4075,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,'const.bool.IN_2'(s(fun(X50,sK0),X51),s(fun(fun(X50,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X50,bool),fun(fun(X50,bool),bool)),X52),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X50,sK0),X51),s(fun(fun(X50,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X50,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X50,bool),fun(fun(X50,bool),bool)),X52))),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X50,bool),fun(fun(X50,bool),bool)),X52)))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(fun(X50,bool),'const.pred_set.INTER_2'(s(fun(X50,bool),'const.util_prob.PREIMAGE_2'(s(fun(X50,sK0),X51),s(fun(sK0,bool),sK8(X51,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X52,X50,sK0)))),s(fun(X50,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X50,bool),fun(fun(X50,bool),bool)),X52))))),s(fun(fun(X50,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X50,bool),fun(fun(X50,bool),bool)),X52))))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f273,f142])).
fof(f4073,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(fun(X38,sK0),X39),s(fun(fun(X38,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X38,bool),fun(fun(X38,bool),bool)),X40),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X38,sK0),X39),s(fun(fun(X38,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X38,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X38,bool),fun(fun(X38,bool),bool)),X40))),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X38,bool),fun(fun(X38,bool),bool)),X40)))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(fun(X38,bool),'const.pred_set.INTER_2'(s(fun(X38,bool),'const.util_prob.PREIMAGE_2'(s(fun(X38,sK0),X39),s(fun(sK0,bool),sK8(X39,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X40,X38,sK0)))),s(fun(X38,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X38,bool),fun(fun(X38,bool),bool)),X40))))),s(fun(fun(X38,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X38,bool),fun(fun(X38,bool),bool)),X40)))))) )),
  inference(resolution,[],[f273,f141])).
fof(f4071,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,'const.bool.IN_2'(s(fun(X26,sK0),X27),s(fun(fun(X26,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X26,bool),fun(fun(X26,bool),bool)),X28),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X26,sK0),X27),s(fun(fun(X26,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X26,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X26,bool),fun(fun(X26,bool),bool)),X28))),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X26,bool),fun(fun(X26,bool),bool)),X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(fun(X26,bool),'const.pred_set.INTER_2'(s(fun(X26,bool),'const.util_prob.PREIMAGE_2'(s(fun(X26,sK0),X27),s(fun(sK0,bool),sK8(X27,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X28,X26,sK0)))),s(fun(X26,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X26,bool),fun(fun(X26,bool),bool)),X28))))),s(fun(fun(X26,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X26,bool),fun(fun(X26,bool),bool)),X28)))))) )),
  inference(resolution,[],[f273,f140])).
fof(f4070,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(fun(X20,sK0),X21),s(fun(fun(X20,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X20,bool),fun(fun(X20,bool),bool)),X22),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X20,sK0),X21),s(fun(fun(X20,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X20,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X20,bool),fun(fun(X20,bool),bool)),X22))),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X20,bool),fun(fun(X20,bool),bool)),X22)))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(X20,bool),'const.pred_set.INTER_2'(s(fun(X20,bool),'const.util_prob.PREIMAGE_2'(s(fun(X20,sK0),X21),s(fun(sK0,bool),sK8(X21,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X22,X20,sK0)))),s(fun(X20,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X20,bool),fun(fun(X20,bool),bool)),X22))))),s(fun(fun(X20,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X20,bool),fun(fun(X20,bool),bool)),X22))))) | s(bool,X23) != s(bool,X25)) )),
  inference(resolution,[],[f273,f138])).
fof(f4068,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X8,sK0),X9),s(fun(fun(X8,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X8,sK0),X9),s(fun(fun(X8,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X8,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10))),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10)))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(X8,bool),'const.pred_set.INTER_2'(s(fun(X8,bool),'const.util_prob.PREIMAGE_2'(s(fun(X8,sK0),X9),s(fun(sK0,bool),sK8(X9,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X10,X8,sK0)))),s(fun(X8,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10))))),s(fun(fun(X8,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X8,bool),fun(fun(X8,bool),bool)),X10)))))) )),
  inference(resolution,[],[f273,f136])).
fof(f4048,plain,(
  ( ! [X14,X12,X13,X16] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X12),X13),s(fun(X12,bool),sK8(X13,X14,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X12)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14))) | ~p(s(bool,X16)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14))))))) ) | $spl106),
  inference(cnf_transformation,[],[f4048_D])).
fof(f4048_D,plain,(
  ( ! [X14,X12,X13,X16] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X12),X13),s(fun(X12,bool),sK8(X13,X14,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X12)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14))) | ~p(s(bool,X16)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14))))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f4037,plain,(
  ( ! [X94,X92,X90,X93,X91] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X90),X91),s(fun(X90,bool),sK8(X91,X92,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X90)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X90),X91),s(fun(fun(sK0,X90),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X90,bool),fun(fun(X90,bool),bool)),X92)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X90),X91),s(fun(fun(sK0,X90),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X90,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X90,bool),fun(fun(X90,bool),bool)),X92)))))))) | p(s(bool,X93)) | ~p(s(bool,X94)) | s(bool,X94) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X90,bool),fun(fun(X90,bool),bool)),X92)))) )),
  inference(resolution,[],[f272,f200])).
fof(f272,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X6),X7),s(fun(X6,bool),sK8(X7,X8,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X6)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X6),X7),s(fun(fun(sK0,X6),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X6),X7),s(fun(fun(sK0,X6),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8))))))))) )),
  inference(forward_demodulation,[],[f271,f196])).
fof(f271,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X6),X7),s(fun(X6,bool),sK8(X7,X8,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X6)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X6),X7),s(fun(fun(sK0,X6),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X6),X7),s(fun(fun(sK0,X6),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8))))) )),
  inference(forward_demodulation,[],[f270,f196])).
fof(f270,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X6),X7),s(fun(X6,bool),sK8(X7,X8,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X6)))),s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X6),X7),s(fun(fun(sK0,X6),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X6),X7),s(fun(fun(sK0,X6),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8))))) )),
  inference(forward_demodulation,[],[f241,f195])).
fof(f241,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X6),X7),s(fun(fun(sK0,X6),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X6),X7),s(fun(X6,bool),sK8(X7,X8,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X6)))),s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))),s(fun(fun(sK0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X6),X7),s(fun(fun(sK0,X6),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X8))))) )),
  inference(resolution,[],[f123,f134])).
fof(f4026,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X54),X55),s(fun(X54,bool),sK8(X55,X56,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X54)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X54),X55),s(fun(fun(sK0,X54),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X54,bool),fun(fun(X54,bool),bool)),X56)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X54),X55),s(fun(fun(sK0,X54),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X54,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X54,bool),fun(fun(X54,bool),bool)),X56)))))))) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X54,bool),fun(fun(X54,bool),bool)),X56))) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f272,f142])).
fof(f4024,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X42),X43),s(fun(X42,bool),sK8(X43,X44,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X42)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X42),X43),s(fun(fun(sK0,X42),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X42,bool),fun(fun(X42,bool),bool)),X44)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X42),X43),s(fun(fun(sK0,X42),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X42,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X42,bool),fun(fun(X42,bool),bool)),X44)))))))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X42,bool),fun(fun(X42,bool),bool)),X44)))) )),
  inference(resolution,[],[f272,f141])).
fof(f4022,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X30),X31),s(fun(X30,bool),sK8(X31,X32,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X30)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X30),X31),s(fun(fun(sK0,X30),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X30,bool),fun(fun(X30,bool),bool)),X32)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X30),X31),s(fun(fun(sK0,X30),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X30,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X30,bool),fun(fun(X30,bool),bool)),X32)))))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X30,bool),fun(fun(X30,bool),bool)),X32)))) )),
  inference(resolution,[],[f272,f140])).
fof(f4021,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X24),X25),s(fun(X24,bool),sK8(X25,X26,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X24)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X24),X25),s(fun(fun(sK0,X24),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X24,bool),fun(fun(X24,bool),bool)),X26)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X24),X25),s(fun(fun(sK0,X24),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X24,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X24,bool),fun(fun(X24,bool),bool)),X26)))))))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X24,bool),fun(fun(X24,bool),bool)),X26))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f272,f138])).
fof(f4019,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,X12),X13),s(fun(X12,bool),sK8(X13,X14,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X12)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))) )),
  inference(resolution,[],[f272,f136])).
fof(f3906,plain,(
  ( ! [X39,X37] : (p(s(bool,X39)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X37),s(fun(sK0,bool),sK2))) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X37))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) | $spl96),
  inference(cnf_transformation,[],[f3906_D])).
fof(f3906_D,plain,(
  ( ! [X39,X37] : (p(s(bool,X39)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X37),s(fun(sK0,bool),sK2))) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X37))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f4004,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X48),s(fun(sK0,bool),sK2))) | p(s(bool,X49)) | p(s(bool,X50)) | s(bool,X50) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X48))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f3887,f200])).
fof(f3887,plain,(
  ( ! [X172] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X172))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X172),s(fun(sK0,bool),sK2)))) )),
  inference(trivial_inequality_removal,[],[f3859])).
fof(f3859,plain,(
  ( ! [X172] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X172))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X172),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f187,f3421])).
fof(f3421,plain,(
  ( ! [X48] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X48))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X48),s(fun(sK0,bool),sK2)))) )),
  inference(resolution,[],[f2651,f182])).
fof(f182,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.EQ_CLAUSES|split|3')).
fof(f187,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.EQ_CLAUSES|split|0')).
fof(f3994,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X35),s(fun(sK0,bool),sK2))) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X35))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X37) != s(bool,X38)) )),
  inference(resolution,[],[f3887,f142])).
fof(f3991,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X23),s(fun(sK0,bool),sK2))) | p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X23))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f3887,f141])).
fof(f3990,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK2))) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X19))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f3887,f140])).
fof(f3989,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK2))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X15))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X17) != s(bool,X18)) )),
  inference(resolution,[],[f3887,f138])).
fof(f3987,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2))) | p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X7))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f3887,f136])).
fof(f3973,plain,(
  ( ! [X14,X12,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))))) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14))) | ~p(s(bool,X16)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X12,bool),sK8(X13,X14,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X12)),s(fun(fun(X12,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14))))))) ) | $spl104),
  inference(cnf_transformation,[],[f3973_D])).
fof(f3973_D,plain,(
  ( ! [X14,X12,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))))) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14))) | ~p(s(bool,X16)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X12,bool),sK8(X13,X14,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X12)),s(fun(fun(X12,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14))))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f3962,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X90),X91),s(fun(fun(sK0,X90),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X90,bool),fun(fun(X90,bool),bool)),X92)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X90,bool),sK8(X91,X92,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X90)),s(fun(fun(X90,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X90,bool),fun(fun(X90,bool),bool)),X92)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X90),X91),s(fun(fun(sK0,X90),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X90,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X90,bool),fun(fun(X90,bool),bool)),X92)))))))) | p(s(bool,X93)) | ~p(s(bool,X94)) | s(bool,X94) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X90,bool),fun(fun(X90,bool),bool)),X92)))) )),
  inference(resolution,[],[f267,f200])).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK8(X1,X2,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2))))))))) )),
  inference(forward_demodulation,[],[f239,f196])).
fof(f239,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK8(X1,X2,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X0),X1),s(fun(fun(sK0,X0),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X2))))) )),
  inference(resolution,[],[f123,f133])).
fof(f133,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK8(X4,X3,X2,X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.util_prob.FUNSET_2'(s(fun(X1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))),s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f3951,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X54),X55),s(fun(fun(sK0,X54),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X54,bool),fun(fun(X54,bool),bool)),X56)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X54,bool),sK8(X55,X56,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X54)),s(fun(fun(X54,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X54,bool),fun(fun(X54,bool),bool)),X56)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X54),X55),s(fun(fun(sK0,X54),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X54,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X54,bool),fun(fun(X54,bool),bool)),X56)))))))) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X54,bool),fun(fun(X54,bool),bool)),X56))) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f267,f142])).
fof(f3949,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X42),X43),s(fun(fun(sK0,X42),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X42,bool),fun(fun(X42,bool),bool)),X44)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X42,bool),sK8(X43,X44,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X42)),s(fun(fun(X42,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X42,bool),fun(fun(X42,bool),bool)),X44)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X42),X43),s(fun(fun(sK0,X42),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X42,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X42,bool),fun(fun(X42,bool),bool)),X44)))))))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X42,bool),fun(fun(X42,bool),bool)),X44)))) )),
  inference(resolution,[],[f267,f141])).
fof(f3947,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X30),X31),s(fun(fun(sK0,X30),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X30,bool),fun(fun(X30,bool),bool)),X32)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X30,bool),sK8(X31,X32,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X30)),s(fun(fun(X30,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X30,bool),fun(fun(X30,bool),bool)),X32)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X30),X31),s(fun(fun(sK0,X30),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X30,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X30,bool),fun(fun(X30,bool),bool)),X32)))))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X30,bool),fun(fun(X30,bool),bool)),X32)))) )),
  inference(resolution,[],[f267,f140])).
fof(f3946,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X24),X25),s(fun(fun(sK0,X24),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X24,bool),fun(fun(X24,bool),bool)),X26)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X24,bool),sK8(X25,X26,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X24)),s(fun(fun(X24,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X24,bool),fun(fun(X24,bool),bool)),X26)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X24),X25),s(fun(fun(sK0,X24),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X24,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X24,bool),fun(fun(X24,bool),bool)),X26)))))))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X24,bool),fun(fun(X24,bool),bool)),X26))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f267,f138])).
fof(f3944,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X12,bool),sK8(X13,X14,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,X12)),s(fun(fun(X12,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,X12),X13),s(fun(fun(sK0,X12),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(X12,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X12,bool),fun(fun(X12,bool),bool)),X14)))) )),
  inference(resolution,[],[f267,f136])).
fof(f3941,plain,(
  ( ! [X223] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X223,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl102),
  inference(cnf_transformation,[],[f3941_D])).
fof(f3941_D,plain,(
  ( ! [X223] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X223,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f3903,plain,(
  ( ! [X13] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X13))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X13),s(fun(sK0,bool),sK2)))) ) | $spl94),
  inference(cnf_transformation,[],[f3903_D])).
fof(f3903_D,plain,(
  ( ! [X13] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X13))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X13),s(fun(sK0,bool),sK2)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f3940,plain,(
  ( ! [X222,X223] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X222))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X223,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X222),s(fun(sK0,bool),sK2)))) ) | $spl56),
  inference(subsumption_resolution,[],[f3882,f2386])).
fof(f2386,plain,(
  ( ! [X41] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X41,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X41),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f1896,f186])).
fof(f1896,plain,(
  ( ! [X40] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X40),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X40,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f1508,f182])).
fof(f1508,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X0,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X0),s(fun(fun(sK0,bool),bool),sK4))))) )),
  inference(resolution,[],[f202,f128])).
fof(f128,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f3882,plain,(
  ( ! [X222,X223] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X223),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X222))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X223,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X222),s(fun(sK0,bool),sK2)))) ) | $spl56),
  inference(superposition,[],[f2398,f3421])).
fof(f3939,plain,(
  ( ! [X221] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X221,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl100),
  inference(cnf_transformation,[],[f3939_D])).
fof(f3939_D,plain,(
  ( ! [X221] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X221,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f3938,plain,(
  ( ! [X220,X221] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X220))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X221,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X220),s(fun(sK0,bool),sK2)))) ) | $spl54),
  inference(subsumption_resolution,[],[f3881,f2325])).
fof(f2325,plain,(
  ( ! [X41] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X41,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X41)))) )),
  inference(resolution,[],[f1597,f186])).
fof(f1597,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X0)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X0,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f1527,f127])).
fof(f1527,plain,(
  ( ! [X49] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X49),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X49),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f202,f186])).
fof(f3881,plain,(
  ( ! [X220,X221] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X221))) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X220))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X221,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X220),s(fun(sK0,bool),sK2)))) ) | $spl54),
  inference(superposition,[],[f2337,f3421])).
fof(f2836,plain,(
  ( ! [X214] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X214),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl64),
  inference(cnf_transformation,[],[f2836_D])).
fof(f2836_D,plain,(
  ( ! [X214] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X214),s(fun(fun(sK0,bool),bool),sK3)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3937,plain,(
  ( ! [X216,X215] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X215))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X216),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X215),s(fun(sK0,bool),sK2)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3879,f1615])).
fof(f1615,plain,(
  ( ! [X48] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X48),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X48),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f1527,f182])).
fof(f3879,plain,(
  ( ! [X216,X215] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X216),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X215))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X216),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X215),s(fun(sK0,bool),sK2)))) ) | $spl36),
  inference(superposition,[],[f1939,f3421])).
fof(f1939,plain,(
  ( ! [X66,X67] : (s(bool,X66) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X67),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X66)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X67),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl36),
  inference(resolution,[],[f1537,f182])).
fof(f3936,plain,(
  ( ! [X214] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X214))))) ) | $spl98),
  inference(cnf_transformation,[],[f3936_D])).
fof(f3936_D,plain,(
  ( ! [X214] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X214))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f3935,plain,(
  ( ! [X213,X214] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X213))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X214)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X213),s(fun(sK0,bool),sK2)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3878,f1597])).
fof(f3878,plain,(
  ( ! [X213,X214] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X214,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X213))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X214)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X213),s(fun(sK0,bool),sK2)))) ) | $spl36),
  inference(superposition,[],[f1921,f3421])).
fof(f1921,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X1,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X0)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X1))))) ) | $spl36),
  inference(resolution,[],[f1537,f127])).
fof(f1776,plain,(
  ( ! [X213] : (s(bool,'F_0') = s(bool,X213)) ) | $spl46),
  inference(cnf_transformation,[],[f1776_D])).
fof(f1776_D,plain,(
  ( ! [X213] : (s(bool,'F_0') = s(bool,X213)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f3933,plain,(
  ( ! [X204,X203] : (p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X203))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,X204) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X203),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f3873,f199])).
fof(f199,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.measure.MEASURABLE_UP_LIFT',aHL_BOOL_CASES)).
fof(f3873,plain,(
  ( ! [X204,X203] : (s(bool,'T_0') != s(bool,X204) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X203))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,X204) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X203),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(superposition,[],[f522,f3421])).
fof(f522,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | p(s(bool,X19)) | s(bool,'F_0') = s(bool,X18)) ) | $spl6),
  inference(subsumption_resolution,[],[f500,f154])).
fof(f154,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.IMP_CLAUSES|split|3')).
fof(f500,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | p(s(bool,X19)) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X18)) ) | $spl6),
  inference(superposition,[],[f230,f199])).
fof(f370,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) | $spl18),
  inference(cnf_transformation,[],[f370_D])).
fof(f370_D,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3931,plain,(
  ( ! [X194,X193] : (s(bool,'T_0') != s(bool,X194) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X193))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X193),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f3868,f185])).
fof(f185,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f116])).
fof(f3868,plain,(
  ( ! [X194,X193] : (s(bool,'T_0') != s(bool,X194) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X193))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | ~p(s(bool,X194)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X193),s(fun(sK0,bool),sK2)))) ) | $spl6),
  inference(superposition,[],[f230,f3421])).
fof(f3930,plain,(
  ( ! [X191,X192,X190] : (p(s(bool,X191)) | p(s(bool,X192)) | s(bool,X192) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X190))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X190),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3867,f154])).
fof(f3867,plain,(
  ( ! [X191,X192,X190] : (~p(s(bool,'T_0')) | p(s(bool,X191)) | p(s(bool,X192)) | s(bool,X192) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X190))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X190),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f200,f3421])).
fof(f3929,plain,(
  ( ! [X182,X183,X181] : (s(bool,'T_0') != s(bool,X182) | p(s(bool,X183)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X181))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X181),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3864,f185])).
fof(f3864,plain,(
  ( ! [X182,X183,X181] : (s(bool,'T_0') != s(bool,X182) | p(s(bool,X183)) | ~p(s(bool,X182)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X181))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X181),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f200,f3421])).
fof(f3922,plain,(
  ( ! [X146,X147,X145,X148] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X146) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X145))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X147) != s(bool,X148) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X145),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3836,f155])).
fof(f155,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f97])).
fof(f3836,plain,(
  ( ! [X146,X147,X145,X148] : (~p(s(bool,'T_0')) | p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X146) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X145))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X147) != s(bool,X148) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X145),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f142,f3421])).
fof(f3918,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,'T_0') != s(bool,X130) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X129))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | ~p(s(bool,X131)) | s(bool,X131) != s(bool,X132) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X129),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3832,f185])).
fof(f3832,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,'T_0') != s(bool,X130) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X129))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | ~p(s(bool,X131)) | ~p(s(bool,X130)) | s(bool,X131) != s(bool,X132) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X129),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f142,f3421])).
fof(f3916,plain,(
  ( ! [X116,X114,X115,X113] : (p(s(bool,X114)) | ~p(s(bool,X115)) | p(s(bool,X116)) | s(bool,X116) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X113))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X113),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3828,f155])).
fof(f3828,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,'T_0')) | p(s(bool,X114)) | ~p(s(bool,X115)) | p(s(bool,X116)) | s(bool,X116) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X113))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X113),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f141,f3421])).
fof(f3915,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'T_0') != s(bool,X102) | p(s(bool,X103)) | ~p(s(bool,X104)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X101))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X101),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3825,f185])).
fof(f3825,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'T_0') != s(bool,X102) | p(s(bool,X103)) | ~p(s(bool,X102)) | ~p(s(bool,X104)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X101))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X101),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f141,f3421])).
fof(f3914,plain,(
  ( ! [X94,X95,X93,X96] : (p(s(bool,X94)) | p(s(bool,X95)) | s(bool,X94) != s(bool,X96) | s(bool,X95) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X93))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X93),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3823,f154])).
fof(f3823,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,'T_0')) | p(s(bool,X94)) | p(s(bool,X95)) | s(bool,X94) != s(bool,X96) | s(bool,X95) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X93))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X93),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f140,f3421])).
fof(f3912,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'T_0') != s(bool,X74) | p(s(bool,X75)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X73))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,X75) != s(bool,X76) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X73),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3818,f185])).
fof(f3818,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'T_0') != s(bool,X74) | p(s(bool,X75)) | ~p(s(bool,X74)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X73))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,X75) != s(bool,X76) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X73),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f140,f3421])).
fof(f3911,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X66) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X65))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X65),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3816,f155])).
fof(f3816,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X66) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X65))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X65),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f138,f3421])).
fof(f3907,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X54) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X53))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X53),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3813,f185])).
fof(f3813,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X54) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X53))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | ~p(s(bool,X55)) | ~p(s(bool,X54)) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X53),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f138,f3421])).
fof(f3905,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X37))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X37),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3809,f154])).
fof(f3809,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X37))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X37),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f136,f3421])).
fof(f3902,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'T_0') != s(bool,X14) | p(s(bool,X15)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X13))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X13),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f3803,f185])).
fof(f3803,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'T_0') != s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X14)) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X13))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X13),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f136,f3421])).
fof(f3694,plain,(
  ( ! [X14,X12,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK8(X12,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X13,X14,sK0)),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X14,bool),fun(fun(X14,bool),bool)),X13))) | ~p(s(bool,X16)) | ~p(s(bool,'const.bool.IN_2'(s(fun(X14,sK0),X12),s(fun(fun(X14,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X14,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X14,bool),fun(fun(X14,bool),bool)),X13))),s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X14,sK0),X12),s(fun(fun(X14,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X14,bool),fun(fun(X14,bool),bool)),X13),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))))) ) | $spl92),
  inference(cnf_transformation,[],[f3694_D])).
fof(f3694_D,plain,(
  ( ! [X14,X12,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK8(X12,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X13,X14,sK0)),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X14,bool),fun(fun(X14,bool),bool)),X13))) | ~p(s(bool,X16)) | ~p(s(bool,'const.bool.IN_2'(s(fun(X14,sK0),X12),s(fun(fun(X14,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X14,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X14,bool),fun(fun(X14,bool),bool)),X13))),s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X14,sK0),X12),s(fun(fun(X14,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X14,bool),fun(fun(X14,bool),bool)),X13),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f3683,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK8(X90,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X91,X92,sK0)),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(X92,sK0),X90),s(fun(fun(X92,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X92,bool),fun(fun(X92,bool),bool)),X91),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X92,sK0),X90),s(fun(fun(X92,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X92,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X92,bool),fun(fun(X92,bool),bool)),X91))),s(fun(sK0,bool),sK2)))))) | p(s(bool,X93)) | ~p(s(bool,X94)) | s(bool,X94) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X92,bool),fun(fun(X92,bool),bool)),X91)))) )),
  inference(resolution,[],[f269,f200])).
fof(f269,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),X5)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK8(X4,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X5,X3,sK0)),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(X3,sK0),X4),s(fun(fun(X3,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),X5),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X3,sK0),X4),s(fun(fun(X3,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X3,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),X5))),s(fun(sK0,bool),sK2))))))) )),
  inference(forward_demodulation,[],[f268,f196])).
fof(f268,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK8(X4,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X5,X3,sK0)),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(X3,sK0),X4),s(fun(fun(X3,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),X5),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X3,sK0),X4),s(fun(fun(X3,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X3,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),X5))),s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),X5))))) )),
  inference(forward_demodulation,[],[f240,f195])).
fof(f240,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X3,sK0),X4),s(fun(fun(X3,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),X5),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK8(X4,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X5,X3,sK0)),s(fun(fun(sK0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X3,sK0),X4),s(fun(fun(X3,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X3,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),X5))),s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X3,bool),fun(fun(X3,bool),bool)),X5))))) )),
  inference(resolution,[],[f123,f133])).
fof(f3672,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK8(X54,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X55,X56,sK0)),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(X56,sK0),X54),s(fun(fun(X56,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X56,bool),fun(fun(X56,bool),bool)),X55),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X56,sK0),X54),s(fun(fun(X56,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X56,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X56,bool),fun(fun(X56,bool),bool)),X55))),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X56,bool),fun(fun(X56,bool),bool)),X55))) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f269,f142])).
fof(f3670,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK8(X42,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X43,X44,sK0)),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(X44,sK0),X42),s(fun(fun(X44,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X44,bool),fun(fun(X44,bool),bool)),X43),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X44,sK0),X42),s(fun(fun(X44,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X44,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X44,bool),fun(fun(X44,bool),bool)),X43))),s(fun(sK0,bool),sK2)))))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X44,bool),fun(fun(X44,bool),bool)),X43)))) )),
  inference(resolution,[],[f269,f141])).
fof(f3668,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK8(X30,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X31,X32,sK0)),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(X32,sK0),X30),s(fun(fun(X32,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X32,bool),fun(fun(X32,bool),bool)),X31),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X32,sK0),X30),s(fun(fun(X32,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X32,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X32,bool),fun(fun(X32,bool),bool)),X31))),s(fun(sK0,bool),sK2)))))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X32,bool),fun(fun(X32,bool),bool)),X31)))) )),
  inference(resolution,[],[f269,f140])).
fof(f3667,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK8(X24,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X25,X26,sK0)),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(X26,sK0),X24),s(fun(fun(X26,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X26,bool),fun(fun(X26,bool),bool)),X25),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X26,sK0),X24),s(fun(fun(X26,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X26,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X26,bool),fun(fun(X26,bool),bool)),X25))),s(fun(sK0,bool),sK2)))))) | ~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X26,bool),fun(fun(X26,bool),bool)),X25))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f269,f138])).
fof(f3665,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK8(X12,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),X13,X14,sK0)),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(fun(X14,sK0),X12),s(fun(fun(X14,sK0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X14,bool),fun(fun(X14,bool),bool)),X13),s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X14,sK0),X12),s(fun(fun(X14,sK0),bool),'const.util_prob.FUNSET_2'(s(fun(X14,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X14,bool),fun(fun(X14,bool),bool)),X13))),s(fun(sK0,bool),sK2)))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X14,bool),fun(fun(X14,bool),bool)),X13)))) )),
  inference(resolution,[],[f269,f136])).
fof(f3474,plain,(
  ( ! [X28,X26,X24,X25] : (p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(fun(sK0,sK1),X24),X25),s(fun(sK0,sK1),sK6))),s(fun(X24,bool),X26)))) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(fun(fun(sK0,sK1),X24),X25),s(fun(fun(fun(sK0,sK1),X24),bool),'const.util_prob.FUNSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(X24,bool),X26))))) | ~p(s(bool,X28))) ) | $spl90),
  inference(cnf_transformation,[],[f3474_D])).
fof(f3474_D,plain,(
  ( ! [X28,X26,X24,X25] : (p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(fun(sK0,sK1),X24),X25),s(fun(sK0,sK1),sK6))),s(fun(X24,bool),X26)))) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(fun(fun(sK0,sK1),X24),X25),s(fun(fun(fun(sK0,sK1),X24),bool),'const.util_prob.FUNSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(X24,bool),X26))))) | ~p(s(bool,X28))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f3463,plain,(
  ( ! [X103,X105,X102,X106,X104] : (p(s(bool,'const.bool.IN_2'(s(X102,happ(s(fun(fun(sK0,sK1),X102),X103),s(fun(sK0,sK1),sK6))),s(fun(X102,bool),X104)))) | p(s(bool,X105)) | ~p(s(bool,X106)) | s(bool,X106) != s(bool,'const.bool.IN_2'(s(fun(fun(sK0,sK1),X102),X103),s(fun(fun(fun(sK0,sK1),X102),bool),'const.util_prob.FUNSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(X102,bool),X104)))))) )),
  inference(resolution,[],[f529,f200])).
fof(f529,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(fun(sK0,sK1),X1),X2),s(fun(fun(fun(sK0,sK1),X1),bool),'const.util_prob.FUNSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(X1,bool),X3)))))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(fun(sK0,sK1),X1),X2),s(fun(sK0,sK1),sK6))),s(fun(X1,bool),X3))))) )),
  inference(resolution,[],[f122,f192])).
fof(f3452,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (p(s(bool,'const.bool.IN_2'(s(X66,happ(s(fun(fun(sK0,sK1),X66),X67),s(fun(sK0,sK1),sK6))),s(fun(X66,bool),X68)))) | ~p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X70) != s(bool,'const.bool.IN_2'(s(fun(fun(sK0,sK1),X66),X67),s(fun(fun(fun(sK0,sK1),X66),bool),'const.util_prob.FUNSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(X66,bool),X68))))) | s(bool,X69) != s(bool,X71)) )),
  inference(resolution,[],[f529,f142])).
fof(f3450,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(X54,happ(s(fun(fun(sK0,sK1),X54),X55),s(fun(sK0,sK1),sK6))),s(fun(X54,bool),X56)))) | p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(fun(fun(sK0,sK1),X54),X55),s(fun(fun(fun(sK0,sK1),X54),bool),'const.util_prob.FUNSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(X54,bool),X56)))))) )),
  inference(resolution,[],[f529,f141])).
fof(f3448,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(X42,happ(s(fun(fun(sK0,sK1),X42),X43),s(fun(sK0,sK1),sK6))),s(fun(X42,bool),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(fun(fun(sK0,sK1),X42),X43),s(fun(fun(fun(sK0,sK1),X42),bool),'const.util_prob.FUNSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(X42,bool),X44)))))) )),
  inference(resolution,[],[f529,f140])).
fof(f3447,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (p(s(bool,'const.bool.IN_2'(s(X36,happ(s(fun(fun(sK0,sK1),X36),X37),s(fun(sK0,sK1),sK6))),s(fun(X36,bool),X38)))) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(fun(fun(sK0,sK1),X36),X37),s(fun(fun(fun(sK0,sK1),X36),bool),'const.util_prob.FUNSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(X36,bool),X38))))) | s(bool,X39) != s(bool,X41)) )),
  inference(resolution,[],[f529,f138])).
fof(f3445,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(X24,happ(s(fun(fun(sK0,sK1),X24),X25),s(fun(sK0,sK1),sK6))),s(fun(X24,bool),X26)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(fun(fun(sK0,sK1),X24),X25),s(fun(fun(fun(sK0,sK1),X24),bool),'const.util_prob.FUNSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(X24,bool),X26)))))) )),
  inference(resolution,[],[f529,f136])).
fof(f3433,plain,(
  ( ! [X12,X10] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X10))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK0,X10),s(fun(sK0,bool),sK2)))) ) | $spl88),
  inference(cnf_transformation,[],[f3433_D])).
fof(f3433_D,plain,(
  ( ! [X12,X10] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X10))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK0,X10),s(fun(sK0,bool),sK2)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3425,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X52))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(sK0,X52),s(fun(sK0,bool),sK2)))) )),
  inference(resolution,[],[f2651,f200])).
fof(f3414,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X38))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK0,X38),s(fun(sK0,bool),sK2))) | s(bool,X39) != s(bool,X41)) )),
  inference(resolution,[],[f2651,f142])).
fof(f3412,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X30))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK0,X30),s(fun(sK0,bool),sK2)))) )),
  inference(resolution,[],[f2651,f141])).
fof(f3410,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X22))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(sK0,X22),s(fun(sK0,bool),sK2)))) )),
  inference(resolution,[],[f2651,f140])).
fof(f3409,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X18))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(sK0,X18),s(fun(sK0,bool),sK2))) | s(bool,X19) != s(bool,X21)) )),
  inference(resolution,[],[f2651,f138])).
fof(f3407,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X10))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(sK0,X10),s(fun(sK0,bool),sK2)))) )),
  inference(resolution,[],[f2651,f136])).
fof(f3386,plain,(
  ( ! [X4,X2] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X2)))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),X2),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) | $spl86),
  inference(cnf_transformation,[],[f3386_D])).
fof(f3386_D,plain,(
  ( ! [X4,X2] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X2)))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),X2),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f3377,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),X44),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f693,f200])).
fof(f693,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),X0),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X0))))) )),
  inference(resolution,[],[f125,f126])).
fof(f125,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))),
  inference(cnf_transformation,[],[f80])).
fof(f3366,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X30)))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),X30),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f693,f142])).
fof(f3364,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X22)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),X22),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f693,f141])).
fof(f3362,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),X14),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f693,f140])).
fof(f3361,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X10)))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),X10),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f693,f138])).
fof(f3359,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X2)))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),X2),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f693,f136])).
fof(f3300,plain,(
  ( ! [X4,X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X2)))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(fun(sK0,sK1),bool),X2)))) ) | $spl84),
  inference(cnf_transformation,[],[f3300_D])).
fof(f3300_D,plain,(
  ( ! [X4,X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X2)))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(fun(sK0,sK1),bool),X2)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f3291,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(fun(sK0,sK1),bool),X44)))) )),
  inference(resolution,[],[f530,f200])).
fof(f530,plain,(
  ( ! [X4] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(fun(sK0,sK1),bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X4))))) )),
  inference(resolution,[],[f122,f126])).
fof(f3280,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X30)))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(fun(sK0,sK1),bool),X30))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f530,f142])).
fof(f3278,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X22)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(fun(sK0,sK1),bool),X22)))) )),
  inference(resolution,[],[f530,f141])).
fof(f3276,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X14)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(fun(sK0,sK1),bool),X14)))) )),
  inference(resolution,[],[f530,f140])).
fof(f3275,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X10)))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(fun(sK0,sK1),bool),X10))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f530,f138])).
fof(f3273,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),X2)))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))),s(fun(fun(sK0,sK1),bool),X2)))) )),
  inference(resolution,[],[f530,f136])).
fof(f3234,plain,(
  ( ! [X12] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X12)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X12,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))))) ) | $spl82),
  inference(cnf_transformation,[],[f3234_D])).
fof(f3234_D,plain,(
  ( ! [X12] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X12)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X12,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3235,plain,(
  ( ! [X14,X15] : (p(s(bool,X15)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X14)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X14,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))))) ) | $spl36),
  inference(subsumption_resolution,[],[f3215,f182])).
fof(f3215,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X15)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X14)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X14,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))))) ) | $spl36),
  inference(superposition,[],[f1921,f182])).
fof(f3233,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X12)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X12,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))))) ) | $spl36),
  inference(subsumption_resolution,[],[f3214,f182])).
fof(f3214,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X13)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X12)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X12,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))))) ) | $spl36),
  inference(superposition,[],[f1921,f184])).
fof(f184,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.EQ_CLAUSES|split|2')).
fof(f3222,plain,(
  ( ! [X28] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X28),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl80),
  inference(cnf_transformation,[],[f3222_D])).
fof(f3222_D,plain,(
  ( ! [X28] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X28),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f3220,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X3,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X3))))) ) | $spl78),
  inference(cnf_transformation,[],[f3220_D])).
fof(f3220_D,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X3,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X3))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f3221,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X29,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X29)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X28),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3202,f3013])).
fof(f3013,plain,(
  ( ! [X169] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X169,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X169),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(trivial_inequality_removal,[],[f2982])).
fof(f2982,plain,(
  ( ! [X169] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X169,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X169),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f183,f2386])).
fof(f183,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f115])).
fof(f3202,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X29,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X28,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X29)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X28),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl36),
  inference(superposition,[],[f1921,f2386])).
fof(f3217,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X11,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X10)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X11))))) ) | $spl36),
  inference(duplicate_literal_removal,[],[f3199])).
fof(f3199,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X11,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X10)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X11)))) | p(s(bool,X10))) ) | $spl36),
  inference(superposition,[],[f1921,f182])).
fof(f3218,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X9,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X8)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X9))))) ) | $spl36),
  inference(duplicate_literal_removal,[],[f3198])).
fof(f3198,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X9,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X8)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X9)))) | p(s(bool,X8))) ) | $spl36),
  inference(superposition,[],[f1921,f184])).
fof(f3219,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X3,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X2)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X3))))) ) | $spl36),
  inference(subsumption_resolution,[],[f3195,f185])).
fof(f3195,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X3,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X2)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X3)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl36),
  inference(superposition,[],[f1921,f199])).
fof(f3149,plain,(
  ( ! [X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),X9),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | p(s(bool,X11)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),X9))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl76),
  inference(cnf_transformation,[],[f3149_D])).
fof(f3149_D,plain,(
  ( ! [X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),X9),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | p(s(bool,X11)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),X9))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK3)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f3141,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),X50),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),X50))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f557,f200])).
fof(f3131,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),X37),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),X37))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X39) != s(bool,X40)) )),
  inference(resolution,[],[f557,f142])).
fof(f3128,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),X25),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),X25))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f557,f141])).
fof(f3127,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),X21),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),X21))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f557,f140])).
fof(f3126,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),X17),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),X17))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f557,f138])).
fof(f3124,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),X9),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),X9))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f557,f136])).
fof(f3028,plain,(
  ( ! [X36,X34] : (s(bool,X36) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X34,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X34),s(fun(fun(sK0,bool),bool),sK4))) | ~p(s(bool,X36))) ) | $spl70),
  inference(cnf_transformation,[],[f3028_D])).
fof(f3028_D,plain,(
  ( ! [X36,X34] : (s(bool,X36) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X34,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X34),s(fun(fun(sK0,bool),bool),sK4))) | ~p(s(bool,X36))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3081,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X44),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X44,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f3013,f200])).
fof(f3070,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X30),s(fun(fun(sK0,bool),bool),sK4))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X30,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f3013,f142])).
fof(f3068,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X22),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X22,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f3013,f141])).
fof(f3066,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X14),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X14,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f3013,f140])).
fof(f3065,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X10),s(fun(fun(sK0,bool),bool),sK4))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X10,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f3013,f138])).
fof(f3063,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X2),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X2,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f3013,f136])).
fof(f3060,plain,(
  ( ! [X215] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X215),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl74),
  inference(cnf_transformation,[],[f3060_D])).
fof(f3060_D,plain,(
  ( ! [X215] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X215),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3025,plain,(
  ( ! [X22] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X22,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X22),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl68),
  inference(cnf_transformation,[],[f3025_D])).
fof(f3025_D,plain,(
  ( ! [X22] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X22,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X22),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3059,plain,(
  ( ! [X215,X214] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X214,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X215),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X214),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl40),
  inference(subsumption_resolution,[],[f3007,f1615])).
fof(f3007,plain,(
  ( ! [X215,X214] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X215),s(fun(fun(sK0,bool),bool),sK3))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X214,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X215),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X214),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl40),
  inference(superposition,[],[f1628,f2386])).
fof(f3056,plain,(
  ( ! [X205] : (s(bool,'T_0') = s(bool,X205)) ) | $spl72),
  inference(cnf_transformation,[],[f3056_D])).
fof(f3056_D,plain,(
  ( ! [X205] : (s(bool,'T_0') = s(bool,X205)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3055,plain,(
  ( ! [X204,X205] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X204,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,'T_0') = s(bool,X205) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X204),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl6),
  inference(subsumption_resolution,[],[f3002,f199])).
fof(f3002,plain,(
  ( ! [X204,X205] : (s(bool,'F_0') != s(bool,X205) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X204,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,'T_0') = s(bool,X205) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X204),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl6),
  inference(superposition,[],[f508,f2386])).
fof(f508,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(subsumption_resolution,[],[f479,f152])).
fof(f152,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X13)))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.IMP_CLAUSES|split|4')).
fof(f479,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,'F_0')) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f230,f199])).
fof(f3054,plain,(
  ( ! [X200,X201] : (p(s(bool,X201)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X200,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X200),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl6),
  inference(subsumption_resolution,[],[f3000,f182])).
fof(f3000,plain,(
  ( ! [X200,X201] : (s(bool,'F_0') != s(bool,X201) | p(s(bool,X201)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X200,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X200),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl6),
  inference(superposition,[],[f230,f2386])).
fof(f3051,plain,(
  ( ! [X191,X189,X190] : (p(s(bool,X190)) | ~p(s(bool,X191)) | s(bool,X191) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X189,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X189),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2996,f168])).
fof(f168,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) & p(s(bool,X13))) <=> p(s(bool,'F_0')))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.AND_CLAUSES|split|2')).
fof(f2996,plain,(
  ( ! [X191,X189,X190] : (p(s(bool,'F_0')) | p(s(bool,X190)) | ~p(s(bool,X191)) | s(bool,X191) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X189,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X189),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f200,f2386])).
fof(f3050,plain,(
  ( ! [X185,X184,X183] : (p(s(bool,X185)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X183,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X184)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X183),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2994,f182])).
fof(f2994,plain,(
  ( ! [X185,X184,X183] : (s(bool,'F_0') != s(bool,X184) | p(s(bool,X185)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X183,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X184)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X183),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f200,f2386])).
fof(f3041,plain,(
  ( ! [X144,X142,X145,X143] : (~p(s(bool,X143)) | ~p(s(bool,X144)) | s(bool,X144) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X142,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X143) != s(bool,X145) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X142),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2965,f152])).
fof(f2965,plain,(
  ( ! [X144,X142,X145,X143] : (p(s(bool,'F_0')) | ~p(s(bool,X143)) | ~p(s(bool,X144)) | s(bool,X144) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X142,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X143) != s(bool,X145) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X142),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f142,f2386])).
fof(f3040,plain,(
  ( ! [X140,X138,X141,X139] : (p(s(bool,X139)) | ~p(s(bool,X140)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X138,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,X140) != s(bool,X141) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X138),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2964,f182])).
fof(f2964,plain,(
  ( ! [X140,X138,X141,X139] : (s(bool,'F_0') != s(bool,X139) | p(s(bool,X139)) | ~p(s(bool,X140)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X138,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,X140) != s(bool,X141) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X138),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f142,f2386])).
fof(f3039,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X111)) | ~p(s(bool,X112)) | ~p(s(bool,X113)) | s(bool,X112) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X110,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X110),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2957,f168])).
fof(f2957,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,'F_0')) | p(s(bool,X111)) | ~p(s(bool,X112)) | ~p(s(bool,X113)) | s(bool,X112) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X110,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X110),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f141,f2386])).
fof(f3038,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X104)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X102,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | ~p(s(bool,X105)) | p(s(bool,X103)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X102),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2955,f182])).
fof(f2955,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'F_0') != s(bool,X103) | p(s(bool,X104)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X102,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | ~p(s(bool,X105)) | p(s(bool,X103)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X102),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f141,f2386])).
fof(f3035,plain,(
  ( ! [X92,X90,X93,X91] : (p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,X93) | s(bool,X92) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X90,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X90),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2952,f168])).
fof(f2952,plain,(
  ( ! [X92,X90,X93,X91] : (p(s(bool,'F_0')) | p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,X93) | s(bool,X92) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X90,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X90),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f140,f2386])).
fof(f3033,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X76)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X74,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X75)) | s(bool,X76) != s(bool,X77) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X74),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2948,f182])).
fof(f2948,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'F_0') != s(bool,X75) | p(s(bool,X76)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X74,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X75)) | s(bool,X76) != s(bool,X77) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X74),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f140,f2386])).
fof(f3032,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X62,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X63) != s(bool,X65) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X62),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2945,f152])).
fof(f2945,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,'F_0')) | ~p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X62,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X63) != s(bool,X65) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X62),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f138,f2386])).
fof(f3031,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X54,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,X56) != s(bool,X57) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X54),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2943,f182])).
fof(f2943,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,'F_0') != s(bool,X55) | p(s(bool,X55)) | ~p(s(bool,X56)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X54,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | s(bool,X56) != s(bool,X57) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X54),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f138,f2386])).
fof(f3027,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X34,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X34),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2938,f168])).
fof(f2938,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'F_0')) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X34,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X34),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f136,f2386])).
fof(f3024,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,X24)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X22,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X23)) | s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X22),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(subsumption_resolution,[],[f2935,f182])).
fof(f2935,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X22,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X23)) | s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X22),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(superposition,[],[f136,f2386])).
fof(f2901,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl66),
  inference(cnf_transformation,[],[f2901_D])).
fof(f2901_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2892,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl31),
  inference(resolution,[],[f733,f200])).
fof(f2881,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X22) != s(bool,X24)) ) | $spl31),
  inference(resolution,[],[f733,f142])).
fof(f2879,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl31),
  inference(resolution,[],[f733,f141])).
fof(f2877,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl31),
  inference(resolution,[],[f733,f140])).
fof(f2876,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X7) != s(bool,X9)) ) | $spl31),
  inference(resolution,[],[f733,f138])).
fof(f2874,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl31),
  inference(resolution,[],[f733,f136])).
fof(f2806,plain,(
  ( ! [X39,X41] : (s(bool,X41) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X39,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X39))) | p(s(bool,X41))) ) | $spl62),
  inference(cnf_transformation,[],[f2806_D])).
fof(f2806_D,plain,(
  ( ! [X39,X41] : (s(bool,X41) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X39,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X39))) | p(s(bool,X41))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2857,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X47))) | p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X49) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X47,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f2789,f200])).
fof(f2789,plain,(
  ( ! [X174] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X174,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X174)))) )),
  inference(trivial_inequality_removal,[],[f2766])).
fof(f2766,plain,(
  ( ! [X174] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X174,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X174)))) )),
  inference(superposition,[],[f187,f2325])).
fof(f2847,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X34))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X34,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X36) != s(bool,X37)) )),
  inference(resolution,[],[f2789,f142])).
fof(f2844,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X22))) | p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X22,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f2789,f141])).
fof(f2843,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X18))) | p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X18,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f2789,f140])).
fof(f2842,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X14))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X14,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f2789,f138])).
fof(f2840,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X6))) | p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X6,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f2789,f136])).
fof(f2803,plain,(
  ( ! [X15] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X15,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X15)))) ) | $spl60),
  inference(cnf_transformation,[],[f2803_D])).
fof(f2803_D,plain,(
  ( ! [X15] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X15,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X15)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2835,plain,(
  ( ! [X213,X214] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X213,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X214),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X213)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2784,f1615])).
fof(f2784,plain,(
  ( ! [X213,X214] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X214),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X213,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X214),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X213)))) ) | $spl36),
  inference(superposition,[],[f1939,f2325])).
fof(f2833,plain,(
  ( ! [X206,X205] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X205,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,X206) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X205)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2780,f199])).
fof(f2780,plain,(
  ( ! [X206,X205] : (s(bool,'T_0') != s(bool,X206) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X205,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,X206) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X205)))) ) | $spl6),
  inference(superposition,[],[f522,f2325])).
fof(f2831,plain,(
  ( ! [X196,X195] : (s(bool,'T_0') != s(bool,X196) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X195,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X195)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2775,f185])).
fof(f2775,plain,(
  ( ! [X196,X195] : (s(bool,'T_0') != s(bool,X196) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X195,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X196)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X195)))) ) | $spl6),
  inference(superposition,[],[f230,f2325])).
fof(f2830,plain,(
  ( ! [X194,X192,X193] : (p(s(bool,X193)) | p(s(bool,X194)) | s(bool,X194) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X192,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X192)))) )),
  inference(subsumption_resolution,[],[f2774,f154])).
fof(f2774,plain,(
  ( ! [X194,X192,X193] : (~p(s(bool,'T_0')) | p(s(bool,X193)) | p(s(bool,X194)) | s(bool,X194) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X192,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X192)))) )),
  inference(superposition,[],[f200,f2325])).
fof(f2829,plain,(
  ( ! [X185,X184,X183] : (s(bool,'T_0') != s(bool,X184) | p(s(bool,X185)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X183,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X183)))) )),
  inference(subsumption_resolution,[],[f2771,f185])).
fof(f2771,plain,(
  ( ! [X185,X184,X183] : (s(bool,'T_0') != s(bool,X184) | p(s(bool,X185)) | ~p(s(bool,X184)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X183,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X183)))) )),
  inference(superposition,[],[f200,f2325])).
fof(f2822,plain,(
  ( ! [X149,X147,X150,X148] : (p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X148) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X147,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X149) != s(bool,X150) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X147)))) )),
  inference(subsumption_resolution,[],[f2743,f155])).
fof(f2743,plain,(
  ( ! [X149,X147,X150,X148] : (~p(s(bool,'T_0')) | p(s(bool,X148)) | ~p(s(bool,X149)) | s(bool,X148) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X147,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X149) != s(bool,X150) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X147)))) )),
  inference(superposition,[],[f142,f2325])).
fof(f2818,plain,(
  ( ! [X134,X132,X133,X131] : (s(bool,'T_0') != s(bool,X132) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X131,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X133)) | s(bool,X133) != s(bool,X134) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X131)))) )),
  inference(subsumption_resolution,[],[f2739,f185])).
fof(f2739,plain,(
  ( ! [X134,X132,X133,X131] : (s(bool,'T_0') != s(bool,X132) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X131,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X133)) | ~p(s(bool,X132)) | s(bool,X133) != s(bool,X134) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X131)))) )),
  inference(superposition,[],[f142,f2325])).
fof(f2816,plain,(
  ( ! [X118,X116,X117,X115] : (p(s(bool,X116)) | ~p(s(bool,X117)) | p(s(bool,X118)) | s(bool,X118) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X115,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X115)))) )),
  inference(subsumption_resolution,[],[f2735,f155])).
fof(f2735,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'T_0')) | p(s(bool,X116)) | ~p(s(bool,X117)) | p(s(bool,X118)) | s(bool,X118) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X115,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X115)))) )),
  inference(superposition,[],[f141,f2325])).
fof(f2815,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'T_0') != s(bool,X104) | p(s(bool,X105)) | ~p(s(bool,X106)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X103,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X103)))) )),
  inference(subsumption_resolution,[],[f2732,f185])).
fof(f2732,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'T_0') != s(bool,X104) | p(s(bool,X105)) | ~p(s(bool,X104)) | ~p(s(bool,X106)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X103,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X103)))) )),
  inference(superposition,[],[f141,f2325])).
fof(f2814,plain,(
  ( ! [X97,X95,X98,X96] : (p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X96) != s(bool,X98) | s(bool,X97) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X95,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X95)))) )),
  inference(subsumption_resolution,[],[f2730,f154])).
fof(f2730,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X96) != s(bool,X98) | s(bool,X97) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X95,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X95)))) )),
  inference(superposition,[],[f140,f2325])).
fof(f2812,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') != s(bool,X76) | p(s(bool,X77)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X75,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X77) != s(bool,X78) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X75)))) )),
  inference(subsumption_resolution,[],[f2725,f185])).
fof(f2725,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') != s(bool,X76) | p(s(bool,X77)) | ~p(s(bool,X76)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X75,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X77) != s(bool,X78) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X75)))) )),
  inference(superposition,[],[f140,f2325])).
fof(f2811,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X67,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X69) != s(bool,X70) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X67)))) )),
  inference(subsumption_resolution,[],[f2723,f155])).
fof(f2723,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X67,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X69) != s(bool,X70) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X67)))) )),
  inference(superposition,[],[f138,f2325])).
fof(f2807,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') != s(bool,X56) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X55,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X57)) | s(bool,X57) != s(bool,X58) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X55)))) )),
  inference(subsumption_resolution,[],[f2720,f185])).
fof(f2720,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') != s(bool,X56) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X55,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X57)) | ~p(s(bool,X56)) | s(bool,X57) != s(bool,X58) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X55)))) )),
  inference(superposition,[],[f138,f2325])).
fof(f2805,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X39,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X39)))) )),
  inference(subsumption_resolution,[],[f2716,f154])).
fof(f2716,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X39,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X39)))) )),
  inference(superposition,[],[f136,f2325])).
fof(f2802,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,X16) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X15,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X15)))) )),
  inference(subsumption_resolution,[],[f2710,f185])).
fof(f2710,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X16)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X15,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X15)))) )),
  inference(superposition,[],[f136,f2325])).
fof(f2660,plain,(
  ( ! [X7,X9] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2)))) | p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X7))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) | $spl58),
  inference(cnf_transformation,[],[f2660_D])).
fof(f2660_D,plain,(
  ( ! [X7,X9] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2)))) | p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X7))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2653,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X48),s(fun(sK0,bool),sK2)))) | p(s(bool,X49)) | p(s(bool,X50)) | s(bool,X50) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X48))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f1553,f200])).
fof(f2643,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X35),s(fun(sK0,bool),sK2)))) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X35))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X37) != s(bool,X38)) )),
  inference(resolution,[],[f1553,f142])).
fof(f2640,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X23),s(fun(sK0,bool),sK2)))) | p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X23))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f1553,f141])).
fof(f2639,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK2)))) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X19))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f1553,f140])).
fof(f2638,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X15),s(fun(sK0,bool),sK2)))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X15))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X17) != s(bool,X18)) )),
  inference(resolution,[],[f1553,f138])).
fof(f2636,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2)))) | p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK6),s(sK0,X7))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f1553,f136])).
fof(f2398,plain,(
  ( ! [X4,X2] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X2),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X2,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl56),
  inference(cnf_transformation,[],[f2398_D])).
fof(f2398_D,plain,(
  ( ! [X4,X2] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X2),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X2,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2388,plain,(
  ( ! [X45,X43,X44] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X43,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X43),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f1896,f200])).
fof(f2378,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X30,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X30),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X32) != s(bool,X33)) )),
  inference(resolution,[],[f1896,f142])).
fof(f2375,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X18,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X18),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f1896,f141])).
fof(f2374,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X14,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X14),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f1896,f140])).
fof(f2373,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X10,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X10),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f1896,f138])).
fof(f2371,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X2,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X2),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f1896,f136])).
fof(f2337,plain,(
  ( ! [X4,X2] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X2))) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X2,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl54),
  inference(cnf_transformation,[],[f2337_D])).
fof(f2337_D,plain,(
  ( ! [X4,X2] : (s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X2))) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X2,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2327,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X43,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X43)))) )),
  inference(resolution,[],[f1597,f200])).
fof(f2317,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X30,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X30))) | s(bool,X32) != s(bool,X33)) )),
  inference(resolution,[],[f1597,f142])).
fof(f2314,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X18,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X18)))) )),
  inference(resolution,[],[f1597,f141])).
fof(f2313,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X14,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X14)))) )),
  inference(resolution,[],[f1597,f140])).
fof(f2312,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X10,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X10))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f1597,f138])).
fof(f2310,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(X2,sK3,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),X2)))) )),
  inference(resolution,[],[f1597,f136])).
fof(f1952,plain,(
  ( ! [X14,X17,X15] : (s(bool,X17) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),sK3))) | ~p(s(bool,X17)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl52),
  inference(cnf_transformation,[],[f1952_D])).
fof(f1952_D,plain,(
  ( ! [X14,X17,X15] : (s(bool,X17) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),sK3))) | ~p(s(bool,X17)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1943,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X74)) | s(bool,X74) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X75),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X76)) | ~p(s(bool,X77)) | s(bool,X77) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X75),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl36),
  inference(resolution,[],[f1537,f200])).
fof(f1932,plain,(
  ( ! [X52,X50,X53,X51,X49] : (p(s(bool,X49)) | s(bool,X49) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X50),s(fun(fun(sK0,bool),bool),sK4))) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X50),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X51) != s(bool,X53)) ) | $spl36),
  inference(resolution,[],[f1537,f142])).
fof(f1930,plain,(
  ( ! [X39,X43,X41,X42,X40] : (p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X40),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X40),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl36),
  inference(resolution,[],[f1537,f141])).
fof(f1928,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X30),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X30),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl36),
  inference(resolution,[],[f1537,f140])).
fof(f1927,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X25),s(fun(fun(sK0,bool),bool),sK4))) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X25),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X26) != s(bool,X28)) ) | $spl36),
  inference(resolution,[],[f1537,f138])).
fof(f1925,plain,(
  ( ! [X14,X17,X15,X18,X16] : (p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl36),
  inference(resolution,[],[f1537,f136])).
fof(f1910,plain,(
  ( ! [X4,X2] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X2,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | ~p(s(bool,X4)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X2),s(fun(fun(sK0,bool),bool),sK4))))) ) | $spl50),
  inference(cnf_transformation,[],[f1910_D])).
fof(f1910_D,plain,(
  ( ! [X4,X2] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X2,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | ~p(s(bool,X4)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X2),s(fun(fun(sK0,bool),bool),sK4))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1900,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X44),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X44,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f1508,f200])).
fof(f1889,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X30),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X30,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X31) != s(bool,X33)) )),
  inference(resolution,[],[f1508,f142])).
fof(f1887,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X22),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X22,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f1508,f141])).
fof(f1885,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X14),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X14,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f1508,f140])).
fof(f1884,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X10),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X10,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X11) != s(bool,X13)) )),
  inference(resolution,[],[f1508,f138])).
fof(f1882,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),X2),s(fun(fun(sK0,bool),bool),sK4)))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK7(sK4,X2,fun(sK0,bool))),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f1508,f136])).
fof(f1841,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) | $spl48),
  inference(cnf_transformation,[],[f1841_D])).
fof(f1841_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1834,plain,(
  ( ! [X31,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) | $spl26),
  inference(resolution,[],[f724,f200])).
fof(f1824,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X29) != s(bool,X30)) ) | $spl26),
  inference(resolution,[],[f724,f142])).
fof(f1821,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) | $spl26),
  inference(resolution,[],[f724,f141])).
fof(f1820,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) | $spl26),
  inference(resolution,[],[f724,f140])).
fof(f1819,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X14) != s(bool,X15)) ) | $spl26),
  inference(resolution,[],[f724,f138])).
fof(f1817,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) | $spl26),
  inference(resolution,[],[f724,f136])).
fof(f1748,plain,(
  ( ! [X48,X46] : (p(s(bool,X48)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X46),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X46),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl44),
  inference(cnf_transformation,[],[f1748_D])).
fof(f1748_D,plain,(
  ( ! [X48,X46] : (p(s(bool,X48)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X46),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X46),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1799,plain,(
  ( ! [X52,X53,X51] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X51),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X53) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X51),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f1727,f200])).
fof(f1727,plain,(
  ( ! [X181] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X181),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X181),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(trivial_inequality_removal,[],[f1706])).
fof(f1706,plain,(
  ( ! [X181] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X181),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X181),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f187,f1615])).
fof(f1789,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X38),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X38),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X40) != s(bool,X41)) )),
  inference(resolution,[],[f1727,f142])).
fof(f1786,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X26),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X26),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f1727,f141])).
fof(f1785,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f1727,f140])).
fof(f1784,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X20) != s(bool,X21)) )),
  inference(resolution,[],[f1727,f138])).
fof(f1782,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK3))) | p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f1727,f136])).
fof(f1745,plain,(
  ( ! [X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl42),
  inference(cnf_transformation,[],[f1745_D])).
fof(f1745_D,plain,(
  ( ! [X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK3)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1775,plain,(
  ( ! [X212,X213] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X212),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,X213) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X212),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl6),
  inference(subsumption_resolution,[],[f1720,f199])).
fof(f1720,plain,(
  ( ! [X212,X213] : (s(bool,'T_0') != s(bool,X213) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X212),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,X213) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X212),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl6),
  inference(superposition,[],[f522,f1615])).
fof(f1773,plain,(
  ( ! [X202,X203] : (s(bool,'T_0') != s(bool,X203) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X202),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X202),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl6),
  inference(subsumption_resolution,[],[f1715,f185])).
fof(f1715,plain,(
  ( ! [X202,X203] : (s(bool,'T_0') != s(bool,X203) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X202),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X203)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X202),s(fun(fun(sK0,bool),bool),sK3)))) ) | $spl6),
  inference(superposition,[],[f230,f1615])).
fof(f1772,plain,(
  ( ! [X200,X199,X201] : (p(s(bool,X200)) | p(s(bool,X201)) | s(bool,X201) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X199),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X199),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1714,f154])).
fof(f1714,plain,(
  ( ! [X200,X199,X201] : (~p(s(bool,'T_0')) | p(s(bool,X200)) | p(s(bool,X201)) | s(bool,X201) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X199),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X199),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f200,f1615])).
fof(f1771,plain,(
  ( ! [X191,X192,X190] : (s(bool,'T_0') != s(bool,X191) | p(s(bool,X192)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X190),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X190),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1711,f185])).
fof(f1711,plain,(
  ( ! [X191,X192,X190] : (s(bool,'T_0') != s(bool,X191) | p(s(bool,X192)) | ~p(s(bool,X191)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X190),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X190),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f200,f1615])).
fof(f1764,plain,(
  ( ! [X156,X154,X157,X155] : (p(s(bool,X155)) | ~p(s(bool,X156)) | s(bool,X155) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X154),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X156) != s(bool,X157) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X154),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1683,f155])).
fof(f1683,plain,(
  ( ! [X156,X154,X157,X155] : (~p(s(bool,'T_0')) | p(s(bool,X155)) | ~p(s(bool,X156)) | s(bool,X155) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X154),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X156) != s(bool,X157) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X154),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f142,f1615])).
fof(f1760,plain,(
  ( ! [X140,X138,X141,X139] : (s(bool,'T_0') != s(bool,X139) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X138),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X140)) | s(bool,X140) != s(bool,X141) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X138),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1679,f185])).
fof(f1679,plain,(
  ( ! [X140,X138,X141,X139] : (s(bool,'T_0') != s(bool,X139) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X138),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X140)) | ~p(s(bool,X139)) | s(bool,X140) != s(bool,X141) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X138),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f142,f1615])).
fof(f1758,plain,(
  ( ! [X125,X123,X124,X122] : (p(s(bool,X123)) | ~p(s(bool,X124)) | p(s(bool,X125)) | s(bool,X125) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X122),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X122),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1675,f155])).
fof(f1675,plain,(
  ( ! [X125,X123,X124,X122] : (~p(s(bool,'T_0')) | p(s(bool,X123)) | ~p(s(bool,X124)) | p(s(bool,X125)) | s(bool,X125) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X122),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X122),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f141,f1615])).
fof(f1757,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') != s(bool,X111) | p(s(bool,X112)) | ~p(s(bool,X113)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X110),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X110),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1672,f185])).
fof(f1672,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') != s(bool,X111) | p(s(bool,X112)) | ~p(s(bool,X111)) | ~p(s(bool,X113)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X110),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X110),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f141,f1615])).
fof(f1756,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X103)) | p(s(bool,X104)) | s(bool,X103) != s(bool,X105) | s(bool,X104) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X102),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X102),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1670,f154])).
fof(f1670,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | p(s(bool,X103)) | p(s(bool,X104)) | s(bool,X103) != s(bool,X105) | s(bool,X104) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X102),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X102),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f140,f1615])).
fof(f1754,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') != s(bool,X83) | p(s(bool,X84)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X82),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X82),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1665,f185])).
fof(f1665,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') != s(bool,X83) | p(s(bool,X84)) | ~p(s(bool,X83)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X82),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X82),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f140,f1615])).
fof(f1753,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X75) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X74),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X76) != s(bool,X77) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X74),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1663,f155])).
fof(f1663,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X75) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X74),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X76) != s(bool,X77) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X74),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f138,f1615])).
fof(f1749,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'T_0') != s(bool,X63) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X62),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,X65) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X62),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1660,f185])).
fof(f1660,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'T_0') != s(bool,X63) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X62),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,X64)) | ~p(s(bool,X63)) | s(bool,X64) != s(bool,X65) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X62),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f138,f1615])).
fof(f1747,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X46),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X46),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1656,f154])).
fof(f1656,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X46),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X46),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f136,f1615])).
fof(f1744,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') != s(bool,X23) | p(s(bool,X24)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(subsumption_resolution,[],[f1650,f185])).
fof(f1650,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') != s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,X23)) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK4)))) | s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(superposition,[],[f136,f1615])).
fof(f1635,plain,(
  $false | $spl29),
  inference(subsumption_resolution,[],[f726,f555])).
fof(f1628,plain,(
  ( ! [X12,X10] : (s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK3))) | ~p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl40),
  inference(cnf_transformation,[],[f1628_D])).
fof(f1628_D,plain,(
  ( ! [X12,X10] : (s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK3))) | ~p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1619,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X52),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X52),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f1527,f200])).
fof(f1608,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X38),s(fun(fun(sK0,bool),bool),sK4))) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X38),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X39) != s(bool,X41)) )),
  inference(resolution,[],[f1527,f142])).
fof(f1606,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X30),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X30),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f1527,f141])).
fof(f1604,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f1527,f140])).
fof(f1603,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),sK4))) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),sK3))) | s(bool,X19) != s(bool,X21)) )),
  inference(resolution,[],[f1527,f138])).
fof(f1601,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK4))) | p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK3)))) )),
  inference(resolution,[],[f1527,f136])).
fof(f1580,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) ) | $spl38),
  inference(cnf_transformation,[],[f1580_D])).
fof(f1580_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1573,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) )),
  inference(resolution,[],[f555,f200])).
fof(f1563,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f555,f142])).
fof(f1560,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) )),
  inference(resolution,[],[f555,f141])).
fof(f1559,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) )),
  inference(resolution,[],[f555,f140])).
fof(f1558,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f555,f138])).
fof(f1556,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) )),
  inference(resolution,[],[f555,f136])).
fof(f1537,plain,(
  ( ! [X12,X10] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl36),
  inference(cnf_transformation,[],[f1537_D])).
fof(f1537_D,plain,(
  ( ! [X12,X10] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1529,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X51),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X53) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X51),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f202,f200])).
fof(f1519,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X38),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X38),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X40) != s(bool,X41)) )),
  inference(resolution,[],[f202,f142])).
fof(f1516,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X26),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X26),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f202,f141])).
fof(f1515,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X22),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f202,f140])).
fof(f1514,plain,(
  ( ! [X21,X19,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X20) != s(bool,X21)) )),
  inference(resolution,[],[f202,f138])).
fof(f1512,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK3)))) | p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X10),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f202,f136])).
fof(f1152,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3)))))) ) | $spl34),
  inference(cnf_transformation,[],[f1152_D])).
fof(f1152_D,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3)))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1138,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X37) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3)))))) )),
  inference(resolution,[],[f525,f200])).
fof(f525,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))))))),
  inference(resolution,[],[f122,f129])).
fof(f129,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f1128,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f525,f142])).
fof(f1125,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3)))))) )),
  inference(resolution,[],[f525,f141])).
fof(f1124,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3)))))) )),
  inference(resolution,[],[f525,f140])).
fof(f1123,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f525,f138])).
fof(f1121,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3)))))) )),
  inference(resolution,[],[f525,f136])).
fof(f1069,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f1060])).
fof(f1060,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(superposition,[],[f229,f218])).
fof(f218,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))),
  inference(resolution,[],[f124,f186])).
fof(f1068,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f229])).
fof(f734,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) | $spl32),
  inference(cnf_transformation,[],[f734_D])).
fof(f734_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f712,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f125,f200])).
fof(f701,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f125,f142])).
fof(f699,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f125,f141])).
fof(f697,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f125,f140])).
fof(f696,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f125,f138])).
fof(f694,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f125,f136])).
fof(f733,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | $spl31),
  inference(cnf_transformation,[],[f733_D])).
fof(f733_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f726,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) | $spl29),
  inference(cnf_transformation,[],[f726_D])).
fof(f726_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f731,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4))))),
  inference(forward_demodulation,[],[f730,f196])).
fof(f730,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),sK2))),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))))),
  inference(forward_demodulation,[],[f729,f196])).
fof(f729,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))),s(fun(fun(sK0,bool),bool),sK4)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))))),
  inference(forward_demodulation,[],[f728,f195])).
fof(f728,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))),s(fun(fun(sK0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))))),
  inference(subsumption_resolution,[],[f727,f123])).
fof(f727,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))),s(fun(fun(sK0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))),
  inference(subsumption_resolution,[],[f692,f526])).
fof(f526,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))),
  inference(resolution,[],[f122,f130])).
fof(f130,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X1,X0),X4),s(fun(fun(X1,X0),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(X1,bool),fun(fun(X1,bool),bool)),X2),s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X3))))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f692,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK6),s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)))),s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))),s(fun(fun(sK0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))),
  inference(resolution,[],[f125,f134])).
fof(f724,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | $spl26),
  inference(cnf_transformation,[],[f724_D])).
fof(f724_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f722,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))),
  inference(forward_demodulation,[],[f721,f196])).
fof(f721,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))))))),
  inference(subsumption_resolution,[],[f720,f123])).
fof(f720,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))),
  inference(subsumption_resolution,[],[f691,f526])).
fof(f691,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK1,bool),sK8(sK6,sK5,'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)),sK0,sK1)),s(fun(fun(sK1,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.util_prob.FUNSET_2'(s(fun(sK0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))),s(fun(sK1,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))))),
  inference(resolution,[],[f125,f133])).
fof(f592,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))) ) | $spl24),
  inference(cnf_transformation,[],[f592_D])).
fof(f592_D,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f585,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X37) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))) )),
  inference(resolution,[],[f526,f200])).
fof(f575,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f526,f142])).
fof(f572,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))) )),
  inference(resolution,[],[f526,f141])).
fof(f571,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))) )),
  inference(resolution,[],[f526,f140])).
fof(f570,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f526,f138])).
fof(f568,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))) )),
  inference(resolution,[],[f526,f136])).
fof(f558,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) | $spl22),
  inference(cnf_transformation,[],[f558_D])).
fof(f558_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f548,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f122,f200])).
fof(f538,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f122,f142])).
fof(f535,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f122,f141])).
fof(f534,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f122,f140])).
fof(f533,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5))))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f122,f138])).
fof(f531,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(fun(sK0,sK1),sK6),s(fun(fun(sK0,sK1),bool),'const.measure.measurable_2'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK3))),s('type.pair.prod'(fun(sK1,bool),fun(fun(sK1,bool),bool)),sK5)))))) )),
  inference(resolution,[],[f122,f136])).
fof(f476,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f452,f232])).
fof(f452,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,'const.bool.IN_2'(s(X111,sK7(X112,X113,X111)),s(fun(X111,bool),X113))))) ) | $spl10),
  inference(resolution,[],[f232,f127])).
fof(f475,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f445,f232])).
fof(f445,plain,(
  ( ! [X80,X83,X81,X84,X82] : (p(s(bool,'const.bool.IN_2'(s(X80,sK9(X81,X82,X83,X84,X80)),s(fun(X80,bool),X82))))) ) | $spl10),
  inference(resolution,[],[f232,f193])).
fof(f193,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,X1),X2),s(fun(fun(X0,X1),bool),'const.util_prob.FUNSET_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK9(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f474,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f440,f232])).
fof(f440,plain,(
  ( ! [X57,X54,X58,X56,X55] : (p(s(bool,'const.bool.IN_2'(s(fun(X54,X55),X56),s(fun(fun(X54,X55),bool),'const.util_prob.FUNSET_2'(s(fun(X54,bool),X57),s(fun(X55,bool),X58))))))) ) | $spl10),
  inference(resolution,[],[f232,f193])).
fof(f473,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f438,f232])).
fof(f438,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X47,bool),X48),s(fun(X47,bool),X49))))) ) | $spl10),
  inference(resolution,[],[f232,f127])).
fof(f472,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f432,f232])).
fof(f432,plain,(
  ( ! [X44] : (p(s(bool,X44))) ) | $spl10),
  inference(resolution,[],[f232,f163])).
fof(f163,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X13] : ((p(s(bool,'T_0')) => p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.IMP_CLAUSES|split|0')).
fof(f471,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f431,f232])).
fof(f431,plain,(
  ( ! [X43] : (p(s(bool,X43))) ) | $spl10),
  inference(resolution,[],[f232,f160])).
fof(f160,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.IMP_CLAUSES|split|1')).
fof(f470,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f429,f232])).
fof(f429,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f232,f157])).
fof(f157,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.IMP_CLAUSES|split|2')).
fof(f469,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f427,f232])).
fof(f427,plain,(
  ( ! [X40] : (p(s(bool,X40))) ) | $spl10),
  inference(resolution,[],[f232,f154])).
fof(f435,plain,(
  $false | $spl10),
  inference(resolution,[],[f232,f190])).
fof(f190,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.measure.MEASURABLE_UP_LIFT','thm.bool.TRUTH')).
fof(f436,plain,(
  $false | $spl10),
  inference(resolution,[],[f232,f191])).
fof(f451,plain,(
  $false | $spl10),
  inference(resolution,[],[f232,f123])).
fof(f454,plain,(
  $false | $spl10),
  inference(resolution,[],[f232,f124])).
fof(f400,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f381])).
fof(f381,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12)) | ~p(s(bool,X12))) ) | $spl8),
  inference(superposition,[],[f231,f186])).
fof(f401,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f380])).
fof(f380,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X10))) ) | $spl8),
  inference(superposition,[],[f231,f188])).
fof(f188,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f117])).
fof(f372,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9)) ) | $spl20),
  inference(cnf_transformation,[],[f372_D])).
fof(f372_D,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f405,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X8))) ) | $spl8),
  inference(subsumption_resolution,[],[f379,f231])).
fof(f379,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8)) ) | $spl8),
  inference(superposition,[],[f231,f199])).
fof(f404,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6))) ) | $spl8),
  inference(subsumption_resolution,[],[f378,f231])).
fof(f378,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6)) | s(bool,'F_0') = s(bool,X6)) ) | $spl8),
  inference(superposition,[],[f231,f199])).
fof(f366,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | p(s(bool,X16))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f349])).
fof(f349,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | p(s(bool,X16)) | p(s(bool,X16))) ) | $spl2),
  inference(superposition,[],[f228,f182])).
fof(f367,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f348])).
fof(f348,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14)) | p(s(bool,X14))) ) | $spl2),
  inference(superposition,[],[f228,f184])).
fof(f371,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8))) ) | $spl2),
  inference(subsumption_resolution,[],[f345,f228])).
fof(f345,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8)) ) | $spl2),
  inference(superposition,[],[f228,f199])).
fof(f369,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl2),
  inference(subsumption_resolution,[],[f344,f228])).
fof(f344,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,X6)) ) | $spl2),
  inference(superposition,[],[f228,f199])).
fof(f338,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f304,f233])).
fof(f304,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) | $spl12),
  inference(resolution,[],[f233,f152])).
fof(f300,plain,(
  $false | $spl12),
  inference(resolution,[],[f233,f189])).
fof(f189,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.measure.MEASURABLE_UP_LIFT',aHL_FALSITY)).
fof(f274,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))) ) | $spl14),
  inference(cnf_transformation,[],[f274_D])).
fof(f274_D,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f260,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X37) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))) )),
  inference(resolution,[],[f123,f200])).
fof(f250,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))) | s(bool,X34) != s(bool,X35)) )),
  inference(resolution,[],[f123,f142])).
fof(f275,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))) ) | $spl16),
  inference(cnf_transformation,[],[f275_D])).
fof(f275_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f230,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl6),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f249,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X31) | s(bool,X32) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))) )),
  inference(resolution,[],[f123,f142])).
fof(f247,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))) )),
  inference(resolution,[],[f123,f141])).
fof(f246,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))) )),
  inference(resolution,[],[f123,f140])).
fof(f245,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4))))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f123,f138])).
fof(f244,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X16) | s(bool,X17) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))) )),
  inference(resolution,[],[f123,f138])).
fof(f243,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK0,bool),fun(fun(sK0,bool),bool)),'const.pair.,_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))) )),
  inference(resolution,[],[f123,f136])).
fof(f227,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl0),
  inference(cnf_transformation,[],[f227_D])).
fof(f227_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f220,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f124,f200])).
fof(f210,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f124,f142])).
fof(f229,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))) ) | $spl4),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f209,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f124,f142])).
fof(f207,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f124,f141])).
fof(f206,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f124,f140])).
fof(f205,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4))) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f124,f138])).
fof(f204,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f124,f138])).
fof(f203,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))) )),
  inference(resolution,[],[f124,f136])).
% SZS output end Proof for thm.measure
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 130
% Passive clauses: 425
% Generated clauses: 4117
% Final active clauses: 110
% Final passive clauses: 217
% Input formulas: 29
% Initial clauses: 78
% 
% Selected by SInE selection: 28
% SInE iterations: 1
% 
% Duplicate literals: 85
% Trivial inequalities: 90
% Fw subsumption resolutions: 857
% Fw demodulations: 335
% Bw demodulations: 15
% 
% Simple tautologies: 159
% Forward subsumptions: 2000
% 
% Binary resolution: 837
% Forward superposition: 807
% Backward superposition: 1025
% Equality resolution: 16
% 
% Split clauses: 274
% Split components: 612
% SAT solver clauses: 295
% SAT solver unit clauses: 17
% SAT solver binary clauses: 168
% 
% Sat splits: 274
% Sat splitting refutations: 18
% 
% Memory used [KB]: 3326
% Time elapsed: 0.116 s
% ------------------------------
% ------------------------------
% Success in time 11.839 s
