Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 44
dis+11_7_41 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f72583,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f72581,f1878,f2862])).
fof(f2862,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f2862_D])).
fof(f2862_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1878,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f523])).
fof(f523,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.extreal.lt_infty|split|2',aHL_FALSITY)).
fof(f72581,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f72580,f2862])).
fof(f72580,plain,(
  p(s(bool,'F_0')) | $spl3),
  inference(forward_demodulation,[],[f72579,f3277])).
fof(f3277,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(resolution,[],[f2045,f1837])).
fof(f1837,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(cnf_transformation,[],[f519])).
fof(f519,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.lt_infty|split|2','thm.extreal.lt_infty|split|2')).
fof(f2045,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f1439])).
fof(f1439,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f663])).
fof(f663,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f662])).
fof(f662,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.extreal.lt_infty|split|2','thm.bool.EQ_CLAUSES|split|3')).
fof(f72579,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl3),
  inference(subsumption_resolution,[],[f72476,f2862])).
fof(f72476,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(superposition,[],[f2519,f1858])).
fof(f1858,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f65])).
fof(f65,axiom,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.lt_infty|split|2','thm.extreal.extreal_le_def|split|4')).
fof(f2519,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f1686])).
fof(f1686,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) & (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f1075])).
fof(f1075,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f1074])).
fof(f1074,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X6] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))))),
  file('thm.extreal.lt_infty|split|2','thm.extreal.extreal_lt_def')).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1669
% Passive clauses: 10594
% Generated clauses: 70646
% Final active clauses: 1405
% Final passive clauses: 8356
% Input formulas: 517
% Initial clauses: 1011
% 
% Duplicate literals: 288
% Trivial inequalities: 112
% Fw subsumption resolutions: 3133
% Fw demodulations: 5931
% Bw demodulations: 218
% 
% Simple tautologies: 426
% Equational tautologies: 56
% Forward subsumptions: 48948
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 13488
% Factoring: 745
% Forward superposition: 21362
% Backward superposition: 24034
% Self superposition: 58
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 1276
% Split components: 2554
% SAT solver clauses: 1654
% SAT solver unit clauses: 38
% SAT solver binary clauses: 93
% SAT solver learnt clauses: 1
% 
% Sat splits: 1276
% Sat splitting refutations: 15
% 
% TWLsolver clauses: 1653
% TWLsolver calls for satisfiability: 555
% 
% Memory used [KB]: 35820
% Time elapsed: 0.801 s
% ------------------------------
% ------------------------------
% Success in time 0.824 s
