Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.numeral
% (105563)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 298
% Passive clauses: 3841
% Generated clauses: 26894
% Final active clauses: 215
% Final passive clauses: 2986
% Input formulas: 69
% Initial clauses: 122
% 
% Duplicate literals: 640
% Trivial inequalities: 467
% Fw subsumption resolutions: 2353
% Fw demodulations: 4139
% Bw demodulations: 80
% 
% Simple tautologies: 450
% Equational tautologies: 48
% Forward subsumptions: 14530
% Fw demodulations to eq. taut.: 63
% 
% Binary resolution: 631
% Factoring: 1
% Forward superposition: 10814
% Backward superposition: 7279
% Self superposition: 136
% Equality factoring: 78
% Equality resolution: 41
% 
% Split clauses: 304
% Split components: 609
% SAT solver clauses: 832
% SAT solver unit clauses: 88
% SAT solver binary clauses: 409
% 
% Sat splits: 304
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 831
% TWLsolver calls for satisfiability: 206
% 
% Memory used [KB]: 14967
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numeral
% (105574)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 313
% Passive clauses: 1655
% Generated clauses: 25696
% Final active clauses: 209
% Final passive clauses: 868
% Input formulas: 69
% Initial clauses: 122
% 
% Duplicate literals: 468
% Trivial inequalities: 159
% Fw subsumption resolutions: 1668
% Fw demodulations: 5150
% Bw demodulations: 505
% Condensations: 497
% Global subsumptions: 17
% 
% Simple tautologies: 182
% Equational tautologies: 38
% Forward subsumptions: 13112
% Fw demodulations to eq. taut.: 56
% 
% Binary resolution: 424
% Factoring: 2
% Forward superposition: 6938
% Backward superposition: 7588
% Self superposition: 94
% Equality factoring: 58
% Equality resolution: 26
% 
% Split clauses: 116
% Split components: 234
% SAT solver clauses: 2535
% SAT solver unit clauses: 164
% SAT solver binary clauses: 57
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 3
% 
% Sat splits: 116
% Sat splitting refutations: 131
% 
% TWLsolver clauses: 2465
% TWLsolver calls for satisfiability: 664
% 
% Memory used [KB]: 18805
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numeral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numeral
% SZS output start Proof for thm.numeral
fof(f49752,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f525,f529,f527,f2362,f2365,f2363,f2477,f2478,f2491,f2504,f2507,f2510,f2513,f2518,f2521,f2539,f2542,f42376,f42982,f42377,f42984,f42981,f43158,f43162,f43195,f43196,f43197,f43199,f43202,f43205,f43967,f43976,f43968,f43969,f43970,f43980,f43973,f43971,f43979,f43974,f43977,f45202,f45205,f45216,f45214,f49687,f49690,f49694,f49698,f49702,f49706,f49710,f49714,f49718,f49726,f49730,f49734,f49738,f49745,f49747,f49749,f49750,f49751])).
fof(f49751,plain,(
  $false | ($spl1 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f49673,f43212])).
fof(f43212,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl3),
  inference(unit_resulting_resolution,[],[f358,f529,f488])).
fof(f488,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f480,f358])).
fof(f480,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f330,f243])).
fof(f243,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0',aHL_BOOL_CASES)).
fof(f330,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f358,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f233,f229])).
fof(f229,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0',aHL_TRUTH)).
fof(f233,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f29])).
fof(f29,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.bool.NOT_CLAUSES|split|2')).
fof(f49673,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6)),
  inference(superposition,[],[f49611,f4821])).
fof(f4821,plain,(
  ( ! [X4,X2,X3] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X2))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(superposition,[],[f346,f312])).
fof(f312,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X8,X9] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X8)))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.ADD_SYM')).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8,X9,X13] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X13))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X13)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.LESS_MONO_ADD_EQ')).
fof(f49611,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | ($spl1 | $spl3 | $spl6)),
  inference(backward_demodulation,[],[f49610,f526])).
fof(f526,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | $spl1),
  inference(cnf_transformation,[],[f526_D])).
fof(f526_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f49610,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f49609,f229])).
fof(f49609,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f49608,f43212])).
fof(f49608,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1)))) | $spl6),
  inference(subsumption_resolution,[],[f49590,f358])).
fof(f49590,plain,(
  p(s(bool,'F_0')) | s('type.num.num',sK1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1)))) | $spl6),
  inference(superposition,[],[f326,f49546])).
fof(f49546,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) | $spl6),
  inference(superposition,[],[f2365,f4659])).
fof(f4659,plain,(
  ( ! [X8,X7] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X8))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X8)))))) )),
  inference(superposition,[],[f319,f360])).
fof(f360,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f316,f314])).
fof(f314,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X9,X8] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f316,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X9,X8] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.ADD_CLAUSES|split|2')).
fof(f319,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X9,X8] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X8))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X8)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.numeral.numeral_9_prove_ltb|split|0')).
fof(f326,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X1) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & ~s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8,X9] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9)))) & ~s('type.num.num',X9) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.LESS_NOT_SUC')).
fof(f49750,plain,(
  $false | ($spl1 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f49672,f43212])).
fof(f49672,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6)),
  inference(superposition,[],[f49611,f346])).
fof(f49749,plain,(
  $false | ($spl1 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f49748,f372])).
fof(f372,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f358,f247])).
fof(f247,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f49748,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49684,f43212])).
fof(f49684,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6)),
  inference(trivial_inequality_removal,[],[f49671])).
fof(f49671,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6)),
  inference(superposition,[],[f49611,f4831])).
fof(f4831,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X11))))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X12)))) )),
  inference(superposition,[],[f243,f346])).
fof(f49747,plain,(
  $false | ($spl1 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f49746,f372])).
fof(f49746,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49634,f43212])).
fof(f49634,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f49611,f4831])).
fof(f49745,plain,(
  $false | ($spl1 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f49744,f229])).
fof(f49744,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49743,f43212])).
fof(f49743,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1)))) | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49636,f346])).
fof(f49636,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))) | ($spl1 | $spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f49611,f248])).
fof(f248,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f181])).
fof(f49738,plain,(
  $false | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f49737,f372])).
fof(f49737,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49736,f43212])).
fof(f49736,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49735,f45235])).
fof(f45235,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl8),
  inference(unit_resulting_resolution,[],[f358,f45195,f330])).
fof(f45195,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl8),
  inference(subsumption_resolution,[],[f45177,f358])).
fof(f45177,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl8),
  inference(superposition,[],[f366,f45121])).
fof(f45121,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl8),
  inference(superposition,[],[f43328,f7878])).
fof(f7878,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X6)))) )),
  inference(superposition,[],[f4843,f238])).
fof(f238,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.ADD_CLAUSES|split|0')).
fof(f4843,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X10))),s('type.num.num',X10)))) )),
  inference(forward_demodulation,[],[f4824,f381])).
fof(f381,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f304,f252])).
fof(f252,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f304,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X9] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0'))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.prim_rec.NOT_LESS_|48|')).
fof(f4824,plain,(
  ( ! [X10,X11] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X10))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f346,f238])).
fof(f43328,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X0)))) ) | $spl8),
  inference(forward_demodulation,[],[f43327,f238])).
fof(f43327,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X0)))) ) | $spl8),
  inference(subsumption_resolution,[],[f43304,f229])).
fof(f43304,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X0)))) ) | $spl8),
  inference(superposition,[],[f5640,f42981])).
fof(f5640,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X12)))) | s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X13))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f5639,f1778])).
fof(f1778,plain,(
  ( ! [X12,X13,X11] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X12))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12)))))))) )),
  inference(superposition,[],[f361,f314])).
fof(f361,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f320,f242])).
fof(f242,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X8] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.ADD|49|')).
fof(f320,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6,X7] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.LE_LT|49|')).
fof(f5639,plain,(
  ( ! [X12,X13,X11] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))),s('type.num.num',X13))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X12))))) )),
  inference(forward_demodulation,[],[f5638,f360])).
fof(f5638,plain,(
  ( ! [X12,X13,X11] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X13))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))),s('type.num.num',X13))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X12))))) )),
  inference(forward_demodulation,[],[f5637,f1196])).
fof(f1196,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f1158,f242])).
fof(f1158,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f242,f238])).
fof(f5637,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X12)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X13))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f5596,f2215])).
fof(f2215,plain,(
  ( ! [X4,X5] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X5))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num',X5)))) )),
  inference(superposition,[],[f317,f361])).
fof(f317,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8,X9] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.prim_rec.LESS_MONO_EQ')).
fof(f5596,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X12)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X13))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X13)))) )),
  inference(superposition,[],[f322,f1196])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ! [X2] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) => ! [X2] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X17,X18] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num',X17)))) => ! [X19] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num',X19))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num',X18))),s('type.num.num',X19))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.LESS_EQ_SUB_LESS')).
fof(f366,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(forward_demodulation,[],[f329,f361])).
fof(f329,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X8,X9] : ((s('type.num.num',X8) = s('type.num.num',X9) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.prim_rec.LESS_LEMMA2')).
fof(f49735,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49643,f346])).
fof(f49643,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f45237,f49611,f427])).
fof(f427,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f243,f243])).
fof(f45237,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl8),
  inference(unit_resulting_resolution,[],[f45195,f247])).
fof(f49734,plain,(
  $false | ($spl1 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f49733,f372])).
fof(f49733,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49732,f43212])).
fof(f49732,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49731,f404])).
fof(f404,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f305,f252])).
fof(f305,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X9] : ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.|48|_0'))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.NOT_SUC_LESS_EQ_|48|')).
fof(f49731,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) ) | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49644,f346])).
fof(f49644,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))) ) | ($spl1 | $spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f405,f49611,f427])).
fof(f405,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f305,f247])).
fof(f49730,plain,(
  $false | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f49729,f372])).
fof(f49729,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49728,f43212])).
fof(f49728,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49727,f43534])).
fof(f43534,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(unit_resulting_resolution,[],[f358,f43520,f330])).
fof(f43520,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0')))) | $spl8),
  inference(unit_resulting_resolution,[],[f43326,f4536])).
fof(f4536,plain,(
  ( ! [X23] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X23),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X23)) )),
  inference(forward_demodulation,[],[f4535,f231])).
fof(f231,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f57])).
fof(f57,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.prim_rec.PRE|split|0')).
fof(f4535,plain,(
  ( ! [X23] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X23),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',X23) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f4534,f358])).
fof(f4534,plain,(
  ( ! [X23] : (p(s(bool,'F_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X23),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',X23) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f4533,f381])).
fof(f4533,plain,(
  ( ! [X23] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X23),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',X23) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f4474,f231])).
fof(f4474,plain,(
  ( ! [X23] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X23),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num',X23) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f364,f231])).
fof(f364,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(forward_demodulation,[],[f324,f361])).
fof(f324,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) => (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X8,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))) => (s('type.num.num',X8) = s('type.num.num',X9) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9))))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.prim_rec.LESS_LEMMA|49|')).
fof(f43326,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) | $spl8),
  inference(subsumption_resolution,[],[f43302,f229])).
fof(f43302,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) | $spl8),
  inference(superposition,[],[f301,f42981])).
fof(f301,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X9] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X9) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.NOT_ZERO_LT_ZERO')).
fof(f49727,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49645,f346])).
fof(f49645,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f43536,f49611,f427])).
fof(f43536,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(unit_resulting_resolution,[],[f43520,f247])).
fof(f49726,plain,(
  $false | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f49725,f372])).
fof(f49725,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49724,f43212])).
fof(f49724,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49723,f45121])).
fof(f49723,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49646,f346])).
fof(f49646,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f48634,f49611,f427])).
fof(f48634,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl8),
  inference(subsumption_resolution,[],[f48440,f373])).
fof(f373,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f304,f247])).
fof(f48440,plain,(
  ( ! [X43] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X43),s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(superposition,[],[f373,f45194])).
fof(f45194,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(subsumption_resolution,[],[f45170,f372])).
fof(f45170,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f45121,f427])).
fof(f49718,plain,(
  $false | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f49717,f372])).
fof(f49717,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49716,f43212])).
fof(f49716,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49715,f45235])).
fof(f49715,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49654,f346])).
fof(f49654,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f45237,f49611,f427])).
fof(f49714,plain,(
  $false | ($spl1 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f49713,f372])).
fof(f49713,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49712,f43212])).
fof(f49712,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49711,f404])).
fof(f49711,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) ) | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49655,f346])).
fof(f49655,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))) ) | ($spl1 | $spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f405,f49611,f427])).
fof(f49710,plain,(
  $false | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f49709,f372])).
fof(f49709,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49708,f43212])).
fof(f49708,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49707,f43534])).
fof(f49707,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49656,f346])).
fof(f49656,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f43536,f49611,f427])).
fof(f49706,plain,(
  $false | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f49705,f372])).
fof(f49705,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49704,f43212])).
fof(f49704,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49703,f45121])).
fof(f49703,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49657,f346])).
fof(f49657,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f48634,f49611,f427])).
fof(f49702,plain,(
  $false | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f49701,f372])).
fof(f49701,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49700,f43212])).
fof(f49700,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49699,f43534])).
fof(f49699,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49661,f346])).
fof(f49661,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f49611,f43547])).
fof(f43547,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(subsumption_resolution,[],[f43538,f229])).
fof(f43538,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f43520,f427])).
fof(f49698,plain,(
  $false | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f49697,f372])).
fof(f49697,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49696,f43212])).
fof(f49696,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49695,f45121])).
fof(f49695,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49662,f346])).
fof(f49662,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f49611,f45194])).
fof(f49694,plain,(
  $false | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f49693,f372])).
fof(f49693,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49692,f43212])).
fof(f49692,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49691,f45235])).
fof(f49691,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f49663,f346])).
fof(f49663,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | ($spl1 | $spl3 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f49611,f45251])).
fof(f45251,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(subsumption_resolution,[],[f45239,f229])).
fof(f45239,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f45195,f427])).
fof(f49690,plain,(
  $false | ($spl1 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f49689,f229])).
fof(f49689,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49688,f43212])).
fof(f49688,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1)))) | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49665,f346])).
fof(f49665,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))) | ($spl1 | $spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f229,f49611,f333])).
fof(f333,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f171])).
fof(f49687,plain,(
  $false | ($spl1 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f49686,f229])).
fof(f49686,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49685,f43212])).
fof(f49685,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1)))) | ($spl1 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f49668,f346])).
fof(f49668,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))) | ($spl1 | $spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f229,f49611,f333])).
fof(f45214,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) | $spl26),
  inference(cnf_transformation,[],[f45214_D])).
fof(f45214_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f45216,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl29),
  inference(cnf_transformation,[],[f45216_D])).
fof(f45216_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f45205,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK5(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f5065,f427])).
fof(f5065,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(X3,sK5(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,g(s(X3,sK5(X4,X5,bool,X3))))) )),
  inference(superposition,[],[f345,f243])).
fof(f345,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0',aHL_EXT)).
fof(f45202,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f5065,f427])).
fof(f43977,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl20),
  inference(cnf_transformation,[],[f43977_D])).
fof(f43977_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f43974,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) | $spl16),
  inference(cnf_transformation,[],[f43974_D])).
fof(f43974_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f43979,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl23),
  inference(cnf_transformation,[],[f43979_D])).
fof(f43979_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f43971,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK5(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f5063,f427])).
fof(f5063,plain,(
  ( ! [X4,X5,X3] : (s(X5,f(s(bool,'F_0'))) != s(X5,g(s(bool,'F_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,'T_0') = s(bool,sK5(X3,X4,X5,bool))) )),
  inference(superposition,[],[f345,f243])).
fof(f43973,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl14),
  inference(cnf_transformation,[],[f43973_D])).
fof(f43973_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f43980,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl24),
  inference(cnf_transformation,[],[f43980_D])).
fof(f43980_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f43970,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK5(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f5063,f427])).
fof(f43969,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK5(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f5063,f427])).
fof(f43968,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK5(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f5063,f427])).
fof(f43976,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl19),
  inference(cnf_transformation,[],[f43976_D])).
fof(f43976_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f43967,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f5063,f243])).
fof(f43205,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f43204,f43171])).
fof(f43171,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl0),
  inference(forward_demodulation,[],[f359,f527])).
fof(f359,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))),
  inference(backward_demodulation,[],[f314,f228])).
fof(f228,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f150])).
fof(f150,plain,(
  ? [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.numeral.numeral_|49||48|_prove_ltb|split|0')).
fof(f43204,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl0),
  inference(condensation,[],[f43203])).
fof(f43203,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(subsumption_resolution,[],[f43193,f248])).
fof(f43193,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f43160,f427])).
fof(f43160,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1)))) | $spl0),
  inference(forward_demodulation,[],[f43159,f319])).
fof(f43159,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))) | $spl0),
  inference(subsumption_resolution,[],[f43132,f229])).
fof(f43132,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))) | $spl0),
  inference(superposition,[],[f1535,f527])).
fof(f1535,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))),s('type.num.num',X7)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6))),s('type.num.num',X7))))) )),
  inference(forward_demodulation,[],[f1515,f314])).
fof(f1515,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X5))))),s('type.num.num',X7)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6))),s('type.num.num',X7))))) )),
  inference(superposition,[],[f323,f312])).
fof(f323,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X8,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.prim_rec.SUC_LESS')).
fof(f43202,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f43201,f43171])).
fof(f43201,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl0),
  inference(condensation,[],[f43200])).
fof(f43200,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1)))) ) | $spl0),
  inference(subsumption_resolution,[],[f43192,f248])).
fof(f43192,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1)))) ) | $spl0),
  inference(superposition,[],[f43160,f427])).
fof(f43199,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f43198,f43171])).
fof(f43198,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl0),
  inference(subsumption_resolution,[],[f43190,f358])).
fof(f43190,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl0),
  inference(superposition,[],[f43160,f243])).
fof(f43197,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f43185,f43171])).
fof(f43185,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl0),
  inference(unit_resulting_resolution,[],[f229,f43160,f333])).
fof(f43196,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f43187,f43171])).
fof(f43187,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl0),
  inference(unit_resulting_resolution,[],[f229,f43160,f333])).
fof(f43195,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f43189,f43171])).
fof(f43189,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl0),
  inference(unit_resulting_resolution,[],[f43160,f248])).
fof(f43162,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f43161,f358])).
fof(f43161,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f43160,f528])).
fof(f528,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl2),
  inference(cnf_transformation,[],[f528_D])).
fof(f528_D,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f43158,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f43157,f358])).
fof(f43157,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f43156,f528])).
fof(f43156,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1)))) | $spl0),
  inference(forward_demodulation,[],[f43155,f319])).
fof(f43155,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))) | $spl0),
  inference(subsumption_resolution,[],[f43131,f229])).
fof(f43131,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))) | $spl0),
  inference(superposition,[],[f1775,f527])).
fof(f1775,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num',X4))))) )),
  inference(superposition,[],[f323,f314])).
fof(f42981,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1))) | $spl8),
  inference(cnf_transformation,[],[f42981_D])).
fof(f42981_D,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f42984,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) | $spl13),
  inference(cnf_transformation,[],[f42984_D])).
fof(f42984_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f42377,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1))) | $spl3),
  inference(superposition,[],[f538,f41079])).
fof(f41079,plain,(
  ( ! [X37,X36] : (s('type.num.num',X37) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X36))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X36)))) ) | $spl3),
  inference(subsumption_resolution,[],[f40983,f358])).
fof(f40983,plain,(
  ( ! [X37,X36] : (p(s(bool,'F_0')) | s('type.num.num',X37) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X36))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X36)))) ) | $spl3),
  inference(superposition,[],[f40590,f243])).
fof(f40590,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl3),
  inference(forward_demodulation,[],[f40589,f4906])).
fof(f4906,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(superposition,[],[f347,f237])).
fof(f237,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.ADD_|48|')).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X8,X9,X13] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num',X8))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.LT_ADD_LCANCEL')).
fof(f40589,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl3),
  inference(subsumption_resolution,[],[f40542,f229])).
fof(f40542,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl3),
  inference(superposition,[],[f4466,f40143])).
fof(f40143,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f39524,f248])).
fof(f39524,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))))) ) | $spl3),
  inference(forward_demodulation,[],[f39301,f238])).
fof(f39301,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))))) ) | $spl3),
  inference(superposition,[],[f38741,f237])).
fof(f38741,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X37))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38))))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f38740,f11376])).
fof(f11376,plain,(
  ( ! [X187,X185,X186] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X185),s('type.num.num',X186))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X187),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X186),s('type.num.num','const.num.SUC_1'(s('type.num.num',X185)))))))) ) | $spl3),
  inference(superposition,[],[f8016,f928])).
fof(f928,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f910,f314])).
fof(f910,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f312,f343])).
fof(f343,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X8,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) <=> s('type.num.num',X8) = s('type.num.num',X9))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.prim_rec.INV_SUC_EQ')).
fof(f8016,plain,(
  ( ! [X8,X7] : (s('type.num.num',X8) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f7983,f358])).
fof(f7983,plain,(
  ( ! [X8,X7] : (p(s(bool,'F_0')) | s('type.num.num',X8) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))) ) | $spl3),
  inference(superposition,[],[f367,f7926])).
fof(f7926,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',X1)))) ) | $spl3),
  inference(forward_demodulation,[],[f7859,f314])).
fof(f7859,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.num.num',X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4843,f5409])).
fof(f5409,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))) ) | $spl3),
  inference(forward_demodulation,[],[f5408,f1196])).
fof(f5408,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X9)))) ) | $spl3),
  inference(forward_demodulation,[],[f5407,f2215])).
fof(f5407,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X9)))) ) | $spl3),
  inference(forward_demodulation,[],[f5373,f361])).
fof(f5373,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X9)))) ) | $spl3),
  inference(superposition,[],[f2189,f1196])).
fof(f2189,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X16),s('type.num.num',X17)))) ) | $spl3),
  inference(subsumption_resolution,[],[f2144,f229])).
fof(f2144,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X16),s('type.num.num',X17)))) ) | $spl3),
  inference(superposition,[],[f251,f1446])).
fof(f1446,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3)))) ) | $spl3),
  inference(forward_demodulation,[],[f1431,f556])).
fof(f556,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | $spl3),
  inference(unit_resulting_resolution,[],[f538,f243])).
fof(f1431,plain,(
  ( ! [X2,X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) ) | $spl3),
  inference(superposition,[],[f361,f897])).
fof(f897,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | s(bool,'T_0') = s(bool,X7)) ) | $spl3),
  inference(subsumption_resolution,[],[f892,f358])).
fof(f892,plain,(
  ( ! [X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | s(bool,'T_0') = s(bool,X7)) ) | $spl3),
  inference(superposition,[],[f488,f556])).
fof(f251,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f183])).
fof(f367,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(forward_demodulation,[],[f328,f361])).
fof(f328,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f169])).
fof(f38740,plain,(
  ( ! [X39,X37,X38] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X37))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X37))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38))))))))) ) | $spl3),
  inference(forward_demodulation,[],[f38739,f314])).
fof(f38739,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X37))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38)))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X37))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38)))))))) ) | $spl3),
  inference(forward_demodulation,[],[f38738,f1778])).
fof(f38738,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X37))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38)))))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X37))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38)))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f38493,f229])).
fof(f38493,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X37))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38)))))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X37))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38)))))))) ) | $spl3),
  inference(superposition,[],[f4455,f10241])).
fof(f10241,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) ) | $spl3),
  inference(forward_demodulation,[],[f10197,f360])).
fof(f10197,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X1)))))) ) | $spl3),
  inference(superposition,[],[f10102,f312])).
fof(f10102,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11)))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f10061,f229])).
fof(f10061,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11)))))))) ) | $spl3),
  inference(superposition,[],[f1089,f772])).
fof(f772,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f313,f248])).
fof(f313,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X8,X9] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.LESS_ADD_SUC')).
fof(f1089,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X4)))) ) | $spl3),
  inference(subsumption_resolution,[],[f1088,f358])).
fof(f1088,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X4)))) ) | $spl3),
  inference(forward_demodulation,[],[f1074,f556])).
fof(f1074,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X4)))) ) | $spl3),
  inference(superposition,[],[f366,f897])).
fof(f4455,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num',X4)))) | s('type.num.num',X4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3)))) )),
  inference(superposition,[],[f364,f312])).
fof(f4466,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3)))))) | s('type.num.num',X4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3)))) )),
  inference(superposition,[],[f364,f312])).
fof(f538,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | $spl3),
  inference(backward_demodulation,[],[f530,f359])).
fof(f530,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl3),
  inference(unit_resulting_resolution,[],[f529,f243])).
fof(f42982,plain,(
  ( ! [X559] : (s(bool,X559) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X559)) ) | $spl10),
  inference(cnf_transformation,[],[f42982_D])).
fof(f42982_D,plain,(
  ( ! [X559] : (s(bool,X559) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X559)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f42376,plain,(
  ( ! [X559] : (s(bool,X559) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X559) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1)))) ) | $spl3),
  inference(superposition,[],[f897,f41079])).
fof(f2542,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2541,f229])).
fof(f2541,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f2540,f2363])).
fof(f2540,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X63),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X61),s('type.num.num',X62))))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2466,f229])).
fof(f2466,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X63),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X61),s('type.num.num',X62))))))) ) | $spl4),
  inference(backward_demodulation,[],[f2363,f1855])).
fof(f1855,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X61),s('type.num.num',X62)))),X63)),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X61),s('type.num.num','const.num.SUC_1'(s('type.num.num',X62)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X63),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X61),s('type.num.num',X62))))))) )),
  inference(forward_demodulation,[],[f1800,f361])).
fof(f1800,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X61),s('type.num.num',X62)))),X63)),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X61),s('type.num.num','const.num.SUC_1'(s('type.num.num',X62)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X61),s('type.num.num',X62))))))))) )),
  inference(superposition,[],[f339,f314])).
fof(f339,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3(X1,X0)),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | (! [X5] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3(X1,X0)),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK3(X1,X0)),s('type.num.num',X5))))) & p(s(bool,happ(s(fun('type.num.num',bool),sK3(X1,X0)),s('type.num.num',X0)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),sK3(X1,X0)),s('type.num.num',X1)))))) & (! [X2] : ((p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X2,X1,X0))))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK4(X2,X1,X0)))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X0)))) | p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4])],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ? [X4] : (! [X5] : (~p(s(bool,happ(s(fun('type.num.num',bool),X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))) | p(s(bool,happ(s(fun('type.num.num',bool),X4),s('type.num.num',X5))))) & p(s(bool,happ(s(fun('type.num.num',bool),X4),s('type.num.num',X0)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X4),s('type.num.num',X1)))))) & (! [X2] : (? [X3] : (p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X3))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X0)))) | p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ? [X2] : (! [X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) | p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X3))))) & p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X0)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X1)))))) & (! [X2] : (? [X3] : (p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X3))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X0)))) | p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> ? [X2] : (! [X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) | p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X3))))) & p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X0)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X1))))))),
  inference(ennf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> ? [X2] : (! [X3] : (p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) => p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X3))))) & p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X0)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X1))))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> ? [X2] : (! [X3] : (p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) => p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X3))))) & p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X0)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X1))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X8,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9)))) <=> ? [X14] : (! [X20] : (p(s(bool,happ(s(fun('type.num.num',bool),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20)))))) => p(s(bool,happ(s(fun('type.num.num',bool),X14),s('type.num.num',X20))))) & p(s(bool,happ(s(fun('type.num.num',bool),X14),s('type.num.num',X8)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X14),s('type.num.num',X9))))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.prim_rec.LESS_DEF')).
fof(f2539,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2538,f229])).
fof(f2538,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f2537,f2363])).
fof(f2537,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',X11))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2465,f229])).
fof(f2465,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',X11))))) ) | $spl4),
  inference(backward_demodulation,[],[f2363,f1692])).
fof(f1692,plain,(
  ( ! [X12,X11] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X12)),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',X11))))) )),
  inference(forward_demodulation,[],[f1691,f361])).
fof(f1691,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11)))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X12)),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))))) )),
  inference(forward_demodulation,[],[f1682,f1196])).
fof(f1682,plain,(
  ( ! [X12,X11] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X12)),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(superposition,[],[f339,f1196])).
fof(f2521,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2520,f229])).
fof(f2520,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f2519,f2363])).
fof(f2519,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34)))),X35)),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X33),s('type.num.num',X34))))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2443,f229])).
fof(f2443,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34)))),X35)),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X33),s('type.num.num',X34))))))) ) | $spl4),
  inference(backward_demodulation,[],[f2363,f2075])).
fof(f2075,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34)))),X35)),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X33),s('type.num.num',X34)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X33),s('type.num.num',X34))))))) )),
  inference(forward_demodulation,[],[f2030,f315])).
fof(f315,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X9,X8] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X9),s('type.num.num',X8))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.SUB_MONO_EQ')).
fof(f2030,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34)))),X35)),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X33),s('type.num.num',X34)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))))))))) )),
  inference(superposition,[],[f339,f315])).
fof(f2518,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2517,f229])).
fof(f2517,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f2516,f2363])).
fof(f2516,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6)),X7)),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X5))))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2438,f229])).
fof(f2438,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6)),X7)),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X5))))))) ) | $spl4),
  inference(backward_demodulation,[],[f2363,f1680])).
fof(f1680,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6)),X7)),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X5)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6))))))) )),
  inference(superposition,[],[f339,f312])).
fof(f2513,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2512,f229])).
fof(f2512,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f2511,f2363])).
fof(f2511,plain,(
  ( ! [X15,X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.-_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')),X16)),s('type.num.num',X15))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2430,f229])).
fof(f2430,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.-_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')),X16)),s('type.num.num',X15))))) ) | $spl4),
  inference(backward_demodulation,[],[f2363,f1694])).
fof(f1694,plain,(
  ( ! [X15,X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.-_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')),X16)),s('type.num.num',X15)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num',X15))))) )),
  inference(forward_demodulation,[],[f1684,f240])).
fof(f240,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.arithmetic.SUB_|48||split|1')).
fof(f1684,plain,(
  ( ! [X15,X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.-_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')),X16)),s('type.num.num',X15)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f339,f240])).
fof(f2510,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2509,f229])).
fof(f2509,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f2508,f2363])).
fof(f2508,plain,(
  ( ! [X14,X13] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)),X14)),s('type.num.num',X13))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2429,f229])).
fof(f2429,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)),X14)),s('type.num.num',X13))))) ) | $spl4),
  inference(backward_demodulation,[],[f2363,f1693])).
fof(f1693,plain,(
  ( ! [X14,X13] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)),X14)),s('type.num.num',X13)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num',X13))))) )),
  inference(forward_demodulation,[],[f1683,f238])).
fof(f1683,plain,(
  ( ! [X14,X13] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)),X14)),s('type.num.num',X13)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))))))) )),
  inference(superposition,[],[f339,f238])).
fof(f2507,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2506,f229])).
fof(f2506,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f2505,f2363])).
fof(f2505,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')),X4)),s('type.num.num',X3))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2428,f229])).
fof(f2428,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')),X4)),s('type.num.num',X3))))) ) | $spl4),
  inference(backward_demodulation,[],[f2363,f1690])).
fof(f1690,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')),X4)),s('type.num.num',X3)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X3))))) )),
  inference(forward_demodulation,[],[f1679,f237])).
fof(f1679,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')),X4)),s('type.num.num',X3)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f339,f237])).
fof(f2504,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2503,f229])).
fof(f2503,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f2502,f2363])).
fof(f2502,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)),s('type.num.num',X1))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2427,f229])).
fof(f2427,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)),s('type.num.num',X1))))) ) | $spl4),
  inference(backward_demodulation,[],[f2363,f1689])).
fof(f1689,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X1))))) )),
  inference(forward_demodulation,[],[f1678,f239])).
fof(f239,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|0','thm.prim_rec.PRE|split|1')).
fof(f1678,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3('const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))),X2)),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(superposition,[],[f339,f239])).
fof(f2491,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2490,f229])).
fof(f2490,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f2489,f2363])).
fof(f2489,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK3(X1,X0)),s('type.num.num',X1))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2406,f229])).
fof(f2406,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK3(X1,X0)),s('type.num.num',X1))))) ) | $spl4),
  inference(backward_demodulation,[],[f2363,f339])).
fof(f2478,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2373,f229])).
fof(f2373,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f2363,f358])).
fof(f2477,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2476,f2363])).
fof(f2476,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl4),
  inference(subsumption_resolution,[],[f2371,f247])).
fof(f2371,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f2363,f251])).
fof(f2363,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5)) ) | $spl4),
  inference(cnf_transformation,[],[f2363_D])).
fof(f2363_D,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f2365,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | $spl6),
  inference(cnf_transformation,[],[f2365_D])).
fof(f2365_D,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2362,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl3),
  inference(forward_demodulation,[],[f2361,f314])).
fof(f2361,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl3),
  inference(subsumption_resolution,[],[f2328,f248])).
fof(f2328,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl3),
  inference(superposition,[],[f2287,f897])).
fof(f2287,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9)))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))) ) | $spl3),
  inference(forward_demodulation,[],[f2286,f1196])).
fof(f2286,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9)))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X9)))) ) | $spl3),
  inference(forward_demodulation,[],[f2250,f2215])).
fof(f2250,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X9)))) ) | $spl3),
  inference(superposition,[],[f2233,f1196])).
fof(f2233,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) ) | $spl3),
  inference(backward_demodulation,[],[f2215,f2187])).
fof(f2187,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7)))) ) | $spl3),
  inference(forward_demodulation,[],[f2186,f361])).
fof(f2186,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7)))) ) | $spl3),
  inference(subsumption_resolution,[],[f2138,f229])).
fof(f2138,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7)))) ) | $spl3),
  inference(superposition,[],[f323,f1446])).
fof(f527,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) | $spl0),
  inference(cnf_transformation,[],[f527_D])).
fof(f527_D,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f529,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl3),
  inference(cnf_transformation,[],[f529_D])).
fof(f529_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f525,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1)))))),
  inference(superposition,[],[f359,f243])).
% SZS output end Proof for thm.numeral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 461
% Passive clauses: 2785
% Generated clauses: 49771
% Final active clauses: 192
% Final passive clauses: 1035
% Input formulas: 69
% Initial clauses: 122
% 
% Duplicate literals: 555
% Trivial inequalities: 291
% Fw subsumption resolutions: 1421
% Fw demodulations: 16428
% Bw demodulations: 196
% Condensations: 532
% 
% Simple tautologies: 148
% Equational tautologies: 179
% Forward subsumptions: 27298
% Fw demodulations to eq. taut.: 86
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 2751
% Factoring: 1
% Forward superposition: 17291
% Backward superposition: 9719
% Self superposition: 88
% Equality factoring: 44
% Equality resolution: 19
% 
% Split clauses: 11
% Split components: 27
% SAT solver clauses: 2846
% SAT solver unit clauses: 289
% SAT solver binary clauses: 1477
% 
% Sat splits: 11
% Sat splitting refutations: 37
% 
% Memory used [KB]: 29807
% Time elapsed: 1.747 s
% ------------------------------
% ------------------------------
% Success in time 3.08 s
