Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.hreal
% (19516)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1215
% Passive clauses: 12258
% Generated clauses: 31807
% Final active clauses: 1028
% Final passive clauses: 10797
% Input formulas: 261
% Initial clauses: 524
% 
% Duplicate literals: 188
% Trivial inequalities: 46
% Fw subsumption resolutions: 2315
% Fw demodulations: 2632
% Bw demodulations: 182
% 
% Simple tautologies: 245
% Equational tautologies: 28
% Forward subsumptions: 13935
% 
% Binary resolution: 8144
% Factoring: 627
% Forward superposition: 7406
% Backward superposition: 9640
% Self superposition: 54
% Equality factoring: 2
% Equality resolution: 17
% 
% Split clauses: 229
% Split components: 460
% SAT solver clauses: 400
% SAT solver unit clauses: 15
% SAT solver binary clauses: 81
% 
% Sat splits: 229
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 400
% TWLsolver calls for satisfiability: 363
% 
% Memory used [KB]: 18293
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.hreal
% (19519)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 671
% Passive clauses: 4170
% Generated clauses: 20227
% Final active clauses: 635
% Final passive clauses: 3357
% Input formulas: 261
% Initial clauses: 524
% 
% Duplicate literals: 101
% Trivial inequalities: 23
% Fw subsumption resolutions: 704
% Fw demodulations: 2225
% Bw demodulations: 157
% Condensations: 58
% Global subsumptions: 19
% 
% Simple tautologies: 243
% Equational tautologies: 10
% Forward subsumptions: 8025
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 17
% 
% Binary resolution: 2738
% Factoring: 57
% Forward superposition: 5436
% Backward superposition: 3905
% Self superposition: 61
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 5184
% SAT solver unit clauses: 650
% SAT solver binary clauses: 164
% SAT solver learnt clauses: 38
% SAT solver learnt literals: 12
% 
% Sat splits: 19
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 5108
% TWLsolver calls for satisfiability: 3482
% 
% Memory used [KB]: 25074
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.hreal
% (19537)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 989
% Passive clauses: 8075
% Generated clauses: 71279
% Final active clauses: 933
% Final passive clauses: 5909
% Input formulas: 261
% Initial clauses: 524
% 
% Duplicate literals: 10
% Trivial inequalities: 27
% Fw subsumption resolutions: 1634
% Fw demodulations: 25205
% Bw demodulations: 317
% Condensations: 34
% Global subsumptions: 1373
% 
% Simple tautologies: 96
% Equational tautologies: 48
% Forward subsumptions: 27732
% Fw demodulations to eq. taut.: 3704
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 26705
% Factoring: 3
% Forward superposition: 8575
% Backward superposition: 6528
% Self superposition: 133
% Equality factoring: 2
% Equality resolution: 16
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 12811
% SAT solver unit clauses: 2813
% SAT solver binary clauses: 5157
% 
% Sat splits: 20
% Sat splitting refutations: 3321
% 
% Memory used [KB]: 305496
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.hreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.hreal
% SZS output start Proof for thm.hreal
fof(f21809,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f587,f884,f883,f586,f886,f882,f887,f888,f1039,f889,f1041,f890,f1043,f1037,f891,f892,f893,f894,f895,f896,f898,f1003,f1350,f1580,f1598,f1599,f1600,f1601,f1602,f1604,f1046,f1605,f1710,f1745,f2069,f2070,f2071,f2072,f2073,f2209,f2074,f2076,f2181,f2208,f2223,f2224,f2225,f2226,f2227,f2229,f2365,f2230,f2335,f2364,f2439,f3254,f3256,f3260,f3272,f3273,f3277,f1047,f3913,f3915,f3916,f3917,f3920,f4025,f4054,f4263,f4261,f4260,f4253,f4250,f4249,f4248,f4245,f4215,f4214,f4213,f4339,f4340,f4341,f4342,f4343,f4344,f5383,f5384,f6180,f6181,f6203,f6204,f6672,f6680,f7465,f7467,f7468,f7469,f7472,f7577,f7600,f7810,f6297,f6298,f7813,f7814,f7812,f7815,f7976,f7972,f7943,f8045,f8046,f8047,f8048,f8049,f8050,f8051,f8052,f8053,f8074,f8076,f8077,f8079,f8081,f8186,f1044,f8631,f8633,f8634,f8635,f8638,f8743,f8778,f8973,f8972,f8976,f9127,f9126,f3255,f9307,f9305,f9304,f9297,f9294,f9291,f9271,f9260,f9259,f9258,f9382,f9383,f9384,f9385,f9386,f9387,f9424,f9426,f9427,f9428,f9431,f9536,f9565,f9573,f9575,f9576,f9578,f9580,f9685,f9712,f9751,f9753,f9754,f9756,f9758,f9863,f9890,f9903,f9905,f9906,f9907,f9910,f10015,f10044,f12908,f12907,f12906,f12902,f12873,f12979,f12980,f12981,f12982,f12983,f12984,f12985,f12986,f12987,f13415,f13417,f13418,f13420,f13422,f13527,f13570,f13583,f13585,f13586,f13587,f13590,f13695,f13724,f13764,f13766,f13767,f13768,f13771,f13876,f13921,f14174,f14176,f14177,f14178,f14181,f14286,f14315,f16393,f16395,f16396,f16397,f16400,f16505,f16550,f18109,f18111,f18112,f18114,f18116,f18221,f18269,f18767,f18769,f18770,f18772,f18774,f18879,f18923,f18931,f18935,f18933,f18943,f18945,f18946,f18948,f18950,f19055,f19233,f19241,f19243,f19244,f19246,f19248,f19353,f19080,f19423,f19425,f19426,f19428,f19430,f19535,f19579,f19684,f19686,f19687,f19689,f19691,f19796,f19818,f19873,f19875,f19876,f19878,f19880,f19985,f20015,f20035,f20037,f20038,f20040,f20042,f20147,f20190,f20332,f20334,f20335,f20337,f20339,f20444,f20488,f21109,f21111,f21112,f21114,f21116,f21221,f21251,f21411,f21413,f21414,f21416,f21418,f21523,f21549,f21636,f21638,f21639,f1045,f21641,f1049,f21643,f1048,f21748,f21793,f1050,f21806])).
fof(f21806,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f21805,f18328])).
fof(f18328,plain,(
  ( ! [X2] : (s('type.hreal.hreal',sK14) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X2)))) ) | $spl4),
  inference(superposition,[],[f631,f8276])).
fof(f8276,plain,(
  ( ! [X0] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X0))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK16),s('type.hreal.hreal',X0)))))) ) | $spl4),
  inference(superposition,[],[f630,f886])).
fof(f630,plain,(
  ( ! [X2,X0,X1] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X2))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.hreal.hreal',X2)))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X2))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.hreal.hreal',X2)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7,X12] : s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X7),s('type.hreal.hreal',X12))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7))),s('type.hreal.hreal',X12)))),
  file('thm.hreal.HREAL_LT','thm.hreal.HREAL_ADD_ASSOC')).
fof(f631,plain,(
  ( ! [X0,X1] : (s('type.hreal.hreal',X0) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1] : s('type.hreal.hreal',X0) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0,X1] : ~s('type.hreal.hreal',X0) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X7] : ~s('type.hreal.hreal',X6) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7)))),
  file('thm.hreal.HREAL_LT','thm.hreal.HREAL_NOZERO')).
fof(f21805,plain,(
  s('type.hreal.hreal',sK14) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))))) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f21777,f879])).
fof(f879,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f792])).
fof(f792,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f547])).
fof(f547,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X29] : ! [X11] : (s(X29,X11) = s(X29,X11) <=> p(s(bool,'T_0')))),
  file('thm.hreal.HREAL_LT','thm.bool.REFL_CLAUSE')).
fof(f21777,plain,(
  ~p(s(bool,'T_0')) | s('type.hreal.hreal',sK14) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))))) | ($spl4 | $spl8)),
  inference(superposition,[],[f8378,f8741])).
fof(f8741,plain,(
  s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14))) | $spl8),
  inference(resolution,[],[f1039,f831])).
fof(f831,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f571])).
fof(f571,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.hreal.HREAL_LT','thm.bool.EQ_CLAUSES|split|0')).
fof(f8378,plain,(
  ( ! [X15] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X15)))) | s('type.hreal.hreal',X15) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X15),s('type.hreal.hreal',sK15)))))) ) | $spl4),
  inference(forward_demodulation,[],[f8377,f886])).
fof(f8377,plain,(
  ( ! [X15] : (s('type.hreal.hreal',X15) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X15),s('type.hreal.hreal',sK15))))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))),s('type.hreal.hreal',X15))))) ) | $spl4),
  inference(forward_demodulation,[],[f8285,f634])).
fof(f634,plain,(
  ( ! [X0,X1] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0)))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1] : s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X7),s('type.hreal.hreal',X6)))),
  file('thm.hreal.HREAL_LT','thm.hreal.HREAL_ADD_SYM')).
fof(f8285,plain,(
  ( ! [X15] : (s('type.hreal.hreal',X15) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X15),s('type.hreal.hreal',sK15))),s('type.hreal.hreal',sK15))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))),s('type.hreal.hreal',X15))))) ) | $spl4),
  inference(superposition,[],[f632,f886])).
fof(f632,plain,(
  ( ! [X0,X1] : (s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))),s('type.hreal.hreal',X0))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) )),
  inference(cnf_transformation,[],[f385])).
fof(f385,plain,(
  ! [X0,X1] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))),s('type.hreal.hreal',X0))))),
  inference(ennf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) => s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))),s('type.hreal.hreal',X0))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7)))) => s('type.hreal.hreal',X7) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X7),s('type.hreal.hreal',X6))),s('type.hreal.hreal',X6))))),
  file('thm.hreal.HREAL_LT','thm.hreal.HREAL_SUB_ADD')).
fof(f1050,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl26),
  inference(cnf_transformation,[],[f1050_D])).
fof(f1050_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f21793,plain,(
  ( ! [X6,X4] : (s('type.hreal.hreal',X4) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X4),s('type.hreal.hreal',sK15))))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X4)))) ) | $spl96),
  inference(cnf_transformation,[],[f21793_D])).
fof(f21793_D,plain,(
  ( ! [X6,X4] : (s('type.hreal.hreal',X4) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X4),s('type.hreal.hreal',sK15))))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X4)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f21748,plain,(
  ( ! [X316,X318,X317] : (s('type.hreal.hreal',X316) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X316),s('type.hreal.hreal',sK15))))) | p(s(bool,X317)) | ~p(s(bool,X318)) | s(bool,X318) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X316)))) ) | $spl4),
  inference(resolution,[],[f8378,f873])).
fof(f873,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f592])).
fof(f592,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f458])).
fof(f458,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f457])).
fof(f457,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f381])).
fof(f381,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f170])).
fof(f170,axiom,(
  ! [X11,X60,X34,X61] : ((s(bool,X11) = s(bool,X60) & (p(s(bool,X60)) => s(bool,X34) = s(bool,X61))) => ((p(s(bool,X11)) => p(s(bool,X34))) <=> (p(s(bool,X60)) => p(s(bool,X61)))))),
  file('thm.hreal.HREAL_LT','thm.bool.IMP_CONG')).
fof(f1048,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) | $spl22),
  inference(cnf_transformation,[],[f1048_D])).
fof(f1048_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f21643,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.hreal.hreal',X32) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',sK15))))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X32))) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(resolution,[],[f8378,f595])).
fof(f595,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f458])).
fof(f1049,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | $spl24),
  inference(cnf_transformation,[],[f1049_D])).
fof(f1049_D,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f21641,plain,(
  ( ! [X26,X24,X27,X25] : (s('type.hreal.hreal',X24) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X24),s('type.hreal.hreal',sK15))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X24)))) ) | $spl4),
  inference(resolution,[],[f8378,f594])).
fof(f594,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f458])).
fof(f1045,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) | $spl16),
  inference(cnf_transformation,[],[f1045_D])).
fof(f1045_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f21639,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.hreal.hreal',X16) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X16),s('type.hreal.hreal',sK15))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X16)))) ) | $spl4),
  inference(resolution,[],[f8378,f593])).
fof(f593,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f458])).
fof(f21638,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.hreal.hreal',X12) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X12),s('type.hreal.hreal',sK15))))) | ~p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X12))) | s(bool,X13) != s(bool,X15)) ) | $spl4),
  inference(resolution,[],[f8378,f591])).
fof(f591,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f458])).
fof(f21636,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.hreal.hreal',X4) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X4),s('type.hreal.hreal',sK15))))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X4)))) ) | $spl4),
  inference(resolution,[],[f8378,f589])).
fof(f589,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f458])).
fof(f21549,plain,(
  ( ! [X14,X16] : (s('type.hreal.hreal',sK15) = s('type.hreal.hreal',sK17('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X14,'type.hreal.hreal')) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun('type.hreal.hreal',bool),X14),s('type.hreal.hreal',sK15)))) ) | $spl94),
  inference(cnf_transformation,[],[f21549_D])).
fof(f21549_D,plain,(
  ( ! [X14,X16] : (s('type.hreal.hreal',sK15) = s('type.hreal.hreal',sK17('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X14,'type.hreal.hreal')) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun('type.hreal.hreal',bool),X14),s('type.hreal.hreal',sK15)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f21523,plain,(
  ( ! [X327,X326,X328] : (s('type.hreal.hreal',sK15) = s('type.hreal.hreal',sK17('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X326,'type.hreal.hreal')) | p(s(bool,X327)) | ~p(s(bool,X328)) | s(bool,X328) != s(bool,happ(s(fun('type.hreal.hreal',bool),X326),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8444,f873])).
fof(f8444,plain,(
  ( ! [X93] : (~p(s(bool,happ(s(fun('type.hreal.hreal',bool),X93),s('type.hreal.hreal',sK15)))) | s('type.hreal.hreal',sK15) = s('type.hreal.hreal',sK17('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X93,'type.hreal.hreal'))) ) | $spl4),
  inference(forward_demodulation,[],[f8353,f886])).
fof(f8353,plain,(
  ( ! [X93] : (~p(s(bool,happ(s(fun('type.hreal.hreal',bool),X93),s('type.hreal.hreal',sK15)))) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))) = s('type.hreal.hreal',sK17('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X93,'type.hreal.hreal'))) ) | $spl4),
  inference(superposition,[],[f625,f886])).
fof(f625,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK17(X2,X1,X0)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f473])).
fof(f473,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK17(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f472])).
fof(f472,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f471])).
fof(f471,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f195])).
fof(f195,axiom,(
  ! [X29] : ! [X45,X30] : (? [X11] : (s(X29,X11) = s(X29,X30) & p(s(bool,happ(s(fun(X29,bool),X45),s(X29,X11))))) <=> p(s(bool,happ(s(fun(X29,bool),X45),s(X29,X30)))))),
  file('thm.hreal.HREAL_LT','thm.bool.UNWIND_THM2')).
fof(f21418,plain,(
  ( ! [X45,X43,X44,X42] : (s('type.hreal.hreal',sK15) = s('type.hreal.hreal',sK17('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X42,'type.hreal.hreal')) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,happ(s(fun('type.hreal.hreal',bool),X42),s('type.hreal.hreal',sK15))) | s(bool,X43) != s(bool,X45)) ) | $spl4),
  inference(resolution,[],[f8444,f595])).
fof(f21416,plain,(
  ( ! [X37,X35,X36,X34] : (s('type.hreal.hreal',sK15) = s('type.hreal.hreal',sK17('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X34,'type.hreal.hreal')) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,happ(s(fun('type.hreal.hreal',bool),X34),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8444,f594])).
fof(f21414,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.hreal.hreal',sK15) = s('type.hreal.hreal',sK17('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X26,'type.hreal.hreal')) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,happ(s(fun('type.hreal.hreal',bool),X26),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8444,f593])).
fof(f21413,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.hreal.hreal',sK15) = s('type.hreal.hreal',sK17('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X22,'type.hreal.hreal')) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,happ(s(fun('type.hreal.hreal',bool),X22),s('type.hreal.hreal',sK15))) | s(bool,X23) != s(bool,X25)) ) | $spl4),
  inference(resolution,[],[f8444,f591])).
fof(f21411,plain,(
  ( ! [X14,X17,X15,X16] : (s('type.hreal.hreal',sK15) = s('type.hreal.hreal',sK17('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X14,'type.hreal.hreal')) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun('type.hreal.hreal',bool),X14),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8444,f589])).
fof(f21251,plain,(
  ( ! [X14,X16] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(X14,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X14)))) ) | $spl92),
  inference(cnf_transformation,[],[f21251_D])).
fof(f21251_D,plain,(
  ( ! [X14,X16] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(X14,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X14)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f21221,plain,(
  ( ! [X327,X326,X328] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(X326,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X327)) | ~p(s(bool,X328)) | s(bool,X328) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X326)))) ) | $spl4),
  inference(resolution,[],[f8438,f873])).
fof(f8438,plain,(
  ( ! [X81] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X81)))) | p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(X81,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))))) ) | $spl4),
  inference(forward_demodulation,[],[f8341,f886])).
fof(f8341,plain,(
  ( ! [X81] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(X81,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))),s('type.hrat.hrat',X81))))) ) | $spl4),
  inference(superposition,[],[f864,f886])).
fof(f864,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK35(X1,X0))))) | ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X1))))) )),
  inference(cnf_transformation,[],[f582])).
fof(f582,plain,(
  ! [X0,X1] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X1)))) | (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK35(X1,X0))))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',sK35(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK35])],[f431])).
fof(f431,plain,(
  ! [X0,X1] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X1)))) | ? [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))))),
  inference(ennf_transformation,[],[f372])).
fof(f372,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X1)))) => ? [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6,X11] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X6),s('type.hrat.hrat',X11)))) => ? [X34] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X6),s('type.hrat.hrat',X34)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X11),s('type.hrat.hrat',X34))))))),
  file('thm.hreal.HREAL_LT','thm.hreal.CUT_UP')).
fof(f21116,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(X42,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X42))) | s(bool,X43) != s(bool,X45)) ) | $spl4),
  inference(resolution,[],[f8438,f595])).
fof(f21114,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(X34,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X34)))) ) | $spl4),
  inference(resolution,[],[f8438,f594])).
fof(f21112,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(X26,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X26)))) ) | $spl4),
  inference(resolution,[],[f8438,f593])).
fof(f21111,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(X22,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X22))) | s(bool,X23) != s(bool,X25)) ) | $spl4),
  inference(resolution,[],[f8438,f591])).
fof(f21109,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(X14,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X14)))) ) | $spl4),
  inference(resolution,[],[f8438,f589])).
fof(f20488,plain,(
  ( ! [X6,X4] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X4))))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X4),s('type.hreal.hreal',sK15)))) ) | $spl90),
  inference(cnf_transformation,[],[f20488_D])).
fof(f20488_D,plain,(
  ( ! [X6,X4] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X4))))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X4),s('type.hreal.hreal',sK15)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f20444,plain,(
  ( ! [X316,X318,X317] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X316))))) | p(s(bool,X317)) | ~p(s(bool,X318)) | s(bool,X318) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X316),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8412,f873])).
fof(f8412,plain,(
  ( ! [X48] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X48),s('type.hreal.hreal',sK15)))) | p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X48)))))) ) | $spl4),
  inference(forward_demodulation,[],[f8314,f886])).
fof(f8314,plain,(
  ( ! [X48] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X48))))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X48),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(superposition,[],[f853,f886])).
fof(f853,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',sK32(X1,X0))))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) )),
  inference(cnf_transformation,[],[f579])).
fof(f579,plain,(
  ! [X0,X1] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK32(X1,X0))))) & p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',sK32(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK32])],[f421])).
fof(f421,plain,(
  ! [X0,X1] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | ? [X2] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X2))))))),
  inference(ennf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) => ? [X2] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X2))))))),
  inference(flattening,[],[f358])).
fof(f358,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) => ? [X2] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X2))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7)))) => ? [X11] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X6),s('type.hrat.hrat',X11)))) & p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X7),s('type.hrat.hrat',X11))))))),
  file('thm.hreal.HREAL_LT','thm.hreal.HREAL_LT_LEMMA')).
fof(f20339,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X32))))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',sK15))) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(resolution,[],[f8412,f595])).
fof(f20337,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X24))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X24),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8412,f594])).
fof(f20335,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X16))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X16),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8412,f593])).
fof(f20334,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X12))))) | ~p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X12),s('type.hreal.hreal',sK15))) | s(bool,X13) != s(bool,X15)) ) | $spl4),
  inference(resolution,[],[f8412,f591])).
fof(f20332,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X4))))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X4),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8412,f589])).
fof(f20190,plain,(
  ( ! [X10,X8] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X8)))) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(X8,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl88),
  inference(cnf_transformation,[],[f20190_D])).
fof(f20190_D,plain,(
  ( ! [X10,X8] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X8)))) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(X8,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f20147,plain,(
  ( ! [X321,X320,X322] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X320)))) | p(s(bool,X321)) | ~p(s(bool,X322)) | s(bool,X322) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(X320,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8410,f873])).
fof(f8410,plain,(
  ( ! [X45] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(X45,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X45))))) ) | $spl4),
  inference(forward_demodulation,[],[f8311,f886])).
fof(f8311,plain,(
  ( ! [X45] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(X45,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))),s('type.hreal.hreal',X45))))) ) | $spl4),
  inference(superposition,[],[f852,f886])).
fof(f852,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK32(X1,X0))))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) )),
  inference(cnf_transformation,[],[f579])).
fof(f20042,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X36)))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(X36,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))) | s(bool,X37) != s(bool,X39)) ) | $spl4),
  inference(resolution,[],[f8410,f595])).
fof(f20040,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X28)))) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(X28,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8410,f594])).
fof(f20038,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X20)))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(X20,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8410,f593])).
fof(f20037,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X16)))) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X18) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(X16,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))) | s(bool,X17) != s(bool,X19)) ) | $spl4),
  inference(resolution,[],[f8410,f591])).
fof(f20035,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X8)))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(X8,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8410,f589])).
fof(f20015,plain,(
  ( ! [X14,X16] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X14),s('type.hrat.hrat',sK35(X14,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X14)))) ) | $spl86),
  inference(cnf_transformation,[],[f20015_D])).
fof(f20015_D,plain,(
  ( ! [X14,X16] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X14),s('type.hrat.hrat',sK35(X14,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X14)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f19985,plain,(
  ( ! [X327,X326,X328] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X326),s('type.hrat.hrat',sK35(X326,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X327)) | ~p(s(bool,X328)) | s(bool,X328) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X326)))) ) | $spl4),
  inference(resolution,[],[f8343,f873])).
fof(f8343,plain,(
  ( ! [X83] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X83)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X83),s('type.hrat.hrat',sK35(X83,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))))) ) | $spl4),
  inference(superposition,[],[f865,f886])).
fof(f865,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',sK35(X1,X0))))) | ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X1))))) )),
  inference(cnf_transformation,[],[f582])).
fof(f19880,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X42),s('type.hrat.hrat',sK35(X42,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X42))) | s(bool,X43) != s(bool,X45)) ) | $spl4),
  inference(resolution,[],[f8343,f595])).
fof(f19878,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X34),s('type.hrat.hrat',sK35(X34,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X34)))) ) | $spl4),
  inference(resolution,[],[f8343,f594])).
fof(f19876,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X26),s('type.hrat.hrat',sK35(X26,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X26)))) ) | $spl4),
  inference(resolution,[],[f8343,f593])).
fof(f19875,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X22),s('type.hrat.hrat',sK35(X22,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X22))) | s(bool,X23) != s(bool,X25)) ) | $spl4),
  inference(resolution,[],[f8343,f591])).
fof(f19873,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X14),s('type.hrat.hrat',sK35(X14,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',X14)))) ) | $spl4),
  inference(resolution,[],[f8343,f589])).
fof(f19818,plain,(
  ( ! [X10,X8] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK34(X8,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X8)))) ) | $spl84),
  inference(cnf_transformation,[],[f19818_D])).
fof(f19818_D,plain,(
  ( ! [X10,X8] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK34(X8,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X8)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f19796,plain,(
  ( ! [X321,X320,X322] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK34(X320,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X321)) | ~p(s(bool,X322)) | s(bool,X322) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X320)))) ) | $spl4),
  inference(resolution,[],[f8339,f873])).
fof(f8339,plain,(
  ( ! [X79] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X79)))) | p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK34(X79,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))))) ) | $spl4),
  inference(superposition,[],[f862,f886])).
fof(f862,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK34(X1,X0))))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X1))))) )),
  inference(cnf_transformation,[],[f581])).
fof(f581,plain,(
  ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X1)))) | (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK34(X1,X0))))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat','const.hrat.hrat_mul_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',sK34(X1,X0)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK34])],[f430])).
fof(f430,plain,(
  ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X1)))) | ? [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat','const.hrat.hrat_mul_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))))))),
  inference(ennf_transformation,[],[f371])).
fof(f371,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X1)))) => ? [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat','const.hrat.hrat_mul_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))))))),
  inference(flattening,[],[f370])).
fof(f370,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X1)))) => ? [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat','const.hrat.hrat_mul_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))))))),
  inference(rectify,[],[f116])).
fof(f116,axiom,(
  ! [X6,X47] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X47)))) => ? [X11] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X6),s('type.hrat.hrat',X11)))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X6),s('type.hrat.hrat','const.hrat.hrat_mul_2'(s('type.hrat.hrat',X47),s('type.hrat.hrat',X11))))))))),
  file('thm.hreal.HREAL_LT','thm.hreal.CUT_NEARTOP_MUL')).
fof(f19691,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK34(X36,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X36))) | s(bool,X37) != s(bool,X39)) ) | $spl4),
  inference(resolution,[],[f8339,f595])).
fof(f19689,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK34(X28,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X28)))) ) | $spl4),
  inference(resolution,[],[f8339,f594])).
fof(f19687,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK34(X20,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X20)))) ) | $spl4),
  inference(resolution,[],[f8339,f593])).
fof(f19686,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK34(X16,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X18) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X16))) | s(bool,X17) != s(bool,X19)) ) | $spl4),
  inference(resolution,[],[f8339,f591])).
fof(f19684,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK34(X8,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X8)))) ) | $spl4),
  inference(resolution,[],[f8339,f589])).
fof(f19579,plain,(
  ( ! [X6,X4] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X4),s('type.hrat.hrat',sK32(X4,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X4)))) ) | $spl82),
  inference(cnf_transformation,[],[f19579_D])).
fof(f19579_D,plain,(
  ( ! [X6,X4] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X4),s('type.hrat.hrat',sK32(X4,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X4)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f19535,plain,(
  ( ! [X316,X318,X317] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X316),s('type.hrat.hrat',sK32(X316,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X317)) | ~p(s(bool,X318)) | s(bool,X318) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X316)))) ) | $spl4),
  inference(resolution,[],[f8315,f873])).
fof(f8315,plain,(
  ( ! [X49] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X49)))) | p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X49),s('type.hrat.hrat',sK32(X49,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))))) ) | $spl4),
  inference(superposition,[],[f853,f886])).
fof(f19430,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X32),s('type.hrat.hrat',sK32(X32,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X32))) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(resolution,[],[f8315,f595])).
fof(f19428,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X24),s('type.hrat.hrat',sK32(X24,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X24)))) ) | $spl4),
  inference(resolution,[],[f8315,f594])).
fof(f19426,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X16),s('type.hrat.hrat',sK32(X16,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X16)))) ) | $spl4),
  inference(resolution,[],[f8315,f593])).
fof(f19425,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X12),s('type.hrat.hrat',sK32(X12,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X12))) | s(bool,X13) != s(bool,X15)) ) | $spl4),
  inference(resolution,[],[f8315,f591])).
fof(f19423,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X4),s('type.hrat.hrat',sK32(X4,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',X4)))) ) | $spl4),
  inference(resolution,[],[f8315,f589])).
fof(f19080,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) ) | $spl80),
  inference(cnf_transformation,[],[f19080_D])).
fof(f19080_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f19353,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X201) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) ) | $spl77),
  inference(resolution,[],[f18933,f873])).
fof(f19248,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl77),
  inference(resolution,[],[f18933,f595])).
fof(f19246,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) ) | $spl77),
  inference(resolution,[],[f18933,f594])).
fof(f19244,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) ) | $spl77),
  inference(resolution,[],[f18933,f593])).
fof(f19243,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl77),
  inference(resolution,[],[f18933,f591])).
fof(f19241,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) ) | $spl77),
  inference(resolution,[],[f18933,f589])).
fof(f19233,plain,(
  $false | $spl76),
  inference(trivial_inequality_removal,[],[f19091])).
fof(f19091,plain,(
  s('type.hreal.hreal',sK15) != s('type.hreal.hreal',sK15) | $spl76),
  inference(resolution,[],[f18932,f848])).
fof(f848,plain,(
  ( ! [X0,X1] : (s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) )),
  inference(cnf_transformation,[],[f578])).
fof(f578,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | (s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) & ! [X3] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X3)))) | p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X3))))))) & (s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1) | (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK31(X1,X0))))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',sK31(X1,X0)))))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31])],[f577])).
fof(f577,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | (s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) & ! [X3] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X3)))) | p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X3))))))) & (s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1) | ? [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X2))))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))))),
  inference(rectify,[],[f576])).
fof(f576,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | (s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) & ! [X2] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) | p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X2))))))) & (s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1) | ? [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X2))))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))))),
  inference(flattening,[],[f575])).
fof(f575,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | (s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) & ! [X2] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) | p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X2))))))) & ((s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1) | ? [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X2)))))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))))),
  inference(nnf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) <=> (s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) & ! [X2] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) | p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X2)))))))),
  inference(ennf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) <=> (s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) & ! [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) => p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X2)))))))),
  inference(flattening,[],[f356])).
fof(f356,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) <=> (~s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1) & ! [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) => p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',X2)))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7)))) <=> (~s('type.hreal.hreal',X6) = s('type.hreal.hreal',X7) & ! [X11] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X6),s('type.hrat.hrat',X11)))) => p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X7),s('type.hrat.hrat',X11)))))))),
  file('thm.hreal.HREAL_LT','thm.hreal.hreal_lt')).
fof(f18932,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) | $spl76),
  inference(cnf_transformation,[],[f18932_D])).
fof(f18932_D,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f19055,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X201) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) ) | $spl77),
  inference(resolution,[],[f18933,f873])).
fof(f18950,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl77),
  inference(resolution,[],[f18933,f595])).
fof(f18948,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) ) | $spl77),
  inference(resolution,[],[f18933,f594])).
fof(f18946,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) ) | $spl77),
  inference(resolution,[],[f18933,f593])).
fof(f18945,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl77),
  inference(resolution,[],[f18933,f591])).
fof(f18943,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) ) | $spl77),
  inference(resolution,[],[f18933,f589])).
fof(f18933,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) | $spl77),
  inference(cnf_transformation,[],[f18933_D])).
fof(f18933_D,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f18935,plain,(
  ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | $spl79),
  inference(cnf_transformation,[],[f18935_D])).
fof(f18935_D,plain,(
  ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f18931,plain,(
  ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) | $spl4),
  inference(forward_demodulation,[],[f18895,f886])).
fof(f18895,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK15)))) | ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | $spl4),
  inference(superposition,[],[f8313,f886])).
fof(f8313,plain,(
  ( ! [X47] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X47),s('type.hreal.hreal',sK15)))) | ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X47),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X47)))))) ) | $spl4),
  inference(superposition,[],[f852,f886])).
fof(f18923,plain,(
  ( ! [X6,X4] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X4),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X4))))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X4),s('type.hreal.hreal',sK15)))) ) | $spl74),
  inference(cnf_transformation,[],[f18923_D])).
fof(f18923_D,plain,(
  ( ! [X6,X4] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X4),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X4))))) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X4),s('type.hreal.hreal',sK15)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f18879,plain,(
  ( ! [X316,X318,X317] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X316),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X316))))) | p(s(bool,X317)) | ~p(s(bool,X318)) | s(bool,X318) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X316),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8313,f873])).
fof(f18774,plain,(
  ( ! [X35,X33,X34,X32] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X32),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X32))))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',sK15))) | s(bool,X33) != s(bool,X35)) ) | $spl4),
  inference(resolution,[],[f8313,f595])).
fof(f18772,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X24),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X24))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X24),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8313,f594])).
fof(f18770,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X16),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X16))))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X16),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8313,f593])).
fof(f18769,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X12),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X12))))) | ~p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X12),s('type.hreal.hreal',sK15))) | s(bool,X13) != s(bool,X15)) ) | $spl4),
  inference(resolution,[],[f8313,f591])).
fof(f18767,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X4),s('type.hrat.hrat',sK32('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)),X4))))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X4),s('type.hreal.hreal',sK15)))) ) | $spl4),
  inference(resolution,[],[f8313,f589])).
fof(f18269,plain,(
  ( ! [X11,X9] : (~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',sK33(X11,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))),s('type.hrat.hrat',X11)))))) ) | $spl72),
  inference(cnf_transformation,[],[f18269_D])).
fof(f18269_D,plain,(
  ( ! [X11,X9] : (~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',sK33(X11,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))),s('type.hrat.hrat',X11)))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f18221,plain,(
  ( ! [X313,X312,X314] : (p(s(bool,X312)) | ~p(s(bool,X313)) | s(bool,X313) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',sK33(X314,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))),s('type.hrat.hrat',X314)))))) ) | $spl4),
  inference(resolution,[],[f8334,f873])).
fof(f8334,plain,(
  ( ! [X74] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',sK33(X74,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))),s('type.hrat.hrat',X74))))))) ) | $spl4),
  inference(superposition,[],[f860,f886])).
fof(f860,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',sK33(X1,X0)),s('type.hrat.hrat',X1))))))) )),
  inference(cnf_transformation,[],[f580])).
fof(f580,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK33(X1,X0))))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',sK33(X1,X0)),s('type.hrat.hrat',X1)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK33])],[f368])).
fof(f368,plain,(
  ! [X0,X1] : ? [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',X1)))))))),
  inference(flattening,[],[f367])).
fof(f367,plain,(
  ! [X0,X1] : ? [X2] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X2)))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',X1)))))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6,X38] : ? [X11] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X6),s('type.hrat.hrat',X11)))) & ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X6),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',X11),s('type.hrat.hrat',X38)))))))),
  file('thm.hreal.HREAL_LT','thm.hreal.CUT_NEARTOP_ADD')).
fof(f18116,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',sK33(X37,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))),s('type.hrat.hrat',X37))))) | s(bool,X35) != s(bool,X38)) ) | $spl4),
  inference(resolution,[],[f8334,f595])).
fof(f18114,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',sK33(X31,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))),s('type.hrat.hrat',X31)))))) ) | $spl4),
  inference(resolution,[],[f8334,f594])).
fof(f18112,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',sK33(X23,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))),s('type.hrat.hrat',X23)))))) ) | $spl4),
  inference(resolution,[],[f8334,f593])).
fof(f18111,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',sK33(X18,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))),s('type.hrat.hrat',X18))))) | s(bool,X16) != s(bool,X19)) ) | $spl4),
  inference(resolution,[],[f8334,f591])).
fof(f18109,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat','const.hrat.hrat_add_2'(s('type.hrat.hrat',sK33(X11,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))),s('type.hrat.hrat',X11)))))) ) | $spl4),
  inference(resolution,[],[f8334,f589])).
fof(f16550,plain,(
  ( ! [X13,X11] : (p(s(bool,X11)) | s(bool,X11) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK33(X13,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl70),
  inference(cnf_transformation,[],[f16550_D])).
fof(f16550_D,plain,(
  ( ! [X13,X11] : (p(s(bool,X11)) | s(bool,X11) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK33(X13,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f16505,plain,(
  ( ! [X321,X320,X322] : (p(s(bool,X320)) | p(s(bool,X321)) | s(bool,X321) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK33(X322,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8333,f873])).
fof(f8333,plain,(
  ( ! [X73] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK33(X73,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))))) ) | $spl4),
  inference(superposition,[],[f859,f886])).
fof(f859,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK33(X1,X0)))))) )),
  inference(cnf_transformation,[],[f580])).
fof(f16400,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK33(X39,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))) | s(bool,X38) != s(bool,X40)) ) | $spl4),
  inference(resolution,[],[f8333,f595])).
fof(f16397,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK33(X29,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8333,f594])).
fof(f16396,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK33(X25,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8333,f593])).
fof(f16395,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK33(X20,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))) | s(bool,X19) != s(bool,X21)) ) | $spl4),
  inference(resolution,[],[f8333,f591])).
fof(f16393,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK33(X13,'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8333,f589])).
fof(f14315,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK32(sK14,sK15),sK14))))) ) | $spl68),
  inference(cnf_transformation,[],[f14315_D])).
fof(f14315_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK32(sK14,sK15),sK14))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f14286,plain,(
  ( ! [X208,X209] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X209) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK32(sK14,sK15),sK14))))) ) | $spl8),
  inference(resolution,[],[f9897,f873])).
fof(f9897,plain,(
  p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK32(sK14,sK15),sK14))))) | $spl8),
  inference(resolution,[],[f8627,f864])).
fof(f8627,plain,(
  p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK14,sK15))))) | $spl8),
  inference(resolution,[],[f1039,f853])).
fof(f14181,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK32(sK14,sK15),sK14)))) | s(bool,X26) != s(bool,X27)) ) | $spl8),
  inference(resolution,[],[f9897,f595])).
fof(f14178,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK32(sK14,sK15),sK14))))) ) | $spl8),
  inference(resolution,[],[f9897,f594])).
fof(f14177,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK32(sK14,sK15),sK14))))) ) | $spl8),
  inference(resolution,[],[f9897,f593])).
fof(f14176,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK32(sK14,sK15),sK14)))) | s(bool,X11) != s(bool,X12)) ) | $spl8),
  inference(resolution,[],[f9897,f591])).
fof(f14174,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK32(sK14,sK15),sK14))))) ) | $spl8),
  inference(resolution,[],[f9897,f589])).
fof(f13921,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK37('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl66),
  inference(cnf_transformation,[],[f13921_D])).
fof(f13921_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK37('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f13876,plain,(
  ( ! [X208,X209] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X209) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK37('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8346,f873])).
fof(f8346,plain,(
  p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK37('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | $spl4),
  inference(superposition,[],[f868,f886])).
fof(f868,plain,(
  ( ! [X0] : (p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK37(X0)))))) )),
  inference(cnf_transformation,[],[f584])).
fof(f584,plain,(
  ! [X0] : p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK37(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK37])],[f376])).
fof(f376,plain,(
  ! [X0] : ? [X1] : p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X1))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X6] : ? [X11] : p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X6),s('type.hrat.hrat',X11))))),
  file('thm.hreal.HREAL_LT','thm.hreal.CUT_NONEMPTY')).
fof(f13771,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK37('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))) | s(bool,X26) != s(bool,X27)) ) | $spl4),
  inference(resolution,[],[f8346,f595])).
fof(f13768,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK37('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8346,f594])).
fof(f13767,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK37('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8346,f593])).
fof(f13766,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK37('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))) | s(bool,X11) != s(bool,X12)) ) | $spl4),
  inference(resolution,[],[f8346,f591])).
fof(f13764,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK37('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8346,f589])).
fof(f13724,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK31(sK15,sK14),sK14))))) ) | $spl64),
  inference(cnf_transformation,[],[f13724_D])).
fof(f13724_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK31(sK15,sK14),sK14))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13695,plain,(
  ( ! [X208,X209] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X209) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK31(sK15,sK14),sK14))))) ) | $spl10),
  inference(resolution,[],[f9418,f873])).
fof(f9418,plain,(
  p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK31(sK15,sK14),sK14))))) | $spl10),
  inference(resolution,[],[f1041,f864])).
fof(f13590,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK31(sK15,sK14),sK14)))) | s(bool,X26) != s(bool,X27)) ) | $spl10),
  inference(resolution,[],[f9418,f595])).
fof(f13587,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK31(sK15,sK14),sK14))))) ) | $spl10),
  inference(resolution,[],[f9418,f594])).
fof(f13586,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK31(sK15,sK14),sK14))))) ) | $spl10),
  inference(resolution,[],[f9418,f593])).
fof(f13585,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK31(sK15,sK14),sK14)))) | s(bool,X11) != s(bool,X12)) ) | $spl10),
  inference(resolution,[],[f9418,f591])).
fof(f13583,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK35(sK31(sK15,sK14),sK14))))) ) | $spl10),
  inference(resolution,[],[f9418,f589])).
fof(f13570,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK36('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl62),
  inference(cnf_transformation,[],[f13570_D])).
fof(f13570_D,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK36('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f13527,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | ~p(s(bool,X205)) | s(bool,X205) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK36('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8345,f873])).
fof(f8345,plain,(
  ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK36('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) | $spl4),
  inference(superposition,[],[f867,f886])).
fof(f867,plain,(
  ( ! [X0] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK36(X0)))))) )),
  inference(cnf_transformation,[],[f583])).
fof(f583,plain,(
  ! [X0] : ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK36(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK36])],[f375])).
fof(f375,plain,(
  ! [X0] : ? [X1] : ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X1))))),
  inference(flattening,[],[f374])).
fof(f374,plain,(
  ! [X0] : ? [X1] : ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',X1))))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X6] : ? [X11] : ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X6),s('type.hrat.hrat',X11))))),
  file('thm.hreal.HREAL_LT','thm.hreal.CUT_BOUNDED')).
fof(f13422,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK36('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))) | s(bool,X25) != s(bool,X27)) ) | $spl4),
  inference(resolution,[],[f8345,f595])).
fof(f13420,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK36('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8345,f594])).
fof(f13418,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK36('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8345,f593])).
fof(f13417,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK36('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))))) | s(bool,X10) != s(bool,X12)) ) | $spl4),
  inference(resolution,[],[f8345,f591])).
fof(f13415,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK36('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))))))) ) | $spl4),
  inference(resolution,[],[f8345,f589])).
fof(f12987,plain,(
  $false | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f12920,f12745])).
fof(f12745,plain,(
  ( ! [X2] : (p(s(bool,X2))) ) | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f12744,f829])).
fof(f829,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f570])).
fof(f570,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f342])).
fof(f342,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.hreal.HREAL_LT','thm.bool.EQ_CLAUSES|split|3')).
fof(f12744,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2))) ) | ($spl1 | $spl28)),
  inference(forward_demodulation,[],[f1745,f8181])).
fof(f8181,plain,(
  s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))) | $spl1),
  inference(resolution,[],[f883,f827])).
fof(f827,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f569])).
fof(f569,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f340])).
fof(f340,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.hreal.HREAL_LT','thm.bool.EQ_CLAUSES|split|2')).
fof(f12920,plain,(
  ( ! [X445,X444] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X444),s('type.hrat.hrat','const.hrat.hrat_mul_2'(s('type.hrat.hrat',X445),s('type.hrat.hrat',sK34(X445,X444)))))))) ) | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f863])).
fof(f863,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat','const.hrat.hrat_mul_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',sK34(X1,X0))))))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X1))))) )),
  inference(cnf_transformation,[],[f581])).
fof(f12986,plain,(
  $false | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f12915,f12745])).
fof(f12915,plain,(
  ( ! [X432,X433] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X432),s('type.hrat.hrat',X433))))) ) | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f768])).
fof(f768,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X0)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1))))) )),
  inference(cnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X0)))))),
  inference(ennf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) => ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X0)))))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) => ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X0)))))),
  inference(rectify,[],[f142])).
fof(f142,axiom,(
  ! [X11,X34] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X11),s('type.hrat.hrat',X34)))) => ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X34),s('type.hrat.hrat',X11)))))),
  file('thm.hreal.HREAL_LT','thm.hreal.HRAT_LT_GT')).
fof(f12985,plain,(
  $false | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f12914,f12745])).
fof(f12914,plain,(
  ( ! [X430,X431] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X430),s('type.hrat.hrat',X431))))) ) | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f768])).
fof(f12984,plain,(
  $false | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f12905,f12745])).
fof(f12905,plain,(
  ( ! [X418,X417] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X417),s('type.hreal.hreal',X418))))) ) | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f852])).
fof(f12983,plain,(
  $false | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f12903,f12745])).
fof(f12903,plain,(
  ( ! [X414] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X414))))) ) | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f863])).
fof(f12982,plain,(
  $false | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f12877,f12745])).
fof(f12877,plain,(
  ( ! [X337,X338] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X337),s('type.hrat.hrat',sK32(X338,X337)))))) ) | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f852])).
fof(f12981,plain,(
  $false | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f12875,f12745])).
fof(f12875,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f845])).
fof(f845,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f574])).
fof(f574,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f355])).
fof(f355,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f55])).
fof(f55,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.hreal.HREAL_LT','thm.bool.NOT_CLAUSES|split|1')).
fof(f12980,plain,(
  $false | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f12874,f12745])).
fof(f12874,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f843])).
fof(f843,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f354])).
fof(f354,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f56])).
fof(f56,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.hreal.HREAL_LT','thm.bool.NOT_CLAUSES|split|2')).
fof(f12979,plain,(
  $false | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f12868,f12745])).
fof(f12868,plain,(
  ( ! [X330] : (~p(s(bool,X330))) ) | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f797])).
fof(f797,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f550])).
fof(f550,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f549])).
fof(f549,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f329])).
fof(f329,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X13)))),
  file('thm.hreal.HREAL_LT','thm.bool.IMP_CLAUSES|split|4')).
fof(f12873,plain,(
  $false | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f841])).
fof(f841,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f353])).
fof(f353,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.hreal.HREAL_LT',aHL_FALSITY)).
fof(f12902,plain,(
  $false | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f860])).
fof(f12906,plain,(
  $false | ($spl1 | $spl28)),
  inference(resolution,[],[f12745,f867])).
fof(f12907,plain,(
  $false | ($spl1 | $spl13 | $spl28)),
  inference(resolution,[],[f12745,f1043])).
fof(f12908,plain,(
  $false | ($spl1 | $spl8 | $spl28)),
  inference(resolution,[],[f12745,f8626])).
fof(f8626,plain,(
  ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK14,sK15))))) | $spl8),
  inference(resolution,[],[f1039,f852])).
fof(f10044,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK14,sK15))))) ) | $spl60),
  inference(cnf_transformation,[],[f10044_D])).
fof(f10044_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK14,sK15))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f10015,plain,(
  ( ! [X208,X209] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X209) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK14,sK15))))) ) | $spl8),
  inference(resolution,[],[f8627,f873])).
fof(f9910,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK14,sK15)))) | s(bool,X26) != s(bool,X27)) ) | $spl8),
  inference(resolution,[],[f8627,f595])).
fof(f9907,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK14,sK15))))) ) | $spl8),
  inference(resolution,[],[f8627,f594])).
fof(f9906,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK14,sK15))))) ) | $spl8),
  inference(resolution,[],[f8627,f593])).
fof(f9905,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK14,sK15)))) | s(bool,X11) != s(bool,X12)) ) | $spl8),
  inference(resolution,[],[f8627,f591])).
fof(f9903,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK14,sK15))))) ) | $spl8),
  inference(resolution,[],[f8627,f589])).
fof(f9890,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK14,sK15))))) ) | $spl58),
  inference(cnf_transformation,[],[f9890_D])).
fof(f9890_D,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK14,sK15))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f9863,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | ~p(s(bool,X205)) | s(bool,X205) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK14,sK15))))) ) | $spl8),
  inference(resolution,[],[f8626,f873])).
fof(f9758,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK14,sK15)))) | s(bool,X25) != s(bool,X27)) ) | $spl8),
  inference(resolution,[],[f8626,f595])).
fof(f9756,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK14,sK15))))) ) | $spl8),
  inference(resolution,[],[f8626,f594])).
fof(f9754,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK14,sK15))))) ) | $spl8),
  inference(resolution,[],[f8626,f593])).
fof(f9753,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK14,sK15)))) | s(bool,X10) != s(bool,X12)) ) | $spl8),
  inference(resolution,[],[f8626,f591])).
fof(f9751,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK14,sK15))))) ) | $spl8),
  inference(resolution,[],[f8626,f589])).
fof(f9712,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK31(sK15,sK14))))) ) | $spl56),
  inference(cnf_transformation,[],[f9712_D])).
fof(f9712_D,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK31(sK15,sK14))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f9685,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | ~p(s(bool,X205)) | s(bool,X205) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK31(sK15,sK14))))) ) | $spl13),
  inference(resolution,[],[f1043,f873])).
fof(f9580,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK31(sK15,sK14)))) | s(bool,X25) != s(bool,X27)) ) | $spl13),
  inference(resolution,[],[f1043,f595])).
fof(f9578,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK31(sK15,sK14))))) ) | $spl13),
  inference(resolution,[],[f1043,f594])).
fof(f9576,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK31(sK15,sK14))))) ) | $spl13),
  inference(resolution,[],[f1043,f593])).
fof(f9575,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK31(sK15,sK14)))) | s(bool,X10) != s(bool,X12)) ) | $spl13),
  inference(resolution,[],[f1043,f591])).
fof(f9573,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK31(sK15,sK14))))) ) | $spl13),
  inference(resolution,[],[f1043,f589])).
fof(f9565,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK31(sK15,sK14))))) ) | $spl54),
  inference(cnf_transformation,[],[f9565_D])).
fof(f9565_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK31(sK15,sK14))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f9536,plain,(
  ( ! [X208,X209] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X209) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK31(sK15,sK14))))) ) | $spl10),
  inference(resolution,[],[f1041,f873])).
fof(f9431,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK31(sK15,sK14)))) | s(bool,X26) != s(bool,X27)) ) | $spl10),
  inference(resolution,[],[f1041,f595])).
fof(f9428,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK31(sK15,sK14))))) ) | $spl10),
  inference(resolution,[],[f1041,f594])).
fof(f9427,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK31(sK15,sK14))))) ) | $spl10),
  inference(resolution,[],[f1041,f593])).
fof(f9426,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK31(sK15,sK14)))) | s(bool,X11) != s(bool,X12)) ) | $spl10),
  inference(resolution,[],[f1041,f591])).
fof(f9424,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK31(sK15,sK14))))) ) | $spl10),
  inference(resolution,[],[f1041,f589])).
fof(f9387,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9257,f1049])).
fof(f9257,plain,(
  p(s(bool,'F_0')) | $spl24),
  inference(resolution,[],[f1049,f844])).
fof(f844,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f574])).
fof(f9386,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9256,f1049])).
fof(f9256,plain,(
  p(s(bool,'F_0')) | $spl24),
  inference(resolution,[],[f1049,f842])).
fof(f842,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f573])).
fof(f9385,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9253,f1049])).
fof(f9253,plain,(
  ( ! [X316] : (p(s(bool,X316))) ) | $spl24),
  inference(resolution,[],[f1049,f808])).
fof(f808,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f558])).
fof(f558,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f557])).
fof(f557,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X13] : ((p(s(bool,'T_0')) => p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.hreal.HREAL_LT','thm.bool.IMP_CLAUSES|split|0')).
fof(f9384,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9251,f1049])).
fof(f9251,plain,(
  p(s(bool,'F_0')) | $spl24),
  inference(resolution,[],[f1049,f805])).
fof(f805,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f556])).
fof(f556,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f555])).
fof(f555,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f413])).
fof(f413,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.hreal.HREAL_LT','thm.bool.IMP_CLAUSES|split|2')).
fof(f9383,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9250,f1049])).
fof(f9250,plain,(
  ( ! [X314] : (p(s(bool,X314))) ) | $spl24),
  inference(resolution,[],[f1049,f802])).
fof(f802,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f554])).
fof(f554,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f553])).
fof(f553,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f412])).
fof(f412,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.hreal.HREAL_LT','thm.bool.IMP_CLAUSES|split|1')).
fof(f9382,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9248,f1049])).
fof(f9248,plain,(
  ( ! [X312] : (p(s(bool,X312))) ) | $spl24),
  inference(resolution,[],[f1049,f799])).
fof(f799,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f552])).
fof(f552,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f551])).
fof(f551,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.hreal.HREAL_LT','thm.bool.IMP_CLAUSES|split|3')).
fof(f9258,plain,(
  $false | $spl24),
  inference(resolution,[],[f1049,f846])).
fof(f846,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f17])).
fof(f17,axiom,(
  p(s(bool,'T_0'))),
  file('thm.hreal.HREAL_LT','thm.bool.TRUTH')).
fof(f9259,plain,(
  $false | $spl24),
  inference(resolution,[],[f1049,f847])).
fof(f847,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.hreal.HREAL_LT',aHL_TRUTH)).
fof(f9260,plain,(
  $false | $spl24),
  inference(resolution,[],[f1049,f879])).
fof(f9271,plain,(
  $false | ($spl8 | $spl24)),
  inference(resolution,[],[f1049,f1039])).
fof(f9291,plain,(
  $false | $spl24),
  inference(resolution,[],[f1049,f859])).
fof(f9294,plain,(
  $false | $spl24),
  inference(resolution,[],[f1049,f868])).
fof(f9297,plain,(
  $false | $spl24),
  inference(resolution,[],[f1049,f872])).
fof(f872,plain,(
  ( ! [X0] : (p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),'const.hreal.cut_1'(s('type.hreal.hreal',X0))))))) )),
  inference(cnf_transformation,[],[f379])).
fof(f379,plain,(
  ! [X0] : p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),'const.hreal.cut_1'(s('type.hreal.hreal',X0))))))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X6] : p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),'const.hreal.cut_1'(s('type.hreal.hreal',X6))))))),
  file('thm.hreal.HREAL_LT','thm.hreal.CUT_ISACUT')).
fof(f9304,plain,(
  $false | $spl24),
  inference(resolution,[],[f1049,f730])).
fof(f730,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK18(X1,X0)),s('type.hrat.hrat',X1))))) )),
  inference(cnf_transformation,[],[f510])).
fof(f510,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK18(X1,X0)),s('type.hrat.hrat',X0)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK18(X1,X0)),s('type.hrat.hrat',X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f293])).
fof(f293,plain,(
  ! [X0,X1] : ? [X2] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',X0)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',X1)))))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X11,X34] : ? [X43] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X43),s('type.hrat.hrat',X11)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X43),s('type.hrat.hrat',X34)))))),
  file('thm.hreal.HREAL_LT','thm.hreal.HRAT_DOWN2')).
fof(f9305,plain,(
  $false | $spl24),
  inference(resolution,[],[f1049,f729])).
fof(f729,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK18(X1,X0)),s('type.hrat.hrat',X0))))) )),
  inference(cnf_transformation,[],[f510])).
fof(f9307,plain,(
  $false | $spl24),
  inference(resolution,[],[f1049,f771])).
fof(f771,plain,(
  ( ! [X0] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK20(X0)),s('type.hrat.hrat',X0))))) )),
  inference(cnf_transformation,[],[f526])).
fof(f526,plain,(
  ! [X0] : p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK20(X0)),s('type.hrat.hrat',X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f313])).
fof(f313,plain,(
  ! [X0] : ? [X1] : p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X0))))),
  inference(rectify,[],[f139])).
fof(f139,axiom,(
  ! [X11] : ? [X34] : p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X34),s('type.hrat.hrat',X11))))),
  file('thm.hreal.HREAL_LT','thm.hreal.HRAT_DOWN')).
fof(f3255,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,X22)) ) | $spl38),
  inference(cnf_transformation,[],[f3255_D])).
fof(f3255_D,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,X22)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f9126,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f9100])).
fof(f9100,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | ~p(s(bool,X34))) ) | $spl22),
  inference(superposition,[],[f1048,f833])).
fof(f833,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f572])).
fof(f572,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.hreal.HREAL_LT','thm.bool.EQ_CLAUSES|split|1')).
fof(f9127,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f9099])).
fof(f9099,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32)) | ~p(s(bool,X32))) ) | $spl22),
  inference(superposition,[],[f1048,f831])).
fof(f8976,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl52),
  inference(cnf_transformation,[],[f8976_D])).
fof(f8976_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f8972,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f8944])).
fof(f8944,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30)) | p(s(bool,X30))) ) | $spl16),
  inference(superposition,[],[f1045,f829])).
fof(f8973,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f8943])).
fof(f8943,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl16),
  inference(superposition,[],[f1045,f827])).
fof(f8778,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14)))) ) | $spl50),
  inference(cnf_transformation,[],[f8778_D])).
fof(f8778_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f8743,plain,(
  ( ! [X206,X205] : (p(s(bool,X205)) | p(s(bool,X206)) | s(bool,X206) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14)))) ) | $spl8),
  inference(resolution,[],[f1039,f873])).
fof(f8638,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl8),
  inference(resolution,[],[f1039,f595])).
fof(f8635,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14)))) ) | $spl8),
  inference(resolution,[],[f1039,f594])).
fof(f8634,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14)))) ) | $spl8),
  inference(resolution,[],[f1039,f593])).
fof(f8633,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl8),
  inference(resolution,[],[f1039,f591])).
fof(f8631,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14)))) ) | $spl8),
  inference(resolution,[],[f1039,f589])).
fof(f1044,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl14),
  inference(cnf_transformation,[],[f1044_D])).
fof(f1044_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f8186,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X201) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl1),
  inference(resolution,[],[f883,f873])).
fof(f8081,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f883,f595])).
fof(f8079,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl1),
  inference(resolution,[],[f883,f594])).
fof(f8077,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl1),
  inference(resolution,[],[f883,f593])).
fof(f8076,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f883,f591])).
fof(f8074,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl1),
  inference(resolution,[],[f883,f589])).
fof(f8053,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f7988,f1050])).
fof(f7988,plain,(
  ( ! [X445,X444] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X444),s('type.hrat.hrat','const.hrat.hrat_mul_2'(s('type.hrat.hrat',X445),s('type.hrat.hrat',sK34(X445,X444)))))))) ) | $spl26),
  inference(resolution,[],[f1050,f863])).
fof(f8052,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f7983,f1050])).
fof(f7983,plain,(
  ( ! [X432,X433] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X432),s('type.hrat.hrat',X433))))) ) | $spl26),
  inference(resolution,[],[f1050,f768])).
fof(f8051,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f7982,f1050])).
fof(f7982,plain,(
  ( ! [X430,X431] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X430),s('type.hrat.hrat',X431))))) ) | $spl26),
  inference(resolution,[],[f1050,f768])).
fof(f8050,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f7975,f1050])).
fof(f7975,plain,(
  ( ! [X418,X417] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X417),s('type.hreal.hreal',X418))))) ) | $spl26),
  inference(resolution,[],[f1050,f852])).
fof(f8049,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f7973,f1050])).
fof(f7973,plain,(
  ( ! [X414] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat','const.hrat.hrat_|49|_0'),s('type.hrat.hrat',X414))))) ) | $spl26),
  inference(resolution,[],[f1050,f863])).
fof(f8048,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f7947,f1050])).
fof(f7947,plain,(
  ( ! [X337,X338] : (~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X337),s('type.hrat.hrat',sK32(X338,X337)))))) ) | $spl26),
  inference(resolution,[],[f1050,f852])).
fof(f8047,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f7945,f1050])).
fof(f7945,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(resolution,[],[f1050,f845])).
fof(f8046,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f7944,f1050])).
fof(f7944,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(resolution,[],[f1050,f843])).
fof(f8045,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f7938,f1050])).
fof(f7938,plain,(
  ( ! [X330] : (~p(s(bool,X330))) ) | $spl26),
  inference(resolution,[],[f1050,f797])).
fof(f7943,plain,(
  $false | $spl26),
  inference(resolution,[],[f1050,f841])).
fof(f7972,plain,(
  $false | $spl26),
  inference(resolution,[],[f1050,f860])).
fof(f7976,plain,(
  $false | $spl26),
  inference(resolution,[],[f1050,f867])).
fof(f7815,plain,(
  $false | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f886,f884])).
fof(f7812,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl49),
  inference(cnf_transformation,[],[f7812_D])).
fof(f7812_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f7814,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X22))) ) | $spl42),
  inference(subsumption_resolution,[],[f6286,f831])).
fof(f6286,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X22) | ~p(s(bool,X22))) ) | $spl42),
  inference(superposition,[],[f6204,f831])).
fof(f7813,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X23))) ) | $spl42),
  inference(subsumption_resolution,[],[f6287,f831])).
fof(f6287,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23))) ) | $spl42),
  inference(superposition,[],[f6204,f833])).
fof(f6298,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X26))) ) | $spl42),
  inference(trivial_inequality_removal,[],[f6288])).
fof(f6288,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,X26) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X26))) ) | $spl42),
  inference(superposition,[],[f6204,f829])).
fof(f6297,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X27))) ) | $spl42),
  inference(trivial_inequality_removal,[],[f6289])).
fof(f6289,plain,(
  ( ! [X27] : (s(bool,X27) != s(bool,X27) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X27))) ) | $spl42),
  inference(superposition,[],[f6204,f827])).
fof(f7810,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1597,f1793])).
fof(f1793,plain,(
  ( ! [X0] : (s('type.hreal.hreal',sK15) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',sK14)))) ) | $spl2),
  inference(superposition,[],[f884,f634])).
fof(f1597,plain,(
  s('type.hreal.hreal',sK15) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14))),s('type.hreal.hreal',sK14))) | $spl0),
  inference(resolution,[],[f882,f632])).
fof(f7600,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK32(sK15,sK14)),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) | $spl46),
  inference(cnf_transformation,[],[f7600_D])).
fof(f7600_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK32(sK15,sK14)),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f7577,plain,(
  ( ! [X208,X209] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X209) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK32(sK15,sK14)),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) | $spl0),
  inference(resolution,[],[f2218,f873])).
fof(f2218,plain,(
  p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK32(sK15,sK14)),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) | $spl0),
  inference(resolution,[],[f1594,f865])).
fof(f1594,plain,(
  p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14))))) | $spl0),
  inference(resolution,[],[f882,f853])).
fof(f7472,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK32(sK15,sK14)),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15)))) | s(bool,X26) != s(bool,X27)) ) | $spl0),
  inference(resolution,[],[f2218,f595])).
fof(f7469,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK32(sK15,sK14)),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) | $spl0),
  inference(resolution,[],[f2218,f594])).
fof(f7468,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK32(sK15,sK14)),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) | $spl0),
  inference(resolution,[],[f2218,f593])).
fof(f7467,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK32(sK15,sK14)),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15)))) | s(bool,X11) != s(bool,X12)) ) | $spl0),
  inference(resolution,[],[f2218,f591])).
fof(f7465,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK32(sK15,sK14)),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) | $spl0),
  inference(resolution,[],[f2218,f589])).
fof(f6680,plain,(
  s(bool,'F_0') != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15)))) | $spl45),
  inference(cnf_transformation,[],[f6680_D])).
fof(f6680_D,plain,(
  s(bool,'F_0') != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f6672,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) | $spl0),
  inference(resolution,[],[f6411,f3932])).
fof(f3932,plain,(
  ( ! [X50,X51] : (sP0('const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))),X50,X51)) ) | $spl0),
  inference(resolution,[],[f2217,f614])).
fof(f614,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f469])).
fof(f469,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f468])).
fof(f468,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f432])).
fof(f432,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f2217,plain,(
  p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) | $spl0),
  inference(resolution,[],[f1594,f864])).
fof(f6411,plain,(
  ( ! [X2,X3] : (~sP0(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f6404])).
fof(f6404,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | ~sP0(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | $spl0),
  inference(resolution,[],[f6394,f618])).
fof(f618,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f470])).
fof(f470,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP3(X2,X1,X0)) & (~sP3(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f436])).
fof(f436,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP3(X2,X1,X0))),
  inference(definition_folding,[],[f266,f435,f434,f433,f432])).
fof(f433,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f434,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f435,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X33,X39,X40] : ((p(s(bool,X40)) <=> s(bool,X33) = s(bool,X39)) <=> ((p(s(bool,X40)) | p(s(bool,X39)) | p(s(bool,X33))) & (p(s(bool,X40)) | ~p(s(bool,X33)) | ~p(s(bool,X39))) & (p(s(bool,X39)) | ~p(s(bool,X33)) | ~p(s(bool,X40))) & (p(s(bool,X33)) | ~p(s(bool,X39)) | ~p(s(bool,X40)))))),
  file('thm.hreal.HREAL_LT','thm.sat.sat_2|49|_prove_ltb')).
fof(f6394,plain,(
  ( ! [X6,X7] : (sP3(X6,'F_0',X7) | p(s(bool,X6)) | ~sP0(X7,'F_0',X6)) ) | $spl0),
  inference(subsumption_resolution,[],[f6391,f606])).
fof(f606,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f463])).
fof(f463,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f462])).
fof(f462,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f434])).
fof(f6391,plain,(
  ( ! [X6,X7] : (sP3(X6,'F_0',X7) | p(s(bool,X6)) | ~sP2(X6,X7,'F_0') | ~sP0(X7,'F_0',X6)) ) | $spl0),
  inference(resolution,[],[f6385,f602])).
fof(f602,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f460])).
fof(f460,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f459])).
fof(f459,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f435])).
fof(f6385,plain,(
  ( ! [X6,X7] : (sP1('F_0',X6,X7)) ) | $spl0),
  inference(subsumption_resolution,[],[f6384,f609])).
fof(f609,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f466])).
fof(f466,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f465])).
fof(f465,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f433])).
fof(f6384,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | sP1('F_0',X6,X7)) ) | $spl0),
  inference(subsumption_resolution,[],[f6378,f608])).
fof(f608,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f467])).
fof(f6378,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | sP1('F_0',X6,X7)) ) | $spl0),
  inference(resolution,[],[f6370,f597])).
fof(f597,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,X2,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f461])).
fof(f6370,plain,(
  ( ! [X4,X5] : (sP3(X5,'F_0',X4) | ~p(s(bool,X4)) | p(s(bool,X5))) ) | $spl0),
  inference(resolution,[],[f6331,f606])).
fof(f6331,plain,(
  ( ! [X2,X3] : (~sP2(X3,X2,'F_0') | ~p(s(bool,X2)) | sP3(X3,'F_0',X2)) ) | $spl0),
  inference(subsumption_resolution,[],[f6330,f831])).
fof(f6330,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | $spl0),
  inference(subsumption_resolution,[],[f6326,f6314])).
fof(f6314,plain,(
  ( ! [X4,X5] : (sP0(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl0),
  inference(subsumption_resolution,[],[f6307,f612])).
fof(f612,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f469])).
fof(f6307,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP0(X4,'F_0',X5)) ) | $spl0),
  inference(resolution,[],[f2897,f596])).
fof(f596,plain,(
  ( ! [X2,X0,X1] : (sP0(X2,X1,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f461])).
fof(f2897,plain,(
  ( ! [X4,X5] : (sP3(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl0),
  inference(superposition,[],[f2892,f615])).
fof(f615,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f470])).
fof(f2892,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl0),
  inference(subsumption_resolution,[],[f2886,f882])).
fof(f2886,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) | $spl0),
  inference(superposition,[],[f1706,f831])).
fof(f1706,plain,(
  s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))) | $spl0),
  inference(resolution,[],[f882,f826])).
fof(f826,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f569])).
fof(f6326,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0') | ~sP0(X2,'F_0',X3)) ) | $spl0),
  inference(resolution,[],[f6315,f600])).
fof(f600,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f461])).
fof(f6315,plain,(
  ( ! [X6,X7] : (sP1('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl0),
  inference(subsumption_resolution,[],[f6308,f608])).
fof(f6308,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP1('F_0',X6,X7)) ) | $spl0),
  inference(resolution,[],[f2897,f597])).
fof(f6204,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f6204_D])).
fof(f6204_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f6203,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f6195])).
fof(f6195,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl0),
  inference(resolution,[],[f2896,f618])).
fof(f2896,plain,(
  ( ! [X2,X3] : (sP3(X3,X2,'F_0') | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl0),
  inference(superposition,[],[f2892,f615])).
fof(f6181,plain,(
  $false | ($spl0 | $spl36)),
  inference(trivial_inequality_removal,[],[f6155])).
fof(f6155,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl0 | $spl36)),
  inference(superposition,[],[f2365,f2333])).
fof(f2333,plain,(
  s(bool,'T_0') = s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14)))) | $spl0),
  inference(resolution,[],[f1594,f831])).
fof(f6180,plain,(
  $false | $spl36),
  inference(equality_resolution,[],[f2365])).
fof(f5384,plain,(
  $false | ($spl0 | $spl32)),
  inference(trivial_inequality_removal,[],[f5358])).
fof(f5358,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl0 | $spl32)),
  inference(superposition,[],[f2209,f2176])).
fof(f2176,plain,(
  s(bool,'F_0') = s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14)))) | $spl0),
  inference(resolution,[],[f1593,f827])).
fof(f1593,plain,(
  ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14))))) | $spl0),
  inference(resolution,[],[f882,f852])).
fof(f5383,plain,(
  $false | $spl32),
  inference(equality_resolution,[],[f2209])).
fof(f4344,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4212,f4087])).
fof(f4087,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4086,f833])).
fof(f4086,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X1))) ) | ($spl0 | $spl14)),
  inference(forward_demodulation,[],[f1044,f1708])).
fof(f1708,plain,(
  s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))) | $spl0),
  inference(resolution,[],[f882,f831])).
fof(f4212,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f844])).
fof(f4343,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4211,f4087])).
fof(f4211,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f842])).
fof(f4342,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4208,f4087])).
fof(f4208,plain,(
  ( ! [X316] : (p(s(bool,X316))) ) | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f808])).
fof(f4341,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4206,f4087])).
fof(f4206,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f805])).
fof(f4340,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4205,f4087])).
fof(f4205,plain,(
  ( ! [X314] : (p(s(bool,X314))) ) | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f802])).
fof(f4339,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4203,f4087])).
fof(f4203,plain,(
  ( ! [X312] : (p(s(bool,X312))) ) | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f799])).
fof(f4213,plain,(
  $false | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f846])).
fof(f4214,plain,(
  $false | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f847])).
fof(f4215,plain,(
  $false | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f879])).
fof(f4245,plain,(
  $false | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f859])).
fof(f4248,plain,(
  $false | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f868])).
fof(f4249,plain,(
  $false | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f1594])).
fof(f4250,plain,(
  $false | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f2217])).
fof(f4253,plain,(
  $false | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f872])).
fof(f4260,plain,(
  $false | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f730])).
fof(f4261,plain,(
  $false | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f729])).
fof(f4263,plain,(
  $false | ($spl0 | $spl14)),
  inference(resolution,[],[f4087,f771])).
fof(f4054,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) | $spl40),
  inference(cnf_transformation,[],[f4054_D])).
fof(f4054_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4025,plain,(
  ( ! [X208,X209] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X209) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) | $spl0),
  inference(resolution,[],[f2217,f873])).
fof(f3920,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15)))) | s(bool,X26) != s(bool,X27)) ) | $spl0),
  inference(resolution,[],[f2217,f595])).
fof(f3917,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) | $spl0),
  inference(resolution,[],[f2217,f594])).
fof(f3916,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) | $spl0),
  inference(resolution,[],[f2217,f593])).
fof(f3915,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15)))) | s(bool,X11) != s(bool,X12)) ) | $spl0),
  inference(resolution,[],[f2217,f591])).
fof(f3913,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK35(sK32(sK15,sK14),sK15))))) ) | $spl0),
  inference(resolution,[],[f2217,f589])).
fof(f1047,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X3)) | p(s(bool,X4))) ) | $spl20),
  inference(cnf_transformation,[],[f1047_D])).
fof(f1047_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X3)) | p(s(bool,X4))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3277,plain,(
  ( ! [X99,X97,X98] : (s(bool,'T_0') != s(bool,X99) | p(s(bool,X97)) | ~p(s(bool,X98)) | s(bool,X97) != s(bool,X98)) ) | $spl0),
  inference(forward_demodulation,[],[f3276,f1708])).
fof(f3276,plain,(
  ( ! [X99,X97,X98] : (p(s(bool,X97)) | ~p(s(bool,X98)) | s(bool,X97) != s(bool,X98) | s(bool,X99) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl0),
  inference(subsumption_resolution,[],[f2980,f873])).
fof(f2980,plain,(
  ( ! [X99,X97,X98] : (~p(s(bool,'T_0')) | p(s(bool,X97)) | ~p(s(bool,X98)) | s(bool,X97) != s(bool,X98) | s(bool,X99) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl0),
  inference(superposition,[],[f595,f1708])).
fof(f3273,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') != s(bool,X85) | p(s(bool,X86)) | ~p(s(bool,X87)) | s(bool,X86) != s(bool,X87)) ) | $spl0),
  inference(subsumption_resolution,[],[f2976,f830])).
fof(f830,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f571])).
fof(f2976,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') != s(bool,X85) | p(s(bool,X86)) | ~p(s(bool,X85)) | ~p(s(bool,X87)) | s(bool,X86) != s(bool,X87)) ) | $spl0),
  inference(superposition,[],[f595,f1708])).
fof(f3272,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') != s(bool,X82) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X84)) ) | $spl0),
  inference(subsumption_resolution,[],[f2975,f873])).
fof(f2975,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') != s(bool,X82) | p(s(bool,X83)) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X84)) ) | $spl0),
  inference(superposition,[],[f595,f1708])).
fof(f3260,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') != s(bool,X39) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38)) ) | $spl0),
  inference(forward_demodulation,[],[f3259,f1708])).
fof(f3259,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38) | s(bool,X39) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl0),
  inference(subsumption_resolution,[],[f2960,f873])).
fof(f2960,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38) | s(bool,X39) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl0),
  inference(superposition,[],[f591,f1708])).
fof(f3256,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X25) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27)) ) | $spl0),
  inference(subsumption_resolution,[],[f2956,f873])).
fof(f2956,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X25) | p(s(bool,X26)) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27)) ) | $spl0),
  inference(superposition,[],[f591,f1708])).
fof(f3254,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X22) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,X24)) ) | $spl0),
  inference(subsumption_resolution,[],[f2955,f830])).
fof(f2955,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X22) | p(s(bool,X23)) | ~p(s(bool,X22)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,X24)) ) | $spl0),
  inference(superposition,[],[f591,f1708])).
fof(f2439,plain,(
  $false | $spl18),
  inference(equality_resolution,[],[f1046])).
fof(f2364,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl34),
  inference(cnf_transformation,[],[f2364_D])).
fof(f2364_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2335,plain,(
  ( ! [X208,X209] : (p(s(bool,X208)) | p(s(bool,X209)) | s(bool,X209) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl0),
  inference(resolution,[],[f1594,f873])).
fof(f2230,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14)))) | s(bool,X26) != s(bool,X27)) ) | $spl0),
  inference(resolution,[],[f1594,f595])).
fof(f2365,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl36),
  inference(cnf_transformation,[],[f2365_D])).
fof(f2365_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2229,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl0),
  inference(resolution,[],[f1594,f595])).
fof(f2227,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl0),
  inference(resolution,[],[f1594,f594])).
fof(f2226,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl0),
  inference(resolution,[],[f1594,f593])).
fof(f2225,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14)))) | s(bool,X11) != s(bool,X12)) ) | $spl0),
  inference(resolution,[],[f1594,f591])).
fof(f2224,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X9) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl0),
  inference(resolution,[],[f1594,f591])).
fof(f2223,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl0),
  inference(resolution,[],[f1594,f589])).
fof(f2208,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl30),
  inference(cnf_transformation,[],[f2208_D])).
fof(f2208_D,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2181,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | ~p(s(bool,X205)) | s(bool,X205) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl0),
  inference(resolution,[],[f1593,f873])).
fof(f2076,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X26) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14)))) | s(bool,X25) != s(bool,X27)) ) | $spl0),
  inference(resolution,[],[f1593,f595])).
fof(f2074,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl0),
  inference(resolution,[],[f1593,f594])).
fof(f2209,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl32),
  inference(cnf_transformation,[],[f2209_D])).
fof(f2209_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2073,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X18) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14)))) | s(bool,X16) != s(bool,X17)) ) | $spl0),
  inference(resolution,[],[f1593,f593])).
fof(f2072,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl0),
  inference(resolution,[],[f1593,f593])).
fof(f2071,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14)))) | s(bool,X10) != s(bool,X12)) ) | $spl0),
  inference(resolution,[],[f1593,f591])).
fof(f2070,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X9) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14)))) | s(bool,X7) != s(bool,X8)) ) | $spl0),
  inference(resolution,[],[f1593,f589])).
fof(f2069,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK32(sK15,sK14))))) ) | $spl0),
  inference(resolution,[],[f1593,f589])).
fof(f1745,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl28),
  inference(cnf_transformation,[],[f1745_D])).
fof(f1745_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1710,plain,(
  ( ! [X206,X205] : (p(s(bool,X205)) | p(s(bool,X206)) | s(bool,X206) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl0),
  inference(resolution,[],[f882,f873])).
fof(f1605,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl0),
  inference(resolution,[],[f882,f595])).
fof(f1046,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl18),
  inference(cnf_transformation,[],[f1046_D])).
fof(f1046_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1604,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl0),
  inference(resolution,[],[f882,f595])).
fof(f1602,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl0),
  inference(resolution,[],[f882,f594])).
fof(f1601,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl0),
  inference(resolution,[],[f882,f593])).
fof(f1600,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl0),
  inference(resolution,[],[f882,f591])).
fof(f1599,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl0),
  inference(resolution,[],[f882,f591])).
fof(f1598,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl0),
  inference(resolution,[],[f882,f589])).
fof(f1580,plain,(
  $false | ($spl0 | $spl6)),
  inference(trivial_inequality_removal,[],[f1438])).
fof(f1438,plain,(
  s('type.hreal.hreal',sK14) != s('type.hreal.hreal',sK14) | ($spl0 | $spl6)),
  inference(resolution,[],[f1434,f848])).
fof(f1434,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK14)))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f882,f1037])).
fof(f1350,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1252,f631])).
fof(f1252,plain,(
  s('type.hreal.hreal',sK14) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))) | ($spl4 | $spl6)),
  inference(backward_demodulation,[],[f1037,f886])).
fof(f1003,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X201) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl1),
  inference(resolution,[],[f883,f873])).
fof(f898,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f883,f595])).
fof(f896,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl1),
  inference(resolution,[],[f883,f594])).
fof(f895,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl1),
  inference(resolution,[],[f883,f593])).
fof(f894,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl1),
  inference(resolution,[],[f883,f593])).
fof(f893,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f883,f591])).
fof(f892,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl1),
  inference(resolution,[],[f883,f589])).
fof(f891,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) ) | $spl1),
  inference(resolution,[],[f883,f589])).
fof(f1037,plain,(
  s('type.hreal.hreal',sK14) = s('type.hreal.hreal',sK15) | $spl6),
  inference(cnf_transformation,[],[f1037_D])).
fof(f1037_D,plain,(
  s('type.hreal.hreal',sK14) = s('type.hreal.hreal',sK15) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1043,plain,(
  ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK31(sK15,sK14))))) | $spl13),
  inference(cnf_transformation,[],[f1043_D])).
fof(f1043_D,plain,(
  ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK31(sK15,sK14))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f890,plain,(
  ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK15),s('type.hrat.hrat',sK31(sK15,sK14))))) | s('type.hreal.hreal',sK14) = s('type.hreal.hreal',sK15) | $spl1),
  inference(resolution,[],[f883,f851])).
fof(f851,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | ~p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X1),s('type.hrat.hrat',sK31(X1,X0))))) | s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1)) )),
  inference(cnf_transformation,[],[f578])).
fof(f1041,plain,(
  p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK31(sK15,sK14))))) | $spl10),
  inference(cnf_transformation,[],[f1041_D])).
fof(f1041_D,plain,(
  p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK31(sK15,sK14))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f889,plain,(
  p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',sK14),s('type.hrat.hrat',sK31(sK15,sK14))))) | s('type.hreal.hreal',sK14) = s('type.hreal.hreal',sK15) | $spl1),
  inference(resolution,[],[f883,f850])).
fof(f850,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | p(s(bool,'const.hreal.cut_2'(s('type.hreal.hreal',X0),s('type.hrat.hrat',sK31(X1,X0))))) | s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1)) )),
  inference(cnf_transformation,[],[f578])).
fof(f1039,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14)))) | $spl8),
  inference(cnf_transformation,[],[f1039_D])).
fof(f1039_D,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f888,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14)))) | s('type.hreal.hreal',sK14) = s('type.hreal.hreal',sK15) | $spl1),
  inference(resolution,[],[f883,f855])).
fof(f855,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0)))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1)) )),
  inference(cnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0,X1] : (s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : (s('type.hreal.hreal',X6) = s('type.hreal.hreal',X7) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7)))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X7),s('type.hreal.hreal',X6)))))),
  file('thm.hreal.HREAL_LT','thm.hreal.HREAL_LT_TOTAL')).
fof(f887,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK15),s('type.hreal.hreal',sK14)))) | s('type.hreal.hreal',sK14) = s('type.hreal.hreal',sK15) | $spl1),
  inference(resolution,[],[f883,f855])).
fof(f882,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) | $spl0),
  inference(cnf_transformation,[],[f882_D])).
fof(f882_D,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f886,plain,(
  s('type.hreal.hreal',sK15) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))) | $spl4),
  inference(cnf_transformation,[],[f886_D])).
fof(f886_D,plain,(
  s('type.hreal.hreal',sK15) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f586,plain,(
  s('type.hreal.hreal',sK15) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))))),
  inference(cnf_transformation,[],[f456])).
fof(f456,plain,(
  (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) | s('type.hreal.hreal',sK15) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK16)))) & (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) | ! [X2] : s('type.hreal.hreal',sK15) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15,sK16])],[f455])).
fof(f455,plain,(
  ? [X0,X1] : ((p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | ? [X3] : s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X3)))) & (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | ! [X2] : s('type.hreal.hreal',X1) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2)))))),
  inference(rectify,[],[f454])).
fof(f454,plain,(
  ? [X0,X1] : ((p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | ? [X2] : s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2)))) & (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | ! [X2] : s('type.hreal.hreal',X1) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2)))))),
  inference(nnf_transformation,[],[f380])).
fof(f380,plain,(
  ? [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) <~> ? [X2] : s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) <=> ? [X2] : s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7)))) <=> ? [X8] : s('type.hreal.hreal',X7) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7)))) <=> ? [X8] : s('type.hreal.hreal',X7) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X8))))),
  file('thm.hreal.HREAL_LT','thm.hreal.HREAL_LT')).
fof(f883,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) | $spl1),
  inference(cnf_transformation,[],[f883_D])).
fof(f883_D,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f884,plain,(
  ( ! [X2] : (s('type.hreal.hreal',sK15) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',X2)))) ) | $spl2),
  inference(cnf_transformation,[],[f884_D])).
fof(f884_D,plain,(
  ( ! [X2] : (s('type.hreal.hreal',sK15) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',X2)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f587,plain,(
  ( ! [X2] : (s('type.hreal.hreal',sK15) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',X2))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK14),s('type.hreal.hreal',sK15))))) )),
  inference(cnf_transformation,[],[f456])).
% SZS output end Proof for thm.hreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 647
% Passive clauses: 9596
% Generated clauses: 20994
% Final active clauses: 456
% Final passive clauses: 6041
% Input formulas: 261
% Initial clauses: 287
% 
% Selected by SInE selection: 83
% SInE iterations: 1
% 
% Duplicate literals: 494
% Trivial inequalities: 61
% Fw subsumption resolutions: 2452
% Fw demodulations: 1305
% Bw demodulations: 17
% 
% Simple tautologies: 953
% Equational tautologies: 10
% Forward subsumptions: 5697
% 
% Binary resolution: 6979
% Forward superposition: 6222
% Backward superposition: 3348
% Self superposition: 1
% Equality resolution: 9
% 
% Split clauses: 205
% Split components: 439
% SAT solver clauses: 276
% SAT solver unit clauses: 31
% SAT solver binary clauses: 39
% 
% Sat splits: 205
% Sat splitting refutations: 70
% 
% Memory used [KB]: 8571
% Time elapsed: 0.386 s
% ------------------------------
% ------------------------------
% Success in time 12.139 s
