Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integral
% (57762)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1272
% Passive clauses: 11270
% Generated clauses: 31304
% Final active clauses: 1183
% Final passive clauses: 9428
% Input formulas: 49
% Initial clauses: 206
% 
% Duplicate literals: 508
% Trivial inequalities: 143
% Fw subsumption resolutions: 3568
% Fw demodulations: 1960
% Bw demodulations: 79
% 
% Simple tautologies: 306
% Equational tautologies: 41
% Forward subsumptions: 12412
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 13354
% Factoring: 1525
% Forward superposition: 4163
% Backward superposition: 5336
% Self superposition: 29
% Equality factoring: 26
% Equality resolution: 22
% 
% Split clauses: 789
% Split components: 1584
% SAT solver clauses: 1357
% SAT solver unit clauses: 20
% SAT solver binary clauses: 279
% 
% Sat splits: 789
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1354
% TWLsolver calls for satisfiability: 692
% 
% Memory used [KB]: 14328
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integral
% (57773)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 461
% Passive clauses: 2626
% Generated clauses: 14055
% Final active clauses: 446
% Final passive clauses: 2128
% Input formulas: 49
% Initial clauses: 206
% 
% Duplicate literals: 266
% Trivial inequalities: 81
% Fw subsumption resolutions: 1110
% Fw demodulations: 1244
% Bw demodulations: 48
% Condensations: 360
% Global subsumptions: 13
% 
% Simple tautologies: 319
% Equational tautologies: 20
% Forward subsumptions: 4851
% 
% Binary resolution: 3437
% Factoring: 154
% Forward superposition: 2037
% Backward superposition: 2377
% Self superposition: 37
% Equality factoring: 30
% Equality resolution: 9
% 
% Split clauses: 77
% Split components: 154
% SAT solver clauses: 3270
% SAT solver unit clauses: 193
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 77
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 3218
% TWLsolver calls for satisfiability: 1842
% 
% Memory used [KB]: 13688
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integral
% (57797)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2355
% Passive clauses: 13800
% Generated clauses: 125478
% Final active clauses: 1870
% Final passive clauses: 4908
% Input formulas: 49
% Initial clauses: 206
% Discarded non-redundant clauses: 5783
% 
% Duplicate literals: 7206
% Trivial inequalities: 106
% Fw subsumption resolutions: 4770
% Fw demodulations: 18207
% Bw demodulations: 2923
% Condensations: 10408
% Global subsumptions: 618
% 
% Simple tautologies: 166
% Equational tautologies: 566
% Forward subsumptions: 66920
% Bw demodulations to eq. taut.: 94
% 
% Unit resulting resolution: 35510
% Factoring: 1
% Forward superposition: 19430
% Backward superposition: 24528
% Self superposition: 177
% Equality factoring: 144
% Equality resolution: 10
% 
% Split clauses: 356
% Split components: 729
% SAT solver clauses: 16934
% SAT solver unit clauses: 1955
% SAT solver binary clauses: 5465
% 
% Sat splits: 356
% Sat splitting refutations: 1802
% 
% Memory used [KB]: 93260
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integral
% (58065)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2097
% Passive clauses: 49217
% Generated clauses: 198937
% Final active clauses: 1945
% Final passive clauses: 44220
% Input formulas: 49
% Initial clauses: 203
% 
% Selected by SInE selection: 46
% SInE iterations: 1
% 
% Duplicate literals: 5503
% Trivial inequalities: 907
% Fw subsumption resolutions: 18271
% Fw demodulations: 1183
% Bw demodulations: 133
% 
% Simple tautologies: 6703
% Equational tautologies: 1302
% Forward subsumptions: 110726
% 
% Binary resolution: 157487
% Factoring: 268
% Forward superposition: 11626
% Backward superposition: 2444
% Self superposition: 7
% Equality resolution: 77
% 
% Split clauses: 7627
% Split components: 17529
% SAT solver clauses: 8137
% SAT solver unit clauses: 36
% SAT solver binary clauses: 2915
% 
% Sat splits: 7627
% Sat splitting refutations: 38
% 
% Memory used [KB]: 47589
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.integral
% (58172)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 494
% Passive clauses: 3654
% Generated clauses: 15604
% Final active clauses: 483
% Final passive clauses: 3052
% Input formulas: 49
% Initial clauses: 206
% 
% Duplicate literals: 391
% Trivial inequalities: 167
% Fw subsumption resolutions: 1136
% Fw demodulations: 1447
% Bw demodulations: 66
% Global subsumptions: 6
% 
% Simple tautologies: 436
% Equational tautologies: 54
% Forward subsumptions: 8215
% 
% Unit resulting resolution: 3738
% Factoring: 2
% Forward superposition: 4094
% Backward superposition: 4191
% Self superposition: 68
% Equality factoring: 6
% Equality resolution: 22
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 3731
% SAT solver unit clauses: 322
% SAT solver binary clauses: 650
% 
% Sat splits: 15
% Sat splitting refutations: 54
% 
% Memory used [KB]: 10106
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.integral
% (58199)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1473
% Passive clauses: 7318
% Generated clauses: 21274
% Final active clauses: 672
% Final passive clauses: 3543
% Input formulas: 49
% Initial clauses: 206
% 
% Duplicate literals: 837
% Trivial inequalities: 56
% Fw subsumption resolutions: 675
% Bw subsumption resolutions: 1478
% Fw demodulations: 745
% Bw demodulations: 140
% Condensations: 1601
% Global subsumptions: 79
% 
% Simple tautologies: 1265
% Equational tautologies: 35
% Forward subsumptions: 8424
% Backward subsumptions: 1545
% 
% Binary resolution: 7275
% Factoring: 70
% Forward superposition: 2582
% Backward superposition: 2661
% Self superposition: 68
% Equality factoring: 4
% Equality resolution: 10
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 7582
% SAT solver unit clauses: 372
% SAT solver binary clauses: 1379
% 
% Sat splits: 20
% Sat splitting refutations: 162
% 
% Memory used [KB]: 12281
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 120 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.integral
% (58234)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 641
% Passive clauses: 3849
% Generated clauses: 19137
% Final active clauses: 542
% Final passive clauses: 1534
% Input formulas: 49
% Initial clauses: 206
% 
% Duplicate literals: 350
% Trivial inequalities: 118
% Fw demodulations: 1790
% Bw demodulations: 10
% Global subsumptions: 172
% 
% Simple tautologies: 510
% Equational tautologies: 354
% Forward subsumptions: 10238
% Backward subsumptions: 1899
% 
% Binary resolution: 5833
% Unit resulting resolution: 4241
% Factoring: 9
% Forward superposition: 764
% Backward superposition: 3490
% Self superposition: 49
% Equality factoring: 52
% Equality resolution: 10
% 
% Split clauses: 255
% Split components: 510
% SAT solver clauses: 6127
% SAT solver unit clauses: 34
% SAT solver binary clauses: 2027
% 
% Sat splits: 255
% Sat splitting refutations: 52
% 
% Memory used [KB]: 14072
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% remaining time: 109 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.integral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integral
% SZS output start Proof for thm.integral
fof(f83952,plain,(
  $false),
  inference(equality_resolution,[],[f83618])).
fof(f83618,plain,(
  ( ! [X0] : (s('type.realax.real',sK5) != s('type.realax.real',X0)) )),
  inference(resolution,[],[f80639,f236])).
fof(f236,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integral.INTEGRABLE_POINT_SPIKE','thm.bool.TRUTH')).
fof(f80639,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s('type.realax.real',sK5) != s('type.realax.real',X2)) )),
  inference(backward_demodulation,[],[f80510,f12327])).
fof(f12327,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s('type.realax.real',X2) != s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X2,sK4,sK3,sK2,sK1))) )),
  inference(superposition,[],[f1702,f2231])).
fof(f2231,plain,(
  s(bool,'T_0') = s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))),s(fun('type.realax.real','type.realax.real'),sK1),s('type.realax.real',sK21(sK1,sK4,sK3))))),
  inference(resolution,[],[f741,f253])).
fof(f253,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.integral.INTEGRABLE_POINT_SPIKE','thm.bool.EQ_CLAUSES|split|0')).
fof(f741,plain,(
  p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))),s(fun('type.realax.real','type.realax.real'),sK1),s('type.realax.real',sK21(sK1,sK4,sK3)))))),
  inference(resolution,[],[f234,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',sK21(X2,X1,X0))))) | ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2)))) | p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',sK21(X2,X1,X0)))))) & (! [X3] : ~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X3)))) | p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21])],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2)))) | ? [X4] : p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X4))))) & (! [X3] : ~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X3)))) | p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2))))))),
  inference(rectify,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2)))) | ? [X3] : p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X3))))) & (! [X3] : ~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X3)))) | p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2))))))),
  inference(nnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2)))) <=> ? [X3] : p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X3)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X8,X9,X6] : (p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s(fun('type.realax.real','type.realax.real'),X6)))) <=> ? [X22] : p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s(fun('type.realax.real','type.realax.real'),X6),s('type.realax.real',X22)))))),
  file('thm.integral.INTEGRABLE_POINT_SPIKE','thm.integral.integrable_def')).
fof(f234,plain,(
  p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))),s(fun('type.realax.real','type.realax.real'),sK1))))),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK4)))) | s('type.realax.real',sK5) = s('type.realax.real',X5) | s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK1),s('type.realax.real',X5))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',X5)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))),s(fun('type.realax.real','type.realax.real'),sK1)))) & ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))),s(fun('type.realax.real','type.realax.real'),sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4,sK5])],[f129])).
fof(f129,plain,(
  ? [X0,X1,X2,X3,X4] : (! [X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X3)))) | s('type.realax.real',X4) = s('type.realax.real',X5) | s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0)))) & ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1)))))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ? [X0,X1,X2,X3,X4] : ((! [X5] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X3)))) | s('type.realax.real',X4) = s('type.realax.real',X5)) | s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0))))) & ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1)))))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ~! [X0,X1,X2,X3,X4] : ((! [X5] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X5)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X3)))) & s('type.realax.real',X4) != s('type.realax.real',X5)) => s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0))))) => p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1)))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ~! [X0,X1,X2,X3,X4] : ((! [X5] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X5)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X3)))) & ~s('type.realax.real',X4) = s('type.realax.real',X5)) => s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0))))) => p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9,X10] : ((! [X11] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X11)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',X9)))) & ~s('type.realax.real',X10) = s('type.realax.real',X11)) => s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X6),s('type.realax.real',X11))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X7),s('type.realax.real',X11)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s(fun('type.realax.real','type.realax.real'),X6))))) => p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s(fun('type.realax.real','type.realax.real'),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9,X10] : ((! [X11] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X11)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',X9)))) & ~s('type.realax.real',X10) = s('type.realax.real',X11)) => s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X6),s('type.realax.real',X11))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X7),s('type.realax.real',X11)))) & p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s(fun('type.realax.real','type.realax.real'),X6))))) => p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s(fun('type.realax.real','type.realax.real'),X7)))))),
  file('thm.integral.INTEGRABLE_POINT_SPIKE','thm.integral.INTEGRABLE_POINT_SPIKE')).
fof(f1702,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))),s(fun('type.realax.real','type.realax.real'),X6),s('type.realax.real',X7)))) | s('type.realax.real',X8) != s('type.realax.real',sK22(X7,X8,sK4,sK3,sK2,X6))) )),
  inference(resolution,[],[f852,f566])).
fof(f566,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))) | ~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5)))) | s('type.realax.real',X4) != s('type.realax.real',sK22(X5,X4,X3,X2,X1,X0))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK22(X5,X4,X3,X2,X1,X0))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X5,X4,X3,X2,X1,X0)),s('type.realax.real',X3)))) & s('type.realax.real',X4) != s('type.realax.real',sK22(X5,X4,X3,X2,X1,X0)) & s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',sK22(X5,X4,X3,X2,X1,X0)))) != s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',sK22(X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5)))) | p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22])],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (? [X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X3)))) & s('type.realax.real',X4) != s('type.realax.real',X6) & s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X6))) != s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X6)))) | ~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5)))) | p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((? [X6] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X3)))) & s('type.realax.real',X4) != s('type.realax.real',X6)) & s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X6))) != s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X6)))) | ~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5))))) | p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))))),
  inference(ennf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((! [X6] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X3)))) & s('type.realax.real',X4) != s('type.realax.real',X6)) => s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X6))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X6)))) & p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5))))) => p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((! [X6] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X3)))) & ~s('type.realax.real',X4) = s('type.realax.real',X6)) => s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X6))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X6)))) & p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5))))) => p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6,X7,X8,X9,X10,X22] : ((! [X11] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X11)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',X9)))) & ~s('type.realax.real',X10) = s('type.realax.real',X11)) => s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X6),s('type.realax.real',X11))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X7),s('type.realax.real',X11)))) & p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s(fun('type.realax.real','type.realax.real'),X6),s('type.realax.real',X22))))) => p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s(fun('type.realax.real','type.realax.real'),X7),s('type.realax.real',X22)))))),
  file('thm.integral.INTEGRABLE_POINT_SPIKE','thm.integral.DINT_POINT_SPIKE')).
fof(f852,plain,(
  ( ! [X0] : (~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))),s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',X0))))) )),
  inference(resolution,[],[f235,f350])).
fof(f350,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2)))) | ~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X3))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f235,plain,(
  ~p(s(bool,'const.integral.integrable_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))),s(fun('type.realax.real','type.realax.real'),sK2))))),
  inference(cnf_transformation,[],[f158])).
fof(f80510,plain,(
  ( ! [X0] : (s('type.realax.real',sK5) = s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X0,sK4,sK3,sK2,sK1))) )),
  inference(resolution,[],[f78762,f236])).
fof(f78762,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',sK5) = s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X0,sK4,sK3,sK2,sK1))) )),
  inference(global_subsumption,[],[f78761,f51192])).
fof(f51192,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK1),s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X2,sK4,sK3,sK2,sK1)))) != s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X2,sK4,sK3,sK2,sK1))))) )),
  inference(superposition,[],[f1703,f2231])).
fof(f1703,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))),s(fun('type.realax.real','type.realax.real'),X9),s('type.realax.real',X10)))) | s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK22(X10,X11,sK4,sK3,sK2,X9)))) != s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X9),s('type.realax.real',sK22(X10,X11,sK4,sK3,sK2,X9))))) )),
  inference(resolution,[],[f852,f567])).
fof(f567,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))) | ~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5)))) | s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',sK22(X5,X4,X3,X2,X1,X0)))) != s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',sK22(X5,X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f78761,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',sK5) = s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X0,sK4,sK3,sK2,sK1)) | s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK1),s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X0,sK4,sK3,sK2,sK1)))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X0,sK4,sK3,sK2,sK1))))) )),
  inference(forward_demodulation,[],[f78654,f75682])).
fof(f75682,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X18,sK4,sK3,sK2,sK1))))) )),
  inference(resolution,[],[f75356,f253])).
fof(f75356,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X0,sK4,sK3,sK2,sK1)))))) )),
  inference(resolution,[],[f12519,f236])).
fof(f12519,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X2,sK4,sK3,sK2,sK1)))))) )),
  inference(superposition,[],[f1700,f2231])).
fof(f1700,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))),s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK22(X1,X2,sK4,sK3,sK2,X0)))))) )),
  inference(resolution,[],[f852,f564])).
fof(f564,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))) | ~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK22(X5,X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f78654,plain,(
  ( ! [X0] : (s('type.realax.real',sK5) = s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X0,sK4,sK3,sK2,sK1)) | s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK1),s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X0,sK4,sK3,sK2,sK1)))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X0,sK4,sK3,sK2,sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X0,sK4,sK3,sK2,sK1)))))) )),
  inference(resolution,[],[f78305,f233])).
fof(f233,plain,(
  ( ! [X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK4)))) | s('type.realax.real',sK5) = s('type.realax.real',X5) | s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK1),s('type.realax.real',X5))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',X5))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',X5))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f78305,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X0,sK4,sK3,sK2,sK1)),s('type.realax.real',sK4))))) )),
  inference(resolution,[],[f12616,f236])).
fof(f12616,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(sK21(sK1,sK4,sK3),X2,sK4,sK3,sK2,sK1)),s('type.realax.real',sK4))))) )),
  inference(superposition,[],[f1701,f2231])).
fof(f1701,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))),s(fun('type.realax.real','type.realax.real'),X3),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X4,X5,sK4,sK3,sK2,X3)),s('type.realax.real',sK4))))) )),
  inference(resolution,[],[f852,f565])).
fof(f565,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X5)))) | ~p(s(bool,'const.transc.Dint_3'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X5)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X5,X4,X3,X2,X1,X0)),s('type.realax.real',X3))))) )),
  inference(cnf_transformation,[],[f232])).
% SZS output end Proof for thm.integral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 593
% Passive clauses: 7595
% Generated clauses: 82994
% Final active clauses: 564
% Final passive clauses: 7028
% Input formulas: 49
% Initial clauses: 335
% 
% Duplicate literals: 5492
% Trivial inequalities: 612
% Fw demodulations: 6654
% Bw demodulations: 74
% Condensations: 14360
% Global subsumptions: 79
% 
% Simple tautologies: 3938
% Equational tautologies: 47
% Forward subsumptions: 44894
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 34243
% Factoring: 41
% Forward superposition: 17812
% Backward superposition: 3938
% Self superposition: 21
% Equality factoring: 24
% Equality resolution: 62
% 
% SAT solver clauses: 7972
% SAT solver unit clauses: 248
% SAT solver binary clauses: 487
% SAT solver learnt clauses: 158
% SAT solver learnt literals: 390
% 
% TWLsolver clauses: 7674
% TWLsolver calls for satisfiability: 282457
% 
% Memory used [KB]: 37483
% Time elapsed: 3.732 s
% ------------------------------
% ------------------------------
% Success in time 22.74 s
