Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.quote
% (102719)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 192
% Passive clauses: 4210
% Generated clauses: 28930
% Final active clauses: 163
% Final passive clauses: 3788
% Input formulas: 69
% Initial clauses: 127
% 
% Duplicate literals: 963
% Trivial inequalities: 499
% Fw subsumption resolutions: 2957
% Fw demodulations: 4556
% Bw demodulations: 27
% 
% Simple tautologies: 446
% Equational tautologies: 57
% Forward subsumptions: 14404
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 477
% Factoring: 1
% Forward superposition: 8964
% Backward superposition: 9925
% Self superposition: 237
% Equality factoring: 113
% Equality resolution: 39
% 
% Split clauses: 463
% Split components: 930
% SAT solver clauses: 1069
% SAT solver unit clauses: 8
% SAT solver binary clauses: 189
% 
% Sat splits: 463
% 
% TWLsolver clauses: 1068
% TWLsolver calls for satisfiability: 212
% 
% Memory used [KB]: 24562
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quote
% (102732)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 233
% Passive clauses: 1810
% Generated clauses: 19009
% Final active clauses: 89
% Final passive clauses: 476
% Input formulas: 69
% Initial clauses: 127
% 
% Duplicate literals: 302
% Trivial inequalities: 148
% Fw subsumption resolutions: 1273
% Fw demodulations: 2555
% Bw demodulations: 1084
% Condensations: 776
% Global subsumptions: 12
% 
% Simple tautologies: 168
% Equational tautologies: 48
% Forward subsumptions: 8898
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 30
% 
% Binary resolution: 289
% Factoring: 1
% Forward superposition: 4299
% Backward superposition: 4638
% Self superposition: 217
% Equality factoring: 58
% Equality resolution: 31
% 
% Split clauses: 257
% Split components: 514
% SAT solver clauses: 2393
% SAT solver unit clauses: 161
% SAT solver binary clauses: 119
% SAT solver learnt clauses: 29
% SAT solver learnt literals: 8
% 
% Sat splits: 257
% Sat splitting refutations: 104
% 
% TWLsolver clauses: 2340
% TWLsolver calls for satisfiability: 514
% 
% Memory used [KB]: 15351
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quote
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quote
% SZS output start Proof for thm.quote
fof(f7254,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f238,f7236])).
fof(f7236,plain,(
  ( ! [X206] : (~p(s(bool,X206))) )),
  inference(forward_demodulation,[],[f7235,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X0),X1),s(fun('type.quote.index',X0),X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : s(X0,X3) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X0),X1),s(fun('type.quote.index',X0),X2),s(X0,X3)))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X3) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X0),X1),s(fun('type.quote.index',X0),X2),s(X0,X3)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X9] : ! [X10,X11,X12] : s(X9,X12) = s(X9,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X9),X10),s(fun('type.quote.index',X9),X11),s(X9,X12)))),
  file('thm.quote.index_distinct|split|2','thm.quote.index_case_def|split|2')).
fof(f7235,plain,(
  ( ! [X206,X204,X205] : (~p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',bool),X204),s(fun('type.quote.index',bool),X205),s(bool,X206))))) )),
  inference(subsumption_resolution,[],[f7139,f7105])).
fof(f7105,plain,(
  ( ! [X6,X4,X7] : (s(X4,X7) = s(X4,happ(s(fun('type.quote.index',X4),X6),s('type.quote.index',sK0)))) )),
  inference(superposition,[],[f2553,f307])).
fof(f2553,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK0))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X1),s(X0,X3)))) )),
  inference(superposition,[],[f312,f237])).
fof(f237,plain,(
  s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f170])).
fof(f170,plain,(
  ? [X0] : s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0] : s('type.quote.index','const.quote.End_idx_0') != s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ~s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X6)))),
  file('thm.quote.index_distinct|split|2','thm.quote.index_distinct|split|2')).
fof(f312,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X3),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.quote.index',X0),X3),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.quote.index',X0),X3),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9] : ! [X6,X10,X11,X12] : s(X9,happ(s(fun('type.quote.index',X9),X11),s('type.quote.index',X6))) = s(X9,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Right_idx_1'(s('type.quote.index',X6))),s(fun('type.quote.index',X9),X10),s(fun('type.quote.index',X9),X11),s(X9,X12)))),
  file('thm.quote.index_distinct|split|2','thm.quote.index_case_def|split|1')).
fof(f7139,plain,(
  ( ! [X206,X204,X205] : (s(bool,'F_0') != s(bool,happ(s(fun('type.quote.index',bool),X205),s('type.quote.index',sK0))) | ~p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',bool),X204),s(fun('type.quote.index',bool),X205),s(bool,X206))))) )),
  inference(superposition,[],[f258,f2553])).
fof(f258,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.quote.index_distinct|split|2','thm.bool.EQ_CLAUSES|split|3')).
fof(f238,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quote.index_distinct|split|2',aHL_TRUTH)).
% SZS output end Proof for thm.quote
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 141
% Passive clauses: 940
% Generated clauses: 7034
% Final active clauses: 118
% Final passive clauses: 708
% Input formulas: 69
% Initial clauses: 127
% 
% Duplicate literals: 4
% Trivial inequalities: 17
% Fw subsumption resolutions: 182
% Fw demodulations: 1258
% Bw demodulations: 111
% Condensations: 30
% 
% Simple tautologies: 46
% Equational tautologies: 69
% Forward subsumptions: 4343
% Fw demodulations to eq. taut.: 48
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 1193
% Factoring: 1
% Forward superposition: 1726
% Backward superposition: 2431
% Self superposition: 99
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 1012
% SAT solver unit clauses: 345
% SAT solver binary clauses: 530
% 
% Sat splits: 10
% Sat splitting refutations: 50
% 
% Memory used [KB]: 16119
% Time elapsed: 0.303 s
% ------------------------------
% ------------------------------
% Success in time 1.631 s
