Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.rich_list
% (91391)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1323
% Passive clauses: 10621
% Generated clauses: 44116
% Final active clauses: 1255
% Final passive clauses: 8802
% Input formulas: 515
% Initial clauses: 917
% 
% Duplicate literals: 230
% Trivial inequalities: 59
% Fw subsumption resolutions: 2122
% Fw demodulations: 5206
% Bw demodulations: 54
% 
% Simple tautologies: 401
% Equational tautologies: 24
% Forward subsumptions: 25168
% 
% Binary resolution: 8537
% Factoring: 725
% Forward superposition: 12711
% Backward superposition: 13260
% Self superposition: 59
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 413
% Split components: 828
% SAT solver clauses: 636
% SAT solver unit clauses: 29
% SAT solver binary clauses: 101
% 
% Sat splits: 413
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 636
% TWLsolver calls for satisfiability: 418
% 
% Memory used [KB]: 23922
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rich_list
% (91411)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 671
% Passive clauses: 5084
% Generated clauses: 31488
% Final active clauses: 648
% Final passive clauses: 4263
% Input formulas: 515
% Initial clauses: 917
% 
% Duplicate literals: 111
% Trivial inequalities: 41
% Fw subsumption resolutions: 1053
% Fw demodulations: 4433
% Bw demodulations: 169
% Condensations: 99
% Global subsumptions: 37
% 
% Simple tautologies: 325
% Equational tautologies: 16
% Forward subsumptions: 14325
% Fw demodulations to eq. taut.: 14
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 3062
% Factoring: 151
% Forward superposition: 8652
% Backward superposition: 7472
% Self superposition: 84
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 6915
% SAT solver unit clauses: 1231
% SAT solver binary clauses: 257
% SAT solver learnt clauses: 74
% SAT solver learnt literals: 14
% 
% Sat splits: 23
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 6767
% TWLsolver calls for satisfiability: 5581
% 
% Memory used [KB]: 35052
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rich_list
% (91430)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1631
% Passive clauses: 17247
% Generated clauses: 278237
% Final active clauses: 1499
% Final passive clauses: 3887
% Input formulas: 515
% Initial clauses: 917
% Discarded non-redundant clauses: 202626
% 
% Duplicate literals: 10
% Trivial inequalities: 107
% Fw subsumption resolutions: 2426
% Fw demodulations: 25136
% Bw demodulations: 3112
% Condensations: 79
% 
% Simple tautologies: 173
% Equational tautologies: 96
% Forward subsumptions: 66526
% Fw demodulations to eq. taut.: 3131
% Bw demodulations to eq. taut.: 37
% 
% Unit resulting resolution: 196483
% Factoring: 1
% Forward superposition: 27871
% Backward superposition: 19118
% Self superposition: 151
% Equality factoring: 2
% Equality resolution: 37
% 
% Split clauses: 43
% Split components: 86
% SAT solver clauses: 19841
% SAT solver unit clauses: 8676
% SAT solver binary clauses: 9216
% 
% Sat splits: 43
% Sat splitting refutations: 2506
% 
% Memory used [KB]: 184858
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.rich_list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rich_list
% SZS output start Proof for thm.rich_list
fof(f33927,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1262,f1419,f1418,f1264,f1266,f1267,f1268,f1271,f1384,f1420,f1435,f1437,f1438,f1439,f1442,f1555,f1589,f1730,f1889,f1890,f2010,f2012,f2013,f2015,f2017,f2130,f2162,f2290,f2289,f2288,f2297,f2495,f2443,f2603,f2604,f2692,f2973,f2972,f2970,f3139,f3142,f3144,f3145,f3146,f3149,f3262,f3296,f3483,f3485,f3486,f3487,f3490,f3603,f3636,f3655,f3657,f3658,f3659,f3662,f3775,f3808,f4869,f4871,f4872,f4874,f4876,f4989,f5110,f5129,f5131,f5132,f5134,f5136,f5249,f5279,f5368,f5370,f5371,f5373,f5375,f5488,f5609,f5493,f5495,f5621,f5623,f5624,f5626,f5628,f5741,f5862,f5746,f5748,f7279,f7278,f7277,f7276,f7275,f7318,f7283,f7317,f7316,f7287,f7315,f7314,f7327,f7329,f7330,f7332,f7334,f7447,f7568,f7451,f7453,f7574,f7589,f7591,f7592,f7594,f7596,f7709,f7831,f7713,f7715,f7837,f7916,f7918,f7919,f7940,f7942,f7943,f7969,f7968,f7967,f7966,f7965,f8011,f7973,f8010,f8009,f7977,f8008,f8007,f8022,f8024,f8025,f8027,f8029,f8142,f8264,f8146,f8148,f8270,f8274,f8276,f8277,f8279,f8281,f8394,f8515,f8398,f8400,f8521,f5615,f8595,f8597,f8598,f8619,f8621,f8622,f8666,f8668,f8669,f8650,f8652,f8678,f8677,f8702,f8704,f8706,f8710,f8713,f8716,f8719,f8722,f8774,f8775,f8776,f8794,f8840,f8842,f8843,f8904,f8905,f8906,f8923,f8940,f8956,f9012,f9014,f9015,f9076,f9077,f9078,f9138,f9140,f9141,f9201,f9203,f9204,f9266,f9267,f9268,f9339,f9340,f9341,f9401,f9402,f9403,f9421,f9465,f9466,f9467,f9528,f9529,f9530,f9547,f9593,f9594,f9595,f9656,f9657,f9658,f9718,f9719,f9720,f9780,f9781,f9782,f9844,f9845,f9846,f9906,f9907,f9908,f9966,f9967,f9968,f9984,f10029,f10030,f10031,f10103,f10104,f10105,f10120,f10166,f10167,f10168,f10227,f10228,f10229,f10287,f10288,f10289,f10347,f10348,f10349,f10421,f10422,f10423,f10481,f10482,f10483,f10541,f10542,f10543,f10559,f10603,f10604,f10605,f10664,f10665,f10666,f10681,f10692,f10694,f10696,f10698,f10886,f10884,f10882,f10844,f10843,f10841,f11058,f11059,f11060,f11119,f11120,f11121,f11179,f11180,f11181,f11239,f11240,f11241,f11265,f11267,f11269,f11271,f11327,f11328,f11329,f11353,f11355,f11357,f11359,f11378,f11380,f11382,f11384,f11405,f11407,f11409,f11411,f11450,f11452,f11454,f11456,f11572,f11577,f11650,f11646,f11682,f11683,f11684,f11685,f11713,f11718,f11741,f11737,f11836,f11841,f11876,f11878,f11880,f11882,f11910,f11906,f11928,f11930,f11931,f11939,f11957,f10707,f11959,f10708,f11961,f10709,f11963,f10710,f11985,f12005,f12007,f12008,f12045,f12047,f12048,f12059,f12079,f12081,f7289,f12082,f7291,f12101,f12147,f12148,f12183,f12188,f12240,f12236,f12322,f12327,f12350,f12346,f12508,f12513,f12568,f12564,f12604,f12606,f12607,f12615,f12635,f12655,f12657,f12658,f12686,f12688,f12689,f12700,f12720,f12722,f7979,f12723,f7981,f12742,f12780,f12781,f12919,f12921,f12922,f12924,f12926,f13039,f13160,f13203,f13205,f13206,f13207,f13210,f13323,f13479,f13336,f13337,f13482,f13338,f13339,f13484,f13340,f13341,f13342,f13343,f13344,f13345,f13347,f13461,f13460,f13459,f13483,f13486,f13490,f13491,f13492,f13497,f13499,f13500,f13502,f13504,f13617,f13739,f13753,f13761,f13770,f13779,f13788,f13794,f13800,f13810,f13496,f13847,f13851,f13852,f13862,f13887,f13891,f13892,f13902,f13919,f13921,f13922,f13924,f13926,f14039,f14161,f14170,f14172,f14173,f14175,f14177,f14290,f14412,f14523,f14522,f14521,f14520,f14471,f14472,f14473,f14474,f14475,f14476,f14477,f14478,f14479,f14480,f14481,f14482,f14483,f14484,f14485,f14486,f14487,f5869,f14621,f14624,f14627,f14630,f14632,f14634,f14801,f14805,f14806,f14812,f14816,f14817,f14820,f14824,f14825,f14830,f14834,f14835,f14868,f14873,f14874,f14877,f14882,f14883,f14886,f14891,f14894,f14899,f14900,f14903,f14908,f14807,f14909,f14811,f14941,f14943,f14944,f14946,f14948,f15061,f15183,f15414,f15413,f15412,f15411,f15360,f15361,f15362,f15363,f15364,f15365,f15366,f15367,f15368,f15369,f15370,f15371,f15372,f15373,f15374,f15375,f15376,f15606,f15605,f15604,f15603,f15583,f15584,f15585,f15586,f15587,f15588,f15589,f15590,f15591,f15592,f15593,f15594,f15595,f15596,f15597,f15598,f15599,f15670,f15672,f15673,f15675,f15677,f15790,f15920,f15929,f15931,f15932,f15934,f15936,f16049,f16171,f16185,f16187,f16189,f16191,f16193,f16195,f16200,f16202,f16241,f16243,f16245,f16247,f16249,f16251,f16254,f16255,f16256,f16258,f16327,f16328,f15439,f16329,f16330,f16320,f16321,f16359,f16361,f16362,f16384,f16386,f16387,f16395,f16397,f16399,f16401,f16403,f16405,f16407,f16409,f16531,f16530,f16529,f16528,f16613,f16615,f16616,f16618,f16620,f16733,f16855,f16864,f16866,f16867,f16869,f16871,f16984,f17114,f17127,f17129,f17131,f17133,f17135,f17137,f17139,f17141,f17180,f17182,f17184,f17186,f17188,f17190,f17192,f17194,f17262,f17263,f14548,f17264,f17265,f17256,f17257,f17290,f17292,f17294,f17296,f17298,f17300,f17302,f17304,f17327,f17329,f17330,f17352,f17354,f17355,f17405,f17406,f17386,f17388,f17390,f17391,f17392,f17393,f17394,f17403,f17404,f17439,f17438,f17494,f17495,f17475,f17477,f17479,f17480,f17481,f17482,f17483,f17490,f17491,f17547,f17548,f17528,f17530,f17532,f17533,f17534,f17535,f17536,f17545,f17546,f17602,f17603,f17582,f17584,f17586,f17587,f17588,f17589,f17590,f17598,f17599,f17655,f17656,f17635,f17637,f17639,f17640,f17641,f17642,f17643,f17653,f17654,f17704,f17705,f17687,f17689,f17691,f17692,f17693,f17694,f17695,f17703,f17734,f17733,f17789,f17790,f17771,f17773,f17775,f17776,f17777,f17778,f17779,f17787,f17841,f17842,f17823,f17825,f17827,f17828,f17829,f17830,f17831,f17840,f13865,f17893,f17894,f17874,f17875,f17877,f17879,f17880,f17407,f17881,f17882,f17409,f17891,f17892,f17945,f17946,f17926,f17927,f17929,f17931,f17932,f17496,f17933,f17934,f17497,f17941,f17942,f17996,f17997,f17977,f17978,f17980,f17982,f17983,f17549,f17984,f17985,f17550,f17994,f17995,f18078,f18079,f18058,f18059,f18061,f18063,f18064,f17604,f18065,f18066,f17605,f18074,f18075,f18129,f18130,f18109,f18110,f18112,f18114,f18115,f17657,f18116,f18117,f17658,f18127,f18128,f18176,f18177,f18159,f18160,f18162,f18164,f18165,f17706,f18166,f18167,f17707,f18175,f18224,f18225,f18206,f18207,f18209,f18211,f18212,f17791,f18213,f18214,f17792,f18222,f18274,f18275,f18256,f18257,f18259,f18261,f18262,f17843,f18263,f18264,f17844,f18273,f13905,f18281,f18282,f18283,f18284,f18285,f18286,f18287,f18288,f18289,f18290,f18308,f18312,f18314,f18363,f18364,f18365,f18368,f18353,f18354,f18374,f18378,f18380,f18427,f18428,f18429,f18417,f18418,f18458,f18462,f18464,f18511,f18512,f18513,f18515,f18503,f18504,f18522,f18526,f18528,f18576,f18577,f18578,f18566,f18567,f18585,f18589,f18591,f18639,f18640,f18641,f18631,f18632,f18647,f18651,f18653,f18699,f18700,f18701,f18691,f18692,f18707,f18711,f18713,f18759,f18760,f18761,f18751,f18769,f18771,f18773,f18775,f18777,f18779,f18781,f18783,f18785,f18787,f18806,f18810,f18812,f18859,f18860,f18861,f18851,f16332,f18852,f18868,f18872,f18874,f18921,f18922,f18923,f18911,f18912,f18932,f18936,f18938,f18985,f18986,f18987,f18990,f18977,f18978,f18996,f19000,f19002,f19049,f19050,f19051,f19039,f19040,f19059,f19063,f19065,f19112,f19113,f19114,f19116,f19104,f19105,f19122,f19124,f19126,f16212,f19128,f16213,f19130,f16214,f19132,f16215,f19134,f19136,f16269,f19138,f16270,f19140,f19142,f16333,f19161,f19165,f19167,f19215,f19216,f19217,f19205,f19206,f19224,f19228,f19230,f19278,f19279,f19280,f19270,f19271,f19286,f19290,f19292,f19338,f19339,f19340,f19330,f19331,f19346,f19350,f19352,f19400,f19401,f19402,f19390,f19391,f19410,f19414,f19416,f19464,f15437,f19465,f15441,f19466,f15443,f19456,f18925,f19457,f18926,f19491,f19489,f19523,f19521,f19591,f19589,f19643,f19641,f19720,f19724,f19726,f14553,f19774,f19775,f19776,f19779,f19767,f19768,f19784,f19788,f19790,f14554,f19838,f19839,f19840,f19829,f19830,f19847,f19851,f19853,f14555,f19901,f19902,f19903,f19905,f19894,f19895,f19911,f19915,f19917,f14556,f19966,f19967,f19968,f19957,f19958,f19974,f19978,f19980,f14557,f20029,f20030,f20031,f20022,f20023,f20036,f20040,f20042,f14558,f20089,f20090,f20091,f20082,f20083,f20096,f20100,f20102,f14559,f20151,f20152,f20153,f20142,f20143,f20160,f20164,f20166,f14560,f20215,f20216,f20217,f20208,f17267,f20209,f17268,f20223,f20227,f20229,f14561,f20277,f20278,f20279,f20268,f20269,f20287,f20291,f20293,f20341,f20342,f20343,f8667,f20346,f20334,f20335,f20351,f20355,f20357,f20405,f20406,f20407,f20396,f20397,f20414,f20418,f20420,f20468,f20469,f20470,f8841,f20472,f20461,f20462,f20478,f20482,f20484,f20533,f20534,f20535,f20524,f20525,f20541,f20545,f20547,f20596,f20597,f20598,f9013,f20589,f20590,f20603,f20607,f20609,f20656,f20657,f20658,f9139,f20649,f20650,f20697,f20701,f20703,f20752,f20753,f20754,f20743,f20744,f20761,f5617,f20765,f20767,f18362,f20816,f14546,f20817,f14550,f20818,f14552,f9202,f20809,f20281,f20810,f20282,f20911,f20913,f20914,f20916,f20944,f20945,f20946,f20947,f20985,f20987,f21021,f21023,f21058,f7285,f21060,f7975,f21103,f21105,f21106,f21108,f21110,f21223,f21376,f21359,f21402,f21360,f21403,f21361,f21404,f21362,f21405,f21363,f21406,f21364,f21407,f21444,f21446,f21447,f21449,f21451,f21564,f21717,f21700,f21743,f21701,f21744,f21702,f21745,f21703,f21746,f21704,f21747,f21705,f21748,f22072,f21783,f21785,f21786,f21787,f21790,f21903,f22073,f22112,f22114,f22115,f22117,f22119,f22232,f22386,f22369,f22412,f22370,f22413,f22371,f22414,f22372,f22415,f22373,f22416,f22374,f22417,f22452,f22454,f22455,f22456,f22459,f22572,f22742,f22781,f22783,f22784,f22786,f22788,f22901,f23056,f23039,f23082,f23040,f23083,f23041,f23084,f23042,f23085,f23043,f23086,f23044,f23087,f23412,f23122,f23124,f23125,f23126,f23129,f23242,f23414,f23439,f23441,f23442,f23444,f23446,f23559,f23655,f23716,f23718,f23719,f23720,f23723,f23836,f24007,f24032,f24034,f24035,f24037,f24039,f24152,f24249,f24310,f24312,f24313,f24314,f24317,f24430,f24601,f24626,f24628,f24629,f24631,f24633,f24746,f24844,f24903,f24905,f24906,f24907,f24910,f25023,f25195,f25220,f25222,f25223,f25225,f25227,f25340,f25439,f25500,f25502,f25503,f25504,f25507,f25620,f25795,f26033,f26035,f26036,f26037,f26040,f26153,f26328,f26365,f26367,f26368,f26369,f26372,f26485,f26659,f26689,f26693,f26694,f26695,f26696,f26710,f26741,f26745,f26746,f26747,f26748,f26762,f26803,f26805,f26806,f26808,f26810,f26923,f27086,f27069,f27112,f27070,f27113,f27071,f27114,f27072,f27115,f27073,f27116,f27074,f27117,f27118,f27120,f27121,f27123,f27125,f27238,f27369,f27378,f27380,f27381,f27383,f27385,f27498,f27629,f27685,f27686,f27666,f27668,f27670,f27672,f27674,f27683,f27684,f27726,f27728,f27729,f27731,f27733,f27846,f28009,f27992,f28035,f27993,f28036,f27994,f28037,f27995,f28038,f27996,f28039,f27997,f28040,f28088,f28089,f28069,f28071,f28073,f28075,f28077,f28084,f28085,f28139,f28140,f28120,f28122,f28124,f28126,f28128,f28137,f28138,f28192,f28193,f28172,f28174,f28176,f28178,f28180,f28188,f28189,f28243,f28244,f28223,f28225,f28227,f28229,f28231,f28241,f28242,f28291,f28292,f28273,f28275,f28277,f28279,f28281,f28289,f28290,f28329,f28331,f28332,f28334,f28336,f28449,f28612,f28595,f28638,f28596,f28639,f28597,f28640,f28598,f28641,f28599,f28642,f28600,f28643,f28692,f28693,f28672,f28674,f28676,f28678,f28680,f28688,f28689,f28744,f28745,f28724,f28726,f28728,f28730,f28732,f28742,f26713,f28743,f26714,f28794,f28795,f28775,f28776,f28778,f28780,f28782,f14562,f28792,f28793,f28846,f28847,f28827,f28828,f28830,f28832,f28834,f14563,f28842,f28843,f28897,f28898,f28878,f28879,f28881,f28883,f28885,f14564,f28895,f28896,f28938,f28940,f28941,f28943,f28945,f29058,f29221,f29204,f29247,f29205,f29248,f29206,f29249,f29207,f29250,f29208,f29251,f29209,f29252,f29301,f29302,f29281,f29282,f29284,f29286,f29288,f14565,f29297,f29298,f29352,f29353,f29332,f29333,f29335,f29337,f29339,f14566,f29350,f29351,f29400,f29401,f29382,f29383,f29385,f29387,f29389,f14567,f29398,f29399,f29450,f29451,f29430,f29431,f29433,f29435,f29437,f14568,f29446,f29447,f29502,f29503,f29482,f29483,f8796,f29485,f9985,f29487,f17408,f29489,f14569,f17410,f29500,f26765,f29501,f26766,f29527,f29529,f29530,f29532,f29534,f29647,f29749,f30111,f8795,f29809,f29811,f29812,f29813,f29816,f29929,f30112,f30137,f30139,f30140,f30142,f30144,f30257,f30360,f30420,f30422,f30423,f30424,f30427,f30540,f30723,f30748,f30750,f30751,f30753,f30755,f30868,f30972,f31335,f8924,f31032,f31034,f31035,f31036,f31039,f31152,f31337,f31362,f31364,f31365,f31367,f31369,f31482,f31587,f31648,f31650,f31651,f31652,f31655,f31768,f31952,f31996,f31998,f31999,f32000,f32003,f32116,f32299,f32395,f32396,f32397,f32398,f32399,f32400,f32499,f32500,f32501,f32614,f32502,f32613,f32503,f32616,f32611,f32504,f32617,f32615,f32612,f32610,f32657,f32659,f32660,f32661,f32664,f32777,f32960,f32999,f33001,f33002,f33003,f33006,f33119,f33304,f33373,f33375,f33376,f1421,f33377,f1423,f33380,f1422,f33493,f33678,f1424,f33888])).
fof(f33888,plain,(
  $false | $spl2),
  inference(trivial_inequality_removal,[],[f33779])).
fof(f33779,plain,(
  s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK17),s('type.list.list'(sK14),sK18))) != s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK17),s('type.list.list'(sK14),sK18))) | $spl2),
  inference(superposition,[],[f918,f33682])).
fof(f33682,plain,(
  ( ! [X0,X1] : (s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X0),s('type.list.list'(sK14),X1))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X0))),s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16),s('type.list.list'(sK14),X1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f33681,f1419])).
fof(f33681,plain,(
  ( ! [X0,X1] : (s(sK14,'const.list.FOLDL_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X0),s('type.list.list'(sK14),X1))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X0))),s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16),s('type.list.list'(sK14),X1)))))) ) | $spl2),
  inference(forward_demodulation,[],[f1588,f1419])).
fof(f1588,plain,(
  ( ! [X0,X1] : (s(sK14,'const.list.FOLDL_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X0),s('type.list.list'(sK14),X1))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X0))),s(sK14,'const.list.FOLDL_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16),s('type.list.list'(sK14),X1)))))) )),
  inference(subsumption_resolution,[],[f1430,f916])).
fof(f916,plain,(
  p(s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))))),
  inference(cnf_transformation,[],[f746])).
fof(f746,plain,(
  p(s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) & (p(s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) & s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK17),s('type.list.list'(sK14),sK18))) != s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK17))),s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16),s('type.list.list'(sK14),sK18))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15,sK16,sK17,sK18])],[f669])).
fof(f669,plain,(
  ? [X0,X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) & ? [X2] : (p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) & ? [X3,X4] : s(X0,'const.list.FOLDR_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X4))) != s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,'const.list.FOLDR_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2),s('type.list.list'(X0),X4)))))))),
  inference(ennf_transformation,[],[f518])).
fof(f518,plain,(
  ~! [X0,X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) => ! [X2] : (p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) => ! [X3,X4] : s(X0,'const.list.FOLDR_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X4))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,'const.list.FOLDR_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2),s('type.list.list'(X0),X4)))))))),
  inference(flattening,[],[f517])).
fof(f517,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) => ! [X2] : (p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) => ! [X3,X4] : s(X0,'const.list.FOLDR_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X4))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,'const.list.FOLDR_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2),s('type.list.list'(X0),X4)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.operator.COMM_1'(s(fun(X6,fun(X6,X6)),X7)))) => ! [X8] : (p(s(bool,'const.operator.MONOID_2'(s(fun(X6,fun(X6,X6)),X7),s(X6,X8)))) => ! [X9,X10] : s(X6,'const.list.FOLDR_3'(s(fun(X6,fun(X6,X6)),X7),s(X6,X9),s('type.list.list'(X6),X10))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X9))),s(X6,'const.list.FOLDR_3'(s(fun(X6,fun(X6,X6)),X7),s(X6,X8),s('type.list.list'(X6),X10)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.operator.COMM_1'(s(fun(X6,fun(X6,X6)),X7)))) => ! [X8] : (p(s(bool,'const.operator.MONOID_2'(s(fun(X6,fun(X6,X6)),X7),s(X6,X8)))) => ! [X9,X10] : s(X6,'const.list.FOLDR_3'(s(fun(X6,fun(X6,X6)),X7),s(X6,X9),s('type.list.list'(X6),X10))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X9))),s(X6,'const.list.FOLDR_3'(s(fun(X6,fun(X6,X6)),X7),s(X6,X8),s('type.list.list'(X6),X10)))))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.rich_list.COMM_MONOID_FOLDR')).
fof(f1430,plain,(
  ( ! [X0,X1] : (s(sK14,'const.list.FOLDL_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X0),s('type.list.list'(sK14),X1))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X0))),s(sK14,'const.list.FOLDL_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16),s('type.list.list'(sK14),X1))))) | ~p(s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))))) )),
  inference(resolution,[],[f917,f1050])).
fof(f1050,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X4))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2),s('type.list.list'(X0),X4))))) | ~p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) | ~p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1))))) )),
  inference(cnf_transformation,[],[f690])).
fof(f690,plain,(
  ! [X0,X1] : (~p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) | ! [X2] : (~p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) | ! [X3,X4] : s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X4))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2),s('type.list.list'(X0),X4)))))))),
  inference(ennf_transformation,[],[f600])).
fof(f600,plain,(
  ! [X0,X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) => ! [X2] : (p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) => ! [X3,X4] : s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X4))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2),s('type.list.list'(X0),X4)))))))),
  inference(flattening,[],[f599])).
fof(f599,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) => ! [X2] : (p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) => ! [X3,X4] : s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X4))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2),s('type.list.list'(X0),X4)))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.operator.COMM_1'(s(fun(X6,fun(X6,X6)),X7)))) => ! [X8] : (p(s(bool,'const.operator.MONOID_2'(s(fun(X6,fun(X6,X6)),X7),s(X6,X8)))) => ! [X9,X10] : s(X6,'const.list.FOLDL_3'(s(fun(X6,fun(X6,X6)),X7),s(X6,X9),s('type.list.list'(X6),X10))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X9))),s(X6,'const.list.FOLDL_3'(s(fun(X6,fun(X6,X6)),X7),s(X6,X8),s('type.list.list'(X6),X10)))))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.rich_list.COMM_MONOID_FOLDL')).
fof(f917,plain,(
  p(s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))))),
  inference(cnf_transformation,[],[f746])).
fof(f918,plain,(
  s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK17),s('type.list.list'(sK14),sK18))) != s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK17))),s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16),s('type.list.list'(sK14),sK18)))))),
  inference(cnf_transformation,[],[f746])).
fof(f1424,plain,(
  ( ! [X19] : (p(s(bool,X19))) ) | $spl12),
  inference(cnf_transformation,[],[f1424_D])).
fof(f1424_D,plain,(
  ( ! [X19] : (p(s(bool,X19))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f33678,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP8(X1,X2,X3)) ) | $spl450),
  inference(cnf_transformation,[],[f33678_D])).
fof(f33678_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP8(X1,X2,X3)) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f33493,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP8(X705,X706,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X704) != s(bool,X709)) )),
  inference(resolution,[],[f9821,f1252])).
fof(f1252,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f923])).
fof(f923,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f748])).
fof(f748,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f747])).
fof(f747,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f671])).
fof(f671,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f670])).
fof(f670,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X12,X27,X13,X28] : ((s(bool,X12) = s(bool,X27) & (p(s(bool,X27)) => s(bool,X13) = s(bool,X28))) => ((p(s(bool,X12)) => p(s(bool,X13))) <=> (p(s(bool,X27)) => p(s(bool,X28)))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.IMP_CONG')).
fof(f9821,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X1,X2,X3) | p(s(bool,X0))) )),
  inference(resolution,[],[f7664,f5228])).
fof(f5228,plain,(
  ( ! [X202] : (~sP12(X202,sK62) | p(s(bool,X202))) )),
  inference(resolution,[],[f5116,f1165])).
fof(f1165,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f875])).
fof(f875,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP12(X0,X1)))),
  inference(rectify,[],[f874])).
fof(f874,plain,(
  ! [X1,X0] : ((~sP12(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP12(X1,X0)))),
  inference(flattening,[],[f873])).
fof(f873,plain,(
  ! [X1,X0] : ((~sP12(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP12(X1,X0)))),
  inference(nnf_transformation,[],[f742])).
fof(f742,plain,(
  ! [X1,X0] : (sP12(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f5116,plain,(
  ~p(s(bool,sK62))),
  inference(subsumption_resolution,[],[f4992,f1381])).
fof(f1381,plain,(
  s(bool,'F_0') != s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))),
  inference(resolution,[],[f916,f1232])).
fof(f1232,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f908])).
fof(f908,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f649])).
fof(f649,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f648])).
fof(f648,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.EQ_CLAUSES|split|3')).
fof(f4992,plain,(
  s(bool,'F_0') = s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | ~p(s(bool,sK62))),
  inference(resolution,[],[f1288,f1249])).
fof(f1249,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,sK62))),
  inference(cnf_transformation,[],[f912])).
fof(f912,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (~p(s(bool,sK62)) | p(s(bool,'F_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK62])],[f911])).
fof(f911,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(rectify,[],[f910])).
fof(f910,plain,(
  (~p(s(bool,'F_0')) | ! [X0] : p(s(bool,X0))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f666])).
fof(f666,plain,(
  p(s(bool,'F_0')) <=> ! [X0] : p(s(bool,X0))),
  inference(rectify,[],[f315])).
fof(f315,axiom,(
  p(s(bool,'F_0')) <=> ! [X14] : p(s(bool,X14))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.F_DEF')).
fof(f1288,plain,(
  ( ! [X61] : (~p(s(bool,X61)) | s(bool,X61) = s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) )),
  inference(resolution,[],[f916,f965])).
fof(f965,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f770])).
fof(f770,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f769])).
fof(f769,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1))))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f533])).
fof(f533,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(flattening,[],[f532])).
fof(f532,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(rectify,[],[f270])).
fof(f270,axiom,(
  ! [X24,X25] : (s(bool,X24) = s(bool,X25) <=> ((p(s(bool,X24)) & p(s(bool,X25))) | (~p(s(bool,X24)) & ~p(s(bool,X25)))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.EQ_EXPAND')).
fof(f7664,plain,(
  ( ! [X304,X306,X305,X307] : (sP12(X304,X305) | sP12(X304,sK62) | sP8(X305,X306,X307)) )),
  inference(resolution,[],[f5470,f1112])).
fof(f1112,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f850])).
fof(f850,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(flattening,[],[f849])).
fof(f849,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(nnf_transformation,[],[f735])).
fof(f735,plain,(
  ! [X0,X1,X2] : (sP8(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f5470,plain,(
  ( ! [X308,X307] : (~p(s(bool,X308)) | sP12(X307,X308) | sP12(X307,sK62)) )),
  inference(resolution,[],[f5230,f1170])).
fof(f1170,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f875])).
fof(f5230,plain,(
  ( ! [X204] : (~p(s(bool,X204)) | sP12(X204,sK62)) )),
  inference(resolution,[],[f5116,f1167])).
fof(f1167,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f875])).
fof(f1422,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | ~p(s(bool,X14))) ) | $spl8),
  inference(cnf_transformation,[],[f1422_D])).
fof(f1422_D,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | ~p(s(bool,X14))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f33380,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP8(X50,X51,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X49) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9821,f926])).
fof(f926,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f748])).
fof(f1423,plain,(
  ( ! [X20] : (~p(s(bool,X20))) ) | $spl10),
  inference(cnf_transformation,[],[f1423_D])).
fof(f1423_D,plain,(
  ( ! [X20] : (~p(s(bool,X20))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f33377,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP8(X29,X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X28) != s(bool,X34)) )),
  inference(resolution,[],[f9821,f925])).
fof(f925,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f748])).
fof(f1421,plain,(
  ( ! [X7,X9] : (s(bool,X7) != s(bool,X9) | p(s(bool,X7))) ) | $spl6),
  inference(cnf_transformation,[],[f1421_D])).
fof(f1421_D,plain,(
  ( ! [X7,X9] : (s(bool,X7) != s(bool,X9) | p(s(bool,X7))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f33376,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP8(X22,X23,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X21) != s(bool,X26)) )),
  inference(resolution,[],[f9821,f924])).
fof(f924,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f748])).
fof(f33375,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP8(X15,X16,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X14) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9821,f922])).
fof(f922,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f748])).
fof(f33373,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP8(X1,X2,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X0) != s(bool,X5)) )),
  inference(resolution,[],[f9821,f920])).
fof(f920,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f748])).
fof(f33304,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP8(X2,X1,X3)) ) | $spl448),
  inference(cnf_transformation,[],[f33304_D])).
fof(f33304_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP8(X2,X1,X3)) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f33119,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP8(X706,X705,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X704) != s(bool,X709)) )),
  inference(resolution,[],[f9757,f1252])).
fof(f9757,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X2,X1,X3) | p(s(bool,X0))) )),
  inference(resolution,[],[f7663,f5228])).
fof(f7663,plain,(
  ( ! [X300,X302,X301,X303] : (sP12(X300,X301) | sP12(X300,sK62) | sP8(X302,X301,X303)) )),
  inference(resolution,[],[f5470,f1111])).
fof(f1111,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f850])).
fof(f33006,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP8(X51,X50,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X49) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9757,f926])).
fof(f33003,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP8(X30,X29,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X28) != s(bool,X34)) )),
  inference(resolution,[],[f9757,f925])).
fof(f33002,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP8(X23,X22,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X21) != s(bool,X26)) )),
  inference(resolution,[],[f9757,f924])).
fof(f33001,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP8(X16,X15,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X14) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9757,f922])).
fof(f32999,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP8(X2,X1,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X0) != s(bool,X5)) )),
  inference(resolution,[],[f9757,f920])).
fof(f32960,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP6(X1,X2,X3)) ) | $spl446),
  inference(cnf_transformation,[],[f32960_D])).
fof(f32960_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP6(X1,X2,X3)) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f32777,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP6(X705,X706,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X704) != s(bool,X709)) )),
  inference(resolution,[],[f9695,f1252])).
fof(f9695,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP6(X1,X2,X3) | p(s(bool,X0))) )),
  inference(resolution,[],[f7650,f5228])).
fof(f7650,plain,(
  ( ! [X251,X249,X250,X248] : (sP12(X248,X249) | sP12(X248,sK62) | sP6(X249,X250,X251)) )),
  inference(resolution,[],[f5470,f1095])).
fof(f1095,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f843])).
fof(f843,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f842])).
fof(f842,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f841])).
fof(f841,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f732])).
fof(f732,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f32664,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP6(X50,X51,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X49) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9695,f926])).
fof(f32661,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP6(X29,X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X28) != s(bool,X34)) )),
  inference(resolution,[],[f9695,f925])).
fof(f32660,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP6(X22,X23,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X21) != s(bool,X26)) )),
  inference(resolution,[],[f9695,f924])).
fof(f32659,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP6(X15,X16,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X14) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9695,f922])).
fof(f32657,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP6(X1,X2,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X0) != s(bool,X5)) )),
  inference(resolution,[],[f9695,f920])).
fof(f32610,plain,(
  s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) != s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14)))) | $spl437),
  inference(cnf_transformation,[],[f32610_D])).
fof(f32610_D,plain,(
  s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) != s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14)))) <=> ~$spl437),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl437])])).
fof(f32612,plain,(
  s(sK14,sK19(sK15,sK14,sK14,sK14)) != s(sK14,sK21(sK15,sK14,sK14,sK14)) | $spl439),
  inference(cnf_transformation,[],[f32612_D])).
fof(f32612_D,plain,(
  s(sK14,sK19(sK15,sK14,sK14,sK14)) != s(sK14,sK21(sK15,sK14,sK14,sK14)) <=> ~$spl439),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl439])])).
fof(f32615,plain,(
  s(sK14,sK20(sK15,sK14,sK14,sK14)) != s(sK14,sK22(sK15,sK14,sK14,sK14)) | $spl443),
  inference(cnf_transformation,[],[f32615_D])).
fof(f32615_D,plain,(
  s(sK14,sK20(sK15,sK14,sK14,sK14)) != s(sK14,sK22(sK15,sK14,sK14,sK14)) <=> ~$spl443),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl443])])).
fof(f32617,plain,(
  ( ! [X10,X9] : (s(sK14,X10) = s(sK14,happ(s(fun(sK14,sK14),sK24(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X10))),s(sK14,X9)))))) ) | $spl444),
  inference(cnf_transformation,[],[f32617_D])).
fof(f32617_D,plain,(
  ( ! [X10,X9] : (s(sK14,X10) = s(sK14,happ(s(fun(sK14,sK14),sK24(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X10))),s(sK14,X9)))))) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f32504,plain,(
  ( ! [X14,X13] : (s(sK14,X14) = s(sK14,happ(s(fun(sK14,sK14),sK24(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X14))),s(sK14,X13))))) | s(sK14,sK20(sK15,sK14,sK14,sK14)) != s(sK14,sK22(sK15,sK14,sK14,sK14)) | s(sK14,sK19(sK15,sK14,sK14,sK14)) != s(sK14,sK21(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) != s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f932,f1263])).
fof(f1263,plain,(
  ( ! [X6,X5] : (s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X5))),s(sK14,X6))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X6))),s(sK14,X5)))) )),
  inference(resolution,[],[f916,f1250])).
fof(f1250,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X5))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X6))),s(X1,X5))) | ~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))) )),
  inference(cnf_transformation,[],[f915])).
fof(f915,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) | ! [X5,X6] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X5))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X6))),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,sK63(X2,X1,X0)))),s(X1,sK64(X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,sK64(X2,X1,X0)))),s(X1,sK63(X2,X1,X0)))) | p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK63,sK64])],[f914])).
fof(f914,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) | ! [X5,X6] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X5))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X6))),s(X1,X5)))) & (? [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))) | p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))))),
  inference(rectify,[],[f913])).
fof(f913,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) | ! [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3)))) & (? [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))) | p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))))),
  inference(nnf_transformation,[],[f668])).
fof(f668,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) <=> ! [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))))),
  inference(flattening,[],[f667])).
fof(f667,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) <=> ! [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11,X6] : ! [X7] : (p(s(bool,'const.operator.COMM_1'(s(fun(X6,fun(X6,X11)),X7)))) <=> ! [X12,X13] : s(X11,happ(s(fun(X6,X11),happ(s(fun(X6,fun(X6,X11)),X7),s(X6,X12))),s(X6,X13))) = s(X11,happ(s(fun(X6,X11),happ(s(fun(X6,fun(X6,X11)),X7),s(X6,X13))),s(X6,X12))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.operator.COMM_DEF')).
fof(f932,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (s(X2,X7) = s(X2,happ(s(fun(X0,X2),sK24(X3,X2,X1,X0)),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))))) | s(X2,sK20(X3,X2,X1,X0)) != s(X2,sK22(X3,X2,X1,X0)) | s(X1,sK19(X3,X2,X1,X0)) != s(X1,sK21(X3,X2,X1,X0)) | s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK19(X3,X2,X1,X0)))),s(X2,sK20(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK21(X3,X2,X1,X0)))),s(X2,sK22(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f752])).
fof(f752,plain,(
  ! [X0,X1,X2,X3] : (((s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK19(X3,X2,X1,X0)))),s(X2,sK20(X3,X2,X1,X0)))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK21(X3,X2,X1,X0)))),s(X2,sK22(X3,X2,X1,X0)))) | (s(X1,sK19(X3,X2,X1,X0)) = s(X1,sK21(X3,X2,X1,X0)) & s(X2,sK20(X3,X2,X1,X0)) = s(X2,sK22(X3,X2,X1,X0)))) & (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK19(X3,X2,X1,X0)))),s(X2,sK20(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK21(X3,X2,X1,X0)))),s(X2,sK22(X3,X2,X1,X0)))) | s(X1,sK19(X3,X2,X1,X0)) != s(X1,sK21(X3,X2,X1,X0)) | s(X2,sK20(X3,X2,X1,X0)) != s(X2,sK22(X3,X2,X1,X0)))) | ! [X6,X7] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),sK23(X3,X2,X1,X0)),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))))) & s(X2,X7) = s(X2,happ(s(fun(X0,X2),sK24(X3,X2,X1,X0)),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20,sK21,sK22,sK23,sK24])],[f751])).
fof(f751,plain,(
  ! [X0,X1,X2,X3] : (? [X8,X9,X10,X11] : ((s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X8))),s(X2,X9))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11))) | (s(X1,X8) = s(X1,X10) & s(X2,X9) = s(X2,X11))) & (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X8))),s(X2,X9))) != s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11))) | s(X1,X8) != s(X1,X10) | s(X2,X9) != s(X2,X11))) | ? [X4,X5] : ! [X6,X7] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))))) & s(X2,X7) = s(X2,happ(s(fun(X0,X2),X5),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7)))))))),
  inference(rectify,[],[f750])).
fof(f750,plain,(
  ! [X0,X1,X2,X3] : (? [X4,X5,X6,X7] : ((s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))) | (s(X1,X4) = s(X1,X6) & s(X2,X5) = s(X2,X7))) & (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) != s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))) | s(X1,X4) != s(X1,X6) | s(X2,X5) != s(X2,X7))) | ? [X8,X9] : ! [X10,X11] : (s(X1,X10) = s(X1,happ(s(fun(X0,X1),X8),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11))))) & s(X2,X11) = s(X2,happ(s(fun(X0,X2),X9),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11)))))))),
  inference(flattening,[],[f749])).
fof(f749,plain,(
  ! [X0,X1,X2,X3] : (? [X4,X5,X6,X7] : ((s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))) | (s(X1,X4) = s(X1,X6) & s(X2,X5) = s(X2,X7))) & (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) != s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))) | (s(X1,X4) != s(X1,X6) | s(X2,X5) != s(X2,X7)))) | ? [X8,X9] : ! [X10,X11] : (s(X1,X10) = s(X1,happ(s(fun(X0,X1),X8),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11))))) & s(X2,X11) = s(X2,happ(s(fun(X0,X2),X9),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11)))))))),
  inference(nnf_transformation,[],[f672])).
fof(f672,plain,(
  ! [X0,X1,X2,X3] : (? [X4,X5,X6,X7] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))) <~> (s(X1,X4) = s(X1,X6) & s(X2,X5) = s(X2,X7))) | ? [X8,X9] : ! [X10,X11] : (s(X1,X10) = s(X1,happ(s(fun(X0,X1),X8),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11))))) & s(X2,X11) = s(X2,happ(s(fun(X0,X2),X9),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11)))))))),
  inference(ennf_transformation,[],[f521])).
fof(f521,plain,(
  ! [X0,X1,X2,X3] : (! [X4,X5,X6,X7] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))) <=> (s(X1,X4) = s(X1,X6) & s(X2,X5) = s(X2,X7))) => ? [X8,X9] : ! [X10,X11] : (s(X1,X10) = s(X1,happ(s(fun(X0,X1),X8),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11))))) & s(X2,X11) = s(X2,happ(s(fun(X0,X2),X9),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11)))))))),
  inference(flattening,[],[f520])).
fof(f520,plain,(
  ! [X0,X1,X2] : ! [X3] : (! [X4,X5,X6,X7] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))) <=> (s(X1,X4) = s(X1,X6) & s(X2,X5) = s(X2,X7))) => ? [X8,X9] : ! [X10,X11] : (s(X1,X10) = s(X1,happ(s(fun(X0,X1),X8),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11))))) & s(X2,X11) = s(X2,happ(s(fun(X0,X2),X9),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X10))),s(X2,X11)))))))),
  inference(rectify,[],[f385])).
fof(f385,axiom,(
  ! [X47,X6,X11] : ! [X16] : (! [X106,X108,X107,X109] : (s(X47,happ(s(fun(X11,X47),happ(s(fun(X6,fun(X11,X47)),X16),s(X6,X106))),s(X11,X108))) = s(X47,happ(s(fun(X11,X47),happ(s(fun(X6,fun(X11,X47)),X16),s(X6,X107))),s(X11,X109))) <=> (s(X6,X106) = s(X6,X107) & s(X11,X108) = s(X11,X109))) => ? [X110,X111] : ! [X12,X13] : (s(X6,X12) = s(X6,happ(s(fun(X47,X6),X110),s(X47,happ(s(fun(X11,X47),happ(s(fun(X6,fun(X11,X47)),X16),s(X6,X12))),s(X11,X13))))) & s(X11,X13) = s(X11,happ(s(fun(X47,X11),X111),s(X47,happ(s(fun(X11,X47),happ(s(fun(X6,fun(X11,X47)),X16),s(X6,X12))),s(X11,X13)))))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.ind_type.INJ_INVERSE2')).
fof(f32611,plain,(
  s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14)))) | $spl436),
  inference(cnf_transformation,[],[f32611_D])).
fof(f32611_D,plain,(
  s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14)))) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f32616,plain,(
  s(sK14,sK20(sK15,sK14,sK14,sK14)) = s(sK14,sK22(sK15,sK14,sK14,sK14)) | $spl442),
  inference(cnf_transformation,[],[f32616_D])).
fof(f32616_D,plain,(
  s(sK14,sK20(sK15,sK14,sK14,sK14)) = s(sK14,sK22(sK15,sK14,sK14,sK14)) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f32503,plain,(
  ( ! [X12,X11] : (s(sK14,X12) = s(sK14,happ(s(fun(sK14,sK14),sK24(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X12))),s(sK14,X11))))) | s(sK14,sK20(sK15,sK14,sK14,sK14)) = s(sK14,sK22(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f931,f1263])).
fof(f931,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (s(X2,X7) = s(X2,happ(s(fun(X0,X2),sK24(X3,X2,X1,X0)),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))))) | s(X2,sK20(X3,X2,X1,X0)) = s(X2,sK22(X3,X2,X1,X0)) | s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK19(X3,X2,X1,X0)))),s(X2,sK20(X3,X2,X1,X0)))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK21(X3,X2,X1,X0)))),s(X2,sK22(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f752])).
fof(f32613,plain,(
  s(sK14,sK19(sK15,sK14,sK14,sK14)) = s(sK14,sK21(sK15,sK14,sK14,sK14)) | $spl438),
  inference(cnf_transformation,[],[f32613_D])).
fof(f32613_D,plain,(
  s(sK14,sK19(sK15,sK14,sK14,sK14)) = s(sK14,sK21(sK15,sK14,sK14,sK14)) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f32502,plain,(
  ( ! [X10,X9] : (s(sK14,X10) = s(sK14,happ(s(fun(sK14,sK14),sK24(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X10))),s(sK14,X9))))) | s(sK14,sK19(sK15,sK14,sK14,sK14)) = s(sK14,sK21(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f930,f1263])).
fof(f930,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (s(X2,X7) = s(X2,happ(s(fun(X0,X2),sK24(X3,X2,X1,X0)),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))))) | s(X1,sK19(X3,X2,X1,X0)) = s(X1,sK21(X3,X2,X1,X0)) | s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK19(X3,X2,X1,X0)))),s(X2,sK20(X3,X2,X1,X0)))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK21(X3,X2,X1,X0)))),s(X2,sK22(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f752])).
fof(f32614,plain,(
  ( ! [X4,X3] : (s(sK14,X3) = s(sK14,happ(s(fun(sK14,sK14),sK23(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X4))),s(sK14,X3)))))) ) | $spl440),
  inference(cnf_transformation,[],[f32614_D])).
fof(f32614_D,plain,(
  ( ! [X4,X3] : (s(sK14,X3) = s(sK14,happ(s(fun(sK14,sK14),sK23(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X4))),s(sK14,X3)))))) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f32501,plain,(
  ( ! [X8,X7] : (s(sK14,X7) = s(sK14,happ(s(fun(sK14,sK14),sK23(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X8))),s(sK14,X7))))) | s(sK14,sK20(sK15,sK14,sK14,sK14)) != s(sK14,sK22(sK15,sK14,sK14,sK14)) | s(sK14,sK19(sK15,sK14,sK14,sK14)) != s(sK14,sK21(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) != s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f929,f1263])).
fof(f929,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),sK23(X3,X2,X1,X0)),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))))) | s(X2,sK20(X3,X2,X1,X0)) != s(X2,sK22(X3,X2,X1,X0)) | s(X1,sK19(X3,X2,X1,X0)) != s(X1,sK21(X3,X2,X1,X0)) | s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK19(X3,X2,X1,X0)))),s(X2,sK20(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK21(X3,X2,X1,X0)))),s(X2,sK22(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f752])).
fof(f32500,plain,(
  ( ! [X6,X5] : (s(sK14,X5) = s(sK14,happ(s(fun(sK14,sK14),sK23(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X6))),s(sK14,X5))))) | s(sK14,sK20(sK15,sK14,sK14,sK14)) = s(sK14,sK22(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f928,f1263])).
fof(f928,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),sK23(X3,X2,X1,X0)),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))))) | s(X2,sK20(X3,X2,X1,X0)) = s(X2,sK22(X3,X2,X1,X0)) | s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK19(X3,X2,X1,X0)))),s(X2,sK20(X3,X2,X1,X0)))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK21(X3,X2,X1,X0)))),s(X2,sK22(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f752])).
fof(f32499,plain,(
  ( ! [X4,X3] : (s(sK14,X3) = s(sK14,happ(s(fun(sK14,sK14),sK23(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X4))),s(sK14,X3))))) | s(sK14,sK19(sK15,sK14,sK14,sK14)) = s(sK14,sK21(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f927,f1263])).
fof(f927,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),sK23(X3,X2,X1,X0)),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X6))),s(X2,X7))))) | s(X1,sK19(X3,X2,X1,X0)) = s(X1,sK21(X3,X2,X1,X0)) | s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK19(X3,X2,X1,X0)))),s(X2,sK20(X3,X2,X1,X0)))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,sK21(X3,X2,X1,X0)))),s(X2,sK22(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f752])).
fof(f32400,plain,(
  ( ! [X14,X13] : (s(sK14,X14) = s(sK14,happ(s(fun(sK14,sK14),sK24(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X14))),s(sK14,X13))))) | s(sK14,sK20(sK15,sK14,sK14,sK14)) != s(sK14,sK22(sK15,sK14,sK14,sK14)) | s(sK14,sK19(sK15,sK14,sK14,sK14)) != s(sK14,sK21(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) != s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f932,f1263])).
fof(f32399,plain,(
  ( ! [X12,X11] : (s(sK14,X12) = s(sK14,happ(s(fun(sK14,sK14),sK24(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X12))),s(sK14,X11))))) | s(sK14,sK20(sK15,sK14,sK14,sK14)) = s(sK14,sK22(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f931,f1263])).
fof(f32398,plain,(
  ( ! [X10,X9] : (s(sK14,X10) = s(sK14,happ(s(fun(sK14,sK14),sK24(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X10))),s(sK14,X9))))) | s(sK14,sK19(sK15,sK14,sK14,sK14)) = s(sK14,sK21(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f930,f1263])).
fof(f32397,plain,(
  ( ! [X8,X7] : (s(sK14,X7) = s(sK14,happ(s(fun(sK14,sK14),sK23(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X8))),s(sK14,X7))))) | s(sK14,sK20(sK15,sK14,sK14,sK14)) != s(sK14,sK22(sK15,sK14,sK14,sK14)) | s(sK14,sK19(sK15,sK14,sK14,sK14)) != s(sK14,sK21(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) != s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f929,f1263])).
fof(f32396,plain,(
  ( ! [X6,X5] : (s(sK14,X5) = s(sK14,happ(s(fun(sK14,sK14),sK23(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X6))),s(sK14,X5))))) | s(sK14,sK20(sK15,sK14,sK14,sK14)) = s(sK14,sK22(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f928,f1263])).
fof(f32395,plain,(
  ( ! [X4,X3] : (s(sK14,X3) = s(sK14,happ(s(fun(sK14,sK14),sK23(sK15,sK14,sK14,sK14)),s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X4))),s(sK14,X3))))) | s(sK14,sK19(sK15,sK14,sK14,sK14)) = s(sK14,sK21(sK15,sK14,sK14,sK14)) | s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK19(sK15,sK14,sK14,sK14)))),s(sK14,sK20(sK15,sK14,sK14,sK14)))) = s(sK14,happ(s(fun(sK14,sK14),happ(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK21(sK15,sK14,sK14,sK14)))),s(sK14,sK22(sK15,sK14,sK14,sK14))))) )),
  inference(superposition,[],[f927,f1263])).
fof(f32299,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP4(X1,X2,X3)) ) | $spl434),
  inference(cnf_transformation,[],[f32299_D])).
fof(f32299_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP4(X1,X2,X3)) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f32116,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP4(X705,X706,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X704) != s(bool,X709)) )),
  inference(resolution,[],[f9633,f1252])).
fof(f9633,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X1,X2,X3) | p(s(bool,X0))) )),
  inference(resolution,[],[f7634,f5228])).
fof(f7634,plain,(
  ( ! [X187,X185,X186,X184] : (sP12(X184,X185) | sP12(X184,sK62) | sP4(X185,X186,X187)) )),
  inference(resolution,[],[f5470,f1078])).
fof(f1078,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f835])).
fof(f835,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f834])).
fof(f834,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f833])).
fof(f833,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f729])).
fof(f729,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f32003,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP4(X50,X51,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X49) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9633,f926])).
fof(f32000,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP4(X29,X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X28) != s(bool,X34)) )),
  inference(resolution,[],[f9633,f925])).
fof(f31999,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP4(X22,X23,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X21) != s(bool,X26)) )),
  inference(resolution,[],[f9633,f924])).
fof(f31998,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP4(X15,X16,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X14) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9633,f922])).
fof(f31996,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP4(X1,X2,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X0) != s(bool,X5)) )),
  inference(resolution,[],[f9633,f920])).
fof(f31952,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP4(X2,X3,X1)) ) | $spl432),
  inference(cnf_transformation,[],[f31952_D])).
fof(f31952_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP4(X2,X3,X1)) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f31768,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP4(X706,X707,X705) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X704) != s(bool,X709)) )),
  inference(resolution,[],[f9570,f1252])).
fof(f9570,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X1) | p(s(bool,X0))) )),
  inference(resolution,[],[f7633,f5228])).
fof(f7633,plain,(
  ( ! [X182,X180,X183,X181] : (sP12(X180,X181) | sP12(X180,sK62) | sP4(X182,X183,X181)) )),
  inference(resolution,[],[f5470,f1076])).
fof(f1076,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f835])).
fof(f31655,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP4(X51,X52,X50) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X49) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9570,f926])).
fof(f31652,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP4(X30,X31,X29) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X28) != s(bool,X34)) )),
  inference(resolution,[],[f9570,f925])).
fof(f31651,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP4(X23,X24,X22) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X21) != s(bool,X26)) )),
  inference(resolution,[],[f9570,f924])).
fof(f31650,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP4(X16,X17,X15) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X14) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9570,f922])).
fof(f31648,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X1) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X0) != s(bool,X5)) )),
  inference(resolution,[],[f9570,f920])).
fof(f31587,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK49(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80)))) | ~p(s(bool,X84))) ) | $spl430),
  inference(cnf_transformation,[],[f31587_D])).
fof(f31587_D,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK49(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80)))) | ~p(s(bool,X84))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f31482,plain,(
  ( ! [X656,X658,X654,X657,X655] : (p(s(bool,happ(s(fun(X654,bool),X655),s(X654,X656)))) | p(s(bool,X657)) | ~p(s(bool,X658)) | s(bool,X658) != s(bool,happ(s(fun(X654,bool),X655),s(X654,sK49(X655,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X654))))) )),
  inference(resolution,[],[f1474,f1252])).
fof(f1474,plain,(
  ( ! [X88,X87,X86] : (~p(s(bool,happ(s(fun(X86,bool),X87),s(X86,sK49(X87,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X86))))) | p(s(bool,happ(s(fun(X86,bool),X87),s(X86,X88))))) )),
  inference(resolution,[],[f917,f1020])).
fof(f1020,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK49(X2,X1,X0))))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f815])).
fof(f815,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK48(X2,X1,X0))))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK49(X2,X1,X0)))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK48,sK49])],[f814])).
fof(f814,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f813])).
fof(f813,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f812])).
fof(f812,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f563])).
fof(f563,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f562])).
fof(f562,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f314])).
fof(f314,axiom,(
  ! [X6] : ! [X16,X29] : ((p(s(bool,X16)) & ! [X12] : p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X12))))) <=> ! [X12] : (p(s(bool,X16)) & p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X12))))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.RIGHT_AND_FORALL_THM')).
fof(f31369,plain,(
  ( ! [X127,X125,X123,X126,X124,X122] : (p(s(bool,happ(s(fun(X122,bool),X123),s(X122,X124)))) | ~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,happ(s(fun(X122,bool),X123),s(X122,sK49(X123,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X122)))) | s(bool,X125) != s(bool,X127)) )),
  inference(resolution,[],[f1474,f926])).
fof(f31367,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) | p(s(bool,X113)) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X114) != s(bool,happ(s(fun(X110,bool),X111),s(X110,sK49(X111,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X110))))) )),
  inference(resolution,[],[f1474,f925])).
fof(f31365,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(X98,bool),X99),s(X98,X100)))) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,happ(s(fun(X98,bool),X99),s(X98,sK49(X99,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X98))))) )),
  inference(resolution,[],[f1474,f924])).
fof(f31364,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,happ(s(fun(X92,bool),X93),s(X92,X94)))) | ~p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X96) != s(bool,happ(s(fun(X92,bool),X93),s(X92,sK49(X93,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X92)))) | s(bool,X95) != s(bool,X97)) )),
  inference(resolution,[],[f1474,f922])).
fof(f31362,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK49(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80))))) )),
  inference(resolution,[],[f1474,f920])).
fof(f31337,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP1(X1,X2,X3)) ) | $spl428),
  inference(cnf_transformation,[],[f31337_D])).
fof(f31337_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP1(X1,X2,X3)) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f31152,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP1(X705,X706,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X704) != s(bool,X709)) )),
  inference(resolution,[],[f9506,f1252])).
fof(f9506,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X1,X2,X3) | p(s(bool,X0))) )),
  inference(resolution,[],[f7605,f5228])).
fof(f7605,plain,(
  ( ! [X74,X72,X75,X73] : (sP12(X72,X73) | sP12(X72,sK62) | sP1(X73,X74,X75)) )),
  inference(resolution,[],[f5470,f947])).
fof(f947,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f761])).
fof(f761,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f760])).
fof(f760,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f759])).
fof(f759,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f725])).
fof(f725,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f31039,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP1(X50,X51,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X49) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9506,f926])).
fof(f31036,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP1(X29,X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X28) != s(bool,X34)) )),
  inference(resolution,[],[f9506,f925])).
fof(f31035,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP1(X22,X23,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X21) != s(bool,X26)) )),
  inference(resolution,[],[f9506,f924])).
fof(f31034,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP1(X15,X16,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X14) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9506,f922])).
fof(f31032,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP1(X1,X2,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X0) != s(bool,X5)) )),
  inference(resolution,[],[f9506,f920])).
fof(f8924,plain,(
  ( ! [X10,X8,X11] : (sP3(X10,X8,X11) | ~sP0(X11,X8,X10) | ~sP2(X10,X11,X8) | p(s(bool,X8))) ) | $spl88),
  inference(cnf_transformation,[],[f8924_D])).
fof(f8924_D,plain,(
  ( ! [X10,X8,X11] : (sP3(X10,X8,X11) | ~sP0(X11,X8,X10) | ~sP2(X10,X11,X8) | p(s(bool,X8))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f31335,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) )),
  inference(subsumption_resolution,[],[f31026,f1165])).
fof(f31026,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,X9) | p(s(bool,X8)) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) )),
  inference(resolution,[],[f9506,f938])).
fof(f938,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f755])).
fof(f755,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f754])).
fof(f754,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f753])).
fof(f753,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f727])).
fof(f727,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f30972,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK48(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80)))) | ~p(s(bool,X84))) ) | $spl426),
  inference(cnf_transformation,[],[f30972_D])).
fof(f30972_D,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK48(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80)))) | ~p(s(bool,X84))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f30868,plain,(
  ( ! [X656,X658,X654,X657,X655] : (p(s(bool,happ(s(fun(X654,bool),X655),s(X654,X656)))) | p(s(bool,X657)) | ~p(s(bool,X658)) | s(bool,X658) != s(bool,happ(s(fun(X654,bool),X655),s(X654,sK48(X655,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X654))))) )),
  inference(resolution,[],[f1473,f1252])).
fof(f1473,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,happ(s(fun(X83,bool),X84),s(X83,sK48(X84,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X83))))) | p(s(bool,happ(s(fun(X83,bool),X84),s(X83,X85))))) )),
  inference(resolution,[],[f917,f1018])).
fof(f1018,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK48(X2,X1,X0))))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f815])).
fof(f30755,plain,(
  ( ! [X127,X125,X123,X126,X124,X122] : (p(s(bool,happ(s(fun(X122,bool),X123),s(X122,X124)))) | ~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,happ(s(fun(X122,bool),X123),s(X122,sK48(X123,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X122)))) | s(bool,X125) != s(bool,X127)) )),
  inference(resolution,[],[f1473,f926])).
fof(f30753,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) | p(s(bool,X113)) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X114) != s(bool,happ(s(fun(X110,bool),X111),s(X110,sK48(X111,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X110))))) )),
  inference(resolution,[],[f1473,f925])).
fof(f30751,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(X98,bool),X99),s(X98,X100)))) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,happ(s(fun(X98,bool),X99),s(X98,sK48(X99,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X98))))) )),
  inference(resolution,[],[f1473,f924])).
fof(f30750,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,happ(s(fun(X92,bool),X93),s(X92,X94)))) | ~p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X96) != s(bool,happ(s(fun(X92,bool),X93),s(X92,sK48(X93,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X92)))) | s(bool,X95) != s(bool,X97)) )),
  inference(resolution,[],[f1473,f922])).
fof(f30748,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK48(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80))))) )),
  inference(resolution,[],[f1473,f920])).
fof(f30723,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP1(X2,X1,X3)) ) | $spl424),
  inference(cnf_transformation,[],[f30723_D])).
fof(f30723_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP1(X2,X1,X3)) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f30540,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP1(X706,X705,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X704) != s(bool,X709)) )),
  inference(resolution,[],[f9443,f1252])).
fof(f9443,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X2,X1,X3) | p(s(bool,X0))) )),
  inference(resolution,[],[f7604,f5228])).
fof(f7604,plain,(
  ( ! [X70,X68,X71,X69] : (sP12(X68,X69) | sP12(X68,sK62) | sP1(X70,X69,X71)) )),
  inference(resolution,[],[f5470,f946])).
fof(f946,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f761])).
fof(f30427,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP1(X51,X50,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X49) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9443,f926])).
fof(f30424,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP1(X30,X29,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X28) != s(bool,X34)) )),
  inference(resolution,[],[f9443,f925])).
fof(f30423,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP1(X23,X22,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X21) != s(bool,X26)) )),
  inference(resolution,[],[f9443,f924])).
fof(f30422,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP1(X16,X15,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X14) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9443,f922])).
fof(f30420,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP1(X2,X1,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X0) != s(bool,X5)) )),
  inference(resolution,[],[f9443,f920])).
fof(f30360,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK38(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80)))) | ~p(s(bool,X84))) ) | $spl422),
  inference(cnf_transformation,[],[f30360_D])).
fof(f30360_D,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK38(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80)))) | ~p(s(bool,X84))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f30257,plain,(
  ( ! [X656,X658,X654,X657,X655] : (p(s(bool,happ(s(fun(X654,bool),X655),s(X654,X656)))) | p(s(bool,X657)) | ~p(s(bool,X658)) | s(bool,X658) != s(bool,happ(s(fun(X654,bool),X655),s(X654,sK38(X655,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X654))))) )),
  inference(resolution,[],[f1472,f1252])).
fof(f1472,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,happ(s(fun(X80,bool),X81),s(X80,sK38(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80))))) | p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82))))) )),
  inference(resolution,[],[f917,f1000])).
fof(f1000,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK38(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f795])).
fof(f795,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0)))))) | ~p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK38(X2,X1,X0)))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK37,sK38])],[f794])).
fof(f794,plain,(
  ! [X0,X1,X2] : ((? [X6] : (p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ~p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f793])).
fof(f793,plain,(
  ! [X0,X1,X2] : ((? [X3] : (p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ~p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f792])).
fof(f792,plain,(
  ! [X0,X1,X2] : ((? [X3] : (p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (~p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f677])).
fof(f677,plain,(
  ! [X0,X1,X2] : (! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (~p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(ennf_transformation,[],[f553])).
fof(f553,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) => ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f552])).
fof(f552,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) => ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X16,X29] : (! [X12] : (p(s(bool,X16)) => p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X12))))) <=> (p(s(bool,X16)) => ! [X12] : p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X12))))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.RIGHT_FORALL_IMP_THM')).
fof(f30144,plain,(
  ( ! [X127,X125,X123,X126,X124,X122] : (p(s(bool,happ(s(fun(X122,bool),X123),s(X122,X124)))) | ~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,happ(s(fun(X122,bool),X123),s(X122,sK38(X123,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X122)))) | s(bool,X125) != s(bool,X127)) )),
  inference(resolution,[],[f1472,f926])).
fof(f30142,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) | p(s(bool,X113)) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X114) != s(bool,happ(s(fun(X110,bool),X111),s(X110,sK38(X111,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X110))))) )),
  inference(resolution,[],[f1472,f925])).
fof(f30140,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(X98,bool),X99),s(X98,X100)))) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,happ(s(fun(X98,bool),X99),s(X98,sK38(X99,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X98))))) )),
  inference(resolution,[],[f1472,f924])).
fof(f30139,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,happ(s(fun(X92,bool),X93),s(X92,X94)))) | ~p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X96) != s(bool,happ(s(fun(X92,bool),X93),s(X92,sK38(X93,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X92)))) | s(bool,X95) != s(bool,X97)) )),
  inference(resolution,[],[f1472,f922])).
fof(f30137,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK38(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80))))) )),
  inference(resolution,[],[f1472,f920])).
fof(f30112,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP2(X1,X2,X3)) ) | $spl420),
  inference(cnf_transformation,[],[f30112_D])).
fof(f30112_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP2(X1,X2,X3)) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f29929,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP2(X705,X706,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X704) != s(bool,X709)) )),
  inference(resolution,[],[f9379,f1252])).
fof(f9379,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X1,X2,X3) | p(s(bool,X0))) )),
  inference(resolution,[],[f7602,f5228])).
fof(f7602,plain,(
  ( ! [X61,X62,X60,X63] : (sP12(X60,X61) | sP12(X60,sK62) | sP2(X61,X62,X63)) )),
  inference(resolution,[],[f5470,f943])).
fof(f943,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f758])).
fof(f758,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f757])).
fof(f757,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f756])).
fof(f756,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f726])).
fof(f726,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f29816,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP2(X50,X51,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X49) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9379,f926])).
fof(f29813,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP2(X29,X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X28) != s(bool,X34)) )),
  inference(resolution,[],[f9379,f925])).
fof(f29812,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP2(X22,X23,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X21) != s(bool,X26)) )),
  inference(resolution,[],[f9379,f924])).
fof(f29811,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP2(X15,X16,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X14) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9379,f922])).
fof(f29809,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP2(X1,X2,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X0) != s(bool,X5)) )),
  inference(resolution,[],[f9379,f920])).
fof(f8795,plain,(
  ( ! [X14,X12,X15] : (sP3(X12,X14,X15) | ~sP0(X15,X14,X12) | ~sP1(X14,X15,X12) | p(s(bool,X12))) ) | $spl82),
  inference(cnf_transformation,[],[f8795_D])).
fof(f8795_D,plain,(
  ( ! [X14,X12,X15] : (sP3(X12,X14,X15) | ~sP0(X15,X14,X12) | ~sP1(X14,X15,X12) | p(s(bool,X12))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f30111,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) )),
  inference(subsumption_resolution,[],[f29804,f1165])).
fof(f29804,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,X13) | p(s(bool,X12)) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) )),
  inference(resolution,[],[f9379,f939])).
fof(f939,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f755])).
fof(f29749,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK37(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80)))) | ~p(s(bool,X84))) ) | $spl418),
  inference(cnf_transformation,[],[f29749_D])).
fof(f29749_D,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK37(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80)))) | ~p(s(bool,X84))) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f29647,plain,(
  ( ! [X656,X658,X654,X657,X655] : (p(s(bool,happ(s(fun(X654,bool),X655),s(X654,X656)))) | p(s(bool,X657)) | ~p(s(bool,X658)) | s(bool,X658) != s(bool,happ(s(fun(X654,bool),X655),s(X654,sK37(X655,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X654))))) )),
  inference(resolution,[],[f1471,f1252])).
fof(f1471,plain,(
  ( ! [X78,X79,X77] : (~p(s(bool,happ(s(fun(X77,bool),X78),s(X77,sK37(X78,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X77))))) | p(s(bool,happ(s(fun(X77,bool),X78),s(X77,X79))))) )),
  inference(resolution,[],[f917,f998])).
fof(f998,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f795])).
fof(f29534,plain,(
  ( ! [X127,X125,X123,X126,X124,X122] : (p(s(bool,happ(s(fun(X122,bool),X123),s(X122,X124)))) | ~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,happ(s(fun(X122,bool),X123),s(X122,sK37(X123,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X122)))) | s(bool,X125) != s(bool,X127)) )),
  inference(resolution,[],[f1471,f926])).
fof(f29532,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) | p(s(bool,X113)) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X114) != s(bool,happ(s(fun(X110,bool),X111),s(X110,sK37(X111,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X110))))) )),
  inference(resolution,[],[f1471,f925])).
fof(f29530,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(X98,bool),X99),s(X98,X100)))) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,happ(s(fun(X98,bool),X99),s(X98,sK37(X99,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X98))))) )),
  inference(resolution,[],[f1471,f924])).
fof(f29529,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,happ(s(fun(X92,bool),X93),s(X92,X94)))) | ~p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X96) != s(bool,happ(s(fun(X92,bool),X93),s(X92,sK37(X93,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X92)))) | s(bool,X95) != s(bool,X97)) )),
  inference(resolution,[],[f1471,f922])).
fof(f29527,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK37(X81,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X80))))) )),
  inference(resolution,[],[f1471,f920])).
fof(f26766,plain,(
  ( ! [X24,X25] : (sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X24,X25)) ) | $spl356),
  inference(cnf_transformation,[],[f26766_D])).
fof(f26766_D,plain,(
  ( ! [X24,X25] : (sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X24,X25)) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f29501,plain,(
  ( ! [X39,X37,X38,X40] : (sP12(X37,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X37,X38) | p(s(bool,X39)) | ~p(s(bool,X40))) ) | $spl352),
  inference(resolution,[],[f27453,f1510])).
fof(f1510,plain,(
  ( ! [X159,X160] : (~sP8('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X159,X160) | p(s(bool,X160)) | ~p(s(bool,X159))) )),
  inference(resolution,[],[f917,f1109])).
fof(f1109,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f850])).
fof(f27453,plain,(
  ( ! [X383,X379,X380,X382,X381] : (sP12(X379,X380) | sP13(X379,X381) | sP8(X380,X382,X383)) ) | $spl352),
  inference(resolution,[],[f26762,f1112])).
fof(f26765,plain,(
  ( ! [X21,X20] : (sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X20,X21)) ) | $spl354),
  inference(cnf_transformation,[],[f26765_D])).
fof(f26765_D,plain,(
  ( ! [X21,X20] : (sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X20,X21)) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f29500,plain,(
  ( ! [X35,X33,X36,X34] : (sP12(X33,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X33,X34) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | $spl352),
  inference(resolution,[],[f27453,f1339])).
fof(f1339,plain,(
  ( ! [X163,X164] : (~sP8('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X163,X164) | p(s(bool,X164)) | ~p(s(bool,X163))) )),
  inference(resolution,[],[f916,f1109])).
fof(f17410,plain,(
  ( ! [X54] : (sP12('F_0',X54)) ) | $spl224),
  inference(cnf_transformation,[],[f17410_D])).
fof(f17410_D,plain,(
  ( ! [X54] : (sP12('F_0',X54)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f14569,plain,(
  ( ! [X243,X241,X242,X240] : (p(s(bool,X240)) | sP8(X241,X242,X243) | sP12(X240,X241)) ) | $spl174),
  inference(cnf_transformation,[],[f14569_D])).
fof(f14569_D,plain,(
  ( ! [X243,X241,X242,X240] : (p(s(bool,X240)) | sP8(X241,X242,X243) | sP12(X240,X241)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f29489,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP8(X46,X47,X48) | p(s(bool,X45)) | sP12('F_0',X49)) ) | ($spl176 | $spl352)),
  inference(resolution,[],[f27453,f15187])).
fof(f15187,plain,(
  ( ! [X318,X317] : (~sP13(X318,X317) | p(s(bool,X318)) | sP12('F_0',X317)) ) | $spl176),
  inference(subsumption_resolution,[],[f15048,f1215])).
fof(f1215,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f897])).
fof(f897,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f896])).
fof(f896,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f642])).
fof(f642,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X14] : ((p(s(bool,X14)) | p(s(bool,'F_0'))) <=> p(s(bool,X14)))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.OR_CLAUSES|split|3')).
fof(f15048,plain,(
  ( ! [X318,X317] : (sP12('F_0',X317) | p(s(bool,'F_0')) | p(s(bool,X318)) | ~sP13(X318,X317)) ) | $spl176),
  inference(resolution,[],[f14807,f1181])).
fof(f1181,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP13(X0,X1)) )),
  inference(cnf_transformation,[],[f880])).
fof(f880,plain,(
  ! [X0,X1] : ((~sP13(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP13(X0,X1)))),
  inference(flattening,[],[f879])).
fof(f879,plain,(
  ! [X0,X1] : ((~sP13(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP13(X0,X1)))),
  inference(nnf_transformation,[],[f744])).
fof(f744,plain,(
  ! [X0,X1] : (sP13(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f17408,plain,(
  ( ! [X44] : (sP12(X44,'F_0')) ) | $spl220),
  inference(cnf_transformation,[],[f17408_D])).
fof(f17408_D,plain,(
  ( ! [X44] : (sP12(X44,'F_0')) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f29487,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP8(X36,X37,X38) | p(s(bool,X35)) | sP12(X39,'F_0')) ) | ($spl116 | $spl352)),
  inference(resolution,[],[f27453,f13741])).
fof(f13741,plain,(
  ( ! [X318,X317] : (~sP13(X318,X317) | p(s(bool,X318)) | sP12(X317,'F_0')) ) | $spl116),
  inference(subsumption_resolution,[],[f13604,f1215])).
fof(f13604,plain,(
  ( ! [X318,X317] : (sP12(X317,'F_0') | p(s(bool,'F_0')) | p(s(bool,X318)) | ~sP13(X318,X317)) ) | $spl116),
  inference(resolution,[],[f13492,f1181])).
fof(f9985,plain,(
  ( ! [X13] : (sP12(sK62,X13)) ) | $spl96),
  inference(cnf_transformation,[],[f9985_D])).
fof(f9985_D,plain,(
  ( ! [X13] : (sP12(sK62,X13)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f29485,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP8(X26,X27,X28) | p(s(bool,X25)) | sP12(sK62,X29)) ) | $spl352),
  inference(resolution,[],[f27453,f5864])).
fof(f5864,plain,(
  ( ! [X318,X317] : (~sP13(X318,X317) | p(s(bool,X318)) | sP12(sK62,X317)) )),
  inference(subsumption_resolution,[],[f5728,f1215])).
fof(f5728,plain,(
  ( ! [X318,X317] : (sP12(sK62,X317) | p(s(bool,'F_0')) | p(s(bool,X318)) | ~sP13(X318,X317)) )),
  inference(resolution,[],[f5231,f1181])).
fof(f5231,plain,(
  ( ! [X205] : (~p(s(bool,X205)) | sP12(sK62,X205)) )),
  inference(resolution,[],[f5116,f1170])).
fof(f8796,plain,(
  ( ! [X13] : (sP12(X13,sK62)) ) | $spl84),
  inference(cnf_transformation,[],[f8796_D])).
fof(f8796_D,plain,(
  ( ! [X13] : (sP12(X13,sK62)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f29483,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP8(X16,X17,X18) | p(s(bool,X15)) | sP12(X19,sK62)) ) | $spl352),
  inference(resolution,[],[f27453,f5611])).
fof(f5611,plain,(
  ( ! [X318,X317] : (~sP13(X318,X317) | p(s(bool,X318)) | sP12(X317,sK62)) )),
  inference(subsumption_resolution,[],[f5475,f1215])).
fof(f5475,plain,(
  ( ! [X318,X317] : (sP12(X317,sK62) | p(s(bool,'F_0')) | p(s(bool,X318)) | ~sP13(X318,X317)) )),
  inference(resolution,[],[f5230,f1181])).
fof(f29482,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP8(X11,X12,X13) | p(s(bool,X14)) | p(s(bool,X10))) ) | $spl352),
  inference(resolution,[],[f27453,f5117])).
fof(f5117,plain,(
  ( ! [X344,X345] : (~sP13(X344,X345) | p(s(bool,X345)) | p(s(bool,X344))) )),
  inference(subsumption_resolution,[],[f4994,f1381])).
fof(f4994,plain,(
  ( ! [X344,X345] : (s(bool,'F_0') = s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X344)) | p(s(bool,X345)) | ~sP13(X344,X345)) )),
  inference(resolution,[],[f1288,f1185])).
fof(f1185,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP13(X0,X1)) )),
  inference(cnf_transformation,[],[f882])).
fof(f882,plain,(
  ! [X0,X1] : ((~sP13(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP13(X0,X1)))),
  inference(flattening,[],[f881])).
fof(f881,plain,(
  ! [X0,X1] : ((~sP13(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP13(X0,X1)))),
  inference(nnf_transformation,[],[f745])).
fof(f745,plain,(
  ! [X0,X1] : (sP13(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f703,f744])).
fof(f703,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f702])).
fof(f702,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f624])).
fof(f624,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f623])).
fof(f623,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f163])).
fof(f163,axiom,(
  ! [X11,X6] : ((~(p(s(bool,X6)) | p(s(bool,X11))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X11)) => p(s(bool,'F_0')))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.sat.sat_9_prove_ltb')).
fof(f29503,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X6,X7,X8) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(subsumption_resolution,[],[f29481,f1215])).
fof(f29481,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X6,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(resolution,[],[f27453,f1185])).
fof(f29502,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X1,X2,X3) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(subsumption_resolution,[],[f29480,f1215])).
fof(f29480,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X1,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(resolution,[],[f27453,f1181])).
fof(f29447,plain,(
  ( ! [X30,X31,X29,X32] : (sP12(X29,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X29,X30) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl352),
  inference(resolution,[],[f27452,f1511])).
fof(f1511,plain,(
  ( ! [X161,X162] : (~sP8(X161,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X162) | p(s(bool,X162)) | ~p(s(bool,X161))) )),
  inference(resolution,[],[f917,f1109])).
fof(f27452,plain,(
  ( ! [X377,X374,X376,X378,X375] : (sP12(X374,X375) | sP13(X374,X376) | sP8(X377,X375,X378)) ) | $spl352),
  inference(resolution,[],[f26762,f1111])).
fof(f29446,plain,(
  ( ! [X28,X26,X27,X25] : (sP12(X25,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X25,X26) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl352),
  inference(resolution,[],[f27452,f1340])).
fof(f1340,plain,(
  ( ! [X165,X166] : (~sP8(X165,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X166) | p(s(bool,X166)) | ~p(s(bool,X165))) )),
  inference(resolution,[],[f916,f1109])).
fof(f14568,plain,(
  ( ! [X239,X237,X238,X236] : (p(s(bool,X236)) | sP8(X238,X237,X239) | sP12(X236,X237)) ) | $spl172),
  inference(cnf_transformation,[],[f14568_D])).
fof(f14568_D,plain,(
  ( ! [X239,X237,X238,X236] : (p(s(bool,X236)) | sP8(X238,X237,X239) | sP12(X236,X237)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f29437,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP8(X47,X46,X48) | p(s(bool,X45)) | sP12('F_0',X49)) ) | ($spl176 | $spl352)),
  inference(resolution,[],[f27452,f15187])).
fof(f29435,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP8(X37,X36,X38) | p(s(bool,X35)) | sP12(X39,'F_0')) ) | ($spl116 | $spl352)),
  inference(resolution,[],[f27452,f13741])).
fof(f29433,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP8(X27,X26,X28) | p(s(bool,X25)) | sP12(sK62,X29)) ) | $spl352),
  inference(resolution,[],[f27452,f5864])).
fof(f29431,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP8(X17,X16,X18) | p(s(bool,X15)) | sP12(X19,sK62)) ) | $spl352),
  inference(resolution,[],[f27452,f5611])).
fof(f29430,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP8(X12,X11,X13) | p(s(bool,X14)) | p(s(bool,X10))) ) | $spl352),
  inference(resolution,[],[f27452,f5117])).
fof(f29451,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X7,X6,X8) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(subsumption_resolution,[],[f29429,f1215])).
fof(f29429,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X7,X6,X8) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(resolution,[],[f27452,f1185])).
fof(f29450,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X2,X1,X3) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(subsumption_resolution,[],[f29428,f1215])).
fof(f29428,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X2,X1,X3) | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(resolution,[],[f27452,f1181])).
fof(f29399,plain,(
  ( ! [X30,X31,X29,X32] : (sP12(X29,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X29,X30) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl352),
  inference(resolution,[],[f27439,f1494])).
fof(f1494,plain,(
  ( ! [X127,X128] : (~sP6('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X127,X128) | p(s(bool,X128)) | p(s(bool,X127))) )),
  inference(resolution,[],[f917,f1092])).
fof(f1092,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f843])).
fof(f27439,plain,(
  ( ! [X313,X310,X312,X309,X311] : (sP12(X309,X310) | sP13(X309,X311) | sP6(X310,X312,X313)) ) | $spl352),
  inference(resolution,[],[f26762,f1095])).
fof(f29398,plain,(
  ( ! [X28,X26,X27,X25] : (sP12(X25,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X25,X26) | p(s(bool,X27)) | p(s(bool,X28))) ) | $spl352),
  inference(resolution,[],[f27439,f1323])).
fof(f1323,plain,(
  ( ! [X132,X131] : (~sP6('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X131,X132) | p(s(bool,X132)) | p(s(bool,X131))) )),
  inference(resolution,[],[f916,f1092])).
fof(f14567,plain,(
  ( ! [X235,X233,X234,X232] : (p(s(bool,X232)) | sP6(X233,X234,X235) | sP12(X232,X233)) ) | $spl170),
  inference(cnf_transformation,[],[f14567_D])).
fof(f14567_D,plain,(
  ( ! [X235,X233,X234,X232] : (p(s(bool,X232)) | sP6(X233,X234,X235) | sP12(X232,X233)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f29389,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP6(X46,X47,X48) | p(s(bool,X45)) | sP12('F_0',X49)) ) | ($spl176 | $spl352)),
  inference(resolution,[],[f27439,f15187])).
fof(f29387,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP6(X36,X37,X38) | p(s(bool,X35)) | sP12(X39,'F_0')) ) | ($spl116 | $spl352)),
  inference(resolution,[],[f27439,f13741])).
fof(f29385,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP6(X26,X27,X28) | p(s(bool,X25)) | sP12(sK62,X29)) ) | $spl352),
  inference(resolution,[],[f27439,f5864])).
fof(f29383,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP6(X16,X17,X18) | p(s(bool,X15)) | sP12(X19,sK62)) ) | $spl352),
  inference(resolution,[],[f27439,f5611])).
fof(f29382,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP6(X11,X12,X13) | p(s(bool,X14)) | p(s(bool,X10))) ) | $spl352),
  inference(resolution,[],[f27439,f5117])).
fof(f29401,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP6(X6,X7,X8) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(subsumption_resolution,[],[f29381,f1215])).
fof(f29381,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP6(X6,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(resolution,[],[f27439,f1185])).
fof(f29400,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP6(X1,X2,X3) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(subsumption_resolution,[],[f29380,f1215])).
fof(f29380,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP6(X1,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(resolution,[],[f27439,f1181])).
fof(f29351,plain,(
  ( ! [X39,X37,X38,X40] : (sP12(X37,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X37,X38) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | $spl352),
  inference(resolution,[],[f27423,f1480])).
fof(f1480,plain,(
  ( ! [X99,X100] : (~sP4('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X99,X100) | ~p(s(bool,X100)) | p(s(bool,X99))) )),
  inference(resolution,[],[f917,f1075])).
fof(f1075,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f835])).
fof(f27423,plain,(
  ( ! [X231,X229,X233,X230,X232] : (sP12(X229,X230) | sP13(X229,X231) | sP4(X230,X232,X233)) ) | $spl352),
  inference(resolution,[],[f26762,f1078])).
fof(f29350,plain,(
  ( ! [X35,X33,X36,X34] : (sP12(X33,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X33,X34) | ~p(s(bool,X35)) | p(s(bool,X36))) ) | $spl352),
  inference(resolution,[],[f27423,f1309])).
fof(f1309,plain,(
  ( ! [X103,X104] : (~sP4('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X103,X104) | ~p(s(bool,X104)) | p(s(bool,X103))) )),
  inference(resolution,[],[f916,f1075])).
fof(f14566,plain,(
  ( ! [X231,X229,X230,X228] : (p(s(bool,X228)) | sP4(X229,X230,X231) | sP12(X228,X229)) ) | $spl168),
  inference(cnf_transformation,[],[f14566_D])).
fof(f14566_D,plain,(
  ( ! [X231,X229,X230,X228] : (p(s(bool,X228)) | sP4(X229,X230,X231) | sP12(X228,X229)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f29339,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP4(X46,X47,X48) | p(s(bool,X45)) | sP12('F_0',X49)) ) | ($spl176 | $spl352)),
  inference(resolution,[],[f27423,f15187])).
fof(f29337,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP4(X36,X37,X38) | p(s(bool,X35)) | sP12(X39,'F_0')) ) | ($spl116 | $spl352)),
  inference(resolution,[],[f27423,f13741])).
fof(f29335,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP4(X26,X27,X28) | p(s(bool,X25)) | sP12(sK62,X29)) ) | $spl352),
  inference(resolution,[],[f27423,f5864])).
fof(f29333,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP4(X16,X17,X18) | p(s(bool,X15)) | sP12(X19,sK62)) ) | $spl352),
  inference(resolution,[],[f27423,f5611])).
fof(f29332,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP4(X11,X12,X13) | p(s(bool,X14)) | p(s(bool,X10))) ) | $spl352),
  inference(resolution,[],[f27423,f5117])).
fof(f29353,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X6,X7,X8) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(subsumption_resolution,[],[f29331,f1215])).
fof(f29331,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X6,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(resolution,[],[f27423,f1185])).
fof(f29352,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X1,X2,X3) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(subsumption_resolution,[],[f29330,f1215])).
fof(f29330,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X1,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(resolution,[],[f27423,f1181])).
fof(f29298,plain,(
  ( ! [X30,X31,X29,X32] : (sP12(X29,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X29,X30) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl352),
  inference(resolution,[],[f27422,f1481])).
fof(f1481,plain,(
  ( ! [X101,X102] : (~sP4(X101,X102,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X102)) | ~p(s(bool,X101))) )),
  inference(resolution,[],[f917,f1075])).
fof(f27422,plain,(
  ( ! [X227,X225,X228,X226,X224] : (sP12(X224,X225) | sP13(X224,X226) | sP4(X227,X228,X225)) ) | $spl352),
  inference(resolution,[],[f26762,f1076])).
fof(f29297,plain,(
  ( ! [X28,X26,X27,X25] : (sP12(X25,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X25,X26) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl352),
  inference(resolution,[],[f27422,f1310])).
fof(f1310,plain,(
  ( ! [X105,X106] : (~sP4(X105,X106,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X106)) | ~p(s(bool,X105))) )),
  inference(resolution,[],[f916,f1075])).
fof(f14565,plain,(
  ( ! [X227,X225,X226,X224] : (p(s(bool,X224)) | sP4(X226,X227,X225) | sP12(X224,X225)) ) | $spl166),
  inference(cnf_transformation,[],[f14565_D])).
fof(f14565_D,plain,(
  ( ! [X227,X225,X226,X224] : (p(s(bool,X224)) | sP4(X226,X227,X225) | sP12(X224,X225)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f29288,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP4(X47,X48,X46) | p(s(bool,X45)) | sP12('F_0',X49)) ) | ($spl176 | $spl352)),
  inference(resolution,[],[f27422,f15187])).
fof(f29286,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP4(X37,X38,X36) | p(s(bool,X35)) | sP12(X39,'F_0')) ) | ($spl116 | $spl352)),
  inference(resolution,[],[f27422,f13741])).
fof(f29284,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP4(X27,X28,X26) | p(s(bool,X25)) | sP12(sK62,X29)) ) | $spl352),
  inference(resolution,[],[f27422,f5864])).
fof(f29282,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP4(X17,X18,X16) | p(s(bool,X15)) | sP12(X19,sK62)) ) | $spl352),
  inference(resolution,[],[f27422,f5611])).
fof(f29281,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP4(X12,X13,X11) | p(s(bool,X14)) | p(s(bool,X10))) ) | $spl352),
  inference(resolution,[],[f27422,f5117])).
fof(f29302,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X7,X8,X6) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(subsumption_resolution,[],[f29280,f1215])).
fof(f29280,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X7,X8,X6) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(resolution,[],[f27422,f1185])).
fof(f29301,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X1) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(subsumption_resolution,[],[f29279,f1215])).
fof(f29279,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X1) | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(resolution,[],[f27422,f1181])).
fof(f29252,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK36(sK24(X142,bool,X143,X141),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) | $spl416),
  inference(cnf_transformation,[],[f29252_D])).
fof(f29252_D,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK36(sK24(X142,bool,X143,X141),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f29209,plain,(
  ( ! [X142,X145,X143,X141] : (~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK36(sK24(X142,bool,X143,X141),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X141))))) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141)) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141))))) )),
  inference(superposition,[],[f1470,f932])).
fof(f1470,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,happ(s(fun(X74,bool),X75),s(X74,X76)))) | p(s(bool,happ(s(fun(X74,bool),X75),s(X74,sK36(X75,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X74)))))) )),
  inference(resolution,[],[f917,f996])).
fof(f996,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f791])).
fof(f791,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK35(X2,X1,X0))))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK35,sK36])],[f790])).
fof(f790,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & ? [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f789])).
fof(f789,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f788])).
fof(f788,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f551])).
fof(f551,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f550])).
fof(f550,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f401])).
fof(f401,axiom,(
  ! [X6] : ! [X16,X29] : (? [X12] : (p(s(bool,X16)) & p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X12))))) <=> (p(s(bool,X16)) & ? [X12] : p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X12))))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.RIGHT_EXISTS_AND_THM')).
fof(f29251,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK36(sK24(X137,bool,X138,X136),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) | $spl414),
  inference(cnf_transformation,[],[f29251_D])).
fof(f29251_D,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK36(sK24(X137,bool,X138,X136),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f29208,plain,(
  ( ! [X140,X138,X136,X137] : (~p(s(bool,X140)) | p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK36(sK24(X137,bool,X138,X136),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X136))))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136)) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136))))) )),
  inference(superposition,[],[f1470,f931])).
fof(f29250,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK36(sK24(X132,bool,X133,X131),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) | $spl412),
  inference(cnf_transformation,[],[f29250_D])).
fof(f29250_D,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK36(sK24(X132,bool,X133,X131),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f29207,plain,(
  ( ! [X132,X135,X133,X131] : (~p(s(bool,X135)) | p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK36(sK24(X132,bool,X133,X131),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X131))))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131)) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131))))) )),
  inference(superposition,[],[f1470,f930])).
fof(f29249,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK36(sK23(X127,X128,bool,X126),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) | $spl410),
  inference(cnf_transformation,[],[f29249_D])).
fof(f29249_D,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK36(sK23(X127,X128,bool,X126),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f29206,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,X129)) | p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK36(sK23(X127,X128,bool,X126),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X126))))) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126)) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126))))) )),
  inference(superposition,[],[f1470,f929])).
fof(f29248,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK36(sK23(X122,X123,bool,X121),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) | $spl408),
  inference(cnf_transformation,[],[f29248_D])).
fof(f29248_D,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK36(sK23(X122,X123,bool,X121),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f29205,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X124)) | p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK36(sK23(X122,X123,bool,X121),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X121))))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121)) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121))))) )),
  inference(superposition,[],[f1470,f928])).
fof(f29247,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK36(sK23(X117,X118,bool,X116),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) | $spl406),
  inference(cnf_transformation,[],[f29247_D])).
fof(f29247_D,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK36(sK23(X117,X118,bool,X116),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f29204,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X119)) | p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK36(sK23(X117,X118,bool,X116),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X116))))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116)) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116))))) )),
  inference(superposition,[],[f1470,f927])).
fof(f29221,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK36(X118,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) | $spl404),
  inference(cnf_transformation,[],[f29221_D])).
fof(f29221_D,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK36(X118,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f29058,plain,(
  ( ! [X685,X687,X688,X684,X686] : (p(s(bool,happ(s(fun(X684,bool),X685),s(X684,sK36(X685,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X684))))) | p(s(bool,X686)) | ~p(s(bool,X687)) | s(bool,X687) != s(bool,happ(s(fun(X684,bool),X685),s(X684,X688)))) )),
  inference(resolution,[],[f1470,f1252])).
fof(f28945,plain,(
  ( ! [X158,X163,X161,X159,X162,X160] : (p(s(bool,happ(s(fun(X158,bool),X159),s(X158,sK36(X159,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X158))))) | ~p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X161) != s(bool,happ(s(fun(X158,bool),X159),s(X158,X162))) | s(bool,X160) != s(bool,X163)) )),
  inference(resolution,[],[f1470,f926])).
fof(f28943,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK36(X148,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X147))))) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,happ(s(fun(X147,bool),X148),s(X147,X152)))) )),
  inference(resolution,[],[f1470,f925])).
fof(f28941,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(X135,bool),X136),s(X135,sK36(X136,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X135))))) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,X139) | s(bool,X138) != s(bool,happ(s(fun(X135,bool),X136),s(X135,X140)))) )),
  inference(resolution,[],[f1470,f924])).
fof(f28940,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (p(s(bool,happ(s(fun(X129,bool),X130),s(X129,sK36(X130,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X129))))) | ~p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,happ(s(fun(X129,bool),X130),s(X129,X133))) | s(bool,X131) != s(bool,X134)) )),
  inference(resolution,[],[f1470,f922])).
fof(f28938,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK36(X118,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X117))))) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122)))) )),
  inference(resolution,[],[f1470,f920])).
fof(f28896,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X32,X33) | p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl352),
  inference(resolution,[],[f27394,f1449])).
fof(f1449,plain,(
  ( ! [X39,X40] : (~sP1('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X39,X40) | p(s(bool,X40)) | ~p(s(bool,X39))) )),
  inference(resolution,[],[f917,f944])).
fof(f944,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f761])).
fof(f27394,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP12(X88,X89) | sP13(X88,X90) | sP1(X89,X91,X92)) ) | $spl352),
  inference(resolution,[],[f26762,f947])).
fof(f28895,plain,(
  ( ! [X30,X28,X31,X29] : (sP12(X28,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X28,X29) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl352),
  inference(resolution,[],[f27394,f1278])).
fof(f1278,plain,(
  ( ! [X43,X44] : (~sP1('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X43,X44) | p(s(bool,X44)) | ~p(s(bool,X43))) )),
  inference(resolution,[],[f916,f944])).
fof(f14564,plain,(
  ( ! [X222,X220,X223,X221] : (p(s(bool,X220)) | sP1(X221,X222,X223) | sP12(X220,X221)) ) | $spl164),
  inference(cnf_transformation,[],[f14564_D])).
fof(f14564_D,plain,(
  ( ! [X222,X220,X223,X221] : (p(s(bool,X220)) | sP1(X221,X222,X223) | sP12(X220,X221)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f28885,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP1(X46,X47,X48) | p(s(bool,X45)) | sP12('F_0',X49)) ) | ($spl176 | $spl352)),
  inference(resolution,[],[f27394,f15187])).
fof(f28883,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP1(X36,X37,X38) | p(s(bool,X35)) | sP12(X39,'F_0')) ) | ($spl116 | $spl352)),
  inference(resolution,[],[f27394,f13741])).
fof(f28881,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP1(X26,X27,X28) | p(s(bool,X25)) | sP12(sK62,X29)) ) | $spl352),
  inference(resolution,[],[f27394,f5864])).
fof(f28879,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP1(X16,X17,X18) | p(s(bool,X15)) | sP12(X19,sK62)) ) | $spl352),
  inference(resolution,[],[f27394,f5611])).
fof(f28878,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP1(X11,X12,X13) | p(s(bool,X14)) | p(s(bool,X10))) ) | $spl352),
  inference(resolution,[],[f27394,f5117])).
fof(f28898,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X6,X7,X8) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(subsumption_resolution,[],[f28877,f1215])).
fof(f28877,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X6,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(resolution,[],[f27394,f1185])).
fof(f28897,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X1,X2,X3) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(subsumption_resolution,[],[f28876,f1215])).
fof(f28876,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X1,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(resolution,[],[f27394,f1181])).
fof(f28843,plain,(
  ( ! [X26,X24,X27,X25] : (sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X24,X25) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl352),
  inference(resolution,[],[f27393,f1450])).
fof(f1450,plain,(
  ( ! [X41,X42] : (~sP1(X41,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X42) | p(s(bool,X42)) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f917,f944])).
fof(f27393,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP12(X83,X84) | sP13(X83,X85) | sP1(X86,X84,X87)) ) | $spl352),
  inference(resolution,[],[f26762,f946])).
fof(f28842,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl352),
  inference(resolution,[],[f27393,f1279])).
fof(f1279,plain,(
  ( ! [X45,X46] : (~sP1(X45,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X46) | p(s(bool,X46)) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f916,f944])).
fof(f14563,plain,(
  ( ! [X218,X216,X219,X217] : (p(s(bool,X216)) | sP1(X218,X217,X219) | sP12(X216,X217)) ) | $spl162),
  inference(cnf_transformation,[],[f14563_D])).
fof(f14563_D,plain,(
  ( ! [X218,X216,X219,X217] : (p(s(bool,X216)) | sP1(X218,X217,X219) | sP12(X216,X217)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f28834,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP1(X47,X46,X48) | p(s(bool,X45)) | sP12('F_0',X49)) ) | ($spl176 | $spl352)),
  inference(resolution,[],[f27393,f15187])).
fof(f28832,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP1(X37,X36,X38) | p(s(bool,X35)) | sP12(X39,'F_0')) ) | ($spl116 | $spl352)),
  inference(resolution,[],[f27393,f13741])).
fof(f28830,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP1(X27,X26,X28) | p(s(bool,X25)) | sP12(sK62,X29)) ) | $spl352),
  inference(resolution,[],[f27393,f5864])).
fof(f28828,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP1(X17,X16,X18) | p(s(bool,X15)) | sP12(X19,sK62)) ) | $spl352),
  inference(resolution,[],[f27393,f5611])).
fof(f28827,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP1(X12,X11,X13) | p(s(bool,X14)) | p(s(bool,X10))) ) | $spl352),
  inference(resolution,[],[f27393,f5117])).
fof(f28847,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X7,X6,X8) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(subsumption_resolution,[],[f28826,f1215])).
fof(f28826,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X7,X6,X8) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(resolution,[],[f27393,f1185])).
fof(f28846,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X2,X1,X3) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(subsumption_resolution,[],[f28825,f1215])).
fof(f28825,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X2,X1,X3) | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(resolution,[],[f27393,f1181])).
fof(f28793,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X32,X33) | p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl352),
  inference(resolution,[],[f27391,f1446])).
fof(f1446,plain,(
  ( ! [X33,X34] : (~sP2('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X33,X34) | p(s(bool,X34)) | ~p(s(bool,X33))) )),
  inference(resolution,[],[f917,f940])).
fof(f940,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f758])).
fof(f27391,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP12(X73,X74) | sP13(X73,X75) | sP2(X74,X76,X77)) ) | $spl352),
  inference(resolution,[],[f26762,f943])).
fof(f28792,plain,(
  ( ! [X30,X28,X31,X29] : (sP12(X28,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X28,X29) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl352),
  inference(resolution,[],[f27391,f1275])).
fof(f1275,plain,(
  ( ! [X37,X38] : (~sP2('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X37,X38) | p(s(bool,X38)) | ~p(s(bool,X37))) )),
  inference(resolution,[],[f916,f940])).
fof(f14562,plain,(
  ( ! [X212,X215,X213,X214] : (p(s(bool,X212)) | sP2(X213,X214,X215) | sP12(X212,X213)) ) | $spl160),
  inference(cnf_transformation,[],[f14562_D])).
fof(f14562_D,plain,(
  ( ! [X212,X215,X213,X214] : (p(s(bool,X212)) | sP2(X213,X214,X215) | sP12(X212,X213)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f28782,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP2(X46,X47,X48) | p(s(bool,X45)) | sP12('F_0',X49)) ) | ($spl176 | $spl352)),
  inference(resolution,[],[f27391,f15187])).
fof(f28780,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP2(X36,X37,X38) | p(s(bool,X35)) | sP12(X39,'F_0')) ) | ($spl116 | $spl352)),
  inference(resolution,[],[f27391,f13741])).
fof(f28778,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP2(X26,X27,X28) | p(s(bool,X25)) | sP12(sK62,X29)) ) | $spl352),
  inference(resolution,[],[f27391,f5864])).
fof(f28776,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP2(X16,X17,X18) | p(s(bool,X15)) | sP12(X19,sK62)) ) | $spl352),
  inference(resolution,[],[f27391,f5611])).
fof(f28775,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP2(X11,X12,X13) | p(s(bool,X14)) | p(s(bool,X10))) ) | $spl352),
  inference(resolution,[],[f27391,f5117])).
fof(f28795,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP2(X6,X7,X8) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(subsumption_resolution,[],[f28774,f1215])).
fof(f28774,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP2(X6,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X9))) ) | $spl352),
  inference(resolution,[],[f27391,f1185])).
fof(f28794,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X1,X2,X3) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(subsumption_resolution,[],[f28773,f1215])).
fof(f28773,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X1,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X4))) ) | $spl352),
  inference(resolution,[],[f27391,f1181])).
fof(f26714,plain,(
  ( ! [X24,X25] : (sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X25,X24)) ) | $spl350),
  inference(cnf_transformation,[],[f26714_D])).
fof(f26714_D,plain,(
  ( ! [X24,X25] : (sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X25,X24)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f28743,plain,(
  ( ! [X39,X37,X38,X40] : (sP12(X37,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X38,X37) | p(s(bool,X39)) | ~p(s(bool,X40))) ) | $spl346),
  inference(resolution,[],[f27193,f1510])).
fof(f27193,plain,(
  ( ! [X383,X379,X380,X382,X381] : (sP12(X379,X380) | sP13(X381,X379) | sP8(X380,X382,X383)) ) | $spl346),
  inference(resolution,[],[f26710,f1112])).
fof(f26713,plain,(
  ( ! [X21,X20] : (sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X21,X20)) ) | $spl348),
  inference(cnf_transformation,[],[f26713_D])).
fof(f26713_D,plain,(
  ( ! [X21,X20] : (sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X21,X20)) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f28742,plain,(
  ( ! [X35,X33,X36,X34] : (sP12(X33,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X34,X33) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | $spl346),
  inference(resolution,[],[f27193,f1339])).
fof(f28732,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP8(X51,X52,X53) | p(s(bool,X50)) | sP12('F_0',X54)) ) | ($spl176 | $spl346)),
  inference(resolution,[],[f27193,f15188])).
fof(f15188,plain,(
  ( ! [X320,X319] : (~sP13(X319,X320) | p(s(bool,X320)) | sP12('F_0',X319)) ) | $spl176),
  inference(subsumption_resolution,[],[f15049,f1215])).
fof(f15049,plain,(
  ( ! [X320,X319] : (sP12('F_0',X319) | p(s(bool,'F_0')) | p(s(bool,X320)) | ~sP13(X319,X320)) ) | $spl176),
  inference(resolution,[],[f14807,f1181])).
fof(f28730,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP8(X41,X42,X43) | p(s(bool,X40)) | sP12(X44,'F_0')) ) | ($spl116 | $spl346)),
  inference(resolution,[],[f27193,f13742])).
fof(f13742,plain,(
  ( ! [X320,X319] : (~sP13(X319,X320) | p(s(bool,X320)) | sP12(X319,'F_0')) ) | $spl116),
  inference(subsumption_resolution,[],[f13605,f1215])).
fof(f13605,plain,(
  ( ! [X320,X319] : (sP12(X319,'F_0') | p(s(bool,'F_0')) | p(s(bool,X320)) | ~sP13(X319,X320)) ) | $spl116),
  inference(resolution,[],[f13492,f1181])).
fof(f28728,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP8(X31,X32,X33) | p(s(bool,X30)) | sP12(sK62,X34)) ) | $spl346),
  inference(resolution,[],[f27193,f5865])).
fof(f5865,plain,(
  ( ! [X320,X319] : (~sP13(X319,X320) | p(s(bool,X320)) | sP12(sK62,X319)) )),
  inference(subsumption_resolution,[],[f5729,f1215])).
fof(f5729,plain,(
  ( ! [X320,X319] : (sP12(sK62,X319) | p(s(bool,'F_0')) | p(s(bool,X320)) | ~sP13(X319,X320)) )),
  inference(resolution,[],[f5231,f1181])).
fof(f28726,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP8(X21,X22,X23) | p(s(bool,X20)) | sP12(X24,sK62)) ) | $spl346),
  inference(resolution,[],[f27193,f5612])).
fof(f5612,plain,(
  ( ! [X320,X319] : (~sP13(X319,X320) | p(s(bool,X320)) | sP12(X319,sK62)) )),
  inference(subsumption_resolution,[],[f5476,f1215])).
fof(f5476,plain,(
  ( ! [X320,X319] : (sP12(X319,sK62) | p(s(bool,'F_0')) | p(s(bool,X320)) | ~sP13(X319,X320)) )),
  inference(resolution,[],[f5230,f1181])).
fof(f28724,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP8(X11,X12,X13) | p(s(bool,X10)) | p(s(bool,X14))) ) | $spl346),
  inference(resolution,[],[f27193,f5117])).
fof(f28745,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X6,X7,X8) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(subsumption_resolution,[],[f28723,f1215])).
fof(f28723,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X6,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(resolution,[],[f27193,f1185])).
fof(f28744,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X1,X2,X3) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(subsumption_resolution,[],[f28722,f1215])).
fof(f28722,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X1,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(resolution,[],[f27193,f1181])).
fof(f28689,plain,(
  ( ! [X30,X31,X29,X32] : (sP12(X29,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X30,X29) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl346),
  inference(resolution,[],[f27192,f1511])).
fof(f27192,plain,(
  ( ! [X377,X374,X376,X378,X375] : (sP12(X374,X375) | sP13(X376,X374) | sP8(X377,X375,X378)) ) | $spl346),
  inference(resolution,[],[f26710,f1111])).
fof(f28688,plain,(
  ( ! [X28,X26,X27,X25] : (sP12(X25,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X26,X25) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl346),
  inference(resolution,[],[f27192,f1340])).
fof(f28680,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP8(X52,X51,X53) | p(s(bool,X50)) | sP12('F_0',X54)) ) | ($spl176 | $spl346)),
  inference(resolution,[],[f27192,f15188])).
fof(f28678,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP8(X42,X41,X43) | p(s(bool,X40)) | sP12(X44,'F_0')) ) | ($spl116 | $spl346)),
  inference(resolution,[],[f27192,f13742])).
fof(f28676,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP8(X32,X31,X33) | p(s(bool,X30)) | sP12(sK62,X34)) ) | $spl346),
  inference(resolution,[],[f27192,f5865])).
fof(f28674,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP8(X22,X21,X23) | p(s(bool,X20)) | sP12(X24,sK62)) ) | $spl346),
  inference(resolution,[],[f27192,f5612])).
fof(f28672,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP8(X12,X11,X13) | p(s(bool,X10)) | p(s(bool,X14))) ) | $spl346),
  inference(resolution,[],[f27192,f5117])).
fof(f28693,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X7,X6,X8) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(subsumption_resolution,[],[f28671,f1215])).
fof(f28671,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X7,X6,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(resolution,[],[f27192,f1185])).
fof(f28692,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X2,X1,X3) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(subsumption_resolution,[],[f28670,f1215])).
fof(f28670,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X2,X1,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(resolution,[],[f27192,f1181])).
fof(f28643,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK35(sK24(X142,bool,X143,X141),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) | $spl402),
  inference(cnf_transformation,[],[f28643_D])).
fof(f28643_D,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK35(sK24(X142,bool,X143,X141),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f28600,plain,(
  ( ! [X142,X145,X143,X141] : (~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK35(sK24(X142,bool,X143,X141),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X141))))) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141)) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141))))) )),
  inference(superposition,[],[f1469,f932])).
fof(f1469,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,happ(s(fun(X71,bool),X72),s(X71,X73)))) | p(s(bool,happ(s(fun(X71,bool),X72),s(X71,sK35(X72,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X71)))))) )),
  inference(resolution,[],[f917,f994])).
fof(f994,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK35(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f791])).
fof(f28642,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK35(sK24(X137,bool,X138,X136),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) | $spl400),
  inference(cnf_transformation,[],[f28642_D])).
fof(f28642_D,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK35(sK24(X137,bool,X138,X136),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f28599,plain,(
  ( ! [X140,X138,X136,X137] : (~p(s(bool,X140)) | p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK35(sK24(X137,bool,X138,X136),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X136))))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136)) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136))))) )),
  inference(superposition,[],[f1469,f931])).
fof(f28641,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK35(sK24(X132,bool,X133,X131),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) | $spl398),
  inference(cnf_transformation,[],[f28641_D])).
fof(f28641_D,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK35(sK24(X132,bool,X133,X131),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f28598,plain,(
  ( ! [X132,X135,X133,X131] : (~p(s(bool,X135)) | p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK35(sK24(X132,bool,X133,X131),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X131))))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131)) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131))))) )),
  inference(superposition,[],[f1469,f930])).
fof(f28640,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK35(sK23(X127,X128,bool,X126),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) | $spl396),
  inference(cnf_transformation,[],[f28640_D])).
fof(f28640_D,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK35(sK23(X127,X128,bool,X126),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f28597,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,X129)) | p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK35(sK23(X127,X128,bool,X126),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X126))))) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126)) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126))))) )),
  inference(superposition,[],[f1469,f929])).
fof(f28639,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK35(sK23(X122,X123,bool,X121),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) | $spl394),
  inference(cnf_transformation,[],[f28639_D])).
fof(f28639_D,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK35(sK23(X122,X123,bool,X121),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f28596,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X124)) | p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK35(sK23(X122,X123,bool,X121),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X121))))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121)) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121))))) )),
  inference(superposition,[],[f1469,f928])).
fof(f28638,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK35(sK23(X117,X118,bool,X116),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) | $spl392),
  inference(cnf_transformation,[],[f28638_D])).
fof(f28638_D,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK35(sK23(X117,X118,bool,X116),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f28595,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X119)) | p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK35(sK23(X117,X118,bool,X116),'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X116))))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116)) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116))))) )),
  inference(superposition,[],[f1469,f927])).
fof(f28612,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK35(X118,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) | $spl390),
  inference(cnf_transformation,[],[f28612_D])).
fof(f28612_D,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK35(X118,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f28449,plain,(
  ( ! [X685,X687,X688,X684,X686] : (p(s(bool,happ(s(fun(X684,bool),X685),s(X684,sK35(X685,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X684))))) | p(s(bool,X686)) | ~p(s(bool,X687)) | s(bool,X687) != s(bool,happ(s(fun(X684,bool),X685),s(X684,X688)))) )),
  inference(resolution,[],[f1469,f1252])).
fof(f28336,plain,(
  ( ! [X158,X163,X161,X159,X162,X160] : (p(s(bool,happ(s(fun(X158,bool),X159),s(X158,sK35(X159,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X158))))) | ~p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X161) != s(bool,happ(s(fun(X158,bool),X159),s(X158,X162))) | s(bool,X160) != s(bool,X163)) )),
  inference(resolution,[],[f1469,f926])).
fof(f28334,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK35(X148,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X147))))) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,happ(s(fun(X147,bool),X148),s(X147,X152)))) )),
  inference(resolution,[],[f1469,f925])).
fof(f28332,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(X135,bool),X136),s(X135,sK35(X136,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X135))))) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,X139) | s(bool,X138) != s(bool,happ(s(fun(X135,bool),X136),s(X135,X140)))) )),
  inference(resolution,[],[f1469,f924])).
fof(f28331,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (p(s(bool,happ(s(fun(X129,bool),X130),s(X129,sK35(X130,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X129))))) | ~p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,happ(s(fun(X129,bool),X130),s(X129,X133))) | s(bool,X131) != s(bool,X134)) )),
  inference(resolution,[],[f1469,f922])).
fof(f28329,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK35(X118,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X117))))) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122)))) )),
  inference(resolution,[],[f1469,f920])).
fof(f28290,plain,(
  ( ! [X30,X31,X29,X32] : (sP12(X29,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X30,X29) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl346),
  inference(resolution,[],[f27179,f1494])).
fof(f27179,plain,(
  ( ! [X313,X310,X312,X309,X311] : (sP12(X309,X310) | sP13(X311,X309) | sP6(X310,X312,X313)) ) | $spl346),
  inference(resolution,[],[f26710,f1095])).
fof(f28289,plain,(
  ( ! [X28,X26,X27,X25] : (sP12(X25,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X26,X25) | p(s(bool,X27)) | p(s(bool,X28))) ) | $spl346),
  inference(resolution,[],[f27179,f1323])).
fof(f28281,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP6(X51,X52,X53) | p(s(bool,X50)) | sP12('F_0',X54)) ) | ($spl176 | $spl346)),
  inference(resolution,[],[f27179,f15188])).
fof(f28279,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP6(X41,X42,X43) | p(s(bool,X40)) | sP12(X44,'F_0')) ) | ($spl116 | $spl346)),
  inference(resolution,[],[f27179,f13742])).
fof(f28277,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP6(X31,X32,X33) | p(s(bool,X30)) | sP12(sK62,X34)) ) | $spl346),
  inference(resolution,[],[f27179,f5865])).
fof(f28275,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP6(X21,X22,X23) | p(s(bool,X20)) | sP12(X24,sK62)) ) | $spl346),
  inference(resolution,[],[f27179,f5612])).
fof(f28273,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP6(X11,X12,X13) | p(s(bool,X10)) | p(s(bool,X14))) ) | $spl346),
  inference(resolution,[],[f27179,f5117])).
fof(f28292,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP6(X6,X7,X8) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(subsumption_resolution,[],[f28272,f1215])).
fof(f28272,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP6(X6,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(resolution,[],[f27179,f1185])).
fof(f28291,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP6(X1,X2,X3) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(subsumption_resolution,[],[f28271,f1215])).
fof(f28271,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP6(X1,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(resolution,[],[f27179,f1181])).
fof(f28242,plain,(
  ( ! [X39,X37,X38,X40] : (sP12(X37,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X38,X37) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | $spl346),
  inference(resolution,[],[f27163,f1480])).
fof(f27163,plain,(
  ( ! [X231,X229,X233,X230,X232] : (sP12(X229,X230) | sP13(X231,X229) | sP4(X230,X232,X233)) ) | $spl346),
  inference(resolution,[],[f26710,f1078])).
fof(f28241,plain,(
  ( ! [X35,X33,X36,X34] : (sP12(X33,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X34,X33) | ~p(s(bool,X35)) | p(s(bool,X36))) ) | $spl346),
  inference(resolution,[],[f27163,f1309])).
fof(f28231,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP4(X51,X52,X53) | p(s(bool,X50)) | sP12('F_0',X54)) ) | ($spl176 | $spl346)),
  inference(resolution,[],[f27163,f15188])).
fof(f28229,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP4(X41,X42,X43) | p(s(bool,X40)) | sP12(X44,'F_0')) ) | ($spl116 | $spl346)),
  inference(resolution,[],[f27163,f13742])).
fof(f28227,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP4(X31,X32,X33) | p(s(bool,X30)) | sP12(sK62,X34)) ) | $spl346),
  inference(resolution,[],[f27163,f5865])).
fof(f28225,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP4(X21,X22,X23) | p(s(bool,X20)) | sP12(X24,sK62)) ) | $spl346),
  inference(resolution,[],[f27163,f5612])).
fof(f28223,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP4(X11,X12,X13) | p(s(bool,X10)) | p(s(bool,X14))) ) | $spl346),
  inference(resolution,[],[f27163,f5117])).
fof(f28244,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X6,X7,X8) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(subsumption_resolution,[],[f28222,f1215])).
fof(f28222,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X6,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(resolution,[],[f27163,f1185])).
fof(f28243,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X1,X2,X3) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(subsumption_resolution,[],[f28221,f1215])).
fof(f28221,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X1,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(resolution,[],[f27163,f1181])).
fof(f28189,plain,(
  ( ! [X30,X31,X29,X32] : (sP12(X29,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X30,X29) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl346),
  inference(resolution,[],[f27162,f1481])).
fof(f27162,plain,(
  ( ! [X227,X225,X228,X226,X224] : (sP12(X224,X225) | sP13(X226,X224) | sP4(X227,X228,X225)) ) | $spl346),
  inference(resolution,[],[f26710,f1076])).
fof(f28188,plain,(
  ( ! [X28,X26,X27,X25] : (sP12(X25,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X26,X25) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl346),
  inference(resolution,[],[f27162,f1310])).
fof(f28180,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP4(X52,X53,X51) | p(s(bool,X50)) | sP12('F_0',X54)) ) | ($spl176 | $spl346)),
  inference(resolution,[],[f27162,f15188])).
fof(f28178,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP4(X42,X43,X41) | p(s(bool,X40)) | sP12(X44,'F_0')) ) | ($spl116 | $spl346)),
  inference(resolution,[],[f27162,f13742])).
fof(f28176,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP4(X32,X33,X31) | p(s(bool,X30)) | sP12(sK62,X34)) ) | $spl346),
  inference(resolution,[],[f27162,f5865])).
fof(f28174,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP4(X22,X23,X21) | p(s(bool,X20)) | sP12(X24,sK62)) ) | $spl346),
  inference(resolution,[],[f27162,f5612])).
fof(f28172,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP4(X12,X13,X11) | p(s(bool,X10)) | p(s(bool,X14))) ) | $spl346),
  inference(resolution,[],[f27162,f5117])).
fof(f28193,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X7,X8,X6) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(subsumption_resolution,[],[f28171,f1215])).
fof(f28171,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X7,X8,X6) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(resolution,[],[f27162,f1185])).
fof(f28192,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X1) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(subsumption_resolution,[],[f28170,f1215])).
fof(f28170,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X1) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(resolution,[],[f27162,f1181])).
fof(f28138,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X33,X32) | p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl346),
  inference(resolution,[],[f27134,f1449])).
fof(f27134,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP12(X88,X89) | sP13(X90,X88) | sP1(X89,X91,X92)) ) | $spl346),
  inference(resolution,[],[f26710,f947])).
fof(f28137,plain,(
  ( ! [X30,X28,X31,X29] : (sP12(X28,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X29,X28) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl346),
  inference(resolution,[],[f27134,f1278])).
fof(f28128,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP1(X51,X52,X53) | p(s(bool,X50)) | sP12('F_0',X54)) ) | ($spl176 | $spl346)),
  inference(resolution,[],[f27134,f15188])).
fof(f28126,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP1(X41,X42,X43) | p(s(bool,X40)) | sP12(X44,'F_0')) ) | ($spl116 | $spl346)),
  inference(resolution,[],[f27134,f13742])).
fof(f28124,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP1(X31,X32,X33) | p(s(bool,X30)) | sP12(sK62,X34)) ) | $spl346),
  inference(resolution,[],[f27134,f5865])).
fof(f28122,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP1(X21,X22,X23) | p(s(bool,X20)) | sP12(X24,sK62)) ) | $spl346),
  inference(resolution,[],[f27134,f5612])).
fof(f28120,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP1(X11,X12,X13) | p(s(bool,X10)) | p(s(bool,X14))) ) | $spl346),
  inference(resolution,[],[f27134,f5117])).
fof(f28140,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X6,X7,X8) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(subsumption_resolution,[],[f28119,f1215])).
fof(f28119,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X6,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(resolution,[],[f27134,f1185])).
fof(f28139,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X1,X2,X3) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(subsumption_resolution,[],[f28118,f1215])).
fof(f28118,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X1,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(resolution,[],[f27134,f1181])).
fof(f28085,plain,(
  ( ! [X26,X24,X27,X25] : (sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X25,X24) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl346),
  inference(resolution,[],[f27133,f1450])).
fof(f27133,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP12(X83,X84) | sP13(X85,X83) | sP1(X86,X84,X87)) ) | $spl346),
  inference(resolution,[],[f26710,f946])).
fof(f28084,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl346),
  inference(resolution,[],[f27133,f1279])).
fof(f28077,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP1(X52,X51,X53) | p(s(bool,X50)) | sP12('F_0',X54)) ) | ($spl176 | $spl346)),
  inference(resolution,[],[f27133,f15188])).
fof(f28075,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP1(X42,X41,X43) | p(s(bool,X40)) | sP12(X44,'F_0')) ) | ($spl116 | $spl346)),
  inference(resolution,[],[f27133,f13742])).
fof(f28073,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP1(X32,X31,X33) | p(s(bool,X30)) | sP12(sK62,X34)) ) | $spl346),
  inference(resolution,[],[f27133,f5865])).
fof(f28071,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP1(X22,X21,X23) | p(s(bool,X20)) | sP12(X24,sK62)) ) | $spl346),
  inference(resolution,[],[f27133,f5612])).
fof(f28069,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP1(X12,X11,X13) | p(s(bool,X10)) | p(s(bool,X14))) ) | $spl346),
  inference(resolution,[],[f27133,f5117])).
fof(f28089,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X7,X6,X8) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(subsumption_resolution,[],[f28068,f1215])).
fof(f28068,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X7,X6,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(resolution,[],[f27133,f1185])).
fof(f28088,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X2,X1,X3) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(subsumption_resolution,[],[f28067,f1215])).
fof(f28067,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X2,X1,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(resolution,[],[f27133,f1181])).
fof(f28040,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X142,bool,X143,X141),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) | $spl388),
  inference(cnf_transformation,[],[f28040_D])).
fof(f28040_D,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X142,bool,X143,X141),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f27997,plain,(
  ( ! [X142,X145,X143,X141] : (~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X142,bool,X143,X141),X141))))) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141)) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141))))) )),
  inference(superposition,[],[f1468,f932])).
fof(f1468,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X70)))) | p(s(bool,happ(s(fun(X68,bool),X69),s(X68,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X69,X68)))))) )),
  inference(resolution,[],[f917,f991])).
fof(f991,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK34(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f787])).
fof(f787,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK33(X2,X1,X0))))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK34(X2,X1,X0))))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK33,sK34])],[f786])).
fof(f786,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f785])).
fof(f785,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f784])).
fof(f784,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f549])).
fof(f549,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f548])).
fof(f548,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f400])).
fof(f400,axiom,(
  ! [X6] : ! [X16,X29] : (? [X12] : (p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X12)))) & p(s(bool,X29))) <=> (? [X12] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X12)))) & p(s(bool,X29))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.LEFT_EXISTS_AND_THM')).
fof(f28039,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X137,bool,X138,X136),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) | $spl386),
  inference(cnf_transformation,[],[f28039_D])).
fof(f28039_D,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X137,bool,X138,X136),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f27996,plain,(
  ( ! [X140,X138,X136,X137] : (~p(s(bool,X140)) | p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X137,bool,X138,X136),X136))))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136)) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136))))) )),
  inference(superposition,[],[f1468,f931])).
fof(f28038,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X132,bool,X133,X131),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) | $spl384),
  inference(cnf_transformation,[],[f28038_D])).
fof(f28038_D,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X132,bool,X133,X131),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f27995,plain,(
  ( ! [X132,X135,X133,X131] : (~p(s(bool,X135)) | p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X132,bool,X133,X131),X131))))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131)) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131))))) )),
  inference(superposition,[],[f1468,f930])).
fof(f28037,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X127,X128,bool,X126),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) | $spl382),
  inference(cnf_transformation,[],[f28037_D])).
fof(f28037_D,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X127,X128,bool,X126),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f27994,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,X129)) | p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X127,X128,bool,X126),X126))))) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126)) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126))))) )),
  inference(superposition,[],[f1468,f929])).
fof(f28036,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X122,X123,bool,X121),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) | $spl380),
  inference(cnf_transformation,[],[f28036_D])).
fof(f28036_D,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X122,X123,bool,X121),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f27993,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X124)) | p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X122,X123,bool,X121),X121))))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121)) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121))))) )),
  inference(superposition,[],[f1468,f928])).
fof(f28035,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X117,X118,bool,X116),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) | $spl378),
  inference(cnf_transformation,[],[f28035_D])).
fof(f28035_D,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X117,X118,bool,X116),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f27992,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X119)) | p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X117,X118,bool,X116),X116))))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116)) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116))))) )),
  inference(superposition,[],[f1468,f927])).
fof(f28009,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X118,X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) | $spl376),
  inference(cnf_transformation,[],[f28009_D])).
fof(f28009_D,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X118,X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f27846,plain,(
  ( ! [X685,X687,X688,X684,X686] : (p(s(bool,happ(s(fun(X684,bool),X685),s(X684,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X685,X684))))) | p(s(bool,X686)) | ~p(s(bool,X687)) | s(bool,X687) != s(bool,happ(s(fun(X684,bool),X685),s(X684,X688)))) )),
  inference(resolution,[],[f1468,f1252])).
fof(f27733,plain,(
  ( ! [X158,X163,X161,X159,X162,X160] : (p(s(bool,happ(s(fun(X158,bool),X159),s(X158,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X159,X158))))) | ~p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X161) != s(bool,happ(s(fun(X158,bool),X159),s(X158,X162))) | s(bool,X160) != s(bool,X163)) )),
  inference(resolution,[],[f1468,f926])).
fof(f27731,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X148,X147))))) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,happ(s(fun(X147,bool),X148),s(X147,X152)))) )),
  inference(resolution,[],[f1468,f925])).
fof(f27729,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(X135,bool),X136),s(X135,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X136,X135))))) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,X139) | s(bool,X138) != s(bool,happ(s(fun(X135,bool),X136),s(X135,X140)))) )),
  inference(resolution,[],[f1468,f924])).
fof(f27728,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (p(s(bool,happ(s(fun(X129,bool),X130),s(X129,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X130,X129))))) | ~p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,happ(s(fun(X129,bool),X130),s(X129,X133))) | s(bool,X131) != s(bool,X134)) )),
  inference(resolution,[],[f1468,f922])).
fof(f27726,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK34('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X118,X117))))) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122)))) )),
  inference(resolution,[],[f1468,f920])).
fof(f27684,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X33,X32) | p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl346),
  inference(resolution,[],[f27131,f1446])).
fof(f27131,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP12(X73,X74) | sP13(X75,X73) | sP2(X74,X76,X77)) ) | $spl346),
  inference(resolution,[],[f26710,f943])).
fof(f27683,plain,(
  ( ! [X30,X28,X31,X29] : (sP12(X28,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X29,X28) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl346),
  inference(resolution,[],[f27131,f1275])).
fof(f27674,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP2(X51,X52,X53) | p(s(bool,X50)) | sP12('F_0',X54)) ) | ($spl176 | $spl346)),
  inference(resolution,[],[f27131,f15188])).
fof(f27672,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP2(X41,X42,X43) | p(s(bool,X40)) | sP12(X44,'F_0')) ) | ($spl116 | $spl346)),
  inference(resolution,[],[f27131,f13742])).
fof(f27670,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP2(X31,X32,X33) | p(s(bool,X30)) | sP12(sK62,X34)) ) | $spl346),
  inference(resolution,[],[f27131,f5865])).
fof(f27668,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP2(X21,X22,X23) | p(s(bool,X20)) | sP12(X24,sK62)) ) | $spl346),
  inference(resolution,[],[f27131,f5612])).
fof(f27666,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP2(X11,X12,X13) | p(s(bool,X10)) | p(s(bool,X14))) ) | $spl346),
  inference(resolution,[],[f27131,f5117])).
fof(f27686,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP2(X6,X7,X8) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(subsumption_resolution,[],[f27665,f1215])).
fof(f27665,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP2(X6,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X5))) ) | $spl346),
  inference(resolution,[],[f27131,f1185])).
fof(f27685,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X1,X2,X3) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(subsumption_resolution,[],[f27664,f1215])).
fof(f27664,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X1,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X0))) ) | $spl346),
  inference(resolution,[],[f27131,f1181])).
fof(f27629,plain,(
  ( ! [X4,X2,X0,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X4) | ~p(s(bool,X4)) | sP13(X0,X2)) ) | $spl374),
  inference(cnf_transformation,[],[f27629_D])).
fof(f27629_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X4) | ~p(s(bool,X4)) | sP13(X0,X2)) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f27498,plain,(
  ( ! [X575,X576,X578,X574,X577] : (sP12(X574,X575) | sP13(X574,X576) | p(s(bool,X577)) | ~p(s(bool,X578)) | s(bool,X575) != s(bool,X578)) ) | $spl352),
  inference(resolution,[],[f26762,f1252])).
fof(f27385,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP12(X42,X43) | sP13(X42,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl352),
  inference(resolution,[],[f26762,f926])).
fof(f27383,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP12(X30,X31) | sP13(X30,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X31) != s(bool,X34)) ) | $spl352),
  inference(resolution,[],[f26762,f925])).
fof(f27381,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP12(X18,X19) | sP13(X18,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl352),
  inference(resolution,[],[f26762,f924])).
fof(f27380,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP12(X12,X13) | sP13(X12,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl352),
  inference(resolution,[],[f26762,f922])).
fof(f27378,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP13(X0,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl352),
  inference(resolution,[],[f26762,f920])).
fof(f27369,plain,(
  ( ! [X4,X2,X0,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X4) | ~p(s(bool,X4)) | sP13(X2,X0)) ) | $spl372),
  inference(cnf_transformation,[],[f27369_D])).
fof(f27369_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X4) | ~p(s(bool,X4)) | sP13(X2,X0)) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f27238,plain,(
  ( ! [X575,X576,X578,X574,X577] : (sP12(X574,X575) | sP13(X576,X574) | p(s(bool,X577)) | ~p(s(bool,X578)) | s(bool,X575) != s(bool,X578)) ) | $spl346),
  inference(resolution,[],[f26710,f1252])).
fof(f27125,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP12(X42,X43) | sP13(X44,X42) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl346),
  inference(resolution,[],[f26710,f926])).
fof(f27123,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP12(X30,X31) | sP13(X32,X30) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X31) != s(bool,X34)) ) | $spl346),
  inference(resolution,[],[f26710,f925])).
fof(f27121,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP12(X18,X19) | sP13(X20,X18) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl346),
  inference(resolution,[],[f26710,f924])).
fof(f27120,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP12(X12,X13) | sP13(X14,X12) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl346),
  inference(resolution,[],[f26710,f922])).
fof(f27118,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP13(X2,X0) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl346),
  inference(resolution,[],[f26710,f920])).
fof(f27117,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X142,bool,X143,X141),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) | $spl370),
  inference(cnf_transformation,[],[f27117_D])).
fof(f27117_D,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X142,bool,X143,X141),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f27074,plain,(
  ( ! [X142,X145,X143,X141] : (~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X142,bool,X143,X141),X141))))) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141)) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141))))) )),
  inference(superposition,[],[f1467,f932])).
fof(f1467,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,happ(s(fun(X65,bool),X66),s(X65,X67)))) | p(s(bool,happ(s(fun(X65,bool),X66),s(X65,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X66,X65)))))) )),
  inference(resolution,[],[f917,f989])).
fof(f989,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK33(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f787])).
fof(f27116,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X137,bool,X138,X136),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) | $spl368),
  inference(cnf_transformation,[],[f27116_D])).
fof(f27116_D,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X137,bool,X138,X136),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f27073,plain,(
  ( ! [X140,X138,X136,X137] : (~p(s(bool,X140)) | p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X137,bool,X138,X136),X136))))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136)) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136))))) )),
  inference(superposition,[],[f1467,f931])).
fof(f27115,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X132,bool,X133,X131),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) | $spl366),
  inference(cnf_transformation,[],[f27115_D])).
fof(f27115_D,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X132,bool,X133,X131),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f27072,plain,(
  ( ! [X132,X135,X133,X131] : (~p(s(bool,X135)) | p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK24(X132,bool,X133,X131),X131))))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131)) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131))))) )),
  inference(superposition,[],[f1467,f930])).
fof(f27114,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X127,X128,bool,X126),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) | $spl364),
  inference(cnf_transformation,[],[f27114_D])).
fof(f27114_D,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X127,X128,bool,X126),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f27071,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,X129)) | p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X127,X128,bool,X126),X126))))) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126)) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126))))) )),
  inference(superposition,[],[f1467,f929])).
fof(f27113,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X122,X123,bool,X121),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) | $spl362),
  inference(cnf_transformation,[],[f27113_D])).
fof(f27113_D,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X122,X123,bool,X121),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f27070,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X124)) | p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X122,X123,bool,X121),X121))))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121)) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121))))) )),
  inference(superposition,[],[f1467,f928])).
fof(f27112,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X117,X118,bool,X116),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) | $spl360),
  inference(cnf_transformation,[],[f27112_D])).
fof(f27112_D,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X117,X118,bool,X116),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f27069,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X119)) | p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK23(X117,X118,bool,X116),X116))))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116)) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116))))) )),
  inference(superposition,[],[f1467,f927])).
fof(f27086,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X118,X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) | $spl358),
  inference(cnf_transformation,[],[f27086_D])).
fof(f27086_D,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X118,X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f26923,plain,(
  ( ! [X685,X687,X688,X684,X686] : (p(s(bool,happ(s(fun(X684,bool),X685),s(X684,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X685,X684))))) | p(s(bool,X686)) | ~p(s(bool,X687)) | s(bool,X687) != s(bool,happ(s(fun(X684,bool),X685),s(X684,X688)))) )),
  inference(resolution,[],[f1467,f1252])).
fof(f26810,plain,(
  ( ! [X158,X163,X161,X159,X162,X160] : (p(s(bool,happ(s(fun(X158,bool),X159),s(X158,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X159,X158))))) | ~p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X161) != s(bool,happ(s(fun(X158,bool),X159),s(X158,X162))) | s(bool,X160) != s(bool,X163)) )),
  inference(resolution,[],[f1467,f926])).
fof(f26808,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X148,X147))))) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,happ(s(fun(X147,bool),X148),s(X147,X152)))) )),
  inference(resolution,[],[f1467,f925])).
fof(f26806,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(X135,bool),X136),s(X135,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X136,X135))))) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,X139) | s(bool,X138) != s(bool,happ(s(fun(X135,bool),X136),s(X135,X140)))) )),
  inference(resolution,[],[f1467,f924])).
fof(f26805,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (p(s(bool,happ(s(fun(X129,bool),X130),s(X129,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X130,X129))))) | ~p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,happ(s(fun(X129,bool),X130),s(X129,X133))) | s(bool,X131) != s(bool,X134)) )),
  inference(resolution,[],[f1467,f922])).
fof(f26803,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK33('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X118,X117))))) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122)))) )),
  inference(resolution,[],[f1467,f920])).
fof(f26762,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP12(X0,X1) | sP13(X0,X2)) ) | $spl352),
  inference(cnf_transformation,[],[f26762_D])).
fof(f26762_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP12(X0,X1) | sP13(X0,X2)) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f26748,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP13(X32,X34) | p(s(bool,X35)) | ~p(s(bool,X33))) )),
  inference(resolution,[],[f26477,f1446])).
fof(f26477,plain,(
  ( ! [X670,X673,X669,X671,X672] : (sP12(X669,X670) | sP2(X671,X670,X672) | sP13(X669,X673)) )),
  inference(resolution,[],[f9317,f1183])).
fof(f1183,plain,(
  ( ! [X0,X1] : (sP13(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f880])).
fof(f9317,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X2,X1,X3) | p(s(bool,X0))) )),
  inference(resolution,[],[f7601,f5228])).
fof(f7601,plain,(
  ( ! [X59,X57,X58,X56] : (sP12(X56,X57) | sP12(X56,sK62) | sP2(X58,X57,X59)) )),
  inference(resolution,[],[f5470,f942])).
fof(f942,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f758])).
fof(f26747,plain,(
  ( ! [X30,X28,X31,X29] : (sP12(X28,X29) | sP13(X28,X30) | p(s(bool,X31)) | ~p(s(bool,X29))) )),
  inference(resolution,[],[f26477,f1275])).
fof(f26746,plain,(
  ( ! [X26,X24,X27,X25] : (sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X24,X25) | p(s(bool,X26)) | ~p(s(bool,X27))) )),
  inference(resolution,[],[f26477,f1447])).
fof(f1447,plain,(
  ( ! [X35,X36] : (~sP2(X35,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X36) | p(s(bool,X36)) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f917,f940])).
fof(f26745,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23))) )),
  inference(resolution,[],[f26477,f1276])).
fof(f1276,plain,(
  ( ! [X39,X40] : (~sP2(X39,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X40) | p(s(bool,X40)) | ~p(s(bool,X39))) )),
  inference(resolution,[],[f916,f940])).
fof(f26741,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP13(X0,X2) | ~p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X4))) )),
  inference(resolution,[],[f26477,f940])).
fof(f26710,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP12(X0,X1) | sP13(X2,X0)) ) | $spl346),
  inference(cnf_transformation,[],[f26710_D])).
fof(f26710_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP12(X0,X1) | sP13(X2,X0)) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f26696,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP13(X34,X32) | p(s(bool,X35)) | ~p(s(bool,X33))) )),
  inference(resolution,[],[f26476,f1446])).
fof(f26476,plain,(
  ( ! [X666,X668,X664,X667,X665] : (sP12(X664,X665) | sP2(X666,X665,X667) | sP13(X668,X664)) )),
  inference(resolution,[],[f9317,f1182])).
fof(f1182,plain,(
  ( ! [X0,X1] : (sP13(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f880])).
fof(f26695,plain,(
  ( ! [X30,X28,X31,X29] : (sP12(X28,X29) | sP13(X30,X28) | p(s(bool,X31)) | ~p(s(bool,X29))) )),
  inference(resolution,[],[f26476,f1275])).
fof(f26694,plain,(
  ( ! [X26,X24,X27,X25] : (sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP13(X25,X24) | p(s(bool,X26)) | ~p(s(bool,X27))) )),
  inference(resolution,[],[f26476,f1447])).
fof(f26693,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP13(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23))) )),
  inference(resolution,[],[f26476,f1276])).
fof(f26689,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP13(X2,X0) | ~p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X4))) )),
  inference(resolution,[],[f26476,f940])).
fof(f26659,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP2(X2,X1,X3)) ) | $spl344),
  inference(cnf_transformation,[],[f26659_D])).
fof(f26659_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X5) | p(s(bool,X5)) | sP2(X2,X1,X3)) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f26485,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP2(X706,X705,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X704) != s(bool,X709)) )),
  inference(resolution,[],[f9317,f1252])).
fof(f26372,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP2(X51,X50,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X49) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9317,f926])).
fof(f26369,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP2(X30,X29,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X28) != s(bool,X34)) )),
  inference(resolution,[],[f9317,f925])).
fof(f26368,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP2(X23,X22,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X21) != s(bool,X26)) )),
  inference(resolution,[],[f9317,f924])).
fof(f26367,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP2(X16,X15,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X14) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9317,f922])).
fof(f26365,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP2(X2,X1,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X0) != s(bool,X5)) )),
  inference(resolution,[],[f9317,f920])).
fof(f26328,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP8(X0,X2,X3)) ) | $spl342),
  inference(cnf_transformation,[],[f26328_D])).
fof(f26328_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP8(X0,X2,X3)) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f26153,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP8(X704,X706,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X705) != s(bool,X709)) )),
  inference(resolution,[],[f9243,f1252])).
fof(f9243,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X0,X2,X3) | p(s(bool,X1))) )),
  inference(resolution,[],[f7402,f5228])).
fof(f7402,plain,(
  ( ! [X304,X306,X305,X307] : (sP12(X304,X305) | sP12(X305,sK62) | sP8(X304,X306,X307)) )),
  inference(resolution,[],[f5469,f1112])).
fof(f5469,plain,(
  ( ! [X306,X305] : (~p(s(bool,X306)) | sP12(X306,X305) | sP12(X305,sK62)) )),
  inference(resolution,[],[f5230,f1167])).
fof(f26040,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP8(X49,X51,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9243,f926])).
fof(f26037,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP8(X28,X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) )),
  inference(resolution,[],[f9243,f925])).
fof(f26036,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP8(X21,X23,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) )),
  inference(resolution,[],[f9243,f924])).
fof(f26035,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP8(X14,X16,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9243,f922])).
fof(f26033,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP8(X0,X2,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) )),
  inference(resolution,[],[f9243,f920])).
fof(f25795,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP8(X2,X0,X3)) ) | $spl340),
  inference(cnf_transformation,[],[f25795_D])).
fof(f25795_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP8(X2,X0,X3)) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f25620,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP8(X706,X704,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X705) != s(bool,X709)) )),
  inference(resolution,[],[f9178,f1252])).
fof(f9178,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X2,X0,X3) | p(s(bool,X1))) )),
  inference(resolution,[],[f7401,f5228])).
fof(f7401,plain,(
  ( ! [X300,X302,X301,X303] : (sP12(X300,X301) | sP12(X301,sK62) | sP8(X302,X300,X303)) )),
  inference(resolution,[],[f5469,f1111])).
fof(f25507,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP8(X51,X49,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9178,f926])).
fof(f25504,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP8(X30,X28,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) )),
  inference(resolution,[],[f9178,f925])).
fof(f25503,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP8(X23,X21,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) )),
  inference(resolution,[],[f9178,f924])).
fof(f25502,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP8(X16,X14,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9178,f922])).
fof(f25500,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP8(X2,X0,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) )),
  inference(resolution,[],[f9178,f920])).
fof(f25439,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK49(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80)))) | ~p(s(bool,X84))) ) | $spl338),
  inference(cnf_transformation,[],[f25439_D])).
fof(f25439_D,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK49(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80)))) | ~p(s(bool,X84))) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f25340,plain,(
  ( ! [X656,X658,X654,X657,X655] : (p(s(bool,happ(s(fun(X654,bool),X655),s(X654,X656)))) | p(s(bool,X657)) | ~p(s(bool,X658)) | s(bool,X658) != s(bool,happ(s(fun(X654,bool),X655),s(X654,sK49(X655,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X654))))) )),
  inference(resolution,[],[f1303,f1252])).
fof(f1303,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,happ(s(fun(X90,bool),X91),s(X90,sK49(X91,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X90))))) | p(s(bool,happ(s(fun(X90,bool),X91),s(X90,X92))))) )),
  inference(resolution,[],[f916,f1020])).
fof(f25227,plain,(
  ( ! [X127,X125,X123,X126,X124,X122] : (p(s(bool,happ(s(fun(X122,bool),X123),s(X122,X124)))) | ~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,happ(s(fun(X122,bool),X123),s(X122,sK49(X123,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X122)))) | s(bool,X125) != s(bool,X127)) )),
  inference(resolution,[],[f1303,f926])).
fof(f25225,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) | p(s(bool,X113)) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X114) != s(bool,happ(s(fun(X110,bool),X111),s(X110,sK49(X111,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X110))))) )),
  inference(resolution,[],[f1303,f925])).
fof(f25223,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(X98,bool),X99),s(X98,X100)))) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,happ(s(fun(X98,bool),X99),s(X98,sK49(X99,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X98))))) )),
  inference(resolution,[],[f1303,f924])).
fof(f25222,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,happ(s(fun(X92,bool),X93),s(X92,X94)))) | ~p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X96) != s(bool,happ(s(fun(X92,bool),X93),s(X92,sK49(X93,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X92)))) | s(bool,X95) != s(bool,X97)) )),
  inference(resolution,[],[f1303,f922])).
fof(f25220,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK49(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80))))) )),
  inference(resolution,[],[f1303,f920])).
fof(f25195,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP6(X0,X2,X3)) ) | $spl336),
  inference(cnf_transformation,[],[f25195_D])).
fof(f25195_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP6(X0,X2,X3)) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f25023,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP6(X704,X706,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X705) != s(bool,X709)) )),
  inference(resolution,[],[f9115,f1252])).
fof(f9115,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP6(X0,X2,X3) | p(s(bool,X1))) )),
  inference(resolution,[],[f7388,f5228])).
fof(f7388,plain,(
  ( ! [X251,X249,X250,X248] : (sP12(X248,X249) | sP12(X249,sK62) | sP6(X248,X250,X251)) )),
  inference(resolution,[],[f5469,f1095])).
fof(f24910,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP6(X49,X51,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9115,f926])).
fof(f24907,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP6(X28,X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) )),
  inference(resolution,[],[f9115,f925])).
fof(f24906,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP6(X21,X23,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) )),
  inference(resolution,[],[f9115,f924])).
fof(f24905,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP6(X14,X16,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9115,f922])).
fof(f24903,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP6(X0,X2,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) )),
  inference(resolution,[],[f9115,f920])).
fof(f24844,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK48(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80)))) | ~p(s(bool,X84))) ) | $spl334),
  inference(cnf_transformation,[],[f24844_D])).
fof(f24844_D,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK48(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80)))) | ~p(s(bool,X84))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f24746,plain,(
  ( ! [X656,X658,X654,X657,X655] : (p(s(bool,happ(s(fun(X654,bool),X655),s(X654,X656)))) | p(s(bool,X657)) | ~p(s(bool,X658)) | s(bool,X658) != s(bool,happ(s(fun(X654,bool),X655),s(X654,sK48(X655,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X654))))) )),
  inference(resolution,[],[f1302,f1252])).
fof(f1302,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,happ(s(fun(X87,bool),X88),s(X87,sK48(X88,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X87))))) | p(s(bool,happ(s(fun(X87,bool),X88),s(X87,X89))))) )),
  inference(resolution,[],[f916,f1018])).
fof(f24633,plain,(
  ( ! [X127,X125,X123,X126,X124,X122] : (p(s(bool,happ(s(fun(X122,bool),X123),s(X122,X124)))) | ~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,happ(s(fun(X122,bool),X123),s(X122,sK48(X123,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X122)))) | s(bool,X125) != s(bool,X127)) )),
  inference(resolution,[],[f1302,f926])).
fof(f24631,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) | p(s(bool,X113)) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X114) != s(bool,happ(s(fun(X110,bool),X111),s(X110,sK48(X111,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X110))))) )),
  inference(resolution,[],[f1302,f925])).
fof(f24629,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(X98,bool),X99),s(X98,X100)))) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,happ(s(fun(X98,bool),X99),s(X98,sK48(X99,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X98))))) )),
  inference(resolution,[],[f1302,f924])).
fof(f24628,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,happ(s(fun(X92,bool),X93),s(X92,X94)))) | ~p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X96) != s(bool,happ(s(fun(X92,bool),X93),s(X92,sK48(X93,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X92)))) | s(bool,X95) != s(bool,X97)) )),
  inference(resolution,[],[f1302,f922])).
fof(f24626,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK48(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80))))) )),
  inference(resolution,[],[f1302,f920])).
fof(f24601,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP4(X0,X2,X3)) ) | $spl332),
  inference(cnf_transformation,[],[f24601_D])).
fof(f24601_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP4(X0,X2,X3)) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f24430,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP4(X704,X706,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X705) != s(bool,X709)) )),
  inference(resolution,[],[f9053,f1252])).
fof(f9053,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X0,X2,X3) | p(s(bool,X1))) )),
  inference(resolution,[],[f7372,f5228])).
fof(f7372,plain,(
  ( ! [X187,X185,X186,X184] : (sP12(X184,X185) | sP12(X185,sK62) | sP4(X184,X186,X187)) )),
  inference(resolution,[],[f5469,f1078])).
fof(f24317,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP4(X49,X51,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f9053,f926])).
fof(f24314,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP4(X28,X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) )),
  inference(resolution,[],[f9053,f925])).
fof(f24313,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP4(X21,X23,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) )),
  inference(resolution,[],[f9053,f924])).
fof(f24312,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP4(X14,X16,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f9053,f922])).
fof(f24310,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP4(X0,X2,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) )),
  inference(resolution,[],[f9053,f920])).
fof(f24249,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK38(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80)))) | ~p(s(bool,X84))) ) | $spl330),
  inference(cnf_transformation,[],[f24249_D])).
fof(f24249_D,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK38(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80)))) | ~p(s(bool,X84))) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f24152,plain,(
  ( ! [X656,X658,X654,X657,X655] : (p(s(bool,happ(s(fun(X654,bool),X655),s(X654,X656)))) | p(s(bool,X657)) | ~p(s(bool,X658)) | s(bool,X658) != s(bool,happ(s(fun(X654,bool),X655),s(X654,sK38(X655,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X654))))) )),
  inference(resolution,[],[f1301,f1252])).
fof(f1301,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,happ(s(fun(X84,bool),X85),s(X84,sK38(X85,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X84))))) | p(s(bool,happ(s(fun(X84,bool),X85),s(X84,X86))))) )),
  inference(resolution,[],[f916,f1000])).
fof(f24039,plain,(
  ( ! [X127,X125,X123,X126,X124,X122] : (p(s(bool,happ(s(fun(X122,bool),X123),s(X122,X124)))) | ~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,happ(s(fun(X122,bool),X123),s(X122,sK38(X123,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X122)))) | s(bool,X125) != s(bool,X127)) )),
  inference(resolution,[],[f1301,f926])).
fof(f24037,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) | p(s(bool,X113)) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X114) != s(bool,happ(s(fun(X110,bool),X111),s(X110,sK38(X111,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X110))))) )),
  inference(resolution,[],[f1301,f925])).
fof(f24035,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(X98,bool),X99),s(X98,X100)))) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,happ(s(fun(X98,bool),X99),s(X98,sK38(X99,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X98))))) )),
  inference(resolution,[],[f1301,f924])).
fof(f24034,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,happ(s(fun(X92,bool),X93),s(X92,X94)))) | ~p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X96) != s(bool,happ(s(fun(X92,bool),X93),s(X92,sK38(X93,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X92)))) | s(bool,X95) != s(bool,X97)) )),
  inference(resolution,[],[f1301,f922])).
fof(f24032,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK38(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80))))) )),
  inference(resolution,[],[f1301,f920])).
fof(f24007,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP4(X2,X3,X0)) ) | $spl328),
  inference(cnf_transformation,[],[f24007_D])).
fof(f24007_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP4(X2,X3,X0)) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f23836,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP4(X706,X707,X704) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X705) != s(bool,X709)) )),
  inference(resolution,[],[f8989,f1252])).
fof(f8989,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X0) | p(s(bool,X1))) )),
  inference(resolution,[],[f7371,f5228])).
fof(f7371,plain,(
  ( ! [X182,X180,X183,X181] : (sP12(X180,X181) | sP12(X181,sK62) | sP4(X182,X183,X180)) )),
  inference(resolution,[],[f5469,f1076])).
fof(f23723,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP4(X51,X52,X49) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f8989,f926])).
fof(f23720,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP4(X30,X31,X28) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) )),
  inference(resolution,[],[f8989,f925])).
fof(f23719,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP4(X23,X24,X21) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) )),
  inference(resolution,[],[f8989,f924])).
fof(f23718,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP4(X16,X17,X14) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f8989,f922])).
fof(f23716,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X0) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) )),
  inference(resolution,[],[f8989,f920])).
fof(f23655,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK37(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80)))) | ~p(s(bool,X84))) ) | $spl326),
  inference(cnf_transformation,[],[f23655_D])).
fof(f23655_D,plain,(
  ( ! [X80,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK37(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80)))) | ~p(s(bool,X84))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f23559,plain,(
  ( ! [X656,X658,X654,X657,X655] : (p(s(bool,happ(s(fun(X654,bool),X655),s(X654,X656)))) | p(s(bool,X657)) | ~p(s(bool,X658)) | s(bool,X658) != s(bool,happ(s(fun(X654,bool),X655),s(X654,sK37(X655,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X654))))) )),
  inference(resolution,[],[f1300,f1252])).
fof(f1300,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,happ(s(fun(X81,bool),X82),s(X81,sK37(X82,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X81))))) | p(s(bool,happ(s(fun(X81,bool),X82),s(X81,X83))))) )),
  inference(resolution,[],[f916,f998])).
fof(f23446,plain,(
  ( ! [X127,X125,X123,X126,X124,X122] : (p(s(bool,happ(s(fun(X122,bool),X123),s(X122,X124)))) | ~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X126) != s(bool,happ(s(fun(X122,bool),X123),s(X122,sK37(X123,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X122)))) | s(bool,X125) != s(bool,X127)) )),
  inference(resolution,[],[f1300,f926])).
fof(f23444,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (p(s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) | p(s(bool,X113)) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | s(bool,X114) != s(bool,happ(s(fun(X110,bool),X111),s(X110,sK37(X111,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X110))))) )),
  inference(resolution,[],[f1300,f925])).
fof(f23442,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (p(s(bool,happ(s(fun(X98,bool),X99),s(X98,X100)))) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,happ(s(fun(X98,bool),X99),s(X98,sK37(X99,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X98))))) )),
  inference(resolution,[],[f1300,f924])).
fof(f23441,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,happ(s(fun(X92,bool),X93),s(X92,X94)))) | ~p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X96) != s(bool,happ(s(fun(X92,bool),X93),s(X92,sK37(X93,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X92)))) | s(bool,X95) != s(bool,X97)) )),
  inference(resolution,[],[f1300,f922])).
fof(f23439,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (p(s(bool,happ(s(fun(X80,bool),X81),s(X80,X82)))) | p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,X85) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,sK37(X81,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X80))))) )),
  inference(resolution,[],[f1300,f920])).
fof(f23414,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP1(X0,X2,X3)) ) | $spl324),
  inference(cnf_transformation,[],[f23414_D])).
fof(f23414_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP1(X0,X2,X3)) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f23242,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP1(X704,X706,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X705) != s(bool,X709)) )),
  inference(resolution,[],[f8882,f1252])).
fof(f8882,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X0,X2,X3) | p(s(bool,X1))) )),
  inference(resolution,[],[f7343,f5228])).
fof(f7343,plain,(
  ( ! [X74,X72,X75,X73] : (sP12(X72,X73) | sP12(X73,sK62) | sP1(X72,X74,X75)) )),
  inference(resolution,[],[f5469,f947])).
fof(f23129,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP1(X49,X51,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f8882,f926])).
fof(f23126,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP1(X28,X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) )),
  inference(resolution,[],[f8882,f925])).
fof(f23125,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP1(X21,X23,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) )),
  inference(resolution,[],[f8882,f924])).
fof(f23124,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP1(X14,X16,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f8882,f922])).
fof(f23122,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP1(X0,X2,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) )),
  inference(resolution,[],[f8882,f920])).
fof(f23412,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) )),
  inference(subsumption_resolution,[],[f23116,f1165])).
fof(f23116,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,X9) | p(s(bool,X9)) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) )),
  inference(resolution,[],[f8882,f938])).
fof(f23087,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK36(sK24(X142,bool,X143,X141),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) | $spl322),
  inference(cnf_transformation,[],[f23087_D])).
fof(f23087_D,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK36(sK24(X142,bool,X143,X141),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f23044,plain,(
  ( ! [X142,X145,X143,X141] : (~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK36(sK24(X142,bool,X143,X141),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X141))))) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141)) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141))))) )),
  inference(superposition,[],[f1299,f932])).
fof(f1299,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,happ(s(fun(X78,bool),X79),s(X78,X80)))) | p(s(bool,happ(s(fun(X78,bool),X79),s(X78,sK36(X79,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X78)))))) )),
  inference(resolution,[],[f916,f996])).
fof(f23086,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK36(sK24(X137,bool,X138,X136),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) | $spl320),
  inference(cnf_transformation,[],[f23086_D])).
fof(f23086_D,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK36(sK24(X137,bool,X138,X136),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f23043,plain,(
  ( ! [X140,X138,X136,X137] : (~p(s(bool,X140)) | p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK36(sK24(X137,bool,X138,X136),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X136))))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136)) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136))))) )),
  inference(superposition,[],[f1299,f931])).
fof(f23085,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK36(sK24(X132,bool,X133,X131),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) | $spl318),
  inference(cnf_transformation,[],[f23085_D])).
fof(f23085_D,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK36(sK24(X132,bool,X133,X131),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f23042,plain,(
  ( ! [X132,X135,X133,X131] : (~p(s(bool,X135)) | p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK36(sK24(X132,bool,X133,X131),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X131))))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131)) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131))))) )),
  inference(superposition,[],[f1299,f930])).
fof(f23084,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK36(sK23(X127,X128,bool,X126),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) | $spl316),
  inference(cnf_transformation,[],[f23084_D])).
fof(f23084_D,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK36(sK23(X127,X128,bool,X126),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f23041,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,X129)) | p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK36(sK23(X127,X128,bool,X126),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X126))))) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126)) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126))))) )),
  inference(superposition,[],[f1299,f929])).
fof(f23083,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK36(sK23(X122,X123,bool,X121),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) | $spl314),
  inference(cnf_transformation,[],[f23083_D])).
fof(f23083_D,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK36(sK23(X122,X123,bool,X121),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f23040,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X124)) | p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK36(sK23(X122,X123,bool,X121),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X121))))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121)) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121))))) )),
  inference(superposition,[],[f1299,f928])).
fof(f23082,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK36(sK23(X117,X118,bool,X116),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) | $spl312),
  inference(cnf_transformation,[],[f23082_D])).
fof(f23082_D,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK36(sK23(X117,X118,bool,X116),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f23039,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X119)) | p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK36(sK23(X117,X118,bool,X116),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X116))))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116)) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116))))) )),
  inference(superposition,[],[f1299,f927])).
fof(f23056,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK36(X118,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) | $spl310),
  inference(cnf_transformation,[],[f23056_D])).
fof(f23056_D,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK36(X118,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f22901,plain,(
  ( ! [X685,X687,X688,X684,X686] : (p(s(bool,happ(s(fun(X684,bool),X685),s(X684,sK36(X685,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X684))))) | p(s(bool,X686)) | ~p(s(bool,X687)) | s(bool,X687) != s(bool,happ(s(fun(X684,bool),X685),s(X684,X688)))) )),
  inference(resolution,[],[f1299,f1252])).
fof(f22788,plain,(
  ( ! [X158,X163,X161,X159,X162,X160] : (p(s(bool,happ(s(fun(X158,bool),X159),s(X158,sK36(X159,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X158))))) | ~p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X161) != s(bool,happ(s(fun(X158,bool),X159),s(X158,X162))) | s(bool,X160) != s(bool,X163)) )),
  inference(resolution,[],[f1299,f926])).
fof(f22786,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK36(X148,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X147))))) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,happ(s(fun(X147,bool),X148),s(X147,X152)))) )),
  inference(resolution,[],[f1299,f925])).
fof(f22784,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(X135,bool),X136),s(X135,sK36(X136,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X135))))) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,X139) | s(bool,X138) != s(bool,happ(s(fun(X135,bool),X136),s(X135,X140)))) )),
  inference(resolution,[],[f1299,f924])).
fof(f22783,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (p(s(bool,happ(s(fun(X129,bool),X130),s(X129,sK36(X130,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X129))))) | ~p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,happ(s(fun(X129,bool),X130),s(X129,X133))) | s(bool,X131) != s(bool,X134)) )),
  inference(resolution,[],[f1299,f922])).
fof(f22781,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK36(X118,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X117))))) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122)))) )),
  inference(resolution,[],[f1299,f920])).
fof(f22742,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP1(X2,X0,X3)) ) | $spl308),
  inference(cnf_transformation,[],[f22742_D])).
fof(f22742_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP1(X2,X0,X3)) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f22572,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP1(X706,X704,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X705) != s(bool,X709)) )),
  inference(resolution,[],[f8818,f1252])).
fof(f8818,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X2,X0,X3) | p(s(bool,X1))) )),
  inference(resolution,[],[f7342,f5228])).
fof(f7342,plain,(
  ( ! [X70,X68,X71,X69] : (sP12(X68,X69) | sP12(X69,sK62) | sP1(X70,X68,X71)) )),
  inference(resolution,[],[f5469,f946])).
fof(f22459,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP1(X51,X49,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f8818,f926])).
fof(f22456,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP1(X30,X28,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) )),
  inference(resolution,[],[f8818,f925])).
fof(f22455,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP1(X23,X21,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) )),
  inference(resolution,[],[f8818,f924])).
fof(f22454,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP1(X16,X14,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f8818,f922])).
fof(f22452,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP1(X2,X0,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) )),
  inference(resolution,[],[f8818,f920])).
fof(f22417,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK35(sK24(X142,bool,X143,X141),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) | $spl306),
  inference(cnf_transformation,[],[f22417_D])).
fof(f22417_D,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK35(sK24(X142,bool,X143,X141),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f22374,plain,(
  ( ! [X142,X145,X143,X141] : (~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK35(sK24(X142,bool,X143,X141),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X141))))) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141)) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141))))) )),
  inference(superposition,[],[f1298,f932])).
fof(f1298,plain,(
  ( ! [X76,X77,X75] : (~p(s(bool,happ(s(fun(X75,bool),X76),s(X75,X77)))) | p(s(bool,happ(s(fun(X75,bool),X76),s(X75,sK35(X76,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X75)))))) )),
  inference(resolution,[],[f916,f994])).
fof(f22416,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK35(sK24(X137,bool,X138,X136),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) | $spl304),
  inference(cnf_transformation,[],[f22416_D])).
fof(f22416_D,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK35(sK24(X137,bool,X138,X136),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f22373,plain,(
  ( ! [X140,X138,X136,X137] : (~p(s(bool,X140)) | p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK35(sK24(X137,bool,X138,X136),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X136))))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136)) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136))))) )),
  inference(superposition,[],[f1298,f931])).
fof(f22415,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK35(sK24(X132,bool,X133,X131),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) | $spl302),
  inference(cnf_transformation,[],[f22415_D])).
fof(f22415_D,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK35(sK24(X132,bool,X133,X131),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f22372,plain,(
  ( ! [X132,X135,X133,X131] : (~p(s(bool,X135)) | p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK35(sK24(X132,bool,X133,X131),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X131))))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131)) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131))))) )),
  inference(superposition,[],[f1298,f930])).
fof(f22414,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK35(sK23(X127,X128,bool,X126),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) | $spl300),
  inference(cnf_transformation,[],[f22414_D])).
fof(f22414_D,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK35(sK23(X127,X128,bool,X126),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f22371,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,X129)) | p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK35(sK23(X127,X128,bool,X126),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X126))))) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126)) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126))))) )),
  inference(superposition,[],[f1298,f929])).
fof(f22413,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK35(sK23(X122,X123,bool,X121),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) | $spl298),
  inference(cnf_transformation,[],[f22413_D])).
fof(f22413_D,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK35(sK23(X122,X123,bool,X121),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f22370,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X124)) | p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK35(sK23(X122,X123,bool,X121),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X121))))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121)) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121))))) )),
  inference(superposition,[],[f1298,f928])).
fof(f22412,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK35(sK23(X117,X118,bool,X116),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) | $spl296),
  inference(cnf_transformation,[],[f22412_D])).
fof(f22412_D,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK35(sK23(X117,X118,bool,X116),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f22369,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X119)) | p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK35(sK23(X117,X118,bool,X116),'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X116))))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116)) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116))))) )),
  inference(superposition,[],[f1298,f927])).
fof(f22386,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK35(X118,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) | $spl294),
  inference(cnf_transformation,[],[f22386_D])).
fof(f22386_D,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK35(X118,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f22232,plain,(
  ( ! [X685,X687,X688,X684,X686] : (p(s(bool,happ(s(fun(X684,bool),X685),s(X684,sK35(X685,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X684))))) | p(s(bool,X686)) | ~p(s(bool,X687)) | s(bool,X687) != s(bool,happ(s(fun(X684,bool),X685),s(X684,X688)))) )),
  inference(resolution,[],[f1298,f1252])).
fof(f22119,plain,(
  ( ! [X158,X163,X161,X159,X162,X160] : (p(s(bool,happ(s(fun(X158,bool),X159),s(X158,sK35(X159,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X158))))) | ~p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X161) != s(bool,happ(s(fun(X158,bool),X159),s(X158,X162))) | s(bool,X160) != s(bool,X163)) )),
  inference(resolution,[],[f1298,f926])).
fof(f22117,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK35(X148,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X147))))) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,happ(s(fun(X147,bool),X148),s(X147,X152)))) )),
  inference(resolution,[],[f1298,f925])).
fof(f22115,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(X135,bool),X136),s(X135,sK35(X136,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X135))))) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,X139) | s(bool,X138) != s(bool,happ(s(fun(X135,bool),X136),s(X135,X140)))) )),
  inference(resolution,[],[f1298,f924])).
fof(f22114,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (p(s(bool,happ(s(fun(X129,bool),X130),s(X129,sK35(X130,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X129))))) | ~p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,happ(s(fun(X129,bool),X130),s(X129,X133))) | s(bool,X131) != s(bool,X134)) )),
  inference(resolution,[],[f1298,f922])).
fof(f22112,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK35(X118,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X117))))) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122)))) )),
  inference(resolution,[],[f1298,f920])).
fof(f22073,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP2(X0,X2,X3)) ) | $spl292),
  inference(cnf_transformation,[],[f22073_D])).
fof(f22073_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP2(X0,X2,X3)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f21903,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP2(X704,X706,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X705) != s(bool,X709)) )),
  inference(resolution,[],[f8752,f1252])).
fof(f8752,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X0,X2,X3) | p(s(bool,X1))) )),
  inference(resolution,[],[f7340,f5228])).
fof(f7340,plain,(
  ( ! [X61,X62,X60,X63] : (sP12(X60,X61) | sP12(X61,sK62) | sP2(X60,X62,X63)) )),
  inference(resolution,[],[f5469,f943])).
fof(f21790,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP2(X49,X51,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f8752,f926])).
fof(f21787,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP2(X28,X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) )),
  inference(resolution,[],[f8752,f925])).
fof(f21786,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP2(X21,X23,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) )),
  inference(resolution,[],[f8752,f924])).
fof(f21785,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP2(X14,X16,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f8752,f922])).
fof(f21783,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP2(X0,X2,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) )),
  inference(resolution,[],[f8752,f920])).
fof(f22072,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) )),
  inference(subsumption_resolution,[],[f21778,f1165])).
fof(f21778,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,X13) | p(s(bool,X13)) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) )),
  inference(resolution,[],[f8752,f939])).
fof(f21748,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X142,bool,X143,X141),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) | $spl290),
  inference(cnf_transformation,[],[f21748_D])).
fof(f21748_D,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X142,bool,X143,X141),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f21705,plain,(
  ( ! [X142,X145,X143,X141] : (~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X142,bool,X143,X141),X141))))) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141)) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141))))) )),
  inference(superposition,[],[f1297,f932])).
fof(f1297,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,happ(s(fun(X72,bool),X73),s(X72,X74)))) | p(s(bool,happ(s(fun(X72,bool),X73),s(X72,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X73,X72)))))) )),
  inference(resolution,[],[f916,f991])).
fof(f21747,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X137,bool,X138,X136),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) | $spl288),
  inference(cnf_transformation,[],[f21747_D])).
fof(f21747_D,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X137,bool,X138,X136),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f21704,plain,(
  ( ! [X140,X138,X136,X137] : (~p(s(bool,X140)) | p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X137,bool,X138,X136),X136))))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136)) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136))))) )),
  inference(superposition,[],[f1297,f931])).
fof(f21746,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X132,bool,X133,X131),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) | $spl286),
  inference(cnf_transformation,[],[f21746_D])).
fof(f21746_D,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X132,bool,X133,X131),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f21703,plain,(
  ( ! [X132,X135,X133,X131] : (~p(s(bool,X135)) | p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X132,bool,X133,X131),X131))))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131)) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131))))) )),
  inference(superposition,[],[f1297,f930])).
fof(f21745,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X127,X128,bool,X126),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) | $spl284),
  inference(cnf_transformation,[],[f21745_D])).
fof(f21745_D,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X127,X128,bool,X126),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f21702,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,X129)) | p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X127,X128,bool,X126),X126))))) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126)) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126))))) )),
  inference(superposition,[],[f1297,f929])).
fof(f21744,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X122,X123,bool,X121),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) | $spl282),
  inference(cnf_transformation,[],[f21744_D])).
fof(f21744_D,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X122,X123,bool,X121),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f21701,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X124)) | p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X122,X123,bool,X121),X121))))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121)) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121))))) )),
  inference(superposition,[],[f1297,f928])).
fof(f21743,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X117,X118,bool,X116),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) | $spl280),
  inference(cnf_transformation,[],[f21743_D])).
fof(f21743_D,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X117,X118,bool,X116),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f21700,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X119)) | p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X117,X118,bool,X116),X116))))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116)) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116))))) )),
  inference(superposition,[],[f1297,f927])).
fof(f21717,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X118,X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) | $spl278),
  inference(cnf_transformation,[],[f21717_D])).
fof(f21717_D,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X118,X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f21564,plain,(
  ( ! [X685,X687,X688,X684,X686] : (p(s(bool,happ(s(fun(X684,bool),X685),s(X684,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X685,X684))))) | p(s(bool,X686)) | ~p(s(bool,X687)) | s(bool,X687) != s(bool,happ(s(fun(X684,bool),X685),s(X684,X688)))) )),
  inference(resolution,[],[f1297,f1252])).
fof(f21451,plain,(
  ( ! [X158,X163,X161,X159,X162,X160] : (p(s(bool,happ(s(fun(X158,bool),X159),s(X158,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X159,X158))))) | ~p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X161) != s(bool,happ(s(fun(X158,bool),X159),s(X158,X162))) | s(bool,X160) != s(bool,X163)) )),
  inference(resolution,[],[f1297,f926])).
fof(f21449,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X148,X147))))) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,happ(s(fun(X147,bool),X148),s(X147,X152)))) )),
  inference(resolution,[],[f1297,f925])).
fof(f21447,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(X135,bool),X136),s(X135,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X136,X135))))) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,X139) | s(bool,X138) != s(bool,happ(s(fun(X135,bool),X136),s(X135,X140)))) )),
  inference(resolution,[],[f1297,f924])).
fof(f21446,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (p(s(bool,happ(s(fun(X129,bool),X130),s(X129,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X130,X129))))) | ~p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,happ(s(fun(X129,bool),X130),s(X129,X133))) | s(bool,X131) != s(bool,X134)) )),
  inference(resolution,[],[f1297,f922])).
fof(f21444,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK34('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X118,X117))))) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122)))) )),
  inference(resolution,[],[f1297,f920])).
fof(f21407,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X142,bool,X143,X141),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) | $spl276),
  inference(cnf_transformation,[],[f21407_D])).
fof(f21407_D,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X142,bool,X143,X141),X141))))) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141)))) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f21364,plain,(
  ( ! [X142,X145,X143,X141] : (~p(s(bool,X145)) | p(s(bool,happ(s(fun(X141,bool),sK24(X142,bool,X143,X141)),s(X141,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X142,bool,X143,X141),X141))))) | s(bool,sK20(X142,bool,X143,X141)) != s(bool,sK22(X142,bool,X143,X141)) | s(X143,sK19(X142,bool,X143,X141)) != s(X143,sK21(X142,bool,X143,X141)) | s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK19(X142,bool,X143,X141)))),s(bool,sK20(X142,bool,X143,X141)))) != s(X141,happ(s(fun(bool,X141),happ(s(fun(X143,fun(bool,X141)),X142),s(X143,sK21(X142,bool,X143,X141)))),s(bool,sK22(X142,bool,X143,X141))))) )),
  inference(superposition,[],[f1296,f932])).
fof(f1296,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,happ(s(fun(X69,bool),X70),s(X69,X71)))) | p(s(bool,happ(s(fun(X69,bool),X70),s(X69,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X70,X69)))))) )),
  inference(resolution,[],[f916,f989])).
fof(f21406,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X137,bool,X138,X136),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) | $spl274),
  inference(cnf_transformation,[],[f21406_D])).
fof(f21406_D,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X137,bool,X138,X136),X136))))) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136)))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f21363,plain,(
  ( ! [X140,X138,X136,X137] : (~p(s(bool,X140)) | p(s(bool,happ(s(fun(X136,bool),sK24(X137,bool,X138,X136)),s(X136,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X137,bool,X138,X136),X136))))) | s(bool,sK20(X137,bool,X138,X136)) = s(bool,sK22(X137,bool,X138,X136)) | s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK19(X137,bool,X138,X136)))),s(bool,sK20(X137,bool,X138,X136)))) = s(X136,happ(s(fun(bool,X136),happ(s(fun(X138,fun(bool,X136)),X137),s(X138,sK21(X137,bool,X138,X136)))),s(bool,sK22(X137,bool,X138,X136))))) )),
  inference(superposition,[],[f1296,f931])).
fof(f21405,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X132,bool,X133,X131),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) | $spl272),
  inference(cnf_transformation,[],[f21405_D])).
fof(f21405_D,plain,(
  ( ! [X132,X133,X131] : (p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X132,bool,X133,X131),X131))))) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131)))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f21362,plain,(
  ( ! [X132,X135,X133,X131] : (~p(s(bool,X135)) | p(s(bool,happ(s(fun(X131,bool),sK24(X132,bool,X133,X131)),s(X131,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK24(X132,bool,X133,X131),X131))))) | s(X133,sK19(X132,bool,X133,X131)) = s(X133,sK21(X132,bool,X133,X131)) | s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK19(X132,bool,X133,X131)))),s(bool,sK20(X132,bool,X133,X131)))) = s(X131,happ(s(fun(bool,X131),happ(s(fun(X133,fun(bool,X131)),X132),s(X133,sK21(X132,bool,X133,X131)))),s(bool,sK22(X132,bool,X133,X131))))) )),
  inference(superposition,[],[f1296,f930])).
fof(f21404,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X127,X128,bool,X126),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) | $spl270),
  inference(cnf_transformation,[],[f21404_D])).
fof(f21404_D,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X127,X128,bool,X126),X126))))) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126)))) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f21361,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,X129)) | p(s(bool,happ(s(fun(X126,bool),sK23(X127,X128,bool,X126)),s(X126,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X127,X128,bool,X126),X126))))) | s(X128,sK20(X127,X128,bool,X126)) != s(X128,sK22(X127,X128,bool,X126)) | s(bool,sK19(X127,X128,bool,X126)) != s(bool,sK21(X127,X128,bool,X126)) | s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK19(X127,X128,bool,X126)))),s(X128,sK20(X127,X128,bool,X126)))) != s(X126,happ(s(fun(X128,X126),happ(s(fun(bool,fun(X128,X126)),X127),s(bool,sK21(X127,X128,bool,X126)))),s(X128,sK22(X127,X128,bool,X126))))) )),
  inference(superposition,[],[f1296,f929])).
fof(f21403,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X122,X123,bool,X121),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) | $spl268),
  inference(cnf_transformation,[],[f21403_D])).
fof(f21403_D,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X122,X123,bool,X121),X121))))) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121)))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f21360,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X124)) | p(s(bool,happ(s(fun(X121,bool),sK23(X122,X123,bool,X121)),s(X121,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X122,X123,bool,X121),X121))))) | s(X123,sK20(X122,X123,bool,X121)) = s(X123,sK22(X122,X123,bool,X121)) | s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK19(X122,X123,bool,X121)))),s(X123,sK20(X122,X123,bool,X121)))) = s(X121,happ(s(fun(X123,X121),happ(s(fun(bool,fun(X123,X121)),X122),s(bool,sK21(X122,X123,bool,X121)))),s(X123,sK22(X122,X123,bool,X121))))) )),
  inference(superposition,[],[f1296,f928])).
fof(f21402,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X117,X118,bool,X116),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) | $spl266),
  inference(cnf_transformation,[],[f21402_D])).
fof(f21402_D,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X117,X118,bool,X116),X116))))) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116)))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f21359,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,X119)) | p(s(bool,happ(s(fun(X116,bool),sK23(X117,X118,bool,X116)),s(X116,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK23(X117,X118,bool,X116),X116))))) | s(bool,sK19(X117,X118,bool,X116)) = s(bool,sK21(X117,X118,bool,X116)) | s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK19(X117,X118,bool,X116)))),s(X118,sK20(X117,X118,bool,X116)))) = s(X116,happ(s(fun(X118,X116),happ(s(fun(bool,fun(X118,X116)),X117),s(bool,sK21(X117,X118,bool,X116)))),s(X118,sK22(X117,X118,bool,X116))))) )),
  inference(superposition,[],[f1296,f927])).
fof(f21376,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X118,X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) | $spl264),
  inference(cnf_transformation,[],[f21376_D])).
fof(f21376_D,plain,(
  ( ! [X118,X122,X120,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X118,X117))))) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122))) | ~p(s(bool,X120))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f21223,plain,(
  ( ! [X685,X687,X688,X684,X686] : (p(s(bool,happ(s(fun(X684,bool),X685),s(X684,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X685,X684))))) | p(s(bool,X686)) | ~p(s(bool,X687)) | s(bool,X687) != s(bool,happ(s(fun(X684,bool),X685),s(X684,X688)))) )),
  inference(resolution,[],[f1296,f1252])).
fof(f21110,plain,(
  ( ! [X158,X163,X161,X159,X162,X160] : (p(s(bool,happ(s(fun(X158,bool),X159),s(X158,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X159,X158))))) | ~p(s(bool,X160)) | ~p(s(bool,X161)) | s(bool,X161) != s(bool,happ(s(fun(X158,bool),X159),s(X158,X162))) | s(bool,X160) != s(bool,X163)) )),
  inference(resolution,[],[f1296,f926])).
fof(f21108,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X148,X147))))) | p(s(bool,X149)) | ~p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,happ(s(fun(X147,bool),X148),s(X147,X152)))) )),
  inference(resolution,[],[f1296,f925])).
fof(f21106,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(X135,bool),X136),s(X135,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X136,X135))))) | p(s(bool,X137)) | ~p(s(bool,X138)) | s(bool,X137) != s(bool,X139) | s(bool,X138) != s(bool,happ(s(fun(X135,bool),X136),s(X135,X140)))) )),
  inference(resolution,[],[f1296,f924])).
fof(f21105,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (p(s(bool,happ(s(fun(X129,bool),X130),s(X129,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X130,X129))))) | ~p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,happ(s(fun(X129,bool),X130),s(X129,X133))) | s(bool,X131) != s(bool,X134)) )),
  inference(resolution,[],[f1296,f922])).
fof(f21103,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,sK33('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X118,X117))))) | p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(X117,bool),X118),s(X117,X122)))) )),
  inference(resolution,[],[f1296,f920])).
fof(f7975,plain,(
  sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | $spl60),
  inference(cnf_transformation,[],[f7975_D])).
fof(f7975_D,plain,(
  sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f21060,plain,(
  ( ! [X18] : (~p(s(bool,X18)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(subsumption_resolution,[],[f21059,f5695])).
fof(f5695,plain,(
  ( ! [X227,X228,X226] : (sP8(X227,X226,X228) | sP12(sK62,X226)) )),
  inference(resolution,[],[f5231,f1111])).
fof(f21059,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | ~sP8(X18,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X19) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(subsumption_resolution,[],[f21039,f1511])).
fof(f21039,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | ~sP8(X18,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X19) | ~p(s(bool,X19)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1507,f5690])).
fof(f5690,plain,(
  ( ! [X212,X213,X211] : (~sP9(X212,X213,X211) | ~p(s(bool,X212)) | sP12(sK62,X211)) )),
  inference(resolution,[],[f5231,f1103])).
fof(f1103,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f848])).
fof(f848,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP8(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP8(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f847])).
fof(f847,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f846])).
fof(f846,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f736])).
fof(f736,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f1507,plain,(
  ( ! [X154,X153] : (sP9(X153,X154,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X154)) | ~sP8(X154,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X153)) )),
  inference(resolution,[],[f917,f1105])).
fof(f1105,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f848])).
fof(f7285,plain,(
  sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | $spl42),
  inference(cnf_transformation,[],[f7285_D])).
fof(f7285_D,plain,(
  sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f21058,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(subsumption_resolution,[],[f21057,f5442])).
fof(f5442,plain,(
  ( ! [X227,X228,X226] : (sP8(X227,X226,X228) | sP12(X226,sK62)) )),
  inference(resolution,[],[f5230,f1111])).
fof(f21057,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | ~sP8(X14,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X15) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(subsumption_resolution,[],[f21037,f1511])).
fof(f21037,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | ~sP8(X14,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X15) | ~p(s(bool,X15)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1507,f5437])).
fof(f5437,plain,(
  ( ! [X212,X213,X211] : (~sP9(X212,X213,X211) | ~p(s(bool,X212)) | sP12(X211,sK62)) )),
  inference(resolution,[],[f5230,f1103])).
fof(f21023,plain,(
  ( ! [X20] : (~p(s(bool,X20)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(subsumption_resolution,[],[f21022,f5696])).
fof(f5696,plain,(
  ( ! [X231,X229,X230] : (sP8(X229,X230,X231) | sP12(sK62,X229)) )),
  inference(resolution,[],[f5231,f1112])).
fof(f21022,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | ~sP8('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X20,X21) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(subsumption_resolution,[],[f21003,f1510])).
fof(f21003,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | ~sP8('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X20,X21) | ~p(s(bool,X21)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1506,f5691])).
fof(f5691,plain,(
  ( ! [X216,X215,X214] : (~sP9(X215,X214,X216) | ~p(s(bool,X215)) | sP12(sK62,X214)) )),
  inference(resolution,[],[f5231,f1104])).
fof(f1104,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f848])).
fof(f1506,plain,(
  ( ! [X152,X151] : (sP9(X151,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X152) | ~p(s(bool,X152)) | ~sP8('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X152,X151)) )),
  inference(resolution,[],[f917,f1105])).
fof(f21021,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(subsumption_resolution,[],[f21020,f5443])).
fof(f5443,plain,(
  ( ! [X231,X229,X230] : (sP8(X229,X230,X231) | sP12(X229,sK62)) )),
  inference(resolution,[],[f5230,f1112])).
fof(f21020,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~sP8('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X16,X17) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(subsumption_resolution,[],[f21001,f1510])).
fof(f21001,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~sP8('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X16,X17) | ~p(s(bool,X17)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1506,f5438])).
fof(f5438,plain,(
  ( ! [X216,X215,X214] : (~sP9(X215,X214,X216) | ~p(s(bool,X215)) | sP12(X214,sK62)) )),
  inference(resolution,[],[f5230,f1104])).
fof(f20987,plain,(
  ( ! [X23,X22] : (p(s(bool,X23)) | p(s(bool,X22)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(subsumption_resolution,[],[f20965,f1078])).
fof(f20965,plain,(
  ( ! [X23,X22] : (~sP4(X22,X23,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X23)) | p(s(bool,X22)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1479,f5660])).
fof(f5660,plain,(
  ( ! [X123,X121,X122] : (~sP5(X122,X123,X121) | p(s(bool,X122)) | sP12(sK62,X121)) )),
  inference(resolution,[],[f5231,f1068])).
fof(f1068,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f832])).
fof(f832,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP4(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP4(X0,X1,X2) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f831])).
fof(f831,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f830])).
fof(f830,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1)) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f730])).
fof(f730,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f1479,plain,(
  ( ! [X97,X98] : (sP5(X97,X98,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~sP4(X97,X98,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X98))) )),
  inference(resolution,[],[f917,f1074])).
fof(f1074,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f832])).
fof(f20985,plain,(
  ( ! [X17,X16] : (p(s(bool,X17)) | p(s(bool,X16)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(subsumption_resolution,[],[f20962,f1078])).
fof(f20962,plain,(
  ( ! [X17,X16] : (~sP4(X16,X17,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X17)) | p(s(bool,X16)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1479,f5407])).
fof(f5407,plain,(
  ( ! [X123,X121,X122] : (~sP5(X122,X123,X121) | p(s(bool,X122)) | sP12(X121,sK62)) )),
  inference(resolution,[],[f5230,f1068])).
fof(f20947,plain,(
  ( ! [X20] : (~p(s(bool,X20)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(subsumption_resolution,[],[f20931,f1094])).
fof(f1094,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f843])).
fof(f20931,plain,(
  ( ! [X21,X20] : (~sP6('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X20,X21) | ~p(s(bool,X20)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1557,f5675])).
fof(f5675,plain,(
  ( ! [X167,X166,X168] : (~sP7(X166,X167,X168) | ~p(s(bool,X167)) | sP12(sK62,X166)) )),
  inference(resolution,[],[f5231,f1086])).
fof(f1086,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f840])).
fof(f840,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP6(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f839])).
fof(f839,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f838])).
fof(f838,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1)) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f733])).
fof(f733,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f1557,plain,(
  ( ! [X231,X232] : (sP7('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X231,X232) | ~sP6('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X231,X232)) )),
  inference(resolution,[],[f917,f1254])).
fof(f1254,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1088])).
fof(f1088,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f840])).
fof(f20946,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(subsumption_resolution,[],[f20930,f1093])).
fof(f1093,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f843])).
fof(f20930,plain,(
  ( ! [X19,X18] : (~sP6('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X18,X19) | ~p(s(bool,X19)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1557,f5674])).
fof(f5674,plain,(
  ( ! [X165,X163,X164] : (~sP7(X163,X165,X164) | ~p(s(bool,X164)) | sP12(sK62,X163)) )),
  inference(resolution,[],[f5231,f1085])).
fof(f1085,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f840])).
fof(f20945,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(subsumption_resolution,[],[f20929,f1094])).
fof(f20929,plain,(
  ( ! [X17,X16] : (~sP6('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X16,X17) | ~p(s(bool,X16)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1557,f5422])).
fof(f5422,plain,(
  ( ! [X167,X166,X168] : (~sP7(X166,X167,X168) | ~p(s(bool,X167)) | sP12(X166,sK62)) )),
  inference(resolution,[],[f5230,f1086])).
fof(f20944,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(subsumption_resolution,[],[f20928,f1093])).
fof(f20928,plain,(
  ( ! [X14,X15] : (~sP6('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X14,X15) | ~p(s(bool,X15)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1557,f5421])).
fof(f5421,plain,(
  ( ! [X165,X163,X164] : (~sP7(X163,X165,X164) | ~p(s(bool,X164)) | sP12(X163,sK62)) )),
  inference(resolution,[],[f5230,f1085])).
fof(f20916,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(subsumption_resolution,[],[f20898,f1077])).
fof(f1077,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f835])).
fof(f20898,plain,(
  ( ! [X24,X25] : (~sP4('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X24,X25) | ~p(s(bool,X24)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1556,f5661])).
fof(f5661,plain,(
  ( ! [X125,X126,X124] : (~sP5(X124,X125,X126) | ~p(s(bool,X125)) | sP12(sK62,X124)) )),
  inference(resolution,[],[f5231,f1069])).
fof(f1069,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f832])).
fof(f1556,plain,(
  ( ! [X229,X230] : (sP5('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X229,X230) | ~sP4('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X229,X230)) )),
  inference(resolution,[],[f917,f1253])).
fof(f1253,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1071])).
fof(f1071,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f832])).
fof(f20914,plain,(
  ( ! [X21] : (p(s(bool,X21)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(subsumption_resolution,[],[f20896,f1076])).
fof(f20896,plain,(
  ( ! [X21,X20] : (~sP4('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X20,X21) | p(s(bool,X21)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1556,f5659])).
fof(f5659,plain,(
  ( ! [X118,X120,X119] : (~sP5(X118,X120,X119) | p(s(bool,X119)) | sP12(sK62,X118)) )),
  inference(resolution,[],[f5231,f1068])).
fof(f20913,plain,(
  ( ! [X18] : (~p(s(bool,X18)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(subsumption_resolution,[],[f20895,f1077])).
fof(f20895,plain,(
  ( ! [X19,X18] : (~sP4('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X18,X19) | ~p(s(bool,X18)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1556,f5408])).
fof(f5408,plain,(
  ( ! [X125,X126,X124] : (~sP5(X124,X125,X126) | ~p(s(bool,X125)) | sP12(X124,sK62)) )),
  inference(resolution,[],[f5230,f1069])).
fof(f20911,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(subsumption_resolution,[],[f20893,f1076])).
fof(f20893,plain,(
  ( ! [X14,X15] : (~sP4('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X14,X15) | p(s(bool,X15)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1556,f5406])).
fof(f5406,plain,(
  ( ! [X118,X120,X119] : (~sP5(X118,X120,X119) | p(s(bool,X119)) | sP12(X118,sK62)) )),
  inference(resolution,[],[f5230,f1068])).
fof(f20282,plain,(
  ( ! [X19] : (sP12(X19,'F_0') | sP12(X19,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl262),
  inference(cnf_transformation,[],[f20282_D])).
fof(f20282_D,plain,(
  ( ! [X19] : (sP12(X19,'F_0') | sP12(X19,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f20810,plain,(
  ( ! [X30,X31,X29] : (sP12(X29,'F_0') | sP12(X29,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl116),
  inference(resolution,[],[f14609,f1510])).
fof(f14609,plain,(
  ( ! [X321,X323,X320,X322] : (sP12(X320,'F_0') | sP12(X320,X321) | sP8(X321,X322,X323)) ) | $spl116),
  inference(subsumption_resolution,[],[f14608,f1247])).
fof(f1247,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f665])).
fof(f665,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.rich_list.COMM_MONOID_FOLDR',aHL_FALSITY)).
fof(f14608,plain,(
  ( ! [X321,X323,X320,X322] : (p(s(bool,'F_0')) | sP12(X320,'F_0') | sP12(X320,X321) | sP8(X321,X322,X323)) ) | $spl116),
  inference(forward_demodulation,[],[f14516,f5240])).
fof(f5240,plain,(
  s(bool,'F_0') = s(bool,sK62)),
  inference(resolution,[],[f5116,f1211])).
fof(f1211,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f895])).
fof(f895,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f894])).
fof(f894,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f714])).
fof(f714,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f641])).
fof(f641,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f217])).
fof(f217,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X14))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.IMP_F_EQ_F')).
fof(f14516,plain,(
  ( ! [X321,X323,X320,X322] : (p(s(bool,sK62)) | sP12(X320,'F_0') | sP12(X320,X321) | sP8(X321,X322,X323)) ) | $spl116),
  inference(resolution,[],[f13596,f8349])).
fof(f8349,plain,(
  ( ! [X304,X306,X305,X307] : (sP12(X304,X305) | sP12(sK62,X304) | sP8(X305,X306,X307)) )),
  inference(resolution,[],[f5723,f1112])).
fof(f5723,plain,(
  ( ! [X308,X307] : (~p(s(bool,X308)) | sP12(X307,X308) | sP12(sK62,X307)) )),
  inference(resolution,[],[f5231,f1170])).
fof(f13596,plain,(
  ( ! [X302,X301] : (~sP12(X302,X301) | p(s(bool,X302)) | sP12(X301,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f1165])).
fof(f20281,plain,(
  ( ! [X16] : (sP12(X16,'F_0') | sP12(X16,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl260),
  inference(cnf_transformation,[],[f20281_D])).
fof(f20281_D,plain,(
  ( ! [X16] : (sP12(X16,'F_0') | sP12(X16,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f20809,plain,(
  ( ! [X28,X26,X27] : (sP12(X26,'F_0') | sP12(X26,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl116),
  inference(resolution,[],[f14609,f1339])).
fof(f9202,plain,(
  ( ! [X37,X35,X36] : (sP8(X35,X36,X37)) ) | $spl94),
  inference(cnf_transformation,[],[f9202_D])).
fof(f9202_D,plain,(
  ( ! [X37,X35,X36] : (sP8(X35,X36,X37)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f14552,plain,(
  sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | $spl140),
  inference(cnf_transformation,[],[f14552_D])).
fof(f14552_D,plain,(
  sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f20818,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X78,X79,X80)) ) | $spl116),
  inference(subsumption_resolution,[],[f20798,f1112])).
fof(f20798,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X78,X79,X80) | ~p(s(bool,X78))) ) | $spl116),
  inference(resolution,[],[f14609,f3758])).
fof(f3758,plain,(
  ( ! [X212] : (~sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X212) | ~p(s(bool,X212))) )),
  inference(resolution,[],[f1433,f1166])).
fof(f1166,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f875])).
fof(f1433,plain,(
  p(s(bool,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))))),
  inference(resolution,[],[f917,f1048])).
fof(f1048,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.operator.LEFT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) | ~p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f829])).
fof(f829,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) | (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) & p(s(bool,'const.operator.RIGHT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) & p(s(bool,'const.operator.LEFT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))))) & (~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | ~p(s(bool,'const.operator.RIGHT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) | ~p(s(bool,'const.operator.LEFT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) | p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))))))),
  inference(flattening,[],[f828])).
fof(f828,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) | (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) & p(s(bool,'const.operator.RIGHT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) & p(s(bool,'const.operator.LEFT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))))) & ((~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | ~p(s(bool,'const.operator.RIGHT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) | ~p(s(bool,'const.operator.LEFT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))))) | p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f598])).
fof(f598,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) <=> (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) & p(s(bool,'const.operator.RIGHT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) & p(s(bool,'const.operator.LEFT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))))))),
  inference(flattening,[],[f597])).
fof(f597,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) <=> (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) & p(s(bool,'const.operator.RIGHT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) & p(s(bool,'const.operator.LEFT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7,X9] : (p(s(bool,'const.operator.MONOID_2'(s(fun(X6,fun(X6,X6)),X7),s(X6,X9)))) <=> (p(s(bool,'const.operator.ASSOC_1'(s(fun(X6,fun(X6,X6)),X7)))) & p(s(bool,'const.operator.RIGHT_ID_2'(s(fun(X6,fun(X6,X6)),X7),s(X6,X9)))) & p(s(bool,'const.operator.LEFT_ID_2'(s(fun(X6,fun(X6,X6)),X7),s(X6,X9))))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.operator.MONOID_DEF')).
fof(f14550,plain,(
  sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | $spl138),
  inference(cnf_transformation,[],[f14550_D])).
fof(f14550_D,plain,(
  sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f20817,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X75,X76,X77)) ) | $spl116),
  inference(subsumption_resolution,[],[f20797,f1112])).
fof(f20797,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X75,X76,X77) | ~p(s(bool,X75))) ) | $spl116),
  inference(resolution,[],[f14609,f3586])).
fof(f3586,plain,(
  ( ! [X212] : (~sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X212) | ~p(s(bool,X212))) )),
  inference(resolution,[],[f1432,f1166])).
fof(f1432,plain,(
  p(s(bool,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))))),
  inference(resolution,[],[f917,f1047])).
fof(f1047,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.operator.RIGHT_ID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2)))) | ~p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f829])).
fof(f14546,plain,(
  sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | $spl134),
  inference(cnf_transformation,[],[f14546_D])).
fof(f14546_D,plain,(
  sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f20816,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X69,X70,X71)) ) | $spl116),
  inference(subsumption_resolution,[],[f20795,f1112])).
fof(f20795,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl116),
  inference(resolution,[],[f14609,f1538])).
fof(f1538,plain,(
  ( ! [X215] : (~sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X215) | ~p(s(bool,X215))) )),
  inference(resolution,[],[f917,f1166])).
fof(f18362,plain,(
  sP12('F_0','F_0') | $spl254),
  inference(cnf_transformation,[],[f18362_D])).
fof(f18362_D,plain,(
  sP12('F_0','F_0') <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f20767,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP8(X41,X42,X43) | p(s(bool,X40)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14609,f15040])).
fof(f15040,plain,(
  ( ! [X302,X301] : (~sP12(X302,X301) | p(s(bool,X302)) | sP12('F_0',X301)) ) | $spl176),
  inference(resolution,[],[f14807,f1165])).
fof(f20765,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP8(X33,X34,X35) | p(s(bool,X32)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14609,f13596])).
fof(f5617,plain,(
  sP12('F_0',sK62) | $spl34),
  inference(cnf_transformation,[],[f5617_D])).
fof(f5617_D,plain,(
  sP12('F_0',sK62) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f20761,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP8(X17,X18,X19) | p(s(bool,X16)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14609,f5467])).
fof(f5467,plain,(
  ( ! [X302,X301] : (~sP12(X302,X301) | p(s(bool,X302)) | sP12(X301,sK62)) )),
  inference(resolution,[],[f5230,f1165])).
fof(f20744,plain,(
  ( ! [X24,X23,X25] : (sP12(X23,'F_0') | sP12(X23,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl116),
  inference(resolution,[],[f14607,f1511])).
fof(f14607,plain,(
  ( ! [X316,X318,X317,X319] : (sP12(X316,'F_0') | sP12(X316,X317) | sP8(X318,X317,X319)) ) | $spl116),
  inference(subsumption_resolution,[],[f14606,f1247])).
fof(f14606,plain,(
  ( ! [X316,X318,X317,X319] : (p(s(bool,'F_0')) | sP12(X316,'F_0') | sP12(X316,X317) | sP8(X318,X317,X319)) ) | $spl116),
  inference(forward_demodulation,[],[f14515,f5240])).
fof(f14515,plain,(
  ( ! [X316,X318,X317,X319] : (p(s(bool,sK62)) | sP12(X316,'F_0') | sP12(X316,X317) | sP8(X318,X317,X319)) ) | $spl116),
  inference(resolution,[],[f13596,f8348])).
fof(f8348,plain,(
  ( ! [X300,X302,X301,X303] : (sP12(X300,X301) | sP12(sK62,X300) | sP8(X302,X301,X303)) )),
  inference(resolution,[],[f5723,f1111])).
fof(f20743,plain,(
  ( ! [X21,X22,X20] : (sP12(X20,'F_0') | sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | $spl116),
  inference(resolution,[],[f14607,f1340])).
fof(f20754,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X78,X79,X80)) ) | $spl116),
  inference(subsumption_resolution,[],[f20734,f1111])).
fof(f20734,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X78,X79,X80) | ~p(s(bool,X79))) ) | $spl116),
  inference(resolution,[],[f14607,f3758])).
fof(f20753,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X75,X76,X77)) ) | $spl116),
  inference(subsumption_resolution,[],[f20733,f1111])).
fof(f20733,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X75,X76,X77) | ~p(s(bool,X76))) ) | $spl116),
  inference(resolution,[],[f14607,f3586])).
fof(f20752,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X69,X70,X71)) ) | $spl116),
  inference(subsumption_resolution,[],[f20731,f1111])).
fof(f20731,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X69,X70,X71) | ~p(s(bool,X70))) ) | $spl116),
  inference(resolution,[],[f14607,f1538])).
fof(f20703,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP8(X42,X41,X43) | p(s(bool,X40)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14607,f15040])).
fof(f20701,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP8(X34,X33,X35) | p(s(bool,X32)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14607,f13596])).
fof(f20697,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP8(X18,X17,X19) | p(s(bool,X16)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14607,f5467])).
fof(f20650,plain,(
  ( ! [X24,X23,X25] : (sP12(X23,'F_0') | sP12(X23,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X24)) | p(s(bool,X25))) ) | $spl116),
  inference(resolution,[],[f14605,f1494])).
fof(f14605,plain,(
  ( ! [X313,X315,X312,X314] : (sP12(X312,'F_0') | sP12(X312,X313) | sP6(X313,X314,X315)) ) | $spl116),
  inference(subsumption_resolution,[],[f14604,f1247])).
fof(f14604,plain,(
  ( ! [X313,X315,X312,X314] : (p(s(bool,'F_0')) | sP12(X312,'F_0') | sP12(X312,X313) | sP6(X313,X314,X315)) ) | $spl116),
  inference(forward_demodulation,[],[f14514,f5240])).
fof(f14514,plain,(
  ( ! [X313,X315,X312,X314] : (p(s(bool,sK62)) | sP12(X312,'F_0') | sP12(X312,X313) | sP6(X313,X314,X315)) ) | $spl116),
  inference(resolution,[],[f13596,f8335])).
fof(f8335,plain,(
  ( ! [X251,X249,X250,X248] : (sP12(X248,X249) | sP12(sK62,X248) | sP6(X249,X250,X251)) )),
  inference(resolution,[],[f5723,f1095])).
fof(f20649,plain,(
  ( ! [X21,X22,X20] : (sP12(X20,'F_0') | sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X21)) | p(s(bool,X22))) ) | $spl116),
  inference(resolution,[],[f14605,f1323])).
fof(f9139,plain,(
  ( ! [X37,X35,X36] : (sP6(X35,X36,X37)) ) | $spl92),
  inference(cnf_transformation,[],[f9139_D])).
fof(f9139_D,plain,(
  ( ! [X37,X35,X36] : (sP6(X35,X36,X37)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f20658,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X78,X79,X80)) ) | $spl116),
  inference(subsumption_resolution,[],[f20640,f1095])).
fof(f20640,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X78,X79,X80) | ~p(s(bool,X78))) ) | $spl116),
  inference(resolution,[],[f14605,f3758])).
fof(f20657,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X75,X76,X77)) ) | $spl116),
  inference(subsumption_resolution,[],[f20639,f1095])).
fof(f20639,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X75,X76,X77) | ~p(s(bool,X75))) ) | $spl116),
  inference(resolution,[],[f14605,f3586])).
fof(f20656,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X69,X70,X71)) ) | $spl116),
  inference(subsumption_resolution,[],[f20637,f1095])).
fof(f20637,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl116),
  inference(resolution,[],[f14605,f1538])).
fof(f20609,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP6(X41,X42,X43) | p(s(bool,X40)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14605,f15040])).
fof(f20607,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP6(X33,X34,X35) | p(s(bool,X32)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14605,f13596])).
fof(f20603,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP6(X17,X18,X19) | p(s(bool,X16)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14605,f5467])).
fof(f20590,plain,(
  ( ! [X30,X31,X29] : (sP12(X29,'F_0') | sP12(X29,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X30)) | p(s(bool,X31))) ) | $spl116),
  inference(resolution,[],[f14603,f1480])).
fof(f14603,plain,(
  ( ! [X308,X310,X309,X311] : (sP12(X308,'F_0') | sP12(X308,X309) | sP4(X309,X310,X311)) ) | $spl116),
  inference(subsumption_resolution,[],[f14602,f1247])).
fof(f14602,plain,(
  ( ! [X308,X310,X309,X311] : (p(s(bool,'F_0')) | sP12(X308,'F_0') | sP12(X308,X309) | sP4(X309,X310,X311)) ) | $spl116),
  inference(forward_demodulation,[],[f14513,f5240])).
fof(f14513,plain,(
  ( ! [X308,X310,X309,X311] : (p(s(bool,sK62)) | sP12(X308,'F_0') | sP12(X308,X309) | sP4(X309,X310,X311)) ) | $spl116),
  inference(resolution,[],[f13596,f8319])).
fof(f8319,plain,(
  ( ! [X187,X185,X186,X184] : (sP12(X184,X185) | sP12(sK62,X184) | sP4(X185,X186,X187)) )),
  inference(resolution,[],[f5723,f1078])).
fof(f20589,plain,(
  ( ! [X28,X26,X27] : (sP12(X26,'F_0') | sP12(X26,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | ~p(s(bool,X27)) | p(s(bool,X28))) ) | $spl116),
  inference(resolution,[],[f14603,f1309])).
fof(f9013,plain,(
  ( ! [X37,X35,X36] : (sP4(X35,X36,X37)) ) | $spl90),
  inference(cnf_transformation,[],[f9013_D])).
fof(f9013_D,plain,(
  ( ! [X37,X35,X36] : (sP4(X35,X36,X37)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f20598,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X78,X79,X80)) ) | $spl116),
  inference(subsumption_resolution,[],[f20578,f1078])).
fof(f20578,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X78,X79,X80) | ~p(s(bool,X78))) ) | $spl116),
  inference(resolution,[],[f14603,f3758])).
fof(f20597,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X75,X76,X77)) ) | $spl116),
  inference(subsumption_resolution,[],[f20577,f1078])).
fof(f20577,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X75,X76,X77) | ~p(s(bool,X75))) ) | $spl116),
  inference(resolution,[],[f14603,f3586])).
fof(f20596,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X69,X70,X71)) ) | $spl116),
  inference(subsumption_resolution,[],[f20575,f1078])).
fof(f20575,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl116),
  inference(resolution,[],[f14603,f1538])).
fof(f20547,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP4(X41,X42,X43) | p(s(bool,X40)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14603,f15040])).
fof(f20545,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP4(X33,X34,X35) | p(s(bool,X32)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14603,f13596])).
fof(f20541,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP4(X17,X18,X19) | p(s(bool,X16)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14603,f5467])).
fof(f20525,plain,(
  ( ! [X24,X23,X25] : (sP12(X23,'F_0') | sP12(X23,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl116),
  inference(resolution,[],[f14601,f1481])).
fof(f14601,plain,(
  ( ! [X304,X306,X305,X307] : (sP12(X304,'F_0') | sP12(X304,X305) | sP4(X306,X307,X305)) ) | $spl116),
  inference(subsumption_resolution,[],[f14600,f1247])).
fof(f14600,plain,(
  ( ! [X304,X306,X305,X307] : (p(s(bool,'F_0')) | sP12(X304,'F_0') | sP12(X304,X305) | sP4(X306,X307,X305)) ) | $spl116),
  inference(forward_demodulation,[],[f14512,f5240])).
fof(f14512,plain,(
  ( ! [X304,X306,X305,X307] : (p(s(bool,sK62)) | sP12(X304,'F_0') | sP12(X304,X305) | sP4(X306,X307,X305)) ) | $spl116),
  inference(resolution,[],[f13596,f8318])).
fof(f8318,plain,(
  ( ! [X182,X180,X183,X181] : (sP12(X180,X181) | sP12(sK62,X180) | sP4(X182,X183,X181)) )),
  inference(resolution,[],[f5723,f1076])).
fof(f20524,plain,(
  ( ! [X21,X22,X20] : (sP12(X20,'F_0') | sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | $spl116),
  inference(resolution,[],[f14601,f1310])).
fof(f20535,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X78,X79,X80)) ) | $spl116),
  inference(subsumption_resolution,[],[f20515,f1076])).
fof(f20515,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X78,X79,X80) | ~p(s(bool,X80))) ) | $spl116),
  inference(resolution,[],[f14601,f3758])).
fof(f20534,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X75,X76,X77)) ) | $spl116),
  inference(subsumption_resolution,[],[f20514,f1076])).
fof(f20514,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X75,X76,X77) | ~p(s(bool,X77))) ) | $spl116),
  inference(resolution,[],[f14601,f3586])).
fof(f20533,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X69,X70,X71)) ) | $spl116),
  inference(subsumption_resolution,[],[f20512,f1076])).
fof(f20512,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X69,X70,X71) | ~p(s(bool,X71))) ) | $spl116),
  inference(resolution,[],[f14601,f1538])).
fof(f20484,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP4(X42,X43,X41) | p(s(bool,X40)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14601,f15040])).
fof(f20482,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP4(X34,X35,X33) | p(s(bool,X32)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14601,f13596])).
fof(f20478,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP4(X18,X19,X17) | p(s(bool,X16)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14601,f5467])).
fof(f20462,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'F_0') | sP12(X25,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl116),
  inference(resolution,[],[f14599,f1449])).
fof(f14599,plain,(
  ( ! [X300,X302,X301,X303] : (sP12(X300,'F_0') | sP12(X300,X301) | sP1(X301,X302,X303)) ) | $spl116),
  inference(subsumption_resolution,[],[f14598,f1247])).
fof(f14598,plain,(
  ( ! [X300,X302,X301,X303] : (p(s(bool,'F_0')) | sP12(X300,'F_0') | sP12(X300,X301) | sP1(X301,X302,X303)) ) | $spl116),
  inference(forward_demodulation,[],[f14511,f5240])).
fof(f14511,plain,(
  ( ! [X300,X302,X301,X303] : (p(s(bool,sK62)) | sP12(X300,'F_0') | sP12(X300,X301) | sP1(X301,X302,X303)) ) | $spl116),
  inference(resolution,[],[f13596,f8290])).
fof(f8290,plain,(
  ( ! [X74,X72,X75,X73] : (sP12(X72,X73) | sP12(sK62,X72) | sP1(X73,X74,X75)) )),
  inference(resolution,[],[f5723,f947])).
fof(f20461,plain,(
  ( ! [X24,X23,X22] : (sP12(X22,'F_0') | sP12(X22,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl116),
  inference(resolution,[],[f14599,f1278])).
fof(f20472,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,'F_0') | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl116),
  inference(subsumption_resolution,[],[f20457,f13597])).
fof(f13597,plain,(
  ( ! [X304,X303] : (~sP12(X303,X304) | p(s(bool,X304)) | sP12(X303,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f1165])).
fof(f20457,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,'F_0') | sP12(X8,X9) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl116),
  inference(resolution,[],[f14599,f938])).
fof(f8841,plain,(
  ( ! [X37,X35,X36] : (sP1(X35,X36,X37)) ) | $spl86),
  inference(cnf_transformation,[],[f8841_D])).
fof(f8841_D,plain,(
  ( ! [X37,X35,X36] : (sP1(X35,X36,X37)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f20470,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X78,X79,X80)) ) | $spl116),
  inference(subsumption_resolution,[],[f20451,f947])).
fof(f20451,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X78,X79,X80) | ~p(s(bool,X78))) ) | $spl116),
  inference(resolution,[],[f14599,f3758])).
fof(f20469,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X75,X76,X77)) ) | $spl116),
  inference(subsumption_resolution,[],[f20450,f947])).
fof(f20450,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X75,X76,X77) | ~p(s(bool,X75))) ) | $spl116),
  inference(resolution,[],[f14599,f3586])).
fof(f20468,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X69,X70,X71)) ) | $spl116),
  inference(subsumption_resolution,[],[f20448,f947])).
fof(f20448,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl116),
  inference(resolution,[],[f14599,f1538])).
fof(f20420,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP1(X41,X42,X43) | p(s(bool,X40)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14599,f15040])).
fof(f20418,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP1(X33,X34,X35) | p(s(bool,X32)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14599,f13596])).
fof(f20414,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP1(X17,X18,X19) | p(s(bool,X16)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14599,f5467])).
fof(f20397,plain,(
  ( ! [X21,X19,X20] : (sP12(X19,'F_0') | sP12(X19,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl116),
  inference(resolution,[],[f14597,f1450])).
fof(f14597,plain,(
  ( ! [X296,X298,X297,X299] : (sP12(X296,'F_0') | sP12(X296,X297) | sP1(X298,X297,X299)) ) | $spl116),
  inference(subsumption_resolution,[],[f14596,f1247])).
fof(f14596,plain,(
  ( ! [X296,X298,X297,X299] : (p(s(bool,'F_0')) | sP12(X296,'F_0') | sP12(X296,X297) | sP1(X298,X297,X299)) ) | $spl116),
  inference(forward_demodulation,[],[f14510,f5240])).
fof(f14510,plain,(
  ( ! [X296,X298,X297,X299] : (p(s(bool,sK62)) | sP12(X296,'F_0') | sP12(X296,X297) | sP1(X298,X297,X299)) ) | $spl116),
  inference(resolution,[],[f13596,f8289])).
fof(f8289,plain,(
  ( ! [X70,X68,X71,X69] : (sP12(X68,X69) | sP12(sK62,X68) | sP1(X70,X69,X71)) )),
  inference(resolution,[],[f5723,f946])).
fof(f20396,plain,(
  ( ! [X17,X18,X16] : (sP12(X16,'F_0') | sP12(X16,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl116),
  inference(resolution,[],[f14597,f1279])).
fof(f20407,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X78,X79,X80)) ) | $spl116),
  inference(subsumption_resolution,[],[f20388,f946])).
fof(f20388,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X78,X79,X80) | ~p(s(bool,X79))) ) | $spl116),
  inference(resolution,[],[f14597,f3758])).
fof(f20406,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X75,X76,X77)) ) | $spl116),
  inference(subsumption_resolution,[],[f20387,f946])).
fof(f20387,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X75,X76,X77) | ~p(s(bool,X76))) ) | $spl116),
  inference(resolution,[],[f14597,f3586])).
fof(f20405,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X69,X70,X71)) ) | $spl116),
  inference(subsumption_resolution,[],[f20385,f946])).
fof(f20385,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X69,X70,X71) | ~p(s(bool,X70))) ) | $spl116),
  inference(resolution,[],[f14597,f1538])).
fof(f20357,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP1(X42,X41,X43) | p(s(bool,X40)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14597,f15040])).
fof(f20355,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP1(X34,X33,X35) | p(s(bool,X32)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14597,f13596])).
fof(f20351,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP1(X18,X17,X19) | p(s(bool,X16)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14597,f5467])).
fof(f20335,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'F_0') | sP12(X25,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl116),
  inference(resolution,[],[f14595,f1446])).
fof(f14595,plain,(
  ( ! [X292,X294,X293,X295] : (sP12(X292,'F_0') | sP12(X292,X293) | sP2(X293,X294,X295)) ) | $spl116),
  inference(subsumption_resolution,[],[f14594,f1247])).
fof(f14594,plain,(
  ( ! [X292,X294,X293,X295] : (p(s(bool,'F_0')) | sP12(X292,'F_0') | sP12(X292,X293) | sP2(X293,X294,X295)) ) | $spl116),
  inference(forward_demodulation,[],[f14509,f5240])).
fof(f14509,plain,(
  ( ! [X292,X294,X293,X295] : (p(s(bool,sK62)) | sP12(X292,'F_0') | sP12(X292,X293) | sP2(X293,X294,X295)) ) | $spl116),
  inference(resolution,[],[f13596,f8287])).
fof(f8287,plain,(
  ( ! [X61,X62,X60,X63] : (sP12(X60,X61) | sP12(sK62,X60) | sP2(X61,X62,X63)) )),
  inference(resolution,[],[f5723,f943])).
fof(f20334,plain,(
  ( ! [X24,X23,X22] : (sP12(X22,'F_0') | sP12(X22,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl116),
  inference(resolution,[],[f14595,f1275])).
fof(f20346,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,'F_0') | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl116),
  inference(subsumption_resolution,[],[f20331,f13597])).
fof(f20331,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,'F_0') | sP12(X12,X13) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl116),
  inference(resolution,[],[f14595,f939])).
fof(f8667,plain,(
  ( ! [X37,X35,X36] : (sP2(X35,X36,X37)) ) | $spl76),
  inference(cnf_transformation,[],[f8667_D])).
fof(f8667_D,plain,(
  ( ! [X37,X35,X36] : (sP2(X35,X36,X37)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f20343,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X78,X79,X80)) ) | $spl116),
  inference(subsumption_resolution,[],[f20324,f943])).
fof(f20324,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X78,X79,X80) | ~p(s(bool,X78))) ) | $spl116),
  inference(resolution,[],[f14595,f3758])).
fof(f20342,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X75,X76,X77)) ) | $spl116),
  inference(subsumption_resolution,[],[f20323,f943])).
fof(f20323,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X75,X76,X77) | ~p(s(bool,X75))) ) | $spl116),
  inference(resolution,[],[f14595,f3586])).
fof(f20341,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X69,X70,X71)) ) | $spl116),
  inference(subsumption_resolution,[],[f20321,f943])).
fof(f20321,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl116),
  inference(resolution,[],[f14595,f1538])).
fof(f20293,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP2(X41,X42,X43) | p(s(bool,X40)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14595,f15040])).
fof(f20291,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP2(X33,X34,X35) | p(s(bool,X32)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14595,f13596])).
fof(f20287,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP2(X17,X18,X19) | p(s(bool,X16)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14595,f5467])).
fof(f20269,plain,(
  ( ! [X21,X19,X20] : (sP12(X19,'F_0') | sP12(X19,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl116),
  inference(resolution,[],[f14593,f1447])).
fof(f14593,plain,(
  ( ! [X288,X290,X289,X291] : (sP12(X288,'F_0') | sP12(X288,X289) | sP2(X290,X289,X291)) ) | $spl116),
  inference(subsumption_resolution,[],[f14592,f1247])).
fof(f14592,plain,(
  ( ! [X288,X290,X289,X291] : (p(s(bool,'F_0')) | sP12(X288,'F_0') | sP12(X288,X289) | sP2(X290,X289,X291)) ) | $spl116),
  inference(forward_demodulation,[],[f14508,f5240])).
fof(f14508,plain,(
  ( ! [X288,X290,X289,X291] : (p(s(bool,sK62)) | sP12(X288,'F_0') | sP12(X288,X289) | sP2(X290,X289,X291)) ) | $spl116),
  inference(resolution,[],[f13596,f8286])).
fof(f8286,plain,(
  ( ! [X59,X57,X58,X56] : (sP12(X56,X57) | sP12(sK62,X56) | sP2(X58,X57,X59)) )),
  inference(resolution,[],[f5723,f942])).
fof(f20268,plain,(
  ( ! [X17,X18,X16] : (sP12(X16,'F_0') | sP12(X16,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl116),
  inference(resolution,[],[f14593,f1276])).
fof(f20279,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X78,X79,X80)) ) | $spl116),
  inference(subsumption_resolution,[],[f20260,f942])).
fof(f20260,plain,(
  ( ! [X80,X78,X79] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X78,X79,X80) | ~p(s(bool,X79))) ) | $spl116),
  inference(resolution,[],[f14593,f3758])).
fof(f20278,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X75,X76,X77)) ) | $spl116),
  inference(subsumption_resolution,[],[f20259,f942])).
fof(f20259,plain,(
  ( ! [X76,X77,X75] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X75,X76,X77) | ~p(s(bool,X76))) ) | $spl116),
  inference(resolution,[],[f14593,f3586])).
fof(f20277,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X69,X70,X71)) ) | $spl116),
  inference(subsumption_resolution,[],[f20257,f942])).
fof(f20257,plain,(
  ( ! [X70,X71,X69] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X69,X70,X71) | ~p(s(bool,X70))) ) | $spl116),
  inference(resolution,[],[f14593,f1538])).
fof(f14561,plain,(
  ( ! [X210,X208,X211,X209] : (p(s(bool,X208)) | sP2(X210,X209,X211) | sP12(X208,X209)) ) | $spl158),
  inference(cnf_transformation,[],[f14561_D])).
fof(f14561_D,plain,(
  ( ! [X210,X208,X211,X209] : (p(s(bool,X208)) | sP2(X210,X209,X211) | sP12(X208,X209)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f20229,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP2(X42,X41,X43) | p(s(bool,X40)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14593,f15040])).
fof(f20227,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP2(X34,X33,X35) | p(s(bool,X32)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14593,f13596])).
fof(f20223,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP2(X18,X17,X19) | p(s(bool,X16)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14593,f5467])).
fof(f17268,plain,(
  ( ! [X19] : (sP12(X19,'F_0') | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X19)) ) | $spl216),
  inference(cnf_transformation,[],[f17268_D])).
fof(f17268_D,plain,(
  ( ! [X19] : (sP12(X19,'F_0') | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X19)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f20209,plain,(
  ( ! [X30,X31,X29] : (sP12(X29,'F_0') | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X29) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl116),
  inference(resolution,[],[f14591,f1510])).
fof(f14591,plain,(
  ( ! [X285,X287,X284,X286] : (sP12(X284,'F_0') | sP12(X285,X284) | sP8(X285,X286,X287)) ) | $spl116),
  inference(subsumption_resolution,[],[f14590,f1247])).
fof(f14590,plain,(
  ( ! [X285,X287,X284,X286] : (p(s(bool,'F_0')) | sP12(X284,'F_0') | sP12(X285,X284) | sP8(X285,X286,X287)) ) | $spl116),
  inference(forward_demodulation,[],[f14507,f5240])).
fof(f14507,plain,(
  ( ! [X285,X287,X284,X286] : (p(s(bool,sK62)) | sP12(X284,'F_0') | sP12(X285,X284) | sP8(X285,X286,X287)) ) | $spl116),
  inference(resolution,[],[f13596,f8097])).
fof(f8097,plain,(
  ( ! [X304,X306,X305,X307] : (sP12(X304,X305) | sP12(sK62,X305) | sP8(X304,X306,X307)) )),
  inference(resolution,[],[f5722,f1112])).
fof(f5722,plain,(
  ( ! [X306,X305] : (~p(s(bool,X306)) | sP12(X306,X305) | sP12(sK62,X305)) )),
  inference(resolution,[],[f5231,f1167])).
fof(f17267,plain,(
  ( ! [X16] : (sP12(X16,'F_0') | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X16)) ) | $spl214),
  inference(cnf_transformation,[],[f17267_D])).
fof(f17267_D,plain,(
  ( ! [X16] : (sP12(X16,'F_0') | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X16)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f20208,plain,(
  ( ! [X28,X26,X27] : (sP12(X26,'F_0') | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X26) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl116),
  inference(resolution,[],[f14591,f1339])).
fof(f20217,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X60,X61,X62)) ) | $spl116),
  inference(subsumption_resolution,[],[f20191,f1112])).
fof(f20191,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl116),
  inference(resolution,[],[f14591,f3759])).
fof(f3759,plain,(
  ( ! [X213] : (~sP12(X213,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X213))) )),
  inference(resolution,[],[f1433,f1166])).
fof(f20216,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X57,X58,X59)) ) | $spl116),
  inference(subsumption_resolution,[],[f20190,f1112])).
fof(f20190,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X57,X58,X59) | ~p(s(bool,X57))) ) | $spl116),
  inference(resolution,[],[f14591,f3587])).
fof(f3587,plain,(
  ( ! [X213] : (~sP12(X213,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X213))) )),
  inference(resolution,[],[f1432,f1166])).
fof(f20215,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X51,X52,X53)) ) | $spl116),
  inference(subsumption_resolution,[],[f20188,f1112])).
fof(f20188,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl116),
  inference(resolution,[],[f14591,f1539])).
fof(f1539,plain,(
  ( ! [X216] : (~sP12(X216,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X216))) )),
  inference(resolution,[],[f917,f1166])).
fof(f14560,plain,(
  ( ! [X206,X204,X207,X205] : (p(s(bool,X204)) | sP8(X205,X206,X207) | sP12(X205,X204)) ) | $spl156),
  inference(cnf_transformation,[],[f14560_D])).
fof(f14560_D,plain,(
  ( ! [X206,X204,X207,X205] : (p(s(bool,X204)) | sP8(X205,X206,X207) | sP12(X205,X204)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f20166,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP8(X40,X42,X43) | p(s(bool,X41)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14591,f15040])).
fof(f20164,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP8(X32,X34,X35) | p(s(bool,X33)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14591,f13596])).
fof(f20160,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP8(X16,X18,X19) | p(s(bool,X17)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14591,f5467])).
fof(f20143,plain,(
  ( ! [X24,X23,X25] : (sP12(X23,'F_0') | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X23) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl116),
  inference(resolution,[],[f14589,f1511])).
fof(f14589,plain,(
  ( ! [X280,X282,X281,X283] : (sP12(X280,'F_0') | sP12(X281,X280) | sP8(X282,X281,X283)) ) | $spl116),
  inference(subsumption_resolution,[],[f14588,f1247])).
fof(f14588,plain,(
  ( ! [X280,X282,X281,X283] : (p(s(bool,'F_0')) | sP12(X280,'F_0') | sP12(X281,X280) | sP8(X282,X281,X283)) ) | $spl116),
  inference(forward_demodulation,[],[f14506,f5240])).
fof(f14506,plain,(
  ( ! [X280,X282,X281,X283] : (p(s(bool,sK62)) | sP12(X280,'F_0') | sP12(X281,X280) | sP8(X282,X281,X283)) ) | $spl116),
  inference(resolution,[],[f13596,f8096])).
fof(f8096,plain,(
  ( ! [X300,X302,X301,X303] : (sP12(X300,X301) | sP12(sK62,X301) | sP8(X302,X300,X303)) )),
  inference(resolution,[],[f5722,f1111])).
fof(f20142,plain,(
  ( ! [X21,X22,X20] : (sP12(X20,'F_0') | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | $spl116),
  inference(resolution,[],[f14589,f1340])).
fof(f20153,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X60,X61,X62)) ) | $spl116),
  inference(subsumption_resolution,[],[f20127,f1111])).
fof(f20127,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X60,X61,X62) | ~p(s(bool,X61))) ) | $spl116),
  inference(resolution,[],[f14589,f3759])).
fof(f20152,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X57,X58,X59)) ) | $spl116),
  inference(subsumption_resolution,[],[f20126,f1111])).
fof(f20126,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X57,X58,X59) | ~p(s(bool,X58))) ) | $spl116),
  inference(resolution,[],[f14589,f3587])).
fof(f20151,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X51,X52,X53)) ) | $spl116),
  inference(subsumption_resolution,[],[f20124,f1111])).
fof(f20124,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP8(X51,X52,X53) | ~p(s(bool,X52))) ) | $spl116),
  inference(resolution,[],[f14589,f1539])).
fof(f14559,plain,(
  ( ! [X202,X200,X203,X201] : (p(s(bool,X200)) | sP8(X202,X201,X203) | sP12(X201,X200)) ) | $spl154),
  inference(cnf_transformation,[],[f14559_D])).
fof(f14559_D,plain,(
  ( ! [X202,X200,X203,X201] : (p(s(bool,X200)) | sP8(X202,X201,X203) | sP12(X201,X200)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f20102,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP8(X42,X40,X43) | p(s(bool,X41)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14589,f15040])).
fof(f20100,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP8(X34,X32,X35) | p(s(bool,X33)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14589,f13596])).
fof(f20096,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP8(X18,X16,X19) | p(s(bool,X17)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14589,f5467])).
fof(f20083,plain,(
  ( ! [X24,X23,X25] : (sP12(X23,'F_0') | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X23) | p(s(bool,X24)) | p(s(bool,X25))) ) | $spl116),
  inference(resolution,[],[f14587,f1494])).
fof(f14587,plain,(
  ( ! [X277,X279,X276,X278] : (sP12(X276,'F_0') | sP12(X277,X276) | sP6(X277,X278,X279)) ) | $spl116),
  inference(subsumption_resolution,[],[f14586,f1247])).
fof(f14586,plain,(
  ( ! [X277,X279,X276,X278] : (p(s(bool,'F_0')) | sP12(X276,'F_0') | sP12(X277,X276) | sP6(X277,X278,X279)) ) | $spl116),
  inference(forward_demodulation,[],[f14505,f5240])).
fof(f14505,plain,(
  ( ! [X277,X279,X276,X278] : (p(s(bool,sK62)) | sP12(X276,'F_0') | sP12(X277,X276) | sP6(X277,X278,X279)) ) | $spl116),
  inference(resolution,[],[f13596,f8083])).
fof(f8083,plain,(
  ( ! [X251,X249,X250,X248] : (sP12(X248,X249) | sP12(sK62,X249) | sP6(X248,X250,X251)) )),
  inference(resolution,[],[f5722,f1095])).
fof(f20082,plain,(
  ( ! [X21,X22,X20] : (sP12(X20,'F_0') | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | p(s(bool,X21)) | p(s(bool,X22))) ) | $spl116),
  inference(resolution,[],[f14587,f1323])).
fof(f20091,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X60,X61,X62)) ) | $spl116),
  inference(subsumption_resolution,[],[f20067,f1095])).
fof(f20067,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl116),
  inference(resolution,[],[f14587,f3759])).
fof(f20090,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X57,X58,X59)) ) | $spl116),
  inference(subsumption_resolution,[],[f20066,f1095])).
fof(f20066,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X57,X58,X59) | ~p(s(bool,X57))) ) | $spl116),
  inference(resolution,[],[f14587,f3587])).
fof(f20089,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X51,X52,X53)) ) | $spl116),
  inference(subsumption_resolution,[],[f20064,f1095])).
fof(f20064,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP6(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl116),
  inference(resolution,[],[f14587,f1539])).
fof(f14558,plain,(
  ( ! [X198,X196,X199,X197] : (p(s(bool,X196)) | sP6(X197,X198,X199) | sP12(X197,X196)) ) | $spl152),
  inference(cnf_transformation,[],[f14558_D])).
fof(f14558_D,plain,(
  ( ! [X198,X196,X199,X197] : (p(s(bool,X196)) | sP6(X197,X198,X199) | sP12(X197,X196)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f20042,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP6(X40,X42,X43) | p(s(bool,X41)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14587,f15040])).
fof(f20040,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP6(X32,X34,X35) | p(s(bool,X33)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14587,f13596])).
fof(f20036,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP6(X16,X18,X19) | p(s(bool,X17)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14587,f5467])).
fof(f20023,plain,(
  ( ! [X30,X31,X29] : (sP12(X29,'F_0') | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X29) | ~p(s(bool,X30)) | p(s(bool,X31))) ) | $spl116),
  inference(resolution,[],[f14585,f1480])).
fof(f14585,plain,(
  ( ! [X273,X275,X272,X274] : (sP12(X272,'F_0') | sP12(X273,X272) | sP4(X273,X274,X275)) ) | $spl116),
  inference(subsumption_resolution,[],[f14584,f1247])).
fof(f14584,plain,(
  ( ! [X273,X275,X272,X274] : (p(s(bool,'F_0')) | sP12(X272,'F_0') | sP12(X273,X272) | sP4(X273,X274,X275)) ) | $spl116),
  inference(forward_demodulation,[],[f14504,f5240])).
fof(f14504,plain,(
  ( ! [X273,X275,X272,X274] : (p(s(bool,sK62)) | sP12(X272,'F_0') | sP12(X273,X272) | sP4(X273,X274,X275)) ) | $spl116),
  inference(resolution,[],[f13596,f8067])).
fof(f8067,plain,(
  ( ! [X187,X185,X186,X184] : (sP12(X184,X185) | sP12(sK62,X185) | sP4(X184,X186,X187)) )),
  inference(resolution,[],[f5722,f1078])).
fof(f20022,plain,(
  ( ! [X28,X26,X27] : (sP12(X26,'F_0') | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X26) | ~p(s(bool,X27)) | p(s(bool,X28))) ) | $spl116),
  inference(resolution,[],[f14585,f1309])).
fof(f20031,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X60,X61,X62)) ) | $spl116),
  inference(subsumption_resolution,[],[f20005,f1078])).
fof(f20005,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl116),
  inference(resolution,[],[f14585,f3759])).
fof(f20030,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X57,X58,X59)) ) | $spl116),
  inference(subsumption_resolution,[],[f20004,f1078])).
fof(f20004,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X57,X58,X59) | ~p(s(bool,X57))) ) | $spl116),
  inference(resolution,[],[f14585,f3587])).
fof(f20029,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X51,X52,X53)) ) | $spl116),
  inference(subsumption_resolution,[],[f20002,f1078])).
fof(f20002,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl116),
  inference(resolution,[],[f14585,f1539])).
fof(f14557,plain,(
  ( ! [X194,X192,X195,X193] : (p(s(bool,X192)) | sP4(X193,X194,X195) | sP12(X193,X192)) ) | $spl150),
  inference(cnf_transformation,[],[f14557_D])).
fof(f14557_D,plain,(
  ( ! [X194,X192,X195,X193] : (p(s(bool,X192)) | sP4(X193,X194,X195) | sP12(X193,X192)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f19980,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP4(X40,X42,X43) | p(s(bool,X41)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14585,f15040])).
fof(f19978,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP4(X32,X34,X35) | p(s(bool,X33)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14585,f13596])).
fof(f19974,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP4(X16,X18,X19) | p(s(bool,X17)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14585,f5467])).
fof(f19958,plain,(
  ( ! [X24,X23,X25] : (sP12(X23,'F_0') | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X23) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl116),
  inference(resolution,[],[f14583,f1481])).
fof(f14583,plain,(
  ( ! [X269,X271,X268,X270] : (sP12(X268,'F_0') | sP12(X269,X268) | sP4(X270,X271,X269)) ) | $spl116),
  inference(subsumption_resolution,[],[f14582,f1247])).
fof(f14582,plain,(
  ( ! [X269,X271,X268,X270] : (p(s(bool,'F_0')) | sP12(X268,'F_0') | sP12(X269,X268) | sP4(X270,X271,X269)) ) | $spl116),
  inference(forward_demodulation,[],[f14503,f5240])).
fof(f14503,plain,(
  ( ! [X269,X271,X268,X270] : (p(s(bool,sK62)) | sP12(X268,'F_0') | sP12(X269,X268) | sP4(X270,X271,X269)) ) | $spl116),
  inference(resolution,[],[f13596,f8066])).
fof(f8066,plain,(
  ( ! [X182,X180,X183,X181] : (sP12(X180,X181) | sP12(sK62,X181) | sP4(X182,X183,X180)) )),
  inference(resolution,[],[f5722,f1076])).
fof(f19957,plain,(
  ( ! [X21,X22,X20] : (sP12(X20,'F_0') | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | $spl116),
  inference(resolution,[],[f14583,f1310])).
fof(f19968,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X60,X61,X62)) ) | $spl116),
  inference(subsumption_resolution,[],[f19942,f1076])).
fof(f19942,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X60,X61,X62) | ~p(s(bool,X62))) ) | $spl116),
  inference(resolution,[],[f14583,f3759])).
fof(f19967,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X57,X58,X59)) ) | $spl116),
  inference(subsumption_resolution,[],[f19941,f1076])).
fof(f19941,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X57,X58,X59) | ~p(s(bool,X59))) ) | $spl116),
  inference(resolution,[],[f14583,f3587])).
fof(f19966,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X51,X52,X53)) ) | $spl116),
  inference(subsumption_resolution,[],[f19939,f1076])).
fof(f19939,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP4(X51,X52,X53) | ~p(s(bool,X53))) ) | $spl116),
  inference(resolution,[],[f14583,f1539])).
fof(f14556,plain,(
  ( ! [X191,X189,X190,X188] : (p(s(bool,X188)) | sP4(X190,X191,X189) | sP12(X189,X188)) ) | $spl148),
  inference(cnf_transformation,[],[f14556_D])).
fof(f14556_D,plain,(
  ( ! [X191,X189,X190,X188] : (p(s(bool,X188)) | sP4(X190,X191,X189) | sP12(X189,X188)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f19917,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP4(X42,X43,X40) | p(s(bool,X41)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14583,f15040])).
fof(f19915,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP4(X34,X35,X32) | p(s(bool,X33)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14583,f13596])).
fof(f19911,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP4(X18,X19,X16) | p(s(bool,X17)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14583,f5467])).
fof(f19895,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'F_0') | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X25) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl116),
  inference(resolution,[],[f14581,f1449])).
fof(f14581,plain,(
  ( ! [X265,X267,X264,X266] : (sP12(X264,'F_0') | sP12(X265,X264) | sP1(X265,X266,X267)) ) | $spl116),
  inference(subsumption_resolution,[],[f14580,f1247])).
fof(f14580,plain,(
  ( ! [X265,X267,X264,X266] : (p(s(bool,'F_0')) | sP12(X264,'F_0') | sP12(X265,X264) | sP1(X265,X266,X267)) ) | $spl116),
  inference(forward_demodulation,[],[f14502,f5240])).
fof(f14502,plain,(
  ( ! [X265,X267,X264,X266] : (p(s(bool,sK62)) | sP12(X264,'F_0') | sP12(X265,X264) | sP1(X265,X266,X267)) ) | $spl116),
  inference(resolution,[],[f13596,f8038])).
fof(f8038,plain,(
  ( ! [X74,X72,X75,X73] : (sP12(X72,X73) | sP12(sK62,X73) | sP1(X72,X74,X75)) )),
  inference(resolution,[],[f5722,f947])).
fof(f19894,plain,(
  ( ! [X24,X23,X22] : (sP12(X22,'F_0') | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X22) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl116),
  inference(resolution,[],[f14581,f1278])).
fof(f19905,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,'F_0') | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl116),
  inference(subsumption_resolution,[],[f19890,f13596])).
fof(f19890,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,'F_0') | sP12(X9,X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl116),
  inference(resolution,[],[f14581,f938])).
fof(f19903,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X60,X61,X62)) ) | $spl116),
  inference(subsumption_resolution,[],[f19878,f947])).
fof(f19878,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl116),
  inference(resolution,[],[f14581,f3759])).
fof(f19902,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X57,X58,X59)) ) | $spl116),
  inference(subsumption_resolution,[],[f19877,f947])).
fof(f19877,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X57,X58,X59) | ~p(s(bool,X57))) ) | $spl116),
  inference(resolution,[],[f14581,f3587])).
fof(f19901,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X51,X52,X53)) ) | $spl116),
  inference(subsumption_resolution,[],[f19875,f947])).
fof(f19875,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl116),
  inference(resolution,[],[f14581,f1539])).
fof(f14555,plain,(
  ( ! [X187,X185,X186,X184] : (p(s(bool,X184)) | sP1(X185,X186,X187) | sP12(X185,X184)) ) | $spl146),
  inference(cnf_transformation,[],[f14555_D])).
fof(f14555_D,plain,(
  ( ! [X187,X185,X186,X184] : (p(s(bool,X184)) | sP1(X185,X186,X187) | sP12(X185,X184)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f19853,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP1(X40,X42,X43) | p(s(bool,X41)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14581,f15040])).
fof(f19851,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP1(X32,X34,X35) | p(s(bool,X33)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14581,f13596])).
fof(f19847,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP1(X16,X18,X19) | p(s(bool,X17)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14581,f5467])).
fof(f19830,plain,(
  ( ! [X21,X19,X20] : (sP12(X19,'F_0') | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X19) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl116),
  inference(resolution,[],[f14579,f1450])).
fof(f14579,plain,(
  ( ! [X261,X263,X260,X262] : (sP12(X260,'F_0') | sP12(X261,X260) | sP1(X262,X261,X263)) ) | $spl116),
  inference(subsumption_resolution,[],[f14578,f1247])).
fof(f14578,plain,(
  ( ! [X261,X263,X260,X262] : (p(s(bool,'F_0')) | sP12(X260,'F_0') | sP12(X261,X260) | sP1(X262,X261,X263)) ) | $spl116),
  inference(forward_demodulation,[],[f14501,f5240])).
fof(f14501,plain,(
  ( ! [X261,X263,X260,X262] : (p(s(bool,sK62)) | sP12(X260,'F_0') | sP12(X261,X260) | sP1(X262,X261,X263)) ) | $spl116),
  inference(resolution,[],[f13596,f8037])).
fof(f8037,plain,(
  ( ! [X70,X68,X71,X69] : (sP12(X68,X69) | sP12(sK62,X69) | sP1(X70,X68,X71)) )),
  inference(resolution,[],[f5722,f946])).
fof(f19829,plain,(
  ( ! [X17,X18,X16] : (sP12(X16,'F_0') | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X16) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl116),
  inference(resolution,[],[f14579,f1279])).
fof(f19840,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X60,X61,X62)) ) | $spl116),
  inference(subsumption_resolution,[],[f19815,f946])).
fof(f19815,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X60,X61,X62) | ~p(s(bool,X61))) ) | $spl116),
  inference(resolution,[],[f14579,f3759])).
fof(f19839,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X57,X58,X59)) ) | $spl116),
  inference(subsumption_resolution,[],[f19814,f946])).
fof(f19814,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X57,X58,X59) | ~p(s(bool,X58))) ) | $spl116),
  inference(resolution,[],[f14579,f3587])).
fof(f19838,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X51,X52,X53)) ) | $spl116),
  inference(subsumption_resolution,[],[f19812,f946])).
fof(f19812,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP1(X51,X52,X53) | ~p(s(bool,X52))) ) | $spl116),
  inference(resolution,[],[f14579,f1539])).
fof(f14554,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,X180)) | sP1(X182,X181,X183) | sP12(X181,X180)) ) | $spl144),
  inference(cnf_transformation,[],[f14554_D])).
fof(f14554_D,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,X180)) | sP1(X182,X181,X183) | sP12(X181,X180)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f19790,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP1(X42,X40,X43) | p(s(bool,X41)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14579,f15040])).
fof(f19788,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP1(X34,X32,X35) | p(s(bool,X33)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14579,f13596])).
fof(f19784,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP1(X18,X16,X19) | p(s(bool,X17)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14579,f5467])).
fof(f19768,plain,(
  ( ! [X26,X27,X25] : (sP12(X25,'F_0') | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X25) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl116),
  inference(resolution,[],[f14577,f1446])).
fof(f14577,plain,(
  ( ! [X257,X259,X256,X258] : (sP12(X256,'F_0') | sP12(X257,X256) | sP2(X257,X258,X259)) ) | $spl116),
  inference(subsumption_resolution,[],[f14576,f1247])).
fof(f14576,plain,(
  ( ! [X257,X259,X256,X258] : (p(s(bool,'F_0')) | sP12(X256,'F_0') | sP12(X257,X256) | sP2(X257,X258,X259)) ) | $spl116),
  inference(forward_demodulation,[],[f14500,f5240])).
fof(f14500,plain,(
  ( ! [X257,X259,X256,X258] : (p(s(bool,sK62)) | sP12(X256,'F_0') | sP12(X257,X256) | sP2(X257,X258,X259)) ) | $spl116),
  inference(resolution,[],[f13596,f8035])).
fof(f8035,plain,(
  ( ! [X61,X62,X60,X63] : (sP12(X60,X61) | sP12(sK62,X61) | sP2(X60,X62,X63)) )),
  inference(resolution,[],[f5722,f943])).
fof(f19767,plain,(
  ( ! [X24,X23,X22] : (sP12(X22,'F_0') | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X22) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl116),
  inference(resolution,[],[f14577,f1275])).
fof(f19779,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,'F_0') | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl116),
  inference(subsumption_resolution,[],[f19764,f13596])).
fof(f19764,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,'F_0') | sP12(X13,X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl116),
  inference(resolution,[],[f14577,f939])).
fof(f19776,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X60,X61,X62)) ) | $spl116),
  inference(subsumption_resolution,[],[f19751,f943])).
fof(f19751,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl116),
  inference(resolution,[],[f14577,f3759])).
fof(f19775,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X57,X58,X59)) ) | $spl116),
  inference(subsumption_resolution,[],[f19750,f943])).
fof(f19750,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X57,X58,X59) | ~p(s(bool,X57))) ) | $spl116),
  inference(resolution,[],[f14577,f3587])).
fof(f19774,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X51,X52,X53)) ) | $spl116),
  inference(subsumption_resolution,[],[f19748,f943])).
fof(f19748,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl116),
  inference(resolution,[],[f14577,f1539])).
fof(f14553,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X176)) | sP2(X177,X178,X179) | sP12(X177,X176)) ) | $spl142),
  inference(cnf_transformation,[],[f14553_D])).
fof(f14553_D,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X176)) | sP2(X177,X178,X179) | sP12(X177,X176)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f19726,plain,(
  ( ! [X43,X41,X42,X40] : (sP12(X40,X41) | sP2(X40,X42,X43) | p(s(bool,X41)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f14577,f15040])).
fof(f19724,plain,(
  ( ! [X35,X33,X34,X32] : (sP12(X32,X33) | sP2(X32,X34,X35) | p(s(bool,X33)) | sP12('F_0','F_0')) ) | $spl116),
  inference(resolution,[],[f14577,f13596])).
fof(f19720,plain,(
  ( ! [X19,X17,X18,X16] : (sP12(X16,X17) | sP2(X16,X18,X19) | p(s(bool,X17)) | sP12('F_0',sK62)) ) | $spl116),
  inference(resolution,[],[f14577,f5467])).
fof(f19641,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f19619])).
fof(f19619,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X27))) )),
  inference(resolution,[],[f1519,f5697])).
fof(f5697,plain,(
  ( ! [X233,X234,X232] : (sP9(X233,X234,X232) | sP12(sK62,X232) | p(s(bool,X233))) )),
  inference(resolution,[],[f5231,f1113])).
fof(f1113,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f852])).
fof(f852,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f851])).
fof(f851,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f737])).
fof(f737,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f609,f736,f735])).
fof(f609,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f608])).
fof(f608,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f167])).
fof(f167,axiom,(
  ! [X38,X77,X78] : ((p(s(bool,X78)) <=> (p(s(bool,X77)) & p(s(bool,X38)))) <=> ((p(s(bool,X78)) | ~p(s(bool,X77)) | ~p(s(bool,X38))) & (p(s(bool,X77)) | ~p(s(bool,X78))) & (p(s(bool,X38)) | ~p(s(bool,X78)))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.sat.sat_23_prove_ltb')).
fof(f1519,plain,(
  ( ! [X177,X178] : (~sP9(X177,X178,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X178)) | p(s(bool,X177))) )),
  inference(resolution,[],[f917,f1118])).
fof(f1118,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f852])).
fof(f19643,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f19614])).
fof(f19614,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X17))) )),
  inference(resolution,[],[f1519,f5444])).
fof(f5444,plain,(
  ( ! [X233,X234,X232] : (sP9(X233,X234,X232) | sP12(X232,sK62) | p(s(bool,X233))) )),
  inference(resolution,[],[f5230,f1113])).
fof(f19589,plain,(
  ( ! [X30,X31] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f19570])).
fof(f19570,plain,(
  ( ! [X30,X31] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X31))) )),
  inference(resolution,[],[f1518,f5699])).
fof(f5699,plain,(
  ( ! [X239,X240,X238] : (sP9(X239,X238,X240) | sP12(sK62,X238) | p(s(bool,X239))) )),
  inference(resolution,[],[f5231,f1114])).
fof(f1114,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f852])).
fof(f1518,plain,(
  ( ! [X175,X176] : (~sP9(X175,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X176) | ~p(s(bool,X176)) | p(s(bool,X175))) )),
  inference(resolution,[],[f917,f1118])).
fof(f19591,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f19565])).
fof(f19565,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X21))) )),
  inference(resolution,[],[f1518,f5446])).
fof(f5446,plain,(
  ( ! [X239,X240,X238] : (sP9(X239,X238,X240) | sP12(X238,sK62) | p(s(bool,X239))) )),
  inference(resolution,[],[f5230,f1114])).
fof(f19521,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f19507])).
fof(f19507,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X17)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1514,f5690])).
fof(f1514,plain,(
  ( ! [X167,X168] : (sP9(X167,X168,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X168)) | ~p(s(bool,X167))) )),
  inference(resolution,[],[f917,f1115])).
fof(f1115,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f852])).
fof(f19523,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f19505])).
fof(f19505,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | ~p(s(bool,X13)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1514,f5437])).
fof(f19489,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f19477])).
fof(f19477,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X19)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1513,f5691])).
fof(f1513,plain,(
  ( ! [X165,X166] : (sP9(X165,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X166) | ~p(s(bool,X166)) | ~p(s(bool,X165))) )),
  inference(resolution,[],[f917,f1115])).
fof(f19491,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | ~p(s(bool,X15)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f19475])).
fof(f19475,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | ~p(s(bool,X15)) | ~p(s(bool,X15)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1513,f5438])).
fof(f18926,plain,(
  ( ! [X19] : (sP12('F_0',X19) | sP12(X19,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl258),
  inference(cnf_transformation,[],[f18926_D])).
fof(f18926_D,plain,(
  ( ! [X19] : (sP12('F_0',X19) | sP12(X19,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f19457,plain,(
  ( ! [X30,X31,X29] : (sP12('F_0',X29) | sP12(X29,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl176),
  inference(resolution,[],[f15483,f1510])).
fof(f15483,plain,(
  ( ! [X321,X323,X324,X322] : (sP12('F_0',X321) | sP12(X321,X322) | sP8(X322,X323,X324)) ) | $spl176),
  inference(subsumption_resolution,[],[f15482,f1247])).
fof(f15482,plain,(
  ( ! [X321,X323,X324,X322] : (p(s(bool,'F_0')) | sP12('F_0',X321) | sP12(X321,X322) | sP8(X322,X323,X324)) ) | $spl176),
  inference(forward_demodulation,[],[f15407,f5240])).
fof(f15407,plain,(
  ( ! [X321,X323,X324,X322] : (p(s(bool,sK62)) | sP12('F_0',X321) | sP12(X321,X322) | sP8(X322,X323,X324)) ) | $spl176),
  inference(resolution,[],[f15040,f8349])).
fof(f18925,plain,(
  ( ! [X16] : (sP12('F_0',X16) | sP12(X16,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl256),
  inference(cnf_transformation,[],[f18925_D])).
fof(f18925_D,plain,(
  ( ! [X16] : (sP12('F_0',X16) | sP12(X16,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f19456,plain,(
  ( ! [X28,X26,X27] : (sP12('F_0',X26) | sP12(X26,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl176),
  inference(resolution,[],[f15483,f1339])).
fof(f15443,plain,(
  sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | $spl188),
  inference(cnf_transformation,[],[f15443_D])).
fof(f15443_D,plain,(
  sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f19466,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X78,X79,X80)) ) | $spl176),
  inference(subsumption_resolution,[],[f19445,f1112])).
fof(f19445,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X78,X79,X80) | ~p(s(bool,X78))) ) | $spl176),
  inference(resolution,[],[f15483,f3758])).
fof(f15441,plain,(
  sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | $spl186),
  inference(cnf_transformation,[],[f15441_D])).
fof(f15441_D,plain,(
  sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f19465,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X75,X76,X77)) ) | $spl176),
  inference(subsumption_resolution,[],[f19444,f1112])).
fof(f19444,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X75,X76,X77) | ~p(s(bool,X75))) ) | $spl176),
  inference(resolution,[],[f15483,f3586])).
fof(f15437,plain,(
  sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | $spl182),
  inference(cnf_transformation,[],[f15437_D])).
fof(f15437_D,plain,(
  sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f19464,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X69,X70,X71)) ) | $spl176),
  inference(subsumption_resolution,[],[f19442,f1112])).
fof(f19442,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl176),
  inference(resolution,[],[f15483,f1538])).
fof(f19416,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP8(X45,X46,X47) | p(s(bool,X44)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15483,f15041])).
fof(f15041,plain,(
  ( ! [X304,X303] : (~sP12(X303,X304) | p(s(bool,X304)) | sP12('F_0',X303)) ) | $spl176),
  inference(resolution,[],[f14807,f1165])).
fof(f19414,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP8(X37,X38,X39) | p(s(bool,X36)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15483,f13597])).
fof(f19410,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP8(X21,X22,X23) | p(s(bool,X20)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15483,f5468])).
fof(f5468,plain,(
  ( ! [X304,X303] : (~sP12(X303,X304) | p(s(bool,X304)) | sP12(X303,sK62)) )),
  inference(resolution,[],[f5230,f1165])).
fof(f19391,plain,(
  ( ! [X24,X23,X25] : (sP12('F_0',X23) | sP12(X23,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl176),
  inference(resolution,[],[f15481,f1511])).
fof(f15481,plain,(
  ( ! [X318,X320,X317,X319] : (sP12('F_0',X317) | sP12(X317,X318) | sP8(X319,X318,X320)) ) | $spl176),
  inference(subsumption_resolution,[],[f15480,f1247])).
fof(f15480,plain,(
  ( ! [X318,X320,X317,X319] : (p(s(bool,'F_0')) | sP12('F_0',X317) | sP12(X317,X318) | sP8(X319,X318,X320)) ) | $spl176),
  inference(forward_demodulation,[],[f15406,f5240])).
fof(f15406,plain,(
  ( ! [X318,X320,X317,X319] : (p(s(bool,sK62)) | sP12('F_0',X317) | sP12(X317,X318) | sP8(X319,X318,X320)) ) | $spl176),
  inference(resolution,[],[f15040,f8348])).
fof(f19390,plain,(
  ( ! [X21,X22,X20] : (sP12('F_0',X20) | sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | $spl176),
  inference(resolution,[],[f15481,f1340])).
fof(f19402,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X78,X79,X80)) ) | $spl176),
  inference(subsumption_resolution,[],[f19381,f1111])).
fof(f19381,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X78,X79,X80) | ~p(s(bool,X79))) ) | $spl176),
  inference(resolution,[],[f15481,f3758])).
fof(f19401,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X75,X76,X77)) ) | $spl176),
  inference(subsumption_resolution,[],[f19380,f1111])).
fof(f19380,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X75,X76,X77) | ~p(s(bool,X76))) ) | $spl176),
  inference(resolution,[],[f15481,f3586])).
fof(f19400,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X69,X70,X71)) ) | $spl176),
  inference(subsumption_resolution,[],[f19378,f1111])).
fof(f19378,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X69,X70,X71) | ~p(s(bool,X70))) ) | $spl176),
  inference(resolution,[],[f15481,f1538])).
fof(f19352,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP8(X46,X45,X47) | p(s(bool,X44)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15481,f15041])).
fof(f19350,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP8(X38,X37,X39) | p(s(bool,X36)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15481,f13597])).
fof(f19346,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP8(X22,X21,X23) | p(s(bool,X20)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15481,f5468])).
fof(f19331,plain,(
  ( ! [X24,X23,X25] : (sP12('F_0',X23) | sP12(X23,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X24)) | p(s(bool,X25))) ) | $spl176),
  inference(resolution,[],[f15479,f1494])).
fof(f15479,plain,(
  ( ! [X313,X315,X316,X314] : (sP12('F_0',X313) | sP12(X313,X314) | sP6(X314,X315,X316)) ) | $spl176),
  inference(subsumption_resolution,[],[f15478,f1247])).
fof(f15478,plain,(
  ( ! [X313,X315,X316,X314] : (p(s(bool,'F_0')) | sP12('F_0',X313) | sP12(X313,X314) | sP6(X314,X315,X316)) ) | $spl176),
  inference(forward_demodulation,[],[f15405,f5240])).
fof(f15405,plain,(
  ( ! [X313,X315,X316,X314] : (p(s(bool,sK62)) | sP12('F_0',X313) | sP12(X313,X314) | sP6(X314,X315,X316)) ) | $spl176),
  inference(resolution,[],[f15040,f8335])).
fof(f19330,plain,(
  ( ! [X21,X22,X20] : (sP12('F_0',X20) | sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X21)) | p(s(bool,X22))) ) | $spl176),
  inference(resolution,[],[f15479,f1323])).
fof(f19340,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X78,X79,X80)) ) | $spl176),
  inference(subsumption_resolution,[],[f19321,f1095])).
fof(f19321,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X78,X79,X80) | ~p(s(bool,X78))) ) | $spl176),
  inference(resolution,[],[f15479,f3758])).
fof(f19339,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X75,X76,X77)) ) | $spl176),
  inference(subsumption_resolution,[],[f19320,f1095])).
fof(f19320,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X75,X76,X77) | ~p(s(bool,X75))) ) | $spl176),
  inference(resolution,[],[f15479,f3586])).
fof(f19338,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X69,X70,X71)) ) | $spl176),
  inference(subsumption_resolution,[],[f19318,f1095])).
fof(f19318,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl176),
  inference(resolution,[],[f15479,f1538])).
fof(f19292,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP6(X45,X46,X47) | p(s(bool,X44)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15479,f15041])).
fof(f19290,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP6(X37,X38,X39) | p(s(bool,X36)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15479,f13597])).
fof(f19286,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP6(X21,X22,X23) | p(s(bool,X20)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15479,f5468])).
fof(f19271,plain,(
  ( ! [X30,X31,X29] : (sP12('F_0',X29) | sP12(X29,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X30)) | p(s(bool,X31))) ) | $spl176),
  inference(resolution,[],[f15477,f1480])).
fof(f15477,plain,(
  ( ! [X310,X312,X309,X311] : (sP12('F_0',X309) | sP12(X309,X310) | sP4(X310,X311,X312)) ) | $spl176),
  inference(subsumption_resolution,[],[f15476,f1247])).
fof(f15476,plain,(
  ( ! [X310,X312,X309,X311] : (p(s(bool,'F_0')) | sP12('F_0',X309) | sP12(X309,X310) | sP4(X310,X311,X312)) ) | $spl176),
  inference(forward_demodulation,[],[f15404,f5240])).
fof(f15404,plain,(
  ( ! [X310,X312,X309,X311] : (p(s(bool,sK62)) | sP12('F_0',X309) | sP12(X309,X310) | sP4(X310,X311,X312)) ) | $spl176),
  inference(resolution,[],[f15040,f8319])).
fof(f19270,plain,(
  ( ! [X28,X26,X27] : (sP12('F_0',X26) | sP12(X26,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | ~p(s(bool,X27)) | p(s(bool,X28))) ) | $spl176),
  inference(resolution,[],[f15477,f1309])).
fof(f19280,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X78,X79,X80)) ) | $spl176),
  inference(subsumption_resolution,[],[f19259,f1078])).
fof(f19259,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X78,X79,X80) | ~p(s(bool,X78))) ) | $spl176),
  inference(resolution,[],[f15477,f3758])).
fof(f19279,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X75,X76,X77)) ) | $spl176),
  inference(subsumption_resolution,[],[f19258,f1078])).
fof(f19258,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X75,X76,X77) | ~p(s(bool,X75))) ) | $spl176),
  inference(resolution,[],[f15477,f3586])).
fof(f19278,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X69,X70,X71)) ) | $spl176),
  inference(subsumption_resolution,[],[f19256,f1078])).
fof(f19256,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl176),
  inference(resolution,[],[f15477,f1538])).
fof(f19230,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP4(X45,X46,X47) | p(s(bool,X44)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15477,f15041])).
fof(f19228,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP4(X37,X38,X39) | p(s(bool,X36)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15477,f13597])).
fof(f19224,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP4(X21,X22,X23) | p(s(bool,X20)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15477,f5468])).
fof(f19206,plain,(
  ( ! [X24,X23,X25] : (sP12('F_0',X23) | sP12(X23,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl176),
  inference(resolution,[],[f15475,f1481])).
fof(f15475,plain,(
  ( ! [X308,X306,X305,X307] : (sP12('F_0',X305) | sP12(X305,X306) | sP4(X307,X308,X306)) ) | $spl176),
  inference(subsumption_resolution,[],[f15474,f1247])).
fof(f15474,plain,(
  ( ! [X308,X306,X305,X307] : (p(s(bool,'F_0')) | sP12('F_0',X305) | sP12(X305,X306) | sP4(X307,X308,X306)) ) | $spl176),
  inference(forward_demodulation,[],[f15403,f5240])).
fof(f15403,plain,(
  ( ! [X308,X306,X305,X307] : (p(s(bool,sK62)) | sP12('F_0',X305) | sP12(X305,X306) | sP4(X307,X308,X306)) ) | $spl176),
  inference(resolution,[],[f15040,f8318])).
fof(f19205,plain,(
  ( ! [X21,X22,X20] : (sP12('F_0',X20) | sP12(X20,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | $spl176),
  inference(resolution,[],[f15475,f1310])).
fof(f19217,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X78,X79,X80)) ) | $spl176),
  inference(subsumption_resolution,[],[f19196,f1076])).
fof(f19196,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X78,X79,X80) | ~p(s(bool,X80))) ) | $spl176),
  inference(resolution,[],[f15475,f3758])).
fof(f19216,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X75,X76,X77)) ) | $spl176),
  inference(subsumption_resolution,[],[f19195,f1076])).
fof(f19195,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X75,X76,X77) | ~p(s(bool,X77))) ) | $spl176),
  inference(resolution,[],[f15475,f3586])).
fof(f19215,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X69,X70,X71)) ) | $spl176),
  inference(subsumption_resolution,[],[f19193,f1076])).
fof(f19193,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X69,X70,X71) | ~p(s(bool,X71))) ) | $spl176),
  inference(resolution,[],[f15475,f1538])).
fof(f19167,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP4(X46,X47,X45) | p(s(bool,X44)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15475,f15041])).
fof(f19165,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP4(X38,X39,X37) | p(s(bool,X36)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15475,f13597])).
fof(f19161,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP4(X22,X23,X21) | p(s(bool,X20)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15475,f5468])).
fof(f16333,plain,(
  ( ! [X19] : (sP12('F_0',X19) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X19)) ) | $spl208),
  inference(cnf_transformation,[],[f16333_D])).
fof(f16333_D,plain,(
  ( ! [X19] : (sP12('F_0',X19) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X19)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f19142,plain,(
  ( ! [X66,X64,X65] : (p(s(bool,X64)) | ~p(s(bool,X65)) | sP12('F_0',X66) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X66)) ) | $spl176),
  inference(resolution,[],[f1511,f15463])).
fof(f15463,plain,(
  ( ! [X282,X281,X283,X284] : (sP12('F_0',X281) | sP12(X282,X281) | sP8(X283,X282,X284)) ) | $spl176),
  inference(subsumption_resolution,[],[f15462,f1247])).
fof(f15462,plain,(
  ( ! [X282,X281,X283,X284] : (p(s(bool,'F_0')) | sP12('F_0',X281) | sP12(X282,X281) | sP8(X283,X282,X284)) ) | $spl176),
  inference(forward_demodulation,[],[f15397,f5240])).
fof(f15397,plain,(
  ( ! [X282,X281,X283,X284] : (p(s(bool,sK62)) | sP12('F_0',X281) | sP12(X282,X281) | sP8(X283,X282,X284)) ) | $spl176),
  inference(resolution,[],[f15040,f8096])).
fof(f19140,plain,(
  ( ! [X61,X60] : (p(s(bool,X60)) | ~p(s(bool,X61)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl39),
  inference(resolution,[],[f1511,f14654])).
fof(f14654,plain,(
  ( ! [X50,X48,X49] : (sP8(X49,X48,X50) | sP12('F_0',X48)) ) | $spl39),
  inference(resolution,[],[f5868,f8348])).
fof(f5868,plain,(
  ~sP12(sK62,'F_0') | $spl39),
  inference(cnf_transformation,[],[f5868_D])).
fof(f5868_D,plain,(
  ~sP12(sK62,'F_0') <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f16270,plain,(
  ( ! [X50,X49] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X49) | sP13(X49,X50)) ) | $spl204),
  inference(cnf_transformation,[],[f16270_D])).
fof(f16270_D,plain,(
  ( ! [X50,X49] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X49) | sP13(X49,X50)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f19138,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,X52)) | ~p(s(bool,X53)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X54) | sP13(X54,X55)) ) | $spl126),
  inference(resolution,[],[f1511,f14244])).
fof(f14244,plain,(
  ( ! [X377,X374,X376,X378,X375] : (sP12(X374,X375) | sP13(X375,X376) | sP8(X377,X374,X378)) ) | $spl126),
  inference(resolution,[],[f13902,f1111])).
fof(f16269,plain,(
  ( ! [X45,X46] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X45) | sP13(X46,X45)) ) | $spl202),
  inference(cnf_transformation,[],[f16269_D])).
fof(f16269_D,plain,(
  ( ! [X45,X46] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X45) | sP13(X46,X45)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f19136,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X46) | sP13(X47,X46)) ) | $spl122),
  inference(resolution,[],[f1511,f13993])).
fof(f13993,plain,(
  ( ! [X377,X374,X376,X378,X375] : (sP12(X374,X375) | sP13(X376,X375) | sP8(X377,X374,X378)) ) | $spl122),
  inference(resolution,[],[f13862,f1111])).
fof(f19134,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f1511,f13571])).
fof(f13571,plain,(
  ( ! [X227,X228,X226] : (sP8(X227,X226,X228) | sP12(X226,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f1111])).
fof(f16215,plain,(
  ( ! [X39] : (sP12(X39,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(sK62,X39)) ) | $spl200),
  inference(cnf_transformation,[],[f16215_D])).
fof(f16215_D,plain,(
  ( ! [X39] : (sP12(X39,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(sK62,X39)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f19132,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | sP12(X36,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(sK62,X36)) )),
  inference(resolution,[],[f1511,f8348])).
fof(f16214,plain,(
  ( ! [X33] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X33) | sP12(sK62,X33)) ) | $spl198),
  inference(cnf_transformation,[],[f16214_D])).
fof(f16214_D,plain,(
  ( ! [X33] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X33) | sP12(sK62,X33)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f19130,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X30) | sP12(sK62,X30)) )),
  inference(resolution,[],[f1511,f8096])).
fof(f16213,plain,(
  ( ! [X27] : (sP12(X27,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(X27,sK62)) ) | $spl196),
  inference(cnf_transformation,[],[f16213_D])).
fof(f16213_D,plain,(
  ( ! [X27] : (sP12(X27,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(X27,sK62)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f19128,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(X24,sK62)) )),
  inference(resolution,[],[f1511,f7663])).
fof(f16212,plain,(
  ( ! [X21] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X21) | sP12(X21,sK62)) ) | $spl194),
  inference(cnf_transformation,[],[f16212_D])).
fof(f16212_D,plain,(
  ( ! [X21] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X21) | sP12(X21,sK62)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f19126,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X18) | sP12(X18,sK62)) )),
  inference(resolution,[],[f1511,f7401])).
fof(f19124,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1511,f5695])).
fof(f19122,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1511,f5442])).
fof(f19105,plain,(
  ( ! [X26,X27,X25] : (sP12('F_0',X25) | sP12(X25,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl176),
  inference(resolution,[],[f15473,f1449])).
fof(f15473,plain,(
  ( ! [X302,X304,X301,X303] : (sP12('F_0',X301) | sP12(X301,X302) | sP1(X302,X303,X304)) ) | $spl176),
  inference(subsumption_resolution,[],[f15472,f1247])).
fof(f15472,plain,(
  ( ! [X302,X304,X301,X303] : (p(s(bool,'F_0')) | sP12('F_0',X301) | sP12(X301,X302) | sP1(X302,X303,X304)) ) | $spl176),
  inference(forward_demodulation,[],[f15402,f5240])).
fof(f15402,plain,(
  ( ! [X302,X304,X301,X303] : (p(s(bool,sK62)) | sP12('F_0',X301) | sP12(X301,X302) | sP1(X302,X303,X304)) ) | $spl176),
  inference(resolution,[],[f15040,f8290])).
fof(f19104,plain,(
  ( ! [X24,X23,X22] : (sP12('F_0',X22) | sP12(X22,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl176),
  inference(resolution,[],[f15473,f1278])).
fof(f19116,plain,(
  ( ! [X10,X8,X11,X9] : (sP12('F_0',X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl176),
  inference(subsumption_resolution,[],[f19100,f15041])).
fof(f19100,plain,(
  ( ! [X10,X8,X11,X9] : (sP12('F_0',X8) | sP12(X8,X9) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl176),
  inference(resolution,[],[f15473,f938])).
fof(f19114,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X78,X79,X80)) ) | $spl176),
  inference(subsumption_resolution,[],[f19094,f947])).
fof(f19094,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X78,X79,X80) | ~p(s(bool,X78))) ) | $spl176),
  inference(resolution,[],[f15473,f3758])).
fof(f19113,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X75,X76,X77)) ) | $spl176),
  inference(subsumption_resolution,[],[f19093,f947])).
fof(f19093,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X75,X76,X77) | ~p(s(bool,X75))) ) | $spl176),
  inference(resolution,[],[f15473,f3586])).
fof(f19112,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X69,X70,X71)) ) | $spl176),
  inference(subsumption_resolution,[],[f19091,f947])).
fof(f19091,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl176),
  inference(resolution,[],[f15473,f1538])).
fof(f19065,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP1(X45,X46,X47) | p(s(bool,X44)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15473,f15041])).
fof(f19063,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP1(X37,X38,X39) | p(s(bool,X36)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15473,f13597])).
fof(f19059,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP1(X21,X22,X23) | p(s(bool,X20)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15473,f5468])).
fof(f19040,plain,(
  ( ! [X21,X19,X20] : (sP12('F_0',X19) | sP12(X19,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl176),
  inference(resolution,[],[f15471,f1450])).
fof(f15471,plain,(
  ( ! [X300,X298,X297,X299] : (sP12('F_0',X297) | sP12(X297,X298) | sP1(X299,X298,X300)) ) | $spl176),
  inference(subsumption_resolution,[],[f15470,f1247])).
fof(f15470,plain,(
  ( ! [X300,X298,X297,X299] : (p(s(bool,'F_0')) | sP12('F_0',X297) | sP12(X297,X298) | sP1(X299,X298,X300)) ) | $spl176),
  inference(forward_demodulation,[],[f15401,f5240])).
fof(f15401,plain,(
  ( ! [X300,X298,X297,X299] : (p(s(bool,sK62)) | sP12('F_0',X297) | sP12(X297,X298) | sP1(X299,X298,X300)) ) | $spl176),
  inference(resolution,[],[f15040,f8289])).
fof(f19039,plain,(
  ( ! [X17,X18,X16] : (sP12('F_0',X16) | sP12(X16,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl176),
  inference(resolution,[],[f15471,f1279])).
fof(f19051,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X78,X79,X80)) ) | $spl176),
  inference(subsumption_resolution,[],[f19031,f946])).
fof(f19031,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X78,X79,X80) | ~p(s(bool,X79))) ) | $spl176),
  inference(resolution,[],[f15471,f3758])).
fof(f19050,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X75,X76,X77)) ) | $spl176),
  inference(subsumption_resolution,[],[f19030,f946])).
fof(f19030,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X75,X76,X77) | ~p(s(bool,X76))) ) | $spl176),
  inference(resolution,[],[f15471,f3586])).
fof(f19049,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X69,X70,X71)) ) | $spl176),
  inference(subsumption_resolution,[],[f19028,f946])).
fof(f19028,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X69,X70,X71) | ~p(s(bool,X70))) ) | $spl176),
  inference(resolution,[],[f15471,f1538])).
fof(f19002,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP1(X46,X45,X47) | p(s(bool,X44)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15471,f15041])).
fof(f19000,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP1(X38,X37,X39) | p(s(bool,X36)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15471,f13597])).
fof(f18996,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP1(X22,X21,X23) | p(s(bool,X20)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15471,f5468])).
fof(f18978,plain,(
  ( ! [X26,X27,X25] : (sP12('F_0',X25) | sP12(X25,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl176),
  inference(resolution,[],[f15469,f1446])).
fof(f15469,plain,(
  ( ! [X294,X296,X293,X295] : (sP12('F_0',X293) | sP12(X293,X294) | sP2(X294,X295,X296)) ) | $spl176),
  inference(subsumption_resolution,[],[f15468,f1247])).
fof(f15468,plain,(
  ( ! [X294,X296,X293,X295] : (p(s(bool,'F_0')) | sP12('F_0',X293) | sP12(X293,X294) | sP2(X294,X295,X296)) ) | $spl176),
  inference(forward_demodulation,[],[f15400,f5240])).
fof(f15400,plain,(
  ( ! [X294,X296,X293,X295] : (p(s(bool,sK62)) | sP12('F_0',X293) | sP12(X293,X294) | sP2(X294,X295,X296)) ) | $spl176),
  inference(resolution,[],[f15040,f8287])).
fof(f18977,plain,(
  ( ! [X24,X23,X22] : (sP12('F_0',X22) | sP12(X22,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl176),
  inference(resolution,[],[f15469,f1275])).
fof(f18990,plain,(
  ( ! [X14,X12,X15,X13] : (sP12('F_0',X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl176),
  inference(subsumption_resolution,[],[f18974,f15041])).
fof(f18974,plain,(
  ( ! [X14,X12,X15,X13] : (sP12('F_0',X12) | sP12(X12,X13) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl176),
  inference(resolution,[],[f15469,f939])).
fof(f18987,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X78,X79,X80)) ) | $spl176),
  inference(subsumption_resolution,[],[f18967,f943])).
fof(f18967,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X78,X79,X80) | ~p(s(bool,X78))) ) | $spl176),
  inference(resolution,[],[f15469,f3758])).
fof(f18986,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X75,X76,X77)) ) | $spl176),
  inference(subsumption_resolution,[],[f18966,f943])).
fof(f18966,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X75,X76,X77) | ~p(s(bool,X75))) ) | $spl176),
  inference(resolution,[],[f15469,f3586])).
fof(f18985,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X69,X70,X71)) ) | $spl176),
  inference(subsumption_resolution,[],[f18964,f943])).
fof(f18964,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X69,X70,X71) | ~p(s(bool,X69))) ) | $spl176),
  inference(resolution,[],[f15469,f1538])).
fof(f18938,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP2(X45,X46,X47) | p(s(bool,X44)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15469,f15041])).
fof(f18936,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP2(X37,X38,X39) | p(s(bool,X36)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15469,f13597])).
fof(f18932,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP2(X21,X22,X23) | p(s(bool,X20)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15469,f5468])).
fof(f18912,plain,(
  ( ! [X21,X19,X20] : (sP12('F_0',X19) | sP12(X19,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl176),
  inference(resolution,[],[f15467,f1447])).
fof(f15467,plain,(
  ( ! [X292,X290,X289,X291] : (sP12('F_0',X289) | sP12(X289,X290) | sP2(X291,X290,X292)) ) | $spl176),
  inference(subsumption_resolution,[],[f15466,f1247])).
fof(f15466,plain,(
  ( ! [X292,X290,X289,X291] : (p(s(bool,'F_0')) | sP12('F_0',X289) | sP12(X289,X290) | sP2(X291,X290,X292)) ) | $spl176),
  inference(forward_demodulation,[],[f15399,f5240])).
fof(f15399,plain,(
  ( ! [X292,X290,X289,X291] : (p(s(bool,sK62)) | sP12('F_0',X289) | sP12(X289,X290) | sP2(X291,X290,X292)) ) | $spl176),
  inference(resolution,[],[f15040,f8286])).
fof(f18911,plain,(
  ( ! [X17,X18,X16] : (sP12('F_0',X16) | sP12(X16,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl176),
  inference(resolution,[],[f15467,f1276])).
fof(f18923,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X78,X79,X80)) ) | $spl176),
  inference(subsumption_resolution,[],[f18903,f942])).
fof(f18903,plain,(
  ( ! [X80,X78,X79] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X78,X79,X80) | ~p(s(bool,X79))) ) | $spl176),
  inference(resolution,[],[f15467,f3758])).
fof(f18922,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X75,X76,X77)) ) | $spl176),
  inference(subsumption_resolution,[],[f18902,f942])).
fof(f18902,plain,(
  ( ! [X76,X77,X75] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X75,X76,X77) | ~p(s(bool,X76))) ) | $spl176),
  inference(resolution,[],[f15467,f3586])).
fof(f18921,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X69,X70,X71)) ) | $spl176),
  inference(subsumption_resolution,[],[f18900,f942])).
fof(f18900,plain,(
  ( ! [X70,X71,X69] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X69,X70,X71) | ~p(s(bool,X70))) ) | $spl176),
  inference(resolution,[],[f15467,f1538])).
fof(f18874,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP2(X46,X45,X47) | p(s(bool,X44)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15467,f15041])).
fof(f18872,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP2(X38,X37,X39) | p(s(bool,X36)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15467,f13597])).
fof(f18868,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP2(X22,X21,X23) | p(s(bool,X20)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15467,f5468])).
fof(f18852,plain,(
  ( ! [X28,X26,X27] : (sP12('F_0',X26) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X26) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl176),
  inference(resolution,[],[f15465,f1510])).
fof(f15465,plain,(
  ( ! [X288,X285,X287,X286] : (sP12('F_0',X285) | sP12(X286,X285) | sP8(X286,X287,X288)) ) | $spl176),
  inference(subsumption_resolution,[],[f15464,f1247])).
fof(f15464,plain,(
  ( ! [X288,X285,X287,X286] : (p(s(bool,'F_0')) | sP12('F_0',X285) | sP12(X286,X285) | sP8(X286,X287,X288)) ) | $spl176),
  inference(forward_demodulation,[],[f15398,f5240])).
fof(f15398,plain,(
  ( ! [X288,X285,X287,X286] : (p(s(bool,sK62)) | sP12('F_0',X285) | sP12(X286,X285) | sP8(X286,X287,X288)) ) | $spl176),
  inference(resolution,[],[f15040,f8097])).
fof(f16332,plain,(
  ( ! [X16] : (sP12('F_0',X16) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X16)) ) | $spl206),
  inference(cnf_transformation,[],[f16332_D])).
fof(f16332_D,plain,(
  ( ! [X16] : (sP12('F_0',X16) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X16)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f18851,plain,(
  ( ! [X24,X23,X25] : (sP12('F_0',X23) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X23) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl176),
  inference(resolution,[],[f15465,f1339])).
fof(f18861,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X60,X61,X62)) ) | $spl176),
  inference(subsumption_resolution,[],[f18835,f1112])).
fof(f18835,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl176),
  inference(resolution,[],[f15465,f3759])).
fof(f18860,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X57,X58,X59)) ) | $spl176),
  inference(subsumption_resolution,[],[f18834,f1112])).
fof(f18834,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X57,X58,X59) | ~p(s(bool,X57))) ) | $spl176),
  inference(resolution,[],[f15465,f3587])).
fof(f18859,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X51,X52,X53)) ) | $spl176),
  inference(subsumption_resolution,[],[f18832,f1112])).
fof(f18832,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl176),
  inference(resolution,[],[f15465,f1539])).
fof(f18812,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP8(X44,X46,X47) | p(s(bool,X45)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15465,f15041])).
fof(f18810,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP8(X36,X38,X39) | p(s(bool,X37)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15465,f13597])).
fof(f18806,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP8(X20,X22,X23) | p(s(bool,X21)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15465,f5468])).
fof(f18787,plain,(
  ( ! [X62,X63] : (p(s(bool,X62)) | ~p(s(bool,X63)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl39),
  inference(resolution,[],[f1510,f14655])).
fof(f14655,plain,(
  ( ! [X52,X53,X51] : (sP8(X51,X52,X53) | sP12('F_0',X51)) ) | $spl39),
  inference(resolution,[],[f5868,f8349])).
fof(f18785,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X58) | sP13(X58,X59)) ) | $spl126),
  inference(resolution,[],[f1510,f14245])).
fof(f14245,plain,(
  ( ! [X383,X379,X380,X382,X381] : (sP12(X379,X380) | sP13(X380,X381) | sP8(X379,X382,X383)) ) | $spl126),
  inference(resolution,[],[f13902,f1112])).
fof(f18783,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,X48)) | ~p(s(bool,X49)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X50) | sP13(X51,X50)) ) | $spl122),
  inference(resolution,[],[f1510,f13994])).
fof(f13994,plain,(
  ( ! [X383,X379,X380,X382,X381] : (sP12(X379,X380) | sP13(X381,X380) | sP8(X379,X382,X383)) ) | $spl122),
  inference(resolution,[],[f13862,f1112])).
fof(f18781,plain,(
  ( ! [X43,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f1510,f13572])).
fof(f13572,plain,(
  ( ! [X231,X229,X230] : (sP8(X229,X230,X231) | sP12(X229,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f1112])).
fof(f18779,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | sP12(X39,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(sK62,X39)) )),
  inference(resolution,[],[f1510,f8349])).
fof(f18777,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X33) | sP12(sK62,X33)) )),
  inference(resolution,[],[f1510,f8097])).
fof(f18775,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | sP12(X27,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(X27,sK62)) )),
  inference(resolution,[],[f1510,f7664])).
fof(f18773,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X21) | sP12(X21,sK62)) )),
  inference(resolution,[],[f1510,f7402])).
fof(f18771,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1510,f5696])).
fof(f18769,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1510,f5443])).
fof(f18751,plain,(
  ( ! [X21,X22,X20] : (sP12('F_0',X20) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | $spl176),
  inference(resolution,[],[f15463,f1340])).
fof(f18761,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X60,X61,X62)) ) | $spl176),
  inference(subsumption_resolution,[],[f18736,f1111])).
fof(f18736,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X60,X61,X62) | ~p(s(bool,X61))) ) | $spl176),
  inference(resolution,[],[f15463,f3759])).
fof(f18760,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X57,X58,X59)) ) | $spl176),
  inference(subsumption_resolution,[],[f18735,f1111])).
fof(f18735,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X57,X58,X59) | ~p(s(bool,X58))) ) | $spl176),
  inference(resolution,[],[f15463,f3587])).
fof(f18759,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X51,X52,X53)) ) | $spl176),
  inference(subsumption_resolution,[],[f18733,f1111])).
fof(f18733,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X51,X52,X53) | ~p(s(bool,X52))) ) | $spl176),
  inference(resolution,[],[f15463,f1539])).
fof(f18713,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP8(X46,X44,X47) | p(s(bool,X45)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15463,f15041])).
fof(f18711,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP8(X38,X36,X39) | p(s(bool,X37)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15463,f13597])).
fof(f18707,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP8(X22,X20,X23) | p(s(bool,X21)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15463,f5468])).
fof(f18692,plain,(
  ( ! [X24,X23,X25] : (sP12('F_0',X23) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X23) | p(s(bool,X24)) | p(s(bool,X25))) ) | $spl176),
  inference(resolution,[],[f15461,f1494])).
fof(f15461,plain,(
  ( ! [X280,X277,X279,X278] : (sP12('F_0',X277) | sP12(X278,X277) | sP6(X278,X279,X280)) ) | $spl176),
  inference(subsumption_resolution,[],[f15460,f1247])).
fof(f15460,plain,(
  ( ! [X280,X277,X279,X278] : (p(s(bool,'F_0')) | sP12('F_0',X277) | sP12(X278,X277) | sP6(X278,X279,X280)) ) | $spl176),
  inference(forward_demodulation,[],[f15396,f5240])).
fof(f15396,plain,(
  ( ! [X280,X277,X279,X278] : (p(s(bool,sK62)) | sP12('F_0',X277) | sP12(X278,X277) | sP6(X278,X279,X280)) ) | $spl176),
  inference(resolution,[],[f15040,f8083])).
fof(f18691,plain,(
  ( ! [X21,X22,X20] : (sP12('F_0',X20) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | p(s(bool,X21)) | p(s(bool,X22))) ) | $spl176),
  inference(resolution,[],[f15461,f1323])).
fof(f18701,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X60,X61,X62)) ) | $spl176),
  inference(subsumption_resolution,[],[f18676,f1095])).
fof(f18676,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl176),
  inference(resolution,[],[f15461,f3759])).
fof(f18700,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X57,X58,X59)) ) | $spl176),
  inference(subsumption_resolution,[],[f18675,f1095])).
fof(f18675,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X57,X58,X59) | ~p(s(bool,X57))) ) | $spl176),
  inference(resolution,[],[f15461,f3587])).
fof(f18699,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X51,X52,X53)) ) | $spl176),
  inference(subsumption_resolution,[],[f18673,f1095])).
fof(f18673,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl176),
  inference(resolution,[],[f15461,f1539])).
fof(f18653,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP6(X44,X46,X47) | p(s(bool,X45)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15461,f15041])).
fof(f18651,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP6(X36,X38,X39) | p(s(bool,X37)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15461,f13597])).
fof(f18647,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP6(X20,X22,X23) | p(s(bool,X21)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15461,f5468])).
fof(f18632,plain,(
  ( ! [X30,X31,X29] : (sP12('F_0',X29) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X29) | ~p(s(bool,X30)) | p(s(bool,X31))) ) | $spl176),
  inference(resolution,[],[f15459,f1480])).
fof(f15459,plain,(
  ( ! [X273,X275,X276,X274] : (sP12('F_0',X273) | sP12(X274,X273) | sP4(X274,X275,X276)) ) | $spl176),
  inference(subsumption_resolution,[],[f15458,f1247])).
fof(f15458,plain,(
  ( ! [X273,X275,X276,X274] : (p(s(bool,'F_0')) | sP12('F_0',X273) | sP12(X274,X273) | sP4(X274,X275,X276)) ) | $spl176),
  inference(forward_demodulation,[],[f15395,f5240])).
fof(f15395,plain,(
  ( ! [X273,X275,X276,X274] : (p(s(bool,sK62)) | sP12('F_0',X273) | sP12(X274,X273) | sP4(X274,X275,X276)) ) | $spl176),
  inference(resolution,[],[f15040,f8067])).
fof(f18631,plain,(
  ( ! [X28,X26,X27] : (sP12('F_0',X26) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X26) | ~p(s(bool,X27)) | p(s(bool,X28))) ) | $spl176),
  inference(resolution,[],[f15459,f1309])).
fof(f18641,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X60,X61,X62)) ) | $spl176),
  inference(subsumption_resolution,[],[f18614,f1078])).
fof(f18614,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl176),
  inference(resolution,[],[f15459,f3759])).
fof(f18640,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X57,X58,X59)) ) | $spl176),
  inference(subsumption_resolution,[],[f18613,f1078])).
fof(f18613,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X57,X58,X59) | ~p(s(bool,X57))) ) | $spl176),
  inference(resolution,[],[f15459,f3587])).
fof(f18639,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X51,X52,X53)) ) | $spl176),
  inference(subsumption_resolution,[],[f18611,f1078])).
fof(f18611,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl176),
  inference(resolution,[],[f15459,f1539])).
fof(f18591,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP4(X44,X46,X47) | p(s(bool,X45)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15459,f15041])).
fof(f18589,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP4(X36,X38,X39) | p(s(bool,X37)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15459,f13597])).
fof(f18585,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP4(X20,X22,X23) | p(s(bool,X21)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15459,f5468])).
fof(f18567,plain,(
  ( ! [X24,X23,X25] : (sP12('F_0',X23) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X23) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl176),
  inference(resolution,[],[f15457,f1481])).
fof(f15457,plain,(
  ( ! [X269,X271,X270,X272] : (sP12('F_0',X269) | sP12(X270,X269) | sP4(X271,X272,X270)) ) | $spl176),
  inference(subsumption_resolution,[],[f15456,f1247])).
fof(f15456,plain,(
  ( ! [X269,X271,X270,X272] : (p(s(bool,'F_0')) | sP12('F_0',X269) | sP12(X270,X269) | sP4(X271,X272,X270)) ) | $spl176),
  inference(forward_demodulation,[],[f15394,f5240])).
fof(f15394,plain,(
  ( ! [X269,X271,X270,X272] : (p(s(bool,sK62)) | sP12('F_0',X269) | sP12(X270,X269) | sP4(X271,X272,X270)) ) | $spl176),
  inference(resolution,[],[f15040,f8066])).
fof(f18566,plain,(
  ( ! [X21,X22,X20] : (sP12('F_0',X20) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | p(s(bool,X21)) | ~p(s(bool,X22))) ) | $spl176),
  inference(resolution,[],[f15457,f1310])).
fof(f18578,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X60,X61,X62)) ) | $spl176),
  inference(subsumption_resolution,[],[f18551,f1076])).
fof(f18551,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X60,X61,X62) | ~p(s(bool,X62))) ) | $spl176),
  inference(resolution,[],[f15457,f3759])).
fof(f18577,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X57,X58,X59)) ) | $spl176),
  inference(subsumption_resolution,[],[f18550,f1076])).
fof(f18550,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X57,X58,X59) | ~p(s(bool,X59))) ) | $spl176),
  inference(resolution,[],[f15457,f3587])).
fof(f18576,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X51,X52,X53)) ) | $spl176),
  inference(subsumption_resolution,[],[f18548,f1076])).
fof(f18548,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X51,X52,X53) | ~p(s(bool,X53))) ) | $spl176),
  inference(resolution,[],[f15457,f1539])).
fof(f18528,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP4(X46,X47,X44) | p(s(bool,X45)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15457,f15041])).
fof(f18526,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP4(X38,X39,X36) | p(s(bool,X37)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15457,f13597])).
fof(f18522,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP4(X22,X23,X20) | p(s(bool,X21)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15457,f5468])).
fof(f18504,plain,(
  ( ! [X26,X27,X25] : (sP12('F_0',X25) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X25) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl176),
  inference(resolution,[],[f15455,f1449])).
fof(f15455,plain,(
  ( ! [X265,X267,X268,X266] : (sP12('F_0',X265) | sP12(X266,X265) | sP1(X266,X267,X268)) ) | $spl176),
  inference(subsumption_resolution,[],[f15454,f1247])).
fof(f15454,plain,(
  ( ! [X265,X267,X268,X266] : (p(s(bool,'F_0')) | sP12('F_0',X265) | sP12(X266,X265) | sP1(X266,X267,X268)) ) | $spl176),
  inference(forward_demodulation,[],[f15393,f5240])).
fof(f15393,plain,(
  ( ! [X265,X267,X268,X266] : (p(s(bool,sK62)) | sP12('F_0',X265) | sP12(X266,X265) | sP1(X266,X267,X268)) ) | $spl176),
  inference(resolution,[],[f15040,f8038])).
fof(f18503,plain,(
  ( ! [X24,X23,X22] : (sP12('F_0',X22) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X22) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl176),
  inference(resolution,[],[f15455,f1278])).
fof(f18515,plain,(
  ( ! [X10,X8,X11,X9] : (sP12('F_0',X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl176),
  inference(subsumption_resolution,[],[f18499,f15040])).
fof(f18499,plain,(
  ( ! [X10,X8,X11,X9] : (sP12('F_0',X8) | sP12(X9,X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl176),
  inference(resolution,[],[f15455,f938])).
fof(f18513,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X60,X61,X62)) ) | $spl176),
  inference(subsumption_resolution,[],[f18487,f947])).
fof(f18487,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl176),
  inference(resolution,[],[f15455,f3759])).
fof(f18512,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X57,X58,X59)) ) | $spl176),
  inference(subsumption_resolution,[],[f18486,f947])).
fof(f18486,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X57,X58,X59) | ~p(s(bool,X57))) ) | $spl176),
  inference(resolution,[],[f15455,f3587])).
fof(f18511,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X51,X52,X53)) ) | $spl176),
  inference(subsumption_resolution,[],[f18484,f947])).
fof(f18484,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl176),
  inference(resolution,[],[f15455,f1539])).
fof(f18464,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP1(X44,X46,X47) | p(s(bool,X45)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15455,f15041])).
fof(f18462,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP1(X36,X38,X39) | p(s(bool,X37)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15455,f13597])).
fof(f18458,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP1(X20,X22,X23) | p(s(bool,X21)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15455,f5468])).
fof(f18418,plain,(
  ( ! [X21,X19,X20] : (sP12('F_0',X19) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X19) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl176),
  inference(resolution,[],[f15453,f1450])).
fof(f15453,plain,(
  ( ! [X261,X263,X262,X264] : (sP12('F_0',X261) | sP12(X262,X261) | sP1(X263,X262,X264)) ) | $spl176),
  inference(subsumption_resolution,[],[f15452,f1247])).
fof(f15452,plain,(
  ( ! [X261,X263,X262,X264] : (p(s(bool,'F_0')) | sP12('F_0',X261) | sP12(X262,X261) | sP1(X263,X262,X264)) ) | $spl176),
  inference(forward_demodulation,[],[f15392,f5240])).
fof(f15392,plain,(
  ( ! [X261,X263,X262,X264] : (p(s(bool,sK62)) | sP12('F_0',X261) | sP12(X262,X261) | sP1(X263,X262,X264)) ) | $spl176),
  inference(resolution,[],[f15040,f8037])).
fof(f18417,plain,(
  ( ! [X17,X18,X16] : (sP12('F_0',X16) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X16) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl176),
  inference(resolution,[],[f15453,f1279])).
fof(f18429,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X60,X61,X62)) ) | $spl176),
  inference(subsumption_resolution,[],[f18403,f946])).
fof(f18403,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X60,X61,X62) | ~p(s(bool,X61))) ) | $spl176),
  inference(resolution,[],[f15453,f3759])).
fof(f18428,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X57,X58,X59)) ) | $spl176),
  inference(subsumption_resolution,[],[f18402,f946])).
fof(f18402,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X57,X58,X59) | ~p(s(bool,X58))) ) | $spl176),
  inference(resolution,[],[f15453,f3587])).
fof(f18427,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X51,X52,X53)) ) | $spl176),
  inference(subsumption_resolution,[],[f18400,f946])).
fof(f18400,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X51,X52,X53) | ~p(s(bool,X52))) ) | $spl176),
  inference(resolution,[],[f15453,f1539])).
fof(f18380,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP1(X46,X44,X47) | p(s(bool,X45)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15453,f15041])).
fof(f18378,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP1(X38,X36,X39) | p(s(bool,X37)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15453,f13597])).
fof(f18374,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP1(X22,X20,X23) | p(s(bool,X21)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15453,f5468])).
fof(f18354,plain,(
  ( ! [X26,X27,X25] : (sP12('F_0',X25) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X25) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl176),
  inference(resolution,[],[f15451,f1446])).
fof(f15451,plain,(
  ( ! [X257,X259,X260,X258] : (sP12('F_0',X257) | sP12(X258,X257) | sP2(X258,X259,X260)) ) | $spl176),
  inference(subsumption_resolution,[],[f15450,f1247])).
fof(f15450,plain,(
  ( ! [X257,X259,X260,X258] : (p(s(bool,'F_0')) | sP12('F_0',X257) | sP12(X258,X257) | sP2(X258,X259,X260)) ) | $spl176),
  inference(forward_demodulation,[],[f15391,f5240])).
fof(f15391,plain,(
  ( ! [X257,X259,X260,X258] : (p(s(bool,sK62)) | sP12('F_0',X257) | sP12(X258,X257) | sP2(X258,X259,X260)) ) | $spl176),
  inference(resolution,[],[f15040,f8035])).
fof(f18353,plain,(
  ( ! [X24,X23,X22] : (sP12('F_0',X22) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X22) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl176),
  inference(resolution,[],[f15451,f1275])).
fof(f18368,plain,(
  ( ! [X14,X12,X15,X13] : (sP12('F_0',X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl176),
  inference(subsumption_resolution,[],[f18350,f15040])).
fof(f18350,plain,(
  ( ! [X14,X12,X15,X13] : (sP12('F_0',X12) | sP12(X13,X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) ) | $spl176),
  inference(resolution,[],[f15451,f939])).
fof(f18365,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X60,X61,X62)) ) | $spl176),
  inference(subsumption_resolution,[],[f18337,f943])).
fof(f18337,plain,(
  ( ! [X61,X62,X60] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X60,X61,X62) | ~p(s(bool,X60))) ) | $spl176),
  inference(resolution,[],[f15451,f3759])).
fof(f18364,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X57,X58,X59)) ) | $spl176),
  inference(subsumption_resolution,[],[f18336,f943])).
fof(f18336,plain,(
  ( ! [X59,X57,X58] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X57,X58,X59) | ~p(s(bool,X57))) ) | $spl176),
  inference(resolution,[],[f15451,f3587])).
fof(f18363,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X51,X52,X53)) ) | $spl176),
  inference(subsumption_resolution,[],[f18334,f943])).
fof(f18334,plain,(
  ( ! [X52,X53,X51] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X51,X52,X53) | ~p(s(bool,X51))) ) | $spl176),
  inference(resolution,[],[f15451,f1539])).
fof(f18314,plain,(
  ( ! [X47,X45,X46,X44] : (sP12(X44,X45) | sP2(X44,X46,X47) | p(s(bool,X45)) | sP12('F_0','F_0')) ) | $spl176),
  inference(resolution,[],[f15451,f15041])).
fof(f18312,plain,(
  ( ! [X39,X37,X38,X36] : (sP12(X36,X37) | sP2(X36,X38,X39) | p(s(bool,X37)) | sP12('F_0','F_0')) ) | ($spl116 | $spl176)),
  inference(resolution,[],[f15451,f13597])).
fof(f18308,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP2(X20,X22,X23) | p(s(bool,X21)) | sP12('F_0',sK62)) ) | $spl176),
  inference(resolution,[],[f15451,f5468])).
fof(f18290,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl39),
  inference(resolution,[],[f1494,f14653])).
fof(f14653,plain,(
  ( ! [X47,X45,X46] : (sP6(X45,X46,X47) | sP12('F_0',X45)) ) | $spl39),
  inference(resolution,[],[f5868,f8335])).
fof(f18289,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X32) | sP13(X32,X33)) ) | $spl126),
  inference(resolution,[],[f1494,f14231])).
fof(f14231,plain,(
  ( ! [X313,X310,X312,X309,X311] : (sP12(X309,X310) | sP13(X310,X311) | sP6(X309,X312,X313)) ) | $spl126),
  inference(resolution,[],[f13902,f1095])).
fof(f18288,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X28) | sP13(X29,X28)) ) | $spl122),
  inference(resolution,[],[f1494,f13980])).
fof(f13980,plain,(
  ( ! [X313,X310,X312,X309,X311] : (sP12(X309,X310) | sP13(X311,X310) | sP6(X309,X312,X313)) ) | $spl122),
  inference(resolution,[],[f13862,f1095])).
fof(f18287,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f1494,f13558])).
fof(f13558,plain,(
  ( ! [X189,X187,X188] : (sP6(X187,X188,X189) | sP12(X187,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f1095])).
fof(f18286,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | sP12(X23,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(sK62,X23)) )),
  inference(resolution,[],[f1494,f8335])).
fof(f18285,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X20) | sP12(sK62,X20)) )),
  inference(resolution,[],[f1494,f8083])).
fof(f18284,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | sP12(X17,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(X17,sK62)) )),
  inference(resolution,[],[f1494,f7650])).
fof(f18283,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X14) | sP12(X14,sK62)) )),
  inference(resolution,[],[f1494,f7388])).
fof(f18282,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1494,f5682])).
fof(f5682,plain,(
  ( ! [X189,X187,X188] : (sP6(X187,X188,X189) | sP12(sK62,X187)) )),
  inference(resolution,[],[f5231,f1095])).
fof(f18281,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1494,f5429])).
fof(f5429,plain,(
  ( ! [X189,X187,X188] : (sP6(X187,X188,X189) | sP12(X187,sK62)) )),
  inference(resolution,[],[f5230,f1095])).
fof(f13905,plain,(
  ( ! [X21,X20] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | sP13(X20,X21)) ) | $spl128),
  inference(cnf_transformation,[],[f13905_D])).
fof(f13905_D,plain,(
  ( ! [X21,X20] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | sP13(X20,X21)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f18273,plain,(
  ( ! [X30,X31,X29,X32] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X29) | sP13(X29,X30) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl126),
  inference(resolution,[],[f14245,f1339])).
fof(f17844,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP8(X45,X47,X48)) ) | $spl252),
  inference(cnf_transformation,[],[f17844_D])).
fof(f17844_D,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP8(X45,X47,X48)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f18264,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP8(X50,X52,X53) | p(s(bool,X54)) | sP12('F_0',X51)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14245,f15188])).
fof(f18263,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP8(X45,X47,X48) | p(s(bool,X46)) | sP12('F_0',X49)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14245,f15187])).
fof(f17843,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP8(X35,X37,X38)) ) | $spl250),
  inference(cnf_transformation,[],[f17843_D])).
fof(f17843_D,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP8(X35,X37,X38)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f18262,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP8(X40,X42,X43) | p(s(bool,X44)) | sP12(X41,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14245,f13742])).
fof(f18261,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP8(X35,X37,X38) | p(s(bool,X36)) | sP12(X39,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14245,f13741])).
fof(f18259,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP8(X25,X27,X28) | p(s(bool,X26)) | sP12(sK62,X29)) ) | $spl126),
  inference(resolution,[],[f14245,f5864])).
fof(f18257,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP8(X15,X17,X18) | p(s(bool,X16)) | sP12(X19,sK62)) ) | $spl126),
  inference(resolution,[],[f14245,f5611])).
fof(f18256,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP8(X10,X12,X13) | p(s(bool,X14)) | p(s(bool,X11))) ) | $spl126),
  inference(resolution,[],[f14245,f5117])).
fof(f18275,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X5,X7,X8) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(subsumption_resolution,[],[f18255,f1215])).
fof(f18255,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X5,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(resolution,[],[f14245,f1185])).
fof(f18274,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X0,X2,X3) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(subsumption_resolution,[],[f18254,f1215])).
fof(f18254,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X0,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(resolution,[],[f14245,f1181])).
fof(f18222,plain,(
  ( ! [X28,X26,X27,X25] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X25) | sP13(X25,X26) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl126),
  inference(resolution,[],[f14244,f1340])).
fof(f17792,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP8(X47,X45,X48)) ) | $spl248),
  inference(cnf_transformation,[],[f17792_D])).
fof(f17792_D,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP8(X47,X45,X48)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f18214,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP8(X52,X50,X53) | p(s(bool,X54)) | sP12('F_0',X51)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14244,f15188])).
fof(f18213,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP8(X47,X45,X48) | p(s(bool,X46)) | sP12('F_0',X49)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14244,f15187])).
fof(f17791,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP8(X37,X35,X38)) ) | $spl246),
  inference(cnf_transformation,[],[f17791_D])).
fof(f17791_D,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP8(X37,X35,X38)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f18212,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP8(X42,X40,X43) | p(s(bool,X44)) | sP12(X41,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14244,f13742])).
fof(f18211,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP8(X37,X35,X38) | p(s(bool,X36)) | sP12(X39,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14244,f13741])).
fof(f18209,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP8(X27,X25,X28) | p(s(bool,X26)) | sP12(sK62,X29)) ) | $spl126),
  inference(resolution,[],[f14244,f5864])).
fof(f18207,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP8(X17,X15,X18) | p(s(bool,X16)) | sP12(X19,sK62)) ) | $spl126),
  inference(resolution,[],[f14244,f5611])).
fof(f18206,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP8(X12,X10,X13) | p(s(bool,X14)) | p(s(bool,X11))) ) | $spl126),
  inference(resolution,[],[f14244,f5117])).
fof(f18225,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X7,X5,X8) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(subsumption_resolution,[],[f18205,f1215])).
fof(f18205,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X7,X5,X8) | p(s(bool,'F_0')) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(resolution,[],[f14244,f1185])).
fof(f18224,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X2,X0,X3) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(subsumption_resolution,[],[f18204,f1215])).
fof(f18204,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X2,X0,X3) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(resolution,[],[f14244,f1181])).
fof(f18175,plain,(
  ( ! [X28,X26,X27,X25] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X25) | sP13(X25,X26) | p(s(bool,X27)) | p(s(bool,X28))) ) | $spl126),
  inference(resolution,[],[f14231,f1323])).
fof(f17707,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP6(X45,X47,X48)) ) | $spl244),
  inference(cnf_transformation,[],[f17707_D])).
fof(f17707_D,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP6(X45,X47,X48)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f18167,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP6(X50,X52,X53) | p(s(bool,X54)) | sP12('F_0',X51)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14231,f15188])).
fof(f18166,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP6(X45,X47,X48) | p(s(bool,X46)) | sP12('F_0',X49)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14231,f15187])).
fof(f17706,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP6(X35,X37,X38)) ) | $spl242),
  inference(cnf_transformation,[],[f17706_D])).
fof(f17706_D,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP6(X35,X37,X38)) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f18165,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP6(X40,X42,X43) | p(s(bool,X44)) | sP12(X41,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14231,f13742])).
fof(f18164,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP6(X35,X37,X38) | p(s(bool,X36)) | sP12(X39,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14231,f13741])).
fof(f18162,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP6(X25,X27,X28) | p(s(bool,X26)) | sP12(sK62,X29)) ) | $spl126),
  inference(resolution,[],[f14231,f5864])).
fof(f18160,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP6(X15,X17,X18) | p(s(bool,X16)) | sP12(X19,sK62)) ) | $spl126),
  inference(resolution,[],[f14231,f5611])).
fof(f18159,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP6(X10,X12,X13) | p(s(bool,X14)) | p(s(bool,X11))) ) | $spl126),
  inference(resolution,[],[f14231,f5117])).
fof(f18177,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP6(X5,X7,X8) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(subsumption_resolution,[],[f18158,f1215])).
fof(f18158,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP6(X5,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(resolution,[],[f14231,f1185])).
fof(f18176,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP6(X0,X2,X3) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(subsumption_resolution,[],[f18157,f1215])).
fof(f18157,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP6(X0,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(resolution,[],[f14231,f1181])).
fof(f18128,plain,(
  ( ! [X39,X37,X38,X40] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X37) | sP13(X37,X38) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | $spl126),
  inference(resolution,[],[f14215,f1480])).
fof(f14215,plain,(
  ( ! [X231,X229,X233,X230,X232] : (sP12(X229,X230) | sP13(X230,X231) | sP4(X229,X232,X233)) ) | $spl126),
  inference(resolution,[],[f13902,f1078])).
fof(f18127,plain,(
  ( ! [X35,X33,X36,X34] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X33) | sP13(X33,X34) | ~p(s(bool,X35)) | p(s(bool,X36))) ) | $spl126),
  inference(resolution,[],[f14215,f1309])).
fof(f17658,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP4(X45,X47,X48)) ) | $spl240),
  inference(cnf_transformation,[],[f17658_D])).
fof(f17658_D,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP4(X45,X47,X48)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f18117,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP4(X50,X52,X53) | p(s(bool,X54)) | sP12('F_0',X51)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14215,f15188])).
fof(f18116,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP4(X45,X47,X48) | p(s(bool,X46)) | sP12('F_0',X49)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14215,f15187])).
fof(f17657,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP4(X35,X37,X38)) ) | $spl238),
  inference(cnf_transformation,[],[f17657_D])).
fof(f17657_D,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP4(X35,X37,X38)) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f18115,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP4(X40,X42,X43) | p(s(bool,X44)) | sP12(X41,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14215,f13742])).
fof(f18114,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP4(X35,X37,X38) | p(s(bool,X36)) | sP12(X39,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14215,f13741])).
fof(f18112,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP4(X25,X27,X28) | p(s(bool,X26)) | sP12(sK62,X29)) ) | $spl126),
  inference(resolution,[],[f14215,f5864])).
fof(f18110,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP4(X15,X17,X18) | p(s(bool,X16)) | sP12(X19,sK62)) ) | $spl126),
  inference(resolution,[],[f14215,f5611])).
fof(f18109,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP4(X10,X12,X13) | p(s(bool,X14)) | p(s(bool,X11))) ) | $spl126),
  inference(resolution,[],[f14215,f5117])).
fof(f18130,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X5,X7,X8) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(subsumption_resolution,[],[f18108,f1215])).
fof(f18108,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X5,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(resolution,[],[f14215,f1185])).
fof(f18129,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X0,X2,X3) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(subsumption_resolution,[],[f18107,f1215])).
fof(f18107,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X0,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(resolution,[],[f14215,f1181])).
fof(f18075,plain,(
  ( ! [X30,X31,X29,X32] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X29) | sP13(X29,X30) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl126),
  inference(resolution,[],[f14214,f1481])).
fof(f14214,plain,(
  ( ! [X227,X225,X228,X226,X224] : (sP12(X224,X225) | sP13(X225,X226) | sP4(X227,X228,X224)) ) | $spl126),
  inference(resolution,[],[f13902,f1076])).
fof(f18074,plain,(
  ( ! [X28,X26,X27,X25] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X25) | sP13(X25,X26) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl126),
  inference(resolution,[],[f14214,f1310])).
fof(f17605,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP4(X47,X48,X45)) ) | $spl236),
  inference(cnf_transformation,[],[f17605_D])).
fof(f17605_D,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP4(X47,X48,X45)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f18066,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP4(X52,X53,X50) | p(s(bool,X54)) | sP12('F_0',X51)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14214,f15188])).
fof(f18065,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP4(X47,X48,X45) | p(s(bool,X46)) | sP12('F_0',X49)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14214,f15187])).
fof(f17604,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP4(X37,X38,X35)) ) | $spl234),
  inference(cnf_transformation,[],[f17604_D])).
fof(f17604_D,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP4(X37,X38,X35)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f18064,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP4(X42,X43,X40) | p(s(bool,X44)) | sP12(X41,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14214,f13742])).
fof(f18063,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP4(X37,X38,X35) | p(s(bool,X36)) | sP12(X39,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14214,f13741])).
fof(f18061,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP4(X27,X28,X25) | p(s(bool,X26)) | sP12(sK62,X29)) ) | $spl126),
  inference(resolution,[],[f14214,f5864])).
fof(f18059,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP4(X17,X18,X15) | p(s(bool,X16)) | sP12(X19,sK62)) ) | $spl126),
  inference(resolution,[],[f14214,f5611])).
fof(f18058,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP4(X12,X13,X10) | p(s(bool,X14)) | p(s(bool,X11))) ) | $spl126),
  inference(resolution,[],[f14214,f5117])).
fof(f18079,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X7,X8,X5) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(subsumption_resolution,[],[f18057,f1215])).
fof(f18057,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X7,X8,X5) | p(s(bool,'F_0')) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(resolution,[],[f14214,f1185])).
fof(f18078,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X0) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(subsumption_resolution,[],[f18056,f1215])).
fof(f18056,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(resolution,[],[f14214,f1181])).
fof(f17995,plain,(
  ( ! [X35,X33,X34,X32] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X32) | sP13(X32,X33) | p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl126),
  inference(resolution,[],[f14186,f1449])).
fof(f14186,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP12(X88,X89) | sP13(X89,X90) | sP1(X88,X91,X92)) ) | $spl126),
  inference(resolution,[],[f13902,f947])).
fof(f17994,plain,(
  ( ! [X30,X28,X31,X29] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X28) | sP13(X28,X29) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl126),
  inference(resolution,[],[f14186,f1278])).
fof(f17550,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP1(X45,X47,X48)) ) | $spl232),
  inference(cnf_transformation,[],[f17550_D])).
fof(f17550_D,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP1(X45,X47,X48)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f17985,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP1(X50,X52,X53) | p(s(bool,X54)) | sP12('F_0',X51)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14186,f15188])).
fof(f17984,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP1(X45,X47,X48) | p(s(bool,X46)) | sP12('F_0',X49)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14186,f15187])).
fof(f17549,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP1(X35,X37,X38)) ) | $spl230),
  inference(cnf_transformation,[],[f17549_D])).
fof(f17549_D,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP1(X35,X37,X38)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f17983,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP1(X40,X42,X43) | p(s(bool,X44)) | sP12(X41,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14186,f13742])).
fof(f17982,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP1(X35,X37,X38) | p(s(bool,X36)) | sP12(X39,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14186,f13741])).
fof(f17980,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP1(X25,X27,X28) | p(s(bool,X26)) | sP12(sK62,X29)) ) | $spl126),
  inference(resolution,[],[f14186,f5864])).
fof(f17978,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP1(X15,X17,X18) | p(s(bool,X16)) | sP12(X19,sK62)) ) | $spl126),
  inference(resolution,[],[f14186,f5611])).
fof(f17977,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP1(X10,X12,X13) | p(s(bool,X14)) | p(s(bool,X11))) ) | $spl126),
  inference(resolution,[],[f14186,f5117])).
fof(f17997,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X5,X7,X8) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(subsumption_resolution,[],[f17976,f1215])).
fof(f17976,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X5,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(resolution,[],[f14186,f1185])).
fof(f17996,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X0,X2,X3) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(subsumption_resolution,[],[f17975,f1215])).
fof(f17975,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X0,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(resolution,[],[f14186,f1181])).
fof(f17942,plain,(
  ( ! [X26,X24,X27,X25] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X24) | sP13(X24,X25) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl126),
  inference(resolution,[],[f14185,f1450])).
fof(f14185,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP12(X83,X84) | sP13(X84,X85) | sP1(X86,X83,X87)) ) | $spl126),
  inference(resolution,[],[f13902,f946])).
fof(f17941,plain,(
  ( ! [X23,X21,X22,X20] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | sP13(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl126),
  inference(resolution,[],[f14185,f1279])).
fof(f17497,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP1(X47,X45,X48)) ) | $spl228),
  inference(cnf_transformation,[],[f17497_D])).
fof(f17497_D,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP1(X47,X45,X48)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f17934,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP1(X52,X50,X53) | p(s(bool,X54)) | sP12('F_0',X51)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14185,f15188])).
fof(f17933,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP1(X47,X45,X48) | p(s(bool,X46)) | sP12('F_0',X49)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14185,f15187])).
fof(f17496,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP1(X37,X35,X38)) ) | $spl226),
  inference(cnf_transformation,[],[f17496_D])).
fof(f17496_D,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP1(X37,X35,X38)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f17932,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP1(X42,X40,X43) | p(s(bool,X44)) | sP12(X41,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14185,f13742])).
fof(f17931,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP1(X37,X35,X38) | p(s(bool,X36)) | sP12(X39,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14185,f13741])).
fof(f17929,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP1(X27,X25,X28) | p(s(bool,X26)) | sP12(sK62,X29)) ) | $spl126),
  inference(resolution,[],[f14185,f5864])).
fof(f17927,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP1(X17,X15,X18) | p(s(bool,X16)) | sP12(X19,sK62)) ) | $spl126),
  inference(resolution,[],[f14185,f5611])).
fof(f17926,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP1(X12,X10,X13) | p(s(bool,X14)) | p(s(bool,X11))) ) | $spl126),
  inference(resolution,[],[f14185,f5117])).
fof(f17946,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X7,X5,X8) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(subsumption_resolution,[],[f17925,f1215])).
fof(f17925,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X7,X5,X8) | p(s(bool,'F_0')) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(resolution,[],[f14185,f1185])).
fof(f17945,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X2,X0,X3) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(subsumption_resolution,[],[f17924,f1215])).
fof(f17924,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X2,X0,X3) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(resolution,[],[f14185,f1181])).
fof(f17892,plain,(
  ( ! [X35,X33,X34,X32] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X32) | sP13(X32,X33) | p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl126),
  inference(resolution,[],[f14183,f1446])).
fof(f14183,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP12(X73,X74) | sP13(X74,X75) | sP2(X73,X76,X77)) ) | $spl126),
  inference(resolution,[],[f13902,f943])).
fof(f17891,plain,(
  ( ! [X30,X28,X31,X29] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X28) | sP13(X28,X29) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl126),
  inference(resolution,[],[f14183,f1275])).
fof(f17409,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP2(X45,X47,X48)) ) | $spl222),
  inference(cnf_transformation,[],[f17409_D])).
fof(f17409_D,plain,(
  ( ! [X47,X45,X48,X46] : (sP12(X45,X46) | sP12('F_0',X46) | sP2(X45,X47,X48)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f17882,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP2(X50,X52,X53) | p(s(bool,X54)) | sP12('F_0',X51)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14183,f15188])).
fof(f17881,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP2(X45,X47,X48) | p(s(bool,X46)) | sP12('F_0',X49)) ) | ($spl126 | $spl176)),
  inference(resolution,[],[f14183,f15187])).
fof(f17407,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP2(X35,X37,X38)) ) | $spl218),
  inference(cnf_transformation,[],[f17407_D])).
fof(f17407_D,plain,(
  ( ! [X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | sP2(X35,X37,X38)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f17880,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP2(X40,X42,X43) | p(s(bool,X44)) | sP12(X41,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14183,f13742])).
fof(f17879,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP2(X35,X37,X38) | p(s(bool,X36)) | sP12(X39,'F_0')) ) | ($spl116 | $spl126)),
  inference(resolution,[],[f14183,f13741])).
fof(f17877,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP2(X25,X27,X28) | p(s(bool,X26)) | sP12(sK62,X29)) ) | $spl126),
  inference(resolution,[],[f14183,f5864])).
fof(f17875,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP2(X15,X17,X18) | p(s(bool,X16)) | sP12(X19,sK62)) ) | $spl126),
  inference(resolution,[],[f14183,f5611])).
fof(f17874,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP2(X10,X12,X13) | p(s(bool,X14)) | p(s(bool,X11))) ) | $spl126),
  inference(resolution,[],[f14183,f5117])).
fof(f17894,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP2(X5,X7,X8) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(subsumption_resolution,[],[f17873,f1215])).
fof(f17873,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP2(X5,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X6)) | p(s(bool,X9))) ) | $spl126),
  inference(resolution,[],[f14183,f1185])).
fof(f17893,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X0,X2,X3) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(subsumption_resolution,[],[f17872,f1215])).
fof(f17872,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X0,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X4))) ) | $spl126),
  inference(resolution,[],[f14183,f1181])).
fof(f13865,plain,(
  ( ! [X21,X20] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | sP13(X21,X20)) ) | $spl124),
  inference(cnf_transformation,[],[f13865_D])).
fof(f13865_D,plain,(
  ( ! [X21,X20] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | sP13(X21,X20)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f17840,plain,(
  ( ! [X30,X31,X29,X32] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X29) | sP13(X30,X29) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl122),
  inference(resolution,[],[f13994,f1339])).
fof(f17831,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP8(X50,X52,X53) | p(s(bool,X51)) | sP12('F_0',X54)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13994,f15188])).
fof(f17830,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP8(X45,X47,X48) | p(s(bool,X49)) | sP12('F_0',X46)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13994,f15187])).
fof(f17829,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP8(X40,X42,X43) | p(s(bool,X41)) | sP12(X44,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13994,f13742])).
fof(f17828,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP8(X35,X37,X38) | p(s(bool,X39)) | sP12(X36,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13994,f13741])).
fof(f17827,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP8(X30,X32,X33) | p(s(bool,X31)) | sP12(sK62,X34)) ) | $spl122),
  inference(resolution,[],[f13994,f5865])).
fof(f17825,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP8(X20,X22,X23) | p(s(bool,X21)) | sP12(X24,sK62)) ) | $spl122),
  inference(resolution,[],[f13994,f5612])).
fof(f17823,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP8(X10,X12,X13) | p(s(bool,X11)) | p(s(bool,X14))) ) | $spl122),
  inference(resolution,[],[f13994,f5117])).
fof(f17842,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X5,X7,X8) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(subsumption_resolution,[],[f17822,f1215])).
fof(f17822,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X5,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(resolution,[],[f13994,f1185])).
fof(f17841,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X0,X2,X3) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(subsumption_resolution,[],[f17821,f1215])).
fof(f17821,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X0,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(resolution,[],[f13994,f1181])).
fof(f17787,plain,(
  ( ! [X28,X26,X27,X25] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X25) | sP13(X26,X25) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl122),
  inference(resolution,[],[f13993,f1340])).
fof(f17779,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP8(X52,X50,X53) | p(s(bool,X51)) | sP12('F_0',X54)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13993,f15188])).
fof(f17778,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP8(X47,X45,X48) | p(s(bool,X49)) | sP12('F_0',X46)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13993,f15187])).
fof(f17777,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP8(X42,X40,X43) | p(s(bool,X41)) | sP12(X44,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13993,f13742])).
fof(f17776,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP8(X37,X35,X38) | p(s(bool,X39)) | sP12(X36,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13993,f13741])).
fof(f17775,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP8(X32,X30,X33) | p(s(bool,X31)) | sP12(sK62,X34)) ) | $spl122),
  inference(resolution,[],[f13993,f5865])).
fof(f17773,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP8(X22,X20,X23) | p(s(bool,X21)) | sP12(X24,sK62)) ) | $spl122),
  inference(resolution,[],[f13993,f5612])).
fof(f17771,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP8(X12,X10,X13) | p(s(bool,X11)) | p(s(bool,X14))) ) | $spl122),
  inference(resolution,[],[f13993,f5117])).
fof(f17790,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X7,X5,X8) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(subsumption_resolution,[],[f17770,f1215])).
fof(f17770,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP8(X7,X5,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(resolution,[],[f13993,f1185])).
fof(f17789,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X2,X0,X3) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(subsumption_resolution,[],[f17769,f1215])).
fof(f17769,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP8(X2,X0,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(resolution,[],[f13993,f1181])).
fof(f17733,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f17717])).
fof(f17717,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f1487,f5669])).
fof(f5669,plain,(
  ( ! [X149,X150,X148] : (sP5(X149,X150,X148) | sP12(sK62,X148) | ~p(s(bool,X149))) )),
  inference(resolution,[],[f5231,f1080])).
fof(f1080,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f837])).
fof(f837,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f836])).
fof(f836,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f731])).
fof(f731,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f695,f730,f729])).
fof(f695,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f605])).
fof(f605,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f604])).
fof(f604,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f174])).
fof(f174,axiom,(
  ! [X38,X77,X78] : ((p(s(bool,X78)) <=> (p(s(bool,X77)) => p(s(bool,X38)))) <=> ((p(s(bool,X78)) | p(s(bool,X77))) & (p(s(bool,X78)) | ~p(s(bool,X38))) & (~p(s(bool,X77)) | p(s(bool,X38)) | ~p(s(bool,X78)))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.sat.sat_27_prove_ltb')).
fof(f1487,plain,(
  ( ! [X114,X113] : (~sP5(X114,X113,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X114)) | p(s(bool,X113))) )),
  inference(resolution,[],[f917,f1082])).
fof(f1082,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f837])).
fof(f17734,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f17716])).
fof(f17716,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f1487,f5416])).
fof(f5416,plain,(
  ( ! [X149,X150,X148] : (sP5(X149,X150,X148) | sP12(X148,sK62) | ~p(s(bool,X149))) )),
  inference(resolution,[],[f5230,f1080])).
fof(f17703,plain,(
  ( ! [X28,X26,X27,X25] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X25) | sP13(X26,X25) | p(s(bool,X27)) | p(s(bool,X28))) ) | $spl122),
  inference(resolution,[],[f13980,f1323])).
fof(f17695,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP6(X50,X52,X53) | p(s(bool,X51)) | sP12('F_0',X54)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13980,f15188])).
fof(f17694,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP6(X45,X47,X48) | p(s(bool,X49)) | sP12('F_0',X46)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13980,f15187])).
fof(f17693,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP6(X40,X42,X43) | p(s(bool,X41)) | sP12(X44,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13980,f13742])).
fof(f17692,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP6(X35,X37,X38) | p(s(bool,X39)) | sP12(X36,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13980,f13741])).
fof(f17691,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP6(X30,X32,X33) | p(s(bool,X31)) | sP12(sK62,X34)) ) | $spl122),
  inference(resolution,[],[f13980,f5865])).
fof(f17689,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP6(X20,X22,X23) | p(s(bool,X21)) | sP12(X24,sK62)) ) | $spl122),
  inference(resolution,[],[f13980,f5612])).
fof(f17687,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP6(X10,X12,X13) | p(s(bool,X11)) | p(s(bool,X14))) ) | $spl122),
  inference(resolution,[],[f13980,f5117])).
fof(f17705,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP6(X5,X7,X8) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(subsumption_resolution,[],[f17686,f1215])).
fof(f17686,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP6(X5,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(resolution,[],[f13980,f1185])).
fof(f17704,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP6(X0,X2,X3) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(subsumption_resolution,[],[f17685,f1215])).
fof(f17685,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP6(X0,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(resolution,[],[f13980,f1181])).
fof(f17654,plain,(
  ( ! [X39,X37,X38,X40] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X37) | sP13(X38,X37) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | $spl122),
  inference(resolution,[],[f13964,f1480])).
fof(f13964,plain,(
  ( ! [X231,X229,X233,X230,X232] : (sP12(X229,X230) | sP13(X231,X230) | sP4(X229,X232,X233)) ) | $spl122),
  inference(resolution,[],[f13862,f1078])).
fof(f17653,plain,(
  ( ! [X35,X33,X36,X34] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X33) | sP13(X34,X33) | ~p(s(bool,X35)) | p(s(bool,X36))) ) | $spl122),
  inference(resolution,[],[f13964,f1309])).
fof(f17643,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP4(X50,X52,X53) | p(s(bool,X51)) | sP12('F_0',X54)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13964,f15188])).
fof(f17642,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP4(X45,X47,X48) | p(s(bool,X49)) | sP12('F_0',X46)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13964,f15187])).
fof(f17641,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP4(X40,X42,X43) | p(s(bool,X41)) | sP12(X44,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13964,f13742])).
fof(f17640,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP4(X35,X37,X38) | p(s(bool,X39)) | sP12(X36,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13964,f13741])).
fof(f17639,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP4(X30,X32,X33) | p(s(bool,X31)) | sP12(sK62,X34)) ) | $spl122),
  inference(resolution,[],[f13964,f5865])).
fof(f17637,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP4(X20,X22,X23) | p(s(bool,X21)) | sP12(X24,sK62)) ) | $spl122),
  inference(resolution,[],[f13964,f5612])).
fof(f17635,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP4(X10,X12,X13) | p(s(bool,X11)) | p(s(bool,X14))) ) | $spl122),
  inference(resolution,[],[f13964,f5117])).
fof(f17656,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X5,X7,X8) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(subsumption_resolution,[],[f17634,f1215])).
fof(f17634,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X5,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(resolution,[],[f13964,f1185])).
fof(f17655,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X0,X2,X3) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(subsumption_resolution,[],[f17633,f1215])).
fof(f17633,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X0,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(resolution,[],[f13964,f1181])).
fof(f17599,plain,(
  ( ! [X30,X31,X29,X32] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X29) | sP13(X30,X29) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl122),
  inference(resolution,[],[f13963,f1481])).
fof(f13963,plain,(
  ( ! [X227,X225,X228,X226,X224] : (sP12(X224,X225) | sP13(X226,X225) | sP4(X227,X228,X224)) ) | $spl122),
  inference(resolution,[],[f13862,f1076])).
fof(f17598,plain,(
  ( ! [X28,X26,X27,X25] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X25) | sP13(X26,X25) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl122),
  inference(resolution,[],[f13963,f1310])).
fof(f17590,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP4(X52,X53,X50) | p(s(bool,X51)) | sP12('F_0',X54)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13963,f15188])).
fof(f17589,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP4(X47,X48,X45) | p(s(bool,X49)) | sP12('F_0',X46)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13963,f15187])).
fof(f17588,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP4(X42,X43,X40) | p(s(bool,X41)) | sP12(X44,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13963,f13742])).
fof(f17587,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP4(X37,X38,X35) | p(s(bool,X39)) | sP12(X36,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13963,f13741])).
fof(f17586,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP4(X32,X33,X30) | p(s(bool,X31)) | sP12(sK62,X34)) ) | $spl122),
  inference(resolution,[],[f13963,f5865])).
fof(f17584,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP4(X22,X23,X20) | p(s(bool,X21)) | sP12(X24,sK62)) ) | $spl122),
  inference(resolution,[],[f13963,f5612])).
fof(f17582,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP4(X12,X13,X10) | p(s(bool,X11)) | p(s(bool,X14))) ) | $spl122),
  inference(resolution,[],[f13963,f5117])).
fof(f17603,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X7,X8,X5) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(subsumption_resolution,[],[f17581,f1215])).
fof(f17581,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP4(X7,X8,X5) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(resolution,[],[f13963,f1185])).
fof(f17602,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X0) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(subsumption_resolution,[],[f17580,f1215])).
fof(f17580,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP4(X2,X3,X0) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(resolution,[],[f13963,f1181])).
fof(f17546,plain,(
  ( ! [X35,X33,X34,X32] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X32) | sP13(X33,X32) | p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl122),
  inference(resolution,[],[f13935,f1449])).
fof(f13935,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP12(X88,X89) | sP13(X90,X89) | sP1(X88,X91,X92)) ) | $spl122),
  inference(resolution,[],[f13862,f947])).
fof(f17545,plain,(
  ( ! [X30,X28,X31,X29] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X28) | sP13(X29,X28) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl122),
  inference(resolution,[],[f13935,f1278])).
fof(f17536,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP1(X50,X52,X53) | p(s(bool,X51)) | sP12('F_0',X54)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13935,f15188])).
fof(f17535,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP1(X45,X47,X48) | p(s(bool,X49)) | sP12('F_0',X46)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13935,f15187])).
fof(f17534,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP1(X40,X42,X43) | p(s(bool,X41)) | sP12(X44,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13935,f13742])).
fof(f17533,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP1(X35,X37,X38) | p(s(bool,X39)) | sP12(X36,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13935,f13741])).
fof(f17532,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP1(X30,X32,X33) | p(s(bool,X31)) | sP12(sK62,X34)) ) | $spl122),
  inference(resolution,[],[f13935,f5865])).
fof(f17530,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP1(X20,X22,X23) | p(s(bool,X21)) | sP12(X24,sK62)) ) | $spl122),
  inference(resolution,[],[f13935,f5612])).
fof(f17528,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP1(X10,X12,X13) | p(s(bool,X11)) | p(s(bool,X14))) ) | $spl122),
  inference(resolution,[],[f13935,f5117])).
fof(f17548,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X5,X7,X8) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(subsumption_resolution,[],[f17527,f1215])).
fof(f17527,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X5,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(resolution,[],[f13935,f1185])).
fof(f17547,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X0,X2,X3) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(subsumption_resolution,[],[f17526,f1215])).
fof(f17526,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X0,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(resolution,[],[f13935,f1181])).
fof(f17491,plain,(
  ( ! [X26,X24,X27,X25] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X24) | sP13(X25,X24) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl122),
  inference(resolution,[],[f13934,f1450])).
fof(f13934,plain,(
  ( ! [X87,X85,X83,X86,X84] : (sP12(X83,X84) | sP13(X85,X84) | sP1(X86,X83,X87)) ) | $spl122),
  inference(resolution,[],[f13862,f946])).
fof(f17490,plain,(
  ( ! [X23,X21,X22,X20] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | sP13(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl122),
  inference(resolution,[],[f13934,f1279])).
fof(f17483,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP1(X52,X50,X53) | p(s(bool,X51)) | sP12('F_0',X54)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13934,f15188])).
fof(f17482,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP1(X47,X45,X48) | p(s(bool,X49)) | sP12('F_0',X46)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13934,f15187])).
fof(f17481,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP1(X42,X40,X43) | p(s(bool,X41)) | sP12(X44,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13934,f13742])).
fof(f17480,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP1(X37,X35,X38) | p(s(bool,X39)) | sP12(X36,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13934,f13741])).
fof(f17479,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP1(X32,X30,X33) | p(s(bool,X31)) | sP12(sK62,X34)) ) | $spl122),
  inference(resolution,[],[f13934,f5865])).
fof(f17477,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP1(X22,X20,X23) | p(s(bool,X21)) | sP12(X24,sK62)) ) | $spl122),
  inference(resolution,[],[f13934,f5612])).
fof(f17475,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP1(X12,X10,X13) | p(s(bool,X11)) | p(s(bool,X14))) ) | $spl122),
  inference(resolution,[],[f13934,f5117])).
fof(f17495,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X7,X5,X8) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(subsumption_resolution,[],[f17474,f1215])).
fof(f17474,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP1(X7,X5,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(resolution,[],[f13934,f1185])).
fof(f17494,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X2,X0,X3) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(subsumption_resolution,[],[f17473,f1215])).
fof(f17473,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP1(X2,X0,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(resolution,[],[f13934,f1181])).
fof(f17438,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f17424])).
fof(f17424,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | p(s(bool,X21)) | p(s(bool,X21)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1483,f5660])).
fof(f1483,plain,(
  ( ! [X105,X106] : (sP5(X105,X106,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X106)) | p(s(bool,X105))) )),
  inference(resolution,[],[f917,f1079])).
fof(f1079,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f837])).
fof(f17439,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,X15)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f17421])).
fof(f17421,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,X15)) | p(s(bool,X15)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1483,f5407])).
fof(f17404,plain,(
  ( ! [X35,X33,X34,X32] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X32) | sP13(X33,X32) | p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl122),
  inference(resolution,[],[f13932,f1446])).
fof(f13932,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP12(X73,X74) | sP13(X75,X74) | sP2(X73,X76,X77)) ) | $spl122),
  inference(resolution,[],[f13862,f943])).
fof(f17403,plain,(
  ( ! [X30,X28,X31,X29] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X28) | sP13(X29,X28) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl122),
  inference(resolution,[],[f13932,f1275])).
fof(f17394,plain,(
  ( ! [X54,X52,X50,X53,X51] : (sP12(X50,X51) | sP2(X50,X52,X53) | p(s(bool,X51)) | sP12('F_0',X54)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13932,f15188])).
fof(f17393,plain,(
  ( ! [X47,X45,X48,X46,X49] : (sP12(X45,X46) | sP2(X45,X47,X48) | p(s(bool,X49)) | sP12('F_0',X46)) ) | ($spl122 | $spl176)),
  inference(resolution,[],[f13932,f15187])).
fof(f17392,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP12(X40,X41) | sP2(X40,X42,X43) | p(s(bool,X41)) | sP12(X44,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13932,f13742])).
fof(f17391,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP2(X35,X37,X38) | p(s(bool,X39)) | sP12(X36,'F_0')) ) | ($spl116 | $spl122)),
  inference(resolution,[],[f13932,f13741])).
fof(f17390,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP12(X30,X31) | sP2(X30,X32,X33) | p(s(bool,X31)) | sP12(sK62,X34)) ) | $spl122),
  inference(resolution,[],[f13932,f5865])).
fof(f17388,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP12(X20,X21) | sP2(X20,X22,X23) | p(s(bool,X21)) | sP12(X24,sK62)) ) | $spl122),
  inference(resolution,[],[f13932,f5612])).
fof(f17386,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP2(X10,X12,X13) | p(s(bool,X11)) | p(s(bool,X14))) ) | $spl122),
  inference(resolution,[],[f13932,f5117])).
fof(f17406,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP2(X5,X7,X8) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(subsumption_resolution,[],[f17385,f1215])).
fof(f17385,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP12(X5,X6) | sP2(X5,X7,X8) | p(s(bool,'F_0')) | p(s(bool,X9)) | p(s(bool,X6))) ) | $spl122),
  inference(resolution,[],[f13932,f1185])).
fof(f17405,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X0,X2,X3) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(subsumption_resolution,[],[f17384,f1215])).
fof(f17384,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X0,X2,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X1))) ) | $spl122),
  inference(resolution,[],[f13932,f1181])).
fof(f17355,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f13742,f3767])).
fof(f3767,plain,(
  ( ! [X221] : (sP13('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X221)) )),
  inference(resolution,[],[f1433,f1183])).
fof(f17354,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f13742,f3595])).
fof(f3595,plain,(
  ( ! [X221] : (sP13('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X221)) )),
  inference(resolution,[],[f1432,f1183])).
fof(f17352,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f13742,f1547])).
fof(f1547,plain,(
  ( ! [X224] : (sP13('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X224)) )),
  inference(resolution,[],[f917,f1183])).
fof(f17330,plain,(
  ( ! [X26] : (p(s(bool,X26)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f13741,f3766])).
fof(f3766,plain,(
  ( ! [X220] : (sP13(X220,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1433,f1182])).
fof(f17329,plain,(
  ( ! [X25] : (p(s(bool,X25)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f13741,f3594])).
fof(f3594,plain,(
  ( ! [X220] : (sP13(X220,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1432,f1182])).
fof(f17327,plain,(
  ( ! [X23] : (p(s(bool,X23)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f13741,f1546])).
fof(f1546,plain,(
  ( ! [X223] : (sP13(X223,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f917,f1182])).
fof(f17304,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl39),
  inference(resolution,[],[f1481,f14651])).
fof(f14651,plain,(
  ( ! [X39,X41,X40] : (sP4(X40,X41,X39) | sP12('F_0',X39)) ) | $spl39),
  inference(resolution,[],[f5868,f8318])).
fof(f17302,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f1481,f13541])).
fof(f13541,plain,(
  ( ! [X138,X136,X137] : (sP4(X137,X138,X136) | sP12(X136,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f1076])).
fof(f17300,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | sP12(X36,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(sK62,X36)) )),
  inference(resolution,[],[f1481,f8318])).
fof(f17298,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X30) | sP12(sK62,X30)) )),
  inference(resolution,[],[f1481,f8066])).
fof(f17296,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(X24,sK62)) )),
  inference(resolution,[],[f1481,f7633])).
fof(f17294,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X18) | sP12(X18,sK62)) )),
  inference(resolution,[],[f1481,f7371])).
fof(f17292,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1481,f5665])).
fof(f5665,plain,(
  ( ! [X138,X136,X137] : (sP4(X137,X138,X136) | sP12(sK62,X136)) )),
  inference(resolution,[],[f5231,f1076])).
fof(f17290,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1481,f5412])).
fof(f5412,plain,(
  ( ! [X138,X136,X137] : (sP4(X137,X138,X136) | sP12(X136,sK62)) )),
  inference(resolution,[],[f5230,f1076])).
fof(f17257,plain,(
  ( ! [X21,X19,X20] : (sP12(X19,'F_0') | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X19) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl116),
  inference(resolution,[],[f13620,f1447])).
fof(f13620,plain,(
  ( ! [X344,X346,X343,X345] : (sP12(X343,'F_0') | sP12(X344,X343) | sP2(X345,X344,X346)) ) | $spl116),
  inference(resolution,[],[f13492,f8645])).
fof(f8645,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,X1) | sP2(X2,X0,X3) | p(s(bool,X1))) )),
  inference(resolution,[],[f7339,f5228])).
fof(f7339,plain,(
  ( ! [X59,X57,X58,X56] : (sP12(X56,X57) | sP12(X57,sK62) | sP2(X58,X56,X59)) )),
  inference(resolution,[],[f5469,f942])).
fof(f17256,plain,(
  ( ! [X17,X18,X16] : (sP12(X16,'F_0') | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X16) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl116),
  inference(resolution,[],[f13620,f1276])).
fof(f17265,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X60,X61,X62)) ) | $spl116),
  inference(subsumption_resolution,[],[f17242,f942])).
fof(f17242,plain,(
  ( ! [X61,X62,X60] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X60,X61,X62) | ~p(s(bool,X61))) ) | $spl116),
  inference(resolution,[],[f13620,f3759])).
fof(f17264,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X57,X58,X59)) ) | $spl116),
  inference(subsumption_resolution,[],[f17241,f942])).
fof(f17241,plain,(
  ( ! [X59,X57,X58] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X57,X58,X59) | ~p(s(bool,X58))) ) | $spl116),
  inference(resolution,[],[f13620,f3587])).
fof(f14548,plain,(
  sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | $spl136),
  inference(cnf_transformation,[],[f14548_D])).
fof(f14548_D,plain,(
  sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f17263,plain,(
  ( ! [X54,X56,X55] : (sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | sP2(X54,X55,X56)) ) | ($spl0 | $spl116)),
  inference(subsumption_resolution,[],[f17240,f942])).
fof(f17240,plain,(
  ( ! [X54,X56,X55] : (sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | sP2(X54,X55,X56) | ~p(s(bool,X55))) ) | ($spl0 | $spl116)),
  inference(resolution,[],[f13620,f3246])).
fof(f3246,plain,(
  ( ! [X216] : (~sP12(X216,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | ~p(s(bool,X216))) ) | $spl0),
  inference(resolution,[],[f1417,f1166])).
fof(f1417,plain,(
  p(s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) | $spl0),
  inference(cnf_transformation,[],[f1417_D])).
fof(f1417_D,plain,(
  p(s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f17262,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X51,X52,X53)) ) | $spl116),
  inference(subsumption_resolution,[],[f17239,f942])).
fof(f17239,plain,(
  ( ! [X52,X53,X51] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | sP2(X51,X52,X53) | ~p(s(bool,X52))) ) | $spl116),
  inference(resolution,[],[f13620,f1539])).
fof(f17194,plain,(
  ( ! [X47,X46] : (~p(s(bool,X46)) | p(s(bool,X47)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl39),
  inference(resolution,[],[f1480,f14652])).
fof(f14652,plain,(
  ( ! [X43,X44,X42] : (sP4(X42,X43,X44) | sP12('F_0',X42)) ) | $spl39),
  inference(resolution,[],[f5868,f8319])).
fof(f17192,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f1480,f13542])).
fof(f13542,plain,(
  ( ! [X140,X141,X139] : (sP4(X139,X140,X141) | sP12(X139,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f1078])).
fof(f17190,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,X37)) | p(s(bool,X38)) | sP12(X39,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(sK62,X39)) )),
  inference(resolution,[],[f1480,f8319])).
fof(f17188,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X33) | sP12(sK62,X33)) )),
  inference(resolution,[],[f1480,f8067])).
fof(f17186,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | sP12(X27,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(X27,sK62)) )),
  inference(resolution,[],[f1480,f7634])).
fof(f17184,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | p(s(bool,X20)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X21) | sP12(X21,sK62)) )),
  inference(resolution,[],[f1480,f7372])).
fof(f17182,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1480,f5666])).
fof(f5666,plain,(
  ( ! [X140,X141,X139] : (sP4(X139,X140,X141) | sP12(sK62,X139)) )),
  inference(resolution,[],[f5231,f1078])).
fof(f17180,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1480,f5413])).
fof(f5413,plain,(
  ( ! [X140,X141,X139] : (sP4(X139,X140,X141) | sP12(X139,sK62)) )),
  inference(resolution,[],[f5230,f1078])).
fof(f17141,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl39),
  inference(resolution,[],[f1450,f14649])).
fof(f14649,plain,(
  ( ! [X35,X33,X34] : (sP1(X34,X33,X35) | sP12('F_0',X33)) ) | $spl39),
  inference(resolution,[],[f5868,f8289])).
fof(f17139,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f1450,f13512])).
fof(f13512,plain,(
  ( ! [X54,X55,X53] : (sP1(X54,X53,X55) | sP12(X53,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f946])).
fof(f17137,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | sP12(X36,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(sK62,X36)) )),
  inference(resolution,[],[f1450,f8289])).
fof(f17135,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X30) | sP12(sK62,X30)) )),
  inference(resolution,[],[f1450,f8037])).
fof(f17133,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(X24,sK62)) )),
  inference(resolution,[],[f1450,f7604])).
fof(f17131,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X18) | sP12(X18,sK62)) )),
  inference(resolution,[],[f1450,f7342])).
fof(f17129,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1450,f5636])).
fof(f5636,plain,(
  ( ! [X54,X55,X53] : (sP1(X54,X53,X55) | sP12(sK62,X53)) )),
  inference(resolution,[],[f5231,f946])).
fof(f17127,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1450,f5383])).
fof(f5383,plain,(
  ( ! [X54,X55,X53] : (sP1(X54,X53,X55) | sP12(X53,sK62)) )),
  inference(resolution,[],[f5230,f946])).
fof(f17114,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP12(X0,'F_0')) ) | $spl212),
  inference(cnf_transformation,[],[f17114_D])).
fof(f17114_D,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP12(X0,'F_0')) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f16984,plain,(
  ( ! [X455,X457,X454,X456] : (sP12(X454,X455) | sP12(X454,'F_0') | p(s(bool,X456)) | ~p(s(bool,X457)) | s(bool,X455) != s(bool,X457)) ) | $spl116),
  inference(resolution,[],[f13599,f1252])).
fof(f13599,plain,(
  ( ! [X308,X307] : (~p(s(bool,X308)) | sP12(X307,X308) | sP12(X307,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f1170])).
fof(f16871,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X35,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl116),
  inference(resolution,[],[f13599,f926])).
fof(f16869,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP12(X25,'F_0') | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) ) | $spl116),
  inference(resolution,[],[f13599,f925])).
fof(f16867,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP12(X15,'F_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl116),
  inference(resolution,[],[f13599,f924])).
fof(f16866,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP12(X10,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl116),
  inference(resolution,[],[f13599,f922])).
fof(f16864,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP12(X0,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl116),
  inference(resolution,[],[f13599,f920])).
fof(f16855,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP12(X1,'F_0')) ) | $spl210),
  inference(cnf_transformation,[],[f16855_D])).
fof(f16855_D,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP12(X1,'F_0')) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f16733,plain,(
  ( ! [X455,X457,X454,X456] : (sP12(X454,X455) | sP12(X455,'F_0') | p(s(bool,X456)) | ~p(s(bool,X457)) | s(bool,X454) != s(bool,X457)) ) | $spl116),
  inference(resolution,[],[f13598,f1252])).
fof(f13598,plain,(
  ( ! [X306,X305] : (~p(s(bool,X306)) | sP12(X306,X305) | sP12(X305,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f1167])).
fof(f16620,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,'F_0') | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl116),
  inference(resolution,[],[f13598,f926])).
fof(f16618,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP12(X26,'F_0') | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) ) | $spl116),
  inference(resolution,[],[f13598,f925])).
fof(f16616,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP12(X16,'F_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | $spl116),
  inference(resolution,[],[f13598,f924])).
fof(f16615,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP12(X11,'F_0') | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl116),
  inference(resolution,[],[f13598,f922])).
fof(f16613,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP12(X1,'F_0') | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl116),
  inference(resolution,[],[f13598,f920])).
fof(f16528,plain,(
  ( ! [X258] : (p(s(bool,X258)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(duplicate_literal_removal,[],[f16503])).
fof(f16503,plain,(
  ( ! [X258] : (p(s(bool,X258)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | p(s(bool,X258))) ) | $spl116),
  inference(resolution,[],[f13597,f3760])).
fof(f3760,plain,(
  ( ! [X214] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X214) | p(s(bool,X214))) )),
  inference(resolution,[],[f1433,f1167])).
fof(f16529,plain,(
  ( ! [X257] : (p(s(bool,X257)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(duplicate_literal_removal,[],[f16502])).
fof(f16502,plain,(
  ( ! [X257] : (p(s(bool,X257)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | p(s(bool,X257))) ) | $spl116),
  inference(resolution,[],[f13597,f3588])).
fof(f3588,plain,(
  ( ! [X214] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X214) | p(s(bool,X214))) )),
  inference(resolution,[],[f1432,f1167])).
fof(f16530,plain,(
  ( ! [X256] : (p(s(bool,X256)) | sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0')) ) | ($spl0 | $spl116)),
  inference(duplicate_literal_removal,[],[f16500])).
fof(f16500,plain,(
  ( ! [X256] : (p(s(bool,X256)) | sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | p(s(bool,X256))) ) | ($spl0 | $spl116)),
  inference(resolution,[],[f13597,f3247])).
fof(f3247,plain,(
  ( ! [X217] : (sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X217) | p(s(bool,X217))) ) | $spl0),
  inference(resolution,[],[f1417,f1167])).
fof(f16531,plain,(
  ( ! [X255] : (p(s(bool,X255)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(duplicate_literal_removal,[],[f16499])).
fof(f16499,plain,(
  ( ! [X255] : (p(s(bool,X255)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | p(s(bool,X255))) ) | $spl116),
  inference(resolution,[],[f13597,f1540])).
fof(f1540,plain,(
  ( ! [X217] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X217) | p(s(bool,X217))) )),
  inference(resolution,[],[f917,f1167])).
fof(f16409,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl39),
  inference(resolution,[],[f1449,f14650])).
fof(f14650,plain,(
  ( ! [X37,X38,X36] : (sP1(X36,X37,X38) | sP12('F_0',X36)) ) | $spl39),
  inference(resolution,[],[f5868,f8290])).
fof(f16407,plain,(
  ( ! [X43,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f1449,f13513])).
fof(f13513,plain,(
  ( ! [X57,X58,X56] : (sP1(X56,X57,X58) | sP12(X56,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f947])).
fof(f16405,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | sP12(X39,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(sK62,X39)) )),
  inference(resolution,[],[f1449,f8290])).
fof(f16403,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X33) | sP12(sK62,X33)) )),
  inference(resolution,[],[f1449,f8038])).
fof(f16401,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | sP12(X27,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(X27,sK62)) )),
  inference(resolution,[],[f1449,f7605])).
fof(f16399,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X21) | sP12(X21,sK62)) )),
  inference(resolution,[],[f1449,f7343])).
fof(f16397,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1449,f5637])).
fof(f5637,plain,(
  ( ! [X57,X58,X56] : (sP1(X56,X57,X58) | sP12(sK62,X56)) )),
  inference(resolution,[],[f5231,f947])).
fof(f16395,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1449,f5384])).
fof(f5384,plain,(
  ( ! [X57,X58,X56] : (sP1(X56,X57,X58) | sP12(X56,sK62)) )),
  inference(resolution,[],[f5230,f947])).
fof(f16387,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(resolution,[],[f15188,f3767])).
fof(f16386,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(resolution,[],[f15188,f3595])).
fof(f16384,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(resolution,[],[f15188,f1547])).
fof(f16362,plain,(
  ( ! [X26] : (p(s(bool,X26)) | sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(resolution,[],[f15187,f3766])).
fof(f16361,plain,(
  ( ! [X25] : (p(s(bool,X25)) | sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(resolution,[],[f15187,f3594])).
fof(f16359,plain,(
  ( ! [X23] : (p(s(bool,X23)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(resolution,[],[f15187,f1546])).
fof(f16321,plain,(
  ( ! [X21,X19,X20] : (sP12('F_0',X19) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X19) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl176),
  inference(resolution,[],[f15064,f1447])).
fof(f15064,plain,(
  ( ! [X344,X346,X343,X345] : (sP12('F_0',X343) | sP12(X344,X343) | sP2(X345,X344,X346)) ) | $spl176),
  inference(resolution,[],[f14807,f8645])).
fof(f16320,plain,(
  ( ! [X17,X18,X16] : (sP12('F_0',X16) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X16) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl176),
  inference(resolution,[],[f15064,f1276])).
fof(f16330,plain,(
  ( ! [X57,X58,X56] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X56,X57,X58)) ) | $spl176),
  inference(subsumption_resolution,[],[f16306,f942])).
fof(f16306,plain,(
  ( ! [X57,X58,X56] : (sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X56,X57,X58) | ~p(s(bool,X57))) ) | $spl176),
  inference(resolution,[],[f15064,f3759])).
fof(f16329,plain,(
  ( ! [X54,X55,X53] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X53,X54,X55)) ) | $spl176),
  inference(subsumption_resolution,[],[f16305,f942])).
fof(f16305,plain,(
  ( ! [X54,X55,X53] : (sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X53,X54,X55) | ~p(s(bool,X54))) ) | $spl176),
  inference(resolution,[],[f15064,f3587])).
fof(f15439,plain,(
  sP12('F_0','const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | $spl184),
  inference(cnf_transformation,[],[f15439_D])).
fof(f15439_D,plain,(
  sP12('F_0','const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f16328,plain,(
  ( ! [X52,X50,X51] : (sP12('F_0','const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP2(X50,X51,X52)) ) | ($spl0 | $spl176)),
  inference(subsumption_resolution,[],[f16304,f942])).
fof(f16304,plain,(
  ( ! [X52,X50,X51] : (sP12('F_0','const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP2(X50,X51,X52) | ~p(s(bool,X51))) ) | ($spl0 | $spl176)),
  inference(resolution,[],[f15064,f3246])).
fof(f16327,plain,(
  ( ! [X47,X48,X49] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X47,X48,X49)) ) | $spl176),
  inference(subsumption_resolution,[],[f16303,f942])).
fof(f16303,plain,(
  ( ! [X47,X48,X49] : (sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X47,X48,X49) | ~p(s(bool,X48))) ) | $spl176),
  inference(resolution,[],[f15064,f1539])).
fof(f16258,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl39),
  inference(resolution,[],[f1447,f14647])).
fof(f14647,plain,(
  ( ! [X28,X29,X27] : (sP2(X28,X27,X29) | sP12('F_0',X27)) ) | $spl39),
  inference(resolution,[],[f5868,f8286])).
fof(f16256,plain,(
  ( ! [X52,X51] : (p(s(bool,X51)) | ~p(s(bool,X52)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) )),
  inference(resolution,[],[f1447,f13346])).
fof(f13346,plain,(
  ( ! [X814,X812,X813] : (sP2(X813,X812,X814) | sP12(X812,'F_0')) )),
  inference(resolution,[],[f8645,f1247])).
fof(f16255,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X47)) | ~p(s(bool,X48)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X49) | sP13(X49,X50)) )),
  inference(resolution,[],[f1447,f13315])).
fof(f13315,plain,(
  ( ! [X670,X673,X669,X671,X672] : (sP12(X669,X670) | sP2(X671,X669,X672) | sP13(X670,X673)) )),
  inference(resolution,[],[f8645,f1183])).
fof(f16254,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,X43)) | ~p(s(bool,X44)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X45) | sP13(X46,X45)) )),
  inference(resolution,[],[f1447,f13314])).
fof(f13314,plain,(
  ( ! [X666,X668,X664,X667,X665] : (sP12(X664,X665) | sP2(X666,X664,X667) | sP13(X668,X665)) )),
  inference(resolution,[],[f8645,f1182])).
fof(f16251,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | sP12(X36,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(sK62,X36)) )),
  inference(resolution,[],[f1447,f8286])).
fof(f16249,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X30) | sP12(sK62,X30)) )),
  inference(resolution,[],[f1447,f8034])).
fof(f8034,plain,(
  ( ! [X59,X57,X58,X56] : (sP12(X56,X57) | sP12(sK62,X57) | sP2(X58,X56,X59)) )),
  inference(resolution,[],[f5722,f942])).
fof(f16247,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | sP12(X24,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(X24,sK62)) )),
  inference(resolution,[],[f1447,f7601])).
fof(f16245,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X18) | sP12(X18,sK62)) )),
  inference(resolution,[],[f1447,f7339])).
fof(f16243,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1447,f5633])).
fof(f5633,plain,(
  ( ! [X45,X46,X44] : (sP2(X45,X44,X46) | sP12(sK62,X44)) )),
  inference(resolution,[],[f5231,f942])).
fof(f16241,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1447,f5380])).
fof(f5380,plain,(
  ( ! [X45,X46,X44] : (sP2(X45,X44,X46) | sP12(X44,sK62)) )),
  inference(resolution,[],[f5230,f942])).
fof(f16202,plain,(
  ( ! [X57,X58] : (p(s(bool,X57)) | ~p(s(bool,X58)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl39),
  inference(resolution,[],[f1446,f14648])).
fof(f14648,plain,(
  ( ! [X30,X31,X32] : (sP2(X30,X31,X32) | sP12('F_0',X30)) ) | $spl39),
  inference(resolution,[],[f5868,f8287])).
fof(f16200,plain,(
  ( ! [X54,X53] : (p(s(bool,X53)) | ~p(s(bool,X54)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(resolution,[],[f1446,f13510])).
fof(f13510,plain,(
  ( ! [X47,X48,X49] : (sP2(X47,X48,X49) | sP12(X47,'F_0')) ) | $spl116),
  inference(resolution,[],[f13492,f943])).
fof(f16195,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | sP12(X39,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(sK62,X39)) )),
  inference(resolution,[],[f1446,f8287])).
fof(f16193,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X33) | sP12(sK62,X33)) )),
  inference(resolution,[],[f1446,f8035])).
fof(f16191,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | sP12(X27,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP12(X27,sK62)) )),
  inference(resolution,[],[f1446,f7602])).
fof(f16189,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X21) | sP12(X21,sK62)) )),
  inference(resolution,[],[f1446,f7340])).
fof(f16187,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1446,f5634])).
fof(f5634,plain,(
  ( ! [X47,X48,X49] : (sP2(X47,X48,X49) | sP12(sK62,X47)) )),
  inference(resolution,[],[f5231,f943])).
fof(f16185,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f1446,f5381])).
fof(f5381,plain,(
  ( ! [X47,X48,X49] : (sP2(X47,X48,X49) | sP12(X47,sK62)) )),
  inference(resolution,[],[f5230,f943])).
fof(f16171,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP12('F_0',X0)) ) | $spl192),
  inference(cnf_transformation,[],[f16171_D])).
fof(f16171_D,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP12('F_0',X0)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f16049,plain,(
  ( ! [X455,X457,X454,X456] : (sP12(X454,X455) | sP12('F_0',X454) | p(s(bool,X456)) | ~p(s(bool,X457)) | s(bool,X455) != s(bool,X457)) ) | $spl176),
  inference(resolution,[],[f15043,f1252])).
fof(f15043,plain,(
  ( ! [X308,X307] : (~p(s(bool,X308)) | sP12(X307,X308) | sP12('F_0',X307)) ) | $spl176),
  inference(resolution,[],[f14807,f1170])).
fof(f15936,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP12('F_0',X35) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl176),
  inference(resolution,[],[f15043,f926])).
fof(f15934,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP12('F_0',X25) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) ) | $spl176),
  inference(resolution,[],[f15043,f925])).
fof(f15932,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP12('F_0',X15) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl176),
  inference(resolution,[],[f15043,f924])).
fof(f15931,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP12('F_0',X10) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl176),
  inference(resolution,[],[f15043,f922])).
fof(f15929,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP12('F_0',X0) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl176),
  inference(resolution,[],[f15043,f920])).
fof(f15920,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP12('F_0',X1)) ) | $spl190),
  inference(cnf_transformation,[],[f15920_D])).
fof(f15920_D,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP12('F_0',X1)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f15790,plain,(
  ( ! [X455,X457,X454,X456] : (sP12(X454,X455) | sP12('F_0',X455) | p(s(bool,X456)) | ~p(s(bool,X457)) | s(bool,X454) != s(bool,X457)) ) | $spl176),
  inference(resolution,[],[f15042,f1252])).
fof(f15042,plain,(
  ( ! [X306,X305] : (~p(s(bool,X306)) | sP12(X306,X305) | sP12('F_0',X305)) ) | $spl176),
  inference(resolution,[],[f14807,f1167])).
fof(f15677,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP12('F_0',X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl176),
  inference(resolution,[],[f15042,f926])).
fof(f15675,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP12('F_0',X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) ) | $spl176),
  inference(resolution,[],[f15042,f925])).
fof(f15673,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP12('F_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | $spl176),
  inference(resolution,[],[f15042,f924])).
fof(f15672,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP12('F_0',X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl176),
  inference(resolution,[],[f15042,f922])).
fof(f15670,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP12('F_0',X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl176),
  inference(resolution,[],[f15042,f920])).
fof(f15599,plain,(
  ( ! [X321,X323,X324,X322] : (p(s(bool,X321)) | sP12('F_0',sK62) | sP12(X321,X322) | sP8(X322,X323,X324)) ) | $spl176),
  inference(resolution,[],[f15041,f8349])).
fof(f15598,plain,(
  ( ! [X318,X320,X317,X319] : (p(s(bool,X317)) | sP12('F_0',sK62) | sP12(X317,X318) | sP8(X319,X318,X320)) ) | $spl176),
  inference(resolution,[],[f15041,f8348])).
fof(f15597,plain,(
  ( ! [X313,X315,X316,X314] : (p(s(bool,X313)) | sP12('F_0',sK62) | sP12(X313,X314) | sP6(X314,X315,X316)) ) | $spl176),
  inference(resolution,[],[f15041,f8335])).
fof(f15596,plain,(
  ( ! [X310,X312,X309,X311] : (p(s(bool,X309)) | sP12('F_0',sK62) | sP12(X309,X310) | sP4(X310,X311,X312)) ) | $spl176),
  inference(resolution,[],[f15041,f8319])).
fof(f15595,plain,(
  ( ! [X308,X306,X305,X307] : (p(s(bool,X305)) | sP12('F_0',sK62) | sP12(X305,X306) | sP4(X307,X308,X306)) ) | $spl176),
  inference(resolution,[],[f15041,f8318])).
fof(f15594,plain,(
  ( ! [X302,X304,X301,X303] : (p(s(bool,X301)) | sP12('F_0',sK62) | sP12(X301,X302) | sP1(X302,X303,X304)) ) | $spl176),
  inference(resolution,[],[f15041,f8290])).
fof(f15593,plain,(
  ( ! [X300,X298,X297,X299] : (p(s(bool,X297)) | sP12('F_0',sK62) | sP12(X297,X298) | sP1(X299,X298,X300)) ) | $spl176),
  inference(resolution,[],[f15041,f8289])).
fof(f15592,plain,(
  ( ! [X294,X296,X293,X295] : (p(s(bool,X293)) | sP12('F_0',sK62) | sP12(X293,X294) | sP2(X294,X295,X296)) ) | $spl176),
  inference(resolution,[],[f15041,f8287])).
fof(f15591,plain,(
  ( ! [X292,X290,X289,X291] : (p(s(bool,X289)) | sP12('F_0',sK62) | sP12(X289,X290) | sP2(X291,X290,X292)) ) | $spl176),
  inference(resolution,[],[f15041,f8286])).
fof(f15590,plain,(
  ( ! [X288,X285,X287,X286] : (p(s(bool,X285)) | sP12('F_0',sK62) | sP12(X286,X285) | sP8(X286,X287,X288)) ) | $spl176),
  inference(resolution,[],[f15041,f8097])).
fof(f15589,plain,(
  ( ! [X282,X281,X283,X284] : (p(s(bool,X281)) | sP12('F_0',sK62) | sP12(X282,X281) | sP8(X283,X282,X284)) ) | $spl176),
  inference(resolution,[],[f15041,f8096])).
fof(f15588,plain,(
  ( ! [X280,X277,X279,X278] : (p(s(bool,X277)) | sP12('F_0',sK62) | sP12(X278,X277) | sP6(X278,X279,X280)) ) | $spl176),
  inference(resolution,[],[f15041,f8083])).
fof(f15587,plain,(
  ( ! [X273,X275,X276,X274] : (p(s(bool,X273)) | sP12('F_0',sK62) | sP12(X274,X273) | sP4(X274,X275,X276)) ) | $spl176),
  inference(resolution,[],[f15041,f8067])).
fof(f15586,plain,(
  ( ! [X269,X271,X270,X272] : (p(s(bool,X269)) | sP12('F_0',sK62) | sP12(X270,X269) | sP4(X271,X272,X270)) ) | $spl176),
  inference(resolution,[],[f15041,f8066])).
fof(f15585,plain,(
  ( ! [X265,X267,X268,X266] : (p(s(bool,X265)) | sP12('F_0',sK62) | sP12(X266,X265) | sP1(X266,X267,X268)) ) | $spl176),
  inference(resolution,[],[f15041,f8038])).
fof(f15584,plain,(
  ( ! [X261,X263,X262,X264] : (p(s(bool,X261)) | sP12('F_0',sK62) | sP12(X262,X261) | sP1(X263,X262,X264)) ) | $spl176),
  inference(resolution,[],[f15041,f8037])).
fof(f15583,plain,(
  ( ! [X257,X259,X260,X258] : (p(s(bool,X257)) | sP12('F_0',sK62) | sP12(X258,X257) | sP2(X258,X259,X260)) ) | $spl176),
  inference(resolution,[],[f15041,f8035])).
fof(f15603,plain,(
  ( ! [X249] : (p(s(bool,X249)) | sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(duplicate_literal_removal,[],[f15578])).
fof(f15578,plain,(
  ( ! [X249] : (p(s(bool,X249)) | sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X249))) ) | $spl176),
  inference(resolution,[],[f15041,f3760])).
fof(f15604,plain,(
  ( ! [X248] : (p(s(bool,X248)) | sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(duplicate_literal_removal,[],[f15577])).
fof(f15577,plain,(
  ( ! [X248] : (p(s(bool,X248)) | sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X248))) ) | $spl176),
  inference(resolution,[],[f15041,f3588])).
fof(f15605,plain,(
  ( ! [X247] : (p(s(bool,X247)) | sP12('F_0','const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | ($spl0 | $spl176)),
  inference(duplicate_literal_removal,[],[f15575])).
fof(f15575,plain,(
  ( ! [X247] : (p(s(bool,X247)) | sP12('F_0','const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X247))) ) | ($spl0 | $spl176)),
  inference(resolution,[],[f15041,f3247])).
fof(f15606,plain,(
  ( ! [X246] : (p(s(bool,X246)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(duplicate_literal_removal,[],[f15574])).
fof(f15574,plain,(
  ( ! [X246] : (p(s(bool,X246)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X246))) ) | $spl176),
  inference(resolution,[],[f15041,f1540])).
fof(f15376,plain,(
  ( ! [X243,X241,X242,X240] : (p(s(bool,X240)) | sP12('F_0',sK62) | sP12(X240,X241) | sP8(X241,X242,X243)) ) | $spl176),
  inference(resolution,[],[f15040,f7664])).
fof(f15375,plain,(
  ( ! [X239,X237,X238,X236] : (p(s(bool,X236)) | sP12('F_0',sK62) | sP12(X236,X237) | sP8(X238,X237,X239)) ) | $spl176),
  inference(resolution,[],[f15040,f7663])).
fof(f15374,plain,(
  ( ! [X235,X233,X234,X232] : (p(s(bool,X232)) | sP12('F_0',sK62) | sP12(X232,X233) | sP6(X233,X234,X235)) ) | $spl176),
  inference(resolution,[],[f15040,f7650])).
fof(f15373,plain,(
  ( ! [X231,X229,X230,X228] : (p(s(bool,X228)) | sP12('F_0',sK62) | sP12(X228,X229) | sP4(X229,X230,X231)) ) | $spl176),
  inference(resolution,[],[f15040,f7634])).
fof(f15372,plain,(
  ( ! [X227,X225,X226,X224] : (p(s(bool,X224)) | sP12('F_0',sK62) | sP12(X224,X225) | sP4(X226,X227,X225)) ) | $spl176),
  inference(resolution,[],[f15040,f7633])).
fof(f15371,plain,(
  ( ! [X222,X220,X223,X221] : (p(s(bool,X220)) | sP12('F_0',sK62) | sP12(X220,X221) | sP1(X221,X222,X223)) ) | $spl176),
  inference(resolution,[],[f15040,f7605])).
fof(f15370,plain,(
  ( ! [X218,X216,X219,X217] : (p(s(bool,X216)) | sP12('F_0',sK62) | sP12(X216,X217) | sP1(X218,X217,X219)) ) | $spl176),
  inference(resolution,[],[f15040,f7604])).
fof(f15369,plain,(
  ( ! [X212,X215,X213,X214] : (p(s(bool,X212)) | sP12('F_0',sK62) | sP12(X212,X213) | sP2(X213,X214,X215)) ) | $spl176),
  inference(resolution,[],[f15040,f7602])).
fof(f15368,plain,(
  ( ! [X210,X208,X211,X209] : (p(s(bool,X208)) | sP12('F_0',sK62) | sP12(X208,X209) | sP2(X210,X209,X211)) ) | $spl176),
  inference(resolution,[],[f15040,f7601])).
fof(f15367,plain,(
  ( ! [X206,X204,X207,X205] : (p(s(bool,X204)) | sP12('F_0',sK62) | sP12(X205,X204) | sP8(X205,X206,X207)) ) | $spl176),
  inference(resolution,[],[f15040,f7402])).
fof(f15366,plain,(
  ( ! [X202,X200,X203,X201] : (p(s(bool,X200)) | sP12('F_0',sK62) | sP12(X201,X200) | sP8(X202,X201,X203)) ) | $spl176),
  inference(resolution,[],[f15040,f7401])).
fof(f15365,plain,(
  ( ! [X198,X196,X199,X197] : (p(s(bool,X196)) | sP12('F_0',sK62) | sP12(X197,X196) | sP6(X197,X198,X199)) ) | $spl176),
  inference(resolution,[],[f15040,f7388])).
fof(f15364,plain,(
  ( ! [X194,X192,X195,X193] : (p(s(bool,X192)) | sP12('F_0',sK62) | sP12(X193,X192) | sP4(X193,X194,X195)) ) | $spl176),
  inference(resolution,[],[f15040,f7372])).
fof(f15363,plain,(
  ( ! [X191,X189,X190,X188] : (p(s(bool,X188)) | sP12('F_0',sK62) | sP12(X189,X188) | sP4(X190,X191,X189)) ) | $spl176),
  inference(resolution,[],[f15040,f7371])).
fof(f15362,plain,(
  ( ! [X187,X185,X186,X184] : (p(s(bool,X184)) | sP12('F_0',sK62) | sP12(X185,X184) | sP1(X185,X186,X187)) ) | $spl176),
  inference(resolution,[],[f15040,f7343])).
fof(f15361,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,X180)) | sP12('F_0',sK62) | sP12(X181,X180) | sP1(X182,X181,X183)) ) | $spl176),
  inference(resolution,[],[f15040,f7342])).
fof(f15360,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X176)) | sP12('F_0',sK62) | sP12(X177,X176) | sP2(X177,X178,X179)) ) | $spl176),
  inference(resolution,[],[f15040,f7340])).
fof(f15411,plain,(
  ( ! [X170] : (p(s(bool,X170)) | sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(duplicate_literal_removal,[],[f15357])).
fof(f15357,plain,(
  ( ! [X170] : (p(s(bool,X170)) | sP12('F_0','const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X170))) ) | $spl176),
  inference(resolution,[],[f15040,f3761])).
fof(f3761,plain,(
  ( ! [X215] : (sP12(X215,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X215))) )),
  inference(resolution,[],[f1433,f1170])).
fof(f15412,plain,(
  ( ! [X169] : (p(s(bool,X169)) | sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(duplicate_literal_removal,[],[f15356])).
fof(f15356,plain,(
  ( ! [X169] : (p(s(bool,X169)) | sP12('F_0','const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X169))) ) | $spl176),
  inference(resolution,[],[f15040,f3589])).
fof(f3589,plain,(
  ( ! [X215] : (sP12(X215,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X215))) )),
  inference(resolution,[],[f1432,f1170])).
fof(f15413,plain,(
  ( ! [X168] : (p(s(bool,X168)) | sP12('F_0','const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | ($spl0 | $spl176)),
  inference(duplicate_literal_removal,[],[f15355])).
fof(f15355,plain,(
  ( ! [X168] : (p(s(bool,X168)) | sP12('F_0','const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X168))) ) | ($spl0 | $spl176)),
  inference(resolution,[],[f15040,f3248])).
fof(f3248,plain,(
  ( ! [X218] : (sP12(X218,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X218))) ) | $spl0),
  inference(resolution,[],[f1417,f1170])).
fof(f15414,plain,(
  ( ! [X167] : (p(s(bool,X167)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl176),
  inference(duplicate_literal_removal,[],[f15354])).
fof(f15354,plain,(
  ( ! [X167] : (p(s(bool,X167)) | sP12('F_0','const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X167))) ) | $spl176),
  inference(resolution,[],[f15040,f1541])).
fof(f1541,plain,(
  ( ! [X218] : (sP12(X218,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X218))) )),
  inference(resolution,[],[f917,f1170])).
fof(f15183,plain,(
  ( ! [X2,X0] : (sP12('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl180),
  inference(cnf_transformation,[],[f15183_D])).
fof(f15183_D,plain,(
  ( ! [X2,X0] : (sP12('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f15061,plain,(
  ( ! [X335,X334,X336] : (sP12('F_0',X334) | p(s(bool,X335)) | ~p(s(bool,X336)) | s(bool,X334) != s(bool,X336)) ) | $spl176),
  inference(resolution,[],[f14807,f1252])).
fof(f14948,plain,(
  ( ! [X30,X28,X31,X29] : (sP12('F_0',X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl176),
  inference(resolution,[],[f14807,f926])).
fof(f14946,plain,(
  ( ! [X23,X21,X22,X20] : (sP12('F_0',X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl176),
  inference(resolution,[],[f14807,f925])).
fof(f14944,plain,(
  ( ! [X14,X12,X15,X13] : (sP12('F_0',X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl176),
  inference(resolution,[],[f14807,f924])).
fof(f14943,plain,(
  ( ! [X10,X8,X11,X9] : (sP12('F_0',X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl176),
  inference(resolution,[],[f14807,f922])).
fof(f14941,plain,(
  ( ! [X2,X0,X3,X1] : (sP12('F_0',X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl176),
  inference(resolution,[],[f14807,f920])).
fof(f14811,plain,(
  sP12('F_0','const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | $spl178),
  inference(cnf_transformation,[],[f14811_D])).
fof(f14811_D,plain,(
  sP12('F_0','const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f14909,plain,(
  ( ! [X17,X18] : (sP12('F_0','const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl39),
  inference(resolution,[],[f14655,f1339])).
fof(f14807,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP12('F_0',X0)) ) | $spl176),
  inference(cnf_transformation,[],[f14807_D])).
fof(f14807_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP12('F_0',X0)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f14908,plain,(
  ( ! [X15,X16] : (sP12('F_0',X15) | p(s(bool,X16)) | ~p(s(bool,X15))) ) | $spl39),
  inference(resolution,[],[f14655,f1340])).
fof(f14903,plain,(
  ( ! [X2,X0,X1] : (sP12('F_0',X0) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl39),
  inference(resolution,[],[f14655,f1109])).
fof(f14900,plain,(
  ( ! [X17,X18] : (sP12('F_0',X17) | p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl39),
  inference(resolution,[],[f14654,f1339])).
fof(f14899,plain,(
  ( ! [X15,X16] : (sP12('F_0','const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl39),
  inference(resolution,[],[f14654,f1340])).
fof(f14894,plain,(
  ( ! [X2,X0,X1] : (sP12('F_0',X0) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))) ) | $spl39),
  inference(resolution,[],[f14654,f1109])).
fof(f14891,plain,(
  ( ! [X15,X16] : (sP12('F_0','const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X15)) | p(s(bool,X16))) ) | $spl39),
  inference(resolution,[],[f14653,f1323])).
fof(f14886,plain,(
  ( ! [X2,X0,X1] : (sP12('F_0',X0) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl39),
  inference(resolution,[],[f14653,f1092])).
fof(f14883,plain,(
  ( ! [X17,X18] : (sP12('F_0','const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | ~p(s(bool,X17)) | p(s(bool,X18))) ) | $spl39),
  inference(resolution,[],[f14652,f1309])).
fof(f14882,plain,(
  ( ! [X15,X16] : (sP12('F_0',X15) | p(s(bool,X16)) | ~p(s(bool,X15))) ) | $spl39),
  inference(resolution,[],[f14652,f1310])).
fof(f14877,plain,(
  ( ! [X2,X0,X1] : (sP12('F_0',X0) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2))) ) | $spl39),
  inference(resolution,[],[f14652,f1075])).
fof(f14874,plain,(
  ( ! [X17,X18] : (sP12('F_0',X17) | ~p(s(bool,X17)) | p(s(bool,X18))) ) | $spl39),
  inference(resolution,[],[f14651,f1309])).
fof(f14873,plain,(
  ( ! [X15,X16] : (sP12('F_0','const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl39),
  inference(resolution,[],[f14651,f1310])).
fof(f14868,plain,(
  ( ! [X2,X0,X1] : (sP12('F_0',X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) ) | $spl39),
  inference(resolution,[],[f14651,f1075])).
fof(f14835,plain,(
  ( ! [X14,X15] : (sP12('F_0','const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X14)) | ~p(s(bool,X15))) ) | $spl39),
  inference(resolution,[],[f14650,f1278])).
fof(f14834,plain,(
  ( ! [X12,X13] : (sP12('F_0',X12) | p(s(bool,X13)) | ~p(s(bool,X12))) ) | $spl39),
  inference(resolution,[],[f14650,f1279])).
fof(f14830,plain,(
  ( ! [X2,X0,X1] : (sP12('F_0',X0) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl39),
  inference(resolution,[],[f14650,f944])).
fof(f14825,plain,(
  ( ! [X14,X15] : (sP12('F_0',X14) | p(s(bool,X15)) | ~p(s(bool,X14))) ) | $spl39),
  inference(resolution,[],[f14649,f1278])).
fof(f14824,plain,(
  ( ! [X12,X13] : (sP12('F_0','const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | $spl39),
  inference(resolution,[],[f14649,f1279])).
fof(f14820,plain,(
  ( ! [X2,X0,X1] : (sP12('F_0',X0) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))) ) | $spl39),
  inference(resolution,[],[f14649,f944])).
fof(f14817,plain,(
  ( ! [X14,X15] : (sP12('F_0','const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X14)) | ~p(s(bool,X15))) ) | $spl39),
  inference(resolution,[],[f14648,f1275])).
fof(f14816,plain,(
  ( ! [X12,X13] : (sP12('F_0',X12) | p(s(bool,X13)) | ~p(s(bool,X12))) ) | $spl39),
  inference(resolution,[],[f14648,f1276])).
fof(f14812,plain,(
  ( ! [X2,X0,X1] : (sP12('F_0',X0) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl39),
  inference(resolution,[],[f14648,f940])).
fof(f14806,plain,(
  ( ! [X14,X15] : (sP12('F_0',X14) | p(s(bool,X15)) | ~p(s(bool,X14))) ) | $spl39),
  inference(resolution,[],[f14647,f1275])).
fof(f14805,plain,(
  ( ! [X12,X13] : (sP12('F_0','const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | $spl39),
  inference(resolution,[],[f14647,f1276])).
fof(f14801,plain,(
  ( ! [X2,X0,X1] : (sP12('F_0',X0) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))) ) | $spl39),
  inference(resolution,[],[f14647,f940])).
fof(f14634,plain,(
  $false | ($spl38 | $spl79)),
  inference(subsumption_resolution,[],[f14633,f8676])).
fof(f8676,plain,(
  ~sP12(sK62,sK62) | $spl79),
  inference(cnf_transformation,[],[f8676_D])).
fof(f8676_D,plain,(
  ~sP12(sK62,sK62) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f14633,plain,(
  sP12(sK62,sK62) | $spl38),
  inference(subsumption_resolution,[],[f14619,f1247])).
fof(f14619,plain,(
  p(s(bool,'F_0')) | sP12(sK62,sK62) | $spl38),
  inference(resolution,[],[f5869,f5721])).
fof(f5721,plain,(
  ( ! [X304,X303] : (~sP12(X303,X304) | p(s(bool,X304)) | sP12(sK62,X303)) )),
  inference(resolution,[],[f5231,f1165])).
fof(f14632,plain,(
  $false | ($spl38 | $spl79)),
  inference(subsumption_resolution,[],[f14631,f8676])).
fof(f14631,plain,(
  sP12(sK62,sK62) | $spl38),
  inference(subsumption_resolution,[],[f14617,f1247])).
fof(f14617,plain,(
  p(s(bool,'F_0')) | sP12(sK62,sK62) | $spl38),
  inference(resolution,[],[f5869,f5468])).
fof(f14630,plain,(
  $false | $spl38),
  inference(subsumption_resolution,[],[f14629,f1247])).
fof(f14629,plain,(
  p(s(bool,'F_0')) | $spl38),
  inference(subsumption_resolution,[],[f14628,f5228])).
fof(f14628,plain,(
  p(s(bool,'F_0')) | sP12('F_0',sK62) | $spl38),
  inference(forward_demodulation,[],[f14616,f5240])).
fof(f14616,plain,(
  p(s(bool,sK62)) | sP12('F_0',sK62) | $spl38),
  inference(resolution,[],[f5869,f5467])).
fof(f14627,plain,(
  $false | $spl38),
  inference(subsumption_resolution,[],[f14626,f1247])).
fof(f14626,plain,(
  p(s(bool,'F_0')) | $spl38),
  inference(forward_demodulation,[],[f14625,f5240])).
fof(f14625,plain,(
  p(s(bool,sK62)) | $spl38),
  inference(subsumption_resolution,[],[f14615,f1248])).
fof(f1248,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f912])).
fof(f14615,plain,(
  p(s(bool,sK62)) | p(s(bool,'F_0')) | $spl38),
  inference(resolution,[],[f5869,f1174])).
fof(f1174,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X1,X0)) )),
  inference(cnf_transformation,[],[f876])).
fof(f876,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP12(X1,X0)) & (~sP12(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f743])).
fof(f743,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP12(X1,X0))),
  inference(definition_folding,[],[f620,f742])).
fof(f620,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f619])).
fof(f619,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f168])).
fof(f168,axiom,(
  ! [X77,X78] : ((p(s(bool,X78)) <=> ~p(s(bool,X77))) <=> ((p(s(bool,X78)) | p(s(bool,X77))) & (~p(s(bool,X77)) | ~p(s(bool,X78)))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.sat.sat_29_prove_ltb')).
fof(f14624,plain,(
  $false | $spl38),
  inference(subsumption_resolution,[],[f14623,f1247])).
fof(f14623,plain,(
  p(s(bool,'F_0')) | $spl38),
  inference(forward_demodulation,[],[f14622,f5240])).
fof(f14622,plain,(
  p(s(bool,sK62)) | $spl38),
  inference(subsumption_resolution,[],[f14612,f1248])).
fof(f14612,plain,(
  p(s(bool,'F_0')) | p(s(bool,sK62)) | $spl38),
  inference(resolution,[],[f5869,f1165])).
fof(f14621,plain,(
  $false | $spl38),
  inference(subsumption_resolution,[],[f14611,f1247])).
fof(f14611,plain,(
  p(s(bool,'F_0')) | $spl38),
  inference(resolution,[],[f5869,f5229])).
fof(f5229,plain,(
  ( ! [X203] : (~sP12(sK62,X203) | p(s(bool,X203))) )),
  inference(resolution,[],[f5116,f1165])).
fof(f5869,plain,(
  sP12(sK62,'F_0') | $spl38),
  inference(cnf_transformation,[],[f5869_D])).
fof(f5869_D,plain,(
  sP12(sK62,'F_0') <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f14487,plain,(
  ( ! [X243,X241,X242,X240] : (p(s(bool,X240)) | sP12(sK62,'F_0') | sP12(X240,X241) | sP8(X241,X242,X243)) ) | $spl116),
  inference(resolution,[],[f13596,f7664])).
fof(f14486,plain,(
  ( ! [X239,X237,X238,X236] : (p(s(bool,X236)) | sP12(sK62,'F_0') | sP12(X236,X237) | sP8(X238,X237,X239)) ) | $spl116),
  inference(resolution,[],[f13596,f7663])).
fof(f14485,plain,(
  ( ! [X235,X233,X234,X232] : (p(s(bool,X232)) | sP12(sK62,'F_0') | sP12(X232,X233) | sP6(X233,X234,X235)) ) | $spl116),
  inference(resolution,[],[f13596,f7650])).
fof(f14484,plain,(
  ( ! [X231,X229,X230,X228] : (p(s(bool,X228)) | sP12(sK62,'F_0') | sP12(X228,X229) | sP4(X229,X230,X231)) ) | $spl116),
  inference(resolution,[],[f13596,f7634])).
fof(f14483,plain,(
  ( ! [X227,X225,X226,X224] : (p(s(bool,X224)) | sP12(sK62,'F_0') | sP12(X224,X225) | sP4(X226,X227,X225)) ) | $spl116),
  inference(resolution,[],[f13596,f7633])).
fof(f14482,plain,(
  ( ! [X222,X220,X223,X221] : (p(s(bool,X220)) | sP12(sK62,'F_0') | sP12(X220,X221) | sP1(X221,X222,X223)) ) | $spl116),
  inference(resolution,[],[f13596,f7605])).
fof(f14481,plain,(
  ( ! [X218,X216,X219,X217] : (p(s(bool,X216)) | sP12(sK62,'F_0') | sP12(X216,X217) | sP1(X218,X217,X219)) ) | $spl116),
  inference(resolution,[],[f13596,f7604])).
fof(f14480,plain,(
  ( ! [X212,X215,X213,X214] : (p(s(bool,X212)) | sP12(sK62,'F_0') | sP12(X212,X213) | sP2(X213,X214,X215)) ) | $spl116),
  inference(resolution,[],[f13596,f7602])).
fof(f14479,plain,(
  ( ! [X210,X208,X211,X209] : (p(s(bool,X208)) | sP12(sK62,'F_0') | sP12(X208,X209) | sP2(X210,X209,X211)) ) | $spl116),
  inference(resolution,[],[f13596,f7601])).
fof(f14478,plain,(
  ( ! [X206,X204,X207,X205] : (p(s(bool,X204)) | sP12(sK62,'F_0') | sP12(X205,X204) | sP8(X205,X206,X207)) ) | $spl116),
  inference(resolution,[],[f13596,f7402])).
fof(f14477,plain,(
  ( ! [X202,X200,X203,X201] : (p(s(bool,X200)) | sP12(sK62,'F_0') | sP12(X201,X200) | sP8(X202,X201,X203)) ) | $spl116),
  inference(resolution,[],[f13596,f7401])).
fof(f14476,plain,(
  ( ! [X198,X196,X199,X197] : (p(s(bool,X196)) | sP12(sK62,'F_0') | sP12(X197,X196) | sP6(X197,X198,X199)) ) | $spl116),
  inference(resolution,[],[f13596,f7388])).
fof(f14475,plain,(
  ( ! [X194,X192,X195,X193] : (p(s(bool,X192)) | sP12(sK62,'F_0') | sP12(X193,X192) | sP4(X193,X194,X195)) ) | $spl116),
  inference(resolution,[],[f13596,f7372])).
fof(f14474,plain,(
  ( ! [X191,X189,X190,X188] : (p(s(bool,X188)) | sP12(sK62,'F_0') | sP12(X189,X188) | sP4(X190,X191,X189)) ) | $spl116),
  inference(resolution,[],[f13596,f7371])).
fof(f14473,plain,(
  ( ! [X187,X185,X186,X184] : (p(s(bool,X184)) | sP12(sK62,'F_0') | sP12(X185,X184) | sP1(X185,X186,X187)) ) | $spl116),
  inference(resolution,[],[f13596,f7343])).
fof(f14472,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,X180)) | sP12(sK62,'F_0') | sP12(X181,X180) | sP1(X182,X181,X183)) ) | $spl116),
  inference(resolution,[],[f13596,f7342])).
fof(f14471,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X176)) | sP12(sK62,'F_0') | sP12(X177,X176) | sP2(X177,X178,X179)) ) | $spl116),
  inference(resolution,[],[f13596,f7340])).
fof(f14520,plain,(
  ( ! [X170] : (p(s(bool,X170)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(duplicate_literal_removal,[],[f14468])).
fof(f14468,plain,(
  ( ! [X170] : (p(s(bool,X170)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | p(s(bool,X170))) ) | $spl116),
  inference(resolution,[],[f13596,f3761])).
fof(f14521,plain,(
  ( ! [X169] : (p(s(bool,X169)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(duplicate_literal_removal,[],[f14467])).
fof(f14467,plain,(
  ( ! [X169] : (p(s(bool,X169)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | p(s(bool,X169))) ) | $spl116),
  inference(resolution,[],[f13596,f3589])).
fof(f14522,plain,(
  ( ! [X168] : (p(s(bool,X168)) | sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0')) ) | ($spl0 | $spl116)),
  inference(duplicate_literal_removal,[],[f14466])).
fof(f14466,plain,(
  ( ! [X168] : (p(s(bool,X168)) | sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | p(s(bool,X168))) ) | ($spl0 | $spl116)),
  inference(resolution,[],[f13596,f3248])).
fof(f14523,plain,(
  ( ! [X167] : (p(s(bool,X167)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0')) ) | $spl116),
  inference(duplicate_literal_removal,[],[f14465])).
fof(f14465,plain,(
  ( ! [X167] : (p(s(bool,X167)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),'F_0') | p(s(bool,X167))) ) | $spl116),
  inference(resolution,[],[f13596,f1541])).
fof(f14412,plain,(
  ( ! [X4,X2,X0,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X4) | ~p(s(bool,X4)) | sP13(X1,X2)) ) | $spl132),
  inference(cnf_transformation,[],[f14412_D])).
fof(f14412_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X4) | ~p(s(bool,X4)) | sP13(X1,X2)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f14290,plain,(
  ( ! [X575,X576,X578,X574,X577] : (sP12(X574,X575) | sP13(X575,X576) | p(s(bool,X577)) | ~p(s(bool,X578)) | s(bool,X574) != s(bool,X578)) ) | $spl126),
  inference(resolution,[],[f13902,f1252])).
fof(f14177,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP12(X42,X43) | sP13(X43,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X42) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl126),
  inference(resolution,[],[f13902,f926])).
fof(f14175,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP12(X30,X31) | sP13(X31,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X30) != s(bool,X34)) ) | $spl126),
  inference(resolution,[],[f13902,f925])).
fof(f14173,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP12(X18,X19) | sP13(X19,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X18) != s(bool,X22)) ) | $spl126),
  inference(resolution,[],[f13902,f924])).
fof(f14172,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP12(X12,X13) | sP13(X13,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X12) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl126),
  inference(resolution,[],[f13902,f922])).
fof(f14170,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP13(X1,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X0) != s(bool,X4)) ) | $spl126),
  inference(resolution,[],[f13902,f920])).
fof(f14161,plain,(
  ( ! [X4,X2,X0,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X4) | ~p(s(bool,X4)) | sP13(X2,X1)) ) | $spl130),
  inference(cnf_transformation,[],[f14161_D])).
fof(f14161_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X4) | ~p(s(bool,X4)) | sP13(X2,X1)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f14039,plain,(
  ( ! [X575,X576,X578,X574,X577] : (sP12(X574,X575) | sP13(X576,X575) | p(s(bool,X577)) | ~p(s(bool,X578)) | s(bool,X574) != s(bool,X578)) ) | $spl122),
  inference(resolution,[],[f13862,f1252])).
fof(f13926,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP12(X42,X43) | sP13(X44,X43) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X42) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl122),
  inference(resolution,[],[f13862,f926])).
fof(f13924,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP12(X30,X31) | sP13(X32,X31) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X30) != s(bool,X34)) ) | $spl122),
  inference(resolution,[],[f13862,f925])).
fof(f13922,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP12(X18,X19) | sP13(X20,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X18) != s(bool,X22)) ) | $spl122),
  inference(resolution,[],[f13862,f924])).
fof(f13921,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP12(X12,X13) | sP13(X14,X13) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X12) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl122),
  inference(resolution,[],[f13862,f922])).
fof(f13919,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP13(X2,X1) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X0) != s(bool,X4)) ) | $spl122),
  inference(resolution,[],[f13862,f920])).
fof(f13902,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP12(X0,X1) | sP13(X1,X2)) ) | $spl126),
  inference(cnf_transformation,[],[f13902_D])).
fof(f13902_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP12(X0,X1) | sP13(X1,X2)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f13892,plain,(
  ( ! [X26,X24,X27,X25] : (sP12(X24,X25) | sP13(X25,X26) | p(s(bool,X27)) | ~p(s(bool,X24))) )),
  inference(resolution,[],[f13315,f1275])).
fof(f13891,plain,(
  ( ! [X23,X21,X22,X20] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | sP13(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23))) )),
  inference(resolution,[],[f13315,f1276])).
fof(f13887,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP13(X1,X2) | ~p(s(bool,X3)) | ~p(s(bool,X0)) | p(s(bool,X4))) )),
  inference(resolution,[],[f13315,f940])).
fof(f13862,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP12(X0,X1) | sP13(X2,X1)) ) | $spl122),
  inference(cnf_transformation,[],[f13862_D])).
fof(f13862_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP12(X0,X1) | sP13(X2,X1)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f13852,plain,(
  ( ! [X26,X24,X27,X25] : (sP12(X24,X25) | sP13(X26,X25) | p(s(bool,X27)) | ~p(s(bool,X24))) )),
  inference(resolution,[],[f13314,f1275])).
fof(f13851,plain,(
  ( ! [X23,X21,X22,X20] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | sP13(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23))) )),
  inference(resolution,[],[f13314,f1276])).
fof(f13847,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP13(X2,X1) | ~p(s(bool,X3)) | ~p(s(bool,X0)) | p(s(bool,X4))) )),
  inference(resolution,[],[f13314,f940])).
fof(f13496,plain,(
  sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | $spl118),
  inference(cnf_transformation,[],[f13496_D])).
fof(f13496_D,plain,(
  sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f13810,plain,(
  ( ! [X17,X18] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl116),
  inference(resolution,[],[f13572,f1339])).
fof(f13800,plain,(
  ( ! [X15,X16] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl116),
  inference(resolution,[],[f13571,f1340])).
fof(f13794,plain,(
  ( ! [X15,X16] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | p(s(bool,X15)) | p(s(bool,X16))) ) | $spl116),
  inference(resolution,[],[f13558,f1323])).
fof(f13788,plain,(
  ( ! [X17,X18] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | ~p(s(bool,X17)) | p(s(bool,X18))) ) | $spl116),
  inference(resolution,[],[f13542,f1309])).
fof(f13779,plain,(
  ( ! [X15,X16] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl116),
  inference(resolution,[],[f13541,f1310])).
fof(f13770,plain,(
  ( ! [X14,X15] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | p(s(bool,X14)) | ~p(s(bool,X15))) ) | $spl116),
  inference(resolution,[],[f13513,f1278])).
fof(f13761,plain,(
  ( ! [X12,X13] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | p(s(bool,X12)) | ~p(s(bool,X13))) ) | $spl116),
  inference(resolution,[],[f13512,f1279])).
fof(f13753,plain,(
  ( ! [X14,X15] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | p(s(bool,X14)) | ~p(s(bool,X15))) ) | $spl116),
  inference(resolution,[],[f13510,f1275])).
fof(f13739,plain,(
  ( ! [X2,X0] : (sP12(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl120),
  inference(cnf_transformation,[],[f13739_D])).
fof(f13739_D,plain,(
  ( ! [X2,X0] : (sP12(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f13617,plain,(
  ( ! [X335,X334,X336] : (sP12(X334,'F_0') | p(s(bool,X335)) | ~p(s(bool,X336)) | s(bool,X334) != s(bool,X336)) ) | $spl116),
  inference(resolution,[],[f13492,f1252])).
fof(f13504,plain,(
  ( ! [X30,X28,X31,X29] : (sP12(X28,'F_0') | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl116),
  inference(resolution,[],[f13492,f926])).
fof(f13502,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,'F_0') | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl116),
  inference(resolution,[],[f13492,f925])).
fof(f13500,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,'F_0') | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl116),
  inference(resolution,[],[f13492,f924])).
fof(f13499,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,'F_0') | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl116),
  inference(resolution,[],[f13492,f922])).
fof(f13497,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,'F_0') | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl116),
  inference(resolution,[],[f13492,f920])).
fof(f13492,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP12(X0,'F_0')) ) | $spl116),
  inference(cnf_transformation,[],[f13492_D])).
fof(f13492_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP12(X0,'F_0')) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f13491,plain,(
  ( ! [X14,X15] : (sP12(X14,'F_0') | p(s(bool,X15)) | ~p(s(bool,X14))) )),
  inference(resolution,[],[f13346,f1275])).
fof(f13490,plain,(
  ( ! [X12,X13] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),'F_0') | p(s(bool,X12)) | ~p(s(bool,X13))) )),
  inference(resolution,[],[f13346,f1276])).
fof(f13486,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,'F_0') | ~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(resolution,[],[f13346,f940])).
fof(f13483,plain,(
  ( ! [X771,X770,X772] : (sP12(X770,'F_0') | sP2(X771,X770,X772)) ) | $spl112),
  inference(cnf_transformation,[],[f13483_D])).
fof(f13483_D,plain,(
  ( ! [X771,X770,X772] : (sP12(X770,'F_0') | sP2(X771,X770,X772)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f13459,plain,(
  ( ! [X111,X109,X110,X108] : (p(s(bool,X108)) | sP12(X109,'F_0') | sP2(X110,X109,X111)) )),
  inference(duplicate_literal_removal,[],[f13422])).
fof(f13422,plain,(
  ( ! [X111,X109,X110,X108] : (p(s(bool,X108)) | sP12(X109,'F_0') | sP2(X110,X109,X111) | p(s(bool,X108))) )),
  inference(superposition,[],[f8645,f1211])).
fof(f13460,plain,(
  ( ! [X107,X105,X106,X104] : (p(s(bool,X104)) | sP12(X105,'F_0') | sP2(X106,X105,X107)) )),
  inference(duplicate_literal_removal,[],[f13421])).
fof(f13421,plain,(
  ( ! [X107,X105,X106,X104] : (p(s(bool,X104)) | sP12(X105,'F_0') | sP2(X106,X105,X107) | p(s(bool,X104))) )),
  inference(superposition,[],[f8645,f1233])).
fof(f1233,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f908])).
fof(f13461,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,X100)) | sP12(X101,'F_0') | sP2(X102,X101,X103)) )),
  inference(duplicate_literal_removal,[],[f13420])).
fof(f13420,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,X100)) | sP12(X101,'F_0') | sP2(X102,X101,X103) | p(s(bool,X100))) )),
  inference(superposition,[],[f8645,f1235])).
fof(f1235,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f909])).
fof(f909,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f651])).
fof(f651,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f650])).
fof(f650,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.EQ_CLAUSES|split|2')).
fof(f13347,plain,(
  ( ! [X817,X818,X816,X815] : (sP12(X815,'F_0') | sP2(X816,X815,X817) | p(s(bool,X818))) )),
  inference(resolution,[],[f8645,f1248])).
fof(f13345,plain,(
  ( ! [X810,X808,X811,X809] : (sP12(X808,'F_0') | sP2(X809,X808,X810) | p(s(bool,X811))) )),
  inference(resolution,[],[f8645,f1245])).
fof(f1245,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f723])).
fof(f723,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f662])).
fof(f662,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X14] : (p(s(bool,'F_0')) => p(s(bool,X14)))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.FALSITY')).
fof(f13344,plain,(
  ( ! [X806,X804,X807,X805] : (sP12(X804,'F_0') | sP2(X805,X804,X806) | ~p(s(bool,X807))) )),
  inference(resolution,[],[f8645,f1243])).
fof(f1243,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f720])).
fof(f720,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f659])).
fof(f659,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f658])).
fof(f658,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f301])).
fof(f301,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,'F_0'))) => ~p(s(bool,X14)))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.IMP_F')).
fof(f13343,plain,(
  ( ! [X802,X800,X803,X801] : (sP12(X800,'F_0') | sP2(X801,X800,X802) | p(s(bool,X803))) )),
  inference(resolution,[],[f8645,f1231])).
fof(f1231,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f907])).
fof(f907,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f906])).
fof(f906,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f647])).
fof(f647,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X14] : ((p(s(bool,'F_0')) & p(s(bool,X14))) <=> p(s(bool,'F_0')))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.AND_CLAUSES|split|2')).
fof(f13342,plain,(
  ( ! [X798,X796,X799,X797] : (sP12(X796,'F_0') | sP2(X797,X796,X798) | p(s(bool,X799))) )),
  inference(resolution,[],[f8645,f1224])).
fof(f1224,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f903])).
fof(f903,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f902])).
fof(f902,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f645])).
fof(f645,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X14] : ((p(s(bool,X14)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.AND_CLAUSES|split|3')).
fof(f13341,plain,(
  ( ! [X795,X793,X794,X792] : (sP12(X792,'F_0') | sP2(X793,X792,X794) | p(s(bool,X795))) )),
  inference(resolution,[],[f8645,f1220])).
fof(f1220,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f901])).
fof(f901,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f900])).
fof(f900,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f644])).
fof(f644,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X14] : ((p(s(bool,'F_0')) | p(s(bool,X14))) <=> p(s(bool,X14)))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.OR_CLAUSES|split|2')).
fof(f13340,plain,(
  ( ! [X790,X788,X791,X789] : (sP12(X788,'F_0') | sP2(X789,X788,X790) | p(s(bool,X791))) )),
  inference(resolution,[],[f8645,f1215])).
fof(f13484,plain,(
  ( ! [X787] : (s(bool,'F_0') = s(bool,X787)) ) | $spl114),
  inference(cnf_transformation,[],[f13484_D])).
fof(f13484_D,plain,(
  ( ! [X787] : (s(bool,'F_0') = s(bool,X787)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f13339,plain,(
  ( ! [X786,X784,X787,X785] : (sP12(X784,'F_0') | sP2(X785,X784,X786) | s(bool,'F_0') = s(bool,X787)) )),
  inference(resolution,[],[f8645,f1212])).
fof(f1212,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f895])).
fof(f13338,plain,(
  ( ! [X783,X781,X782,X780] : (sP12(X780,'F_0') | sP2(X781,X780,X782) | ~p(s(bool,X783))) )),
  inference(resolution,[],[f8645,f1209])).
fof(f1209,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f893])).
fof(f893,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f892])).
fof(f892,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f713])).
fof(f713,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f640])).
fof(f640,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f639])).
fof(f639,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X14)))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.IMP_CLAUSES|split|4')).
fof(f13482,plain,(
  ( ! [X773,X774] : (sP13(X773,X774)) ) | $spl110),
  inference(cnf_transformation,[],[f13482_D])).
fof(f13482_D,plain,(
  ( ! [X773,X774] : (sP13(X773,X774)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f13337,plain,(
  ( ! [X779,X777,X775,X778,X776] : (sP12(X775,'F_0') | sP2(X776,X775,X777) | sP13(X778,X779)) )),
  inference(resolution,[],[f8645,f1189])).
fof(f1189,plain,(
  ( ! [X0,X1] : (sP13(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f882])).
fof(f13336,plain,(
  ( ! [X771,X773,X770,X774,X772] : (sP12(X770,'F_0') | sP2(X771,X770,X772) | sP13(X773,X774)) )),
  inference(resolution,[],[f8645,f1184])).
fof(f1184,plain,(
  ( ! [X0,X1] : (sP13(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f880])).
fof(f13479,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP2(X2,X0,X3)) ) | $spl108),
  inference(cnf_transformation,[],[f13479_D])).
fof(f13479_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X5) | p(s(bool,X5)) | sP2(X2,X0,X3)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f13323,plain,(
  ( ! [X707,X708,X704,X706,X709,X705] : (sP12(X704,X705) | sP2(X706,X704,X707) | p(s(bool,X708)) | p(s(bool,X709)) | s(bool,X705) != s(bool,X709)) )),
  inference(resolution,[],[f8645,f1252])).
fof(f13210,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (sP12(X49,X50) | sP2(X51,X49,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f8645,f926])).
fof(f13207,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (sP12(X28,X29) | sP2(X30,X28,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) )),
  inference(resolution,[],[f8645,f925])).
fof(f13206,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (sP12(X21,X22) | sP2(X23,X21,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) )),
  inference(resolution,[],[f8645,f924])).
fof(f13205,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (sP12(X14,X15) | sP2(X16,X14,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f8645,f922])).
fof(f13203,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP12(X0,X1) | sP2(X2,X0,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) )),
  inference(resolution,[],[f8645,f920])).
fof(f13160,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl106),
  inference(cnf_transformation,[],[f13160_D])).
fof(f13160_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f13039,plain,(
  ( ! [X335,X334,X336] : (s(bool,X334) = s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X335)) | ~p(s(bool,X336)) | s(bool,X334) != s(bool,X336)) )),
  inference(resolution,[],[f1459,f1252])).
fof(f1459,plain,(
  ( ! [X57] : (~p(s(bool,X57)) | s(bool,X57) = s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f917,f965])).
fof(f12926,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f1459,f926])).
fof(f12924,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f1459,f925])).
fof(f12922,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1459,f924])).
fof(f12921,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f1459,f922])).
fof(f12919,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1459,f920])).
fof(f12781,plain,(
  ( ! [X45] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X45))) )),
  inference(subsumption_resolution,[],[f12773,f1111])).
fof(f12773,plain,(
  ( ! [X45,X44] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~sP8(X44,X45,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X45))) )),
  inference(resolution,[],[f5743,f1504])).
fof(f1504,plain,(
  ( ! [X147,X148] : (~sP9('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X148,X147) | p(s(bool,X147))) )),
  inference(resolution,[],[f917,f1103])).
fof(f5743,plain,(
  ( ! [X341,X340,X342] : (sP9(X340,X341,X342) | sP12(sK62,X340) | ~sP8(X341,X342,X340)) )),
  inference(resolution,[],[f5231,f1255])).
fof(f1255,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f1108])).
fof(f1108,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f848])).
fof(f12780,plain,(
  ( ! [X42] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X42))) )),
  inference(subsumption_resolution,[],[f12772,f1112])).
fof(f12772,plain,(
  ( ! [X43,X42] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~sP8(X42,X43,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X42))) )),
  inference(resolution,[],[f5743,f1505])).
fof(f1505,plain,(
  ( ! [X149,X150] : (~sP9('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X149,X150) | p(s(bool,X149))) )),
  inference(resolution,[],[f917,f1104])).
fof(f12742,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X29))) )),
  inference(resolution,[],[f5715,f1533])).
fof(f1533,plain,(
  ( ! [X206,X205] : (sP11('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X205,X206) | ~p(s(bool,X206))) )),
  inference(resolution,[],[f917,f1140])).
fof(f1140,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f862])).
fof(f862,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f861])).
fof(f861,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f860])).
fof(f860,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f740])).
fof(f740,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f5715,plain,(
  ( ! [X288,X287,X286] : (~sP11(X286,X287,X288) | p(s(bool,X287)) | sP12(sK62,X286)) )),
  inference(resolution,[],[f5231,f1136])).
fof(f1136,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f862])).
fof(f7981,plain,(
  sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | $spl66),
  inference(cnf_transformation,[],[f7981_D])).
fof(f7981_D,plain,(
  sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f12723,plain,(
  ( ! [X25] : (p(s(bool,X25)) | sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5714,f3755])).
fof(f3755,plain,(
  ( ! [X206,X207] : (sP11(X206,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X207)) )),
  inference(resolution,[],[f1433,f1141])).
fof(f1141,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f864])).
fof(f864,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))))),
  inference(flattening,[],[f863])).
fof(f863,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | (p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f741])).
fof(f741,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f612,f740])).
fof(f612,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(rectify,[],[f268])).
fof(f268,axiom,(
  ! [X6,X11,X47] : ((p(s(bool,X6)) | (p(s(bool,X11)) & p(s(bool,X47)))) <=> ((p(s(bool,X6)) | p(s(bool,X11))) & (p(s(bool,X6)) | p(s(bool,X47)))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.LEFT_OR_OVER_AND')).
fof(f5714,plain,(
  ( ! [X285,X283,X284] : (~sP11(X284,X283,X285) | p(s(bool,X284)) | sP12(sK62,X283)) )),
  inference(resolution,[],[f5231,f1136])).
fof(f7979,plain,(
  sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | $spl64),
  inference(cnf_transformation,[],[f7979_D])).
fof(f7979_D,plain,(
  sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f12722,plain,(
  ( ! [X24] : (p(s(bool,X24)) | sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5714,f3583])).
fof(f3583,plain,(
  ( ! [X206,X207] : (sP11(X206,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X207)) )),
  inference(resolution,[],[f1432,f1141])).
fof(f12720,plain,(
  ( ! [X22] : (p(s(bool,X22)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5714,f1535])).
fof(f1535,plain,(
  ( ! [X210,X209] : (sP11(X209,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X210)) )),
  inference(resolution,[],[f917,f1141])).
fof(f12700,plain,(
  ( ! [X17,X18] : (p(s(bool,X17)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f5713,f1534])).
fof(f1534,plain,(
  ( ! [X208,X207] : (sP11(X207,X208,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X207))) )),
  inference(resolution,[],[f917,f1140])).
fof(f5713,plain,(
  ( ! [X280,X282,X281] : (~sP11(X282,X281,X280) | p(s(bool,X281)) | sP12(sK62,X280)) )),
  inference(resolution,[],[f5231,f1135])).
fof(f1135,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X1)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f862])).
fof(f12689,plain,(
  ( ! [X25] : (p(s(bool,X25)) | sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5712,f3755])).
fof(f5712,plain,(
  ( ! [X277,X279,X278] : (~sP11(X279,X277,X278) | p(s(bool,X278)) | sP12(sK62,X277)) )),
  inference(resolution,[],[f5231,f1135])).
fof(f12688,plain,(
  ( ! [X24] : (p(s(bool,X24)) | sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5712,f3583])).
fof(f12686,plain,(
  ( ! [X22] : (p(s(bool,X22)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5712,f1535])).
fof(f12658,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5707,f3748])).
fof(f3748,plain,(
  ( ! [X192,X193] : (sP10('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X192,X193)) )),
  inference(resolution,[],[f1433,f1132])).
fof(f1132,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X2,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f859])).
fof(f859,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP10(X0,X2,X1)) & (~sP10(X0,X2,X1) | (p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))))),
  inference(flattening,[],[f858])).
fof(f858,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP10(X0,X2,X1)) & (~sP10(X0,X2,X1) | ((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f739])).
fof(f739,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> sP10(X0,X2,X1))),
  inference(definition_folding,[],[f611,f738])).
fof(f738,plain,(
  ! [X0,X2,X1] : (sP10(X0,X2,X1) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f611,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  inference(rectify,[],[f212])).
fof(f212,axiom,(
  ! [X6,X11,X47] : (((p(s(bool,X11)) & p(s(bool,X47))) | p(s(bool,X6))) <=> ((p(s(bool,X11)) | p(s(bool,X6))) & (p(s(bool,X47)) | p(s(bool,X6)))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.bool.RIGHT_OR_OVER_AND')).
fof(f5707,plain,(
  ( ! [X263,X262,X264] : (~sP10(X262,X263,X264) | p(s(bool,X263)) | sP12(sK62,X262)) )),
  inference(resolution,[],[f5231,f1126])).
fof(f1126,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f857])).
fof(f857,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X2)) | p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X2)) & ~p(s(bool,X0))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f856])).
fof(f856,plain,(
  ! [X0,X2,X1] : ((~sP10(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (~p(s(bool,X2)) & ~p(s(bool,X0))) | sP10(X0,X2,X1)))),
  inference(flattening,[],[f855])).
fof(f855,plain,(
  ! [X0,X2,X1] : ((~sP10(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (~p(s(bool,X2)) & ~p(s(bool,X0)))) | sP10(X0,X2,X1)))),
  inference(nnf_transformation,[],[f738])).
fof(f12657,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5707,f3576])).
fof(f3576,plain,(
  ( ! [X192,X193] : (sP10('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X192,X193)) )),
  inference(resolution,[],[f1432,f1132])).
fof(f12655,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5707,f1528])).
fof(f1528,plain,(
  ( ! [X196,X195] : (sP10('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X195,X196)) )),
  inference(resolution,[],[f917,f1132])).
fof(f12635,plain,(
  ( ! [X24,X23] : (p(s(bool,X23)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X24))) )),
  inference(resolution,[],[f5706,f1520])).
fof(f1520,plain,(
  ( ! [X180,X179] : (sP10(X179,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X180) | ~p(s(bool,X180))) )),
  inference(resolution,[],[f917,f1127])).
fof(f1127,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f857])).
fof(f5706,plain,(
  ( ! [X261,X259,X260] : (~sP10(X260,X259,X261) | p(s(bool,X260)) | sP12(sK62,X259)) )),
  inference(resolution,[],[f5231,f1126])).
fof(f12615,plain,(
  ( ! [X17,X18] : (p(s(bool,X17)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f5705,f1521])).
fof(f1521,plain,(
  ( ! [X182,X181] : (sP10(X181,X182,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X182))) )),
  inference(resolution,[],[f917,f1127])).
fof(f5705,plain,(
  ( ! [X257,X256,X258] : (~sP10(X257,X258,X256) | p(s(bool,X257)) | sP12(sK62,X256)) )),
  inference(resolution,[],[f5231,f1125])).
fof(f1125,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X2)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f857])).
fof(f12607,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5704,f3748])).
fof(f5704,plain,(
  ( ! [X255,X253,X254] : (~sP10(X253,X255,X254) | p(s(bool,X254)) | sP12(sK62,X253)) )),
  inference(resolution,[],[f5231,f1125])).
fof(f12606,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5704,f3576])).
fof(f12604,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5704,f1528])).
fof(f12564,plain,(
  ( ! [X37,X38] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X37)) | p(s(bool,X38))) )),
  inference(resolution,[],[f5700,f1504])).
fof(f5700,plain,(
  ( ! [X243,X241,X242] : (sP9(X241,X242,X243) | sP12(sK62,X241) | p(s(bool,X242))) )),
  inference(resolution,[],[f5231,f1114])).
fof(f12568,plain,(
  ( ! [X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X36))) )),
  inference(duplicate_literal_removal,[],[f12563])).
fof(f12563,plain,(
  ( ! [X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X36)) | p(s(bool,X36))) )),
  inference(resolution,[],[f5700,f1505])).
fof(f12513,plain,(
  ( ! [X38] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X38))) )),
  inference(duplicate_literal_removal,[],[f12509])).
fof(f12509,plain,(
  ( ! [X38] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X38)) | p(s(bool,X38))) )),
  inference(resolution,[],[f5698,f1504])).
fof(f5698,plain,(
  ( ! [X237,X235,X236] : (sP9(X235,X236,X237) | sP12(sK62,X235) | p(s(bool,X237))) )),
  inference(resolution,[],[f5231,f1113])).
fof(f12508,plain,(
  ( ! [X37,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X36)) | p(s(bool,X37))) )),
  inference(resolution,[],[f5698,f1505])).
fof(f12346,plain,(
  ( ! [X33,X34] : (~p(s(bool,X33)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X34))) )),
  inference(resolution,[],[f5675,f1498])).
fof(f1498,plain,(
  ( ! [X136,X135] : (sP7('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X135,X136) | ~p(s(bool,X136))) )),
  inference(resolution,[],[f917,f1097])).
fof(f1097,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f845])).
fof(f845,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f844])).
fof(f844,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f734])).
fof(f734,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f607,f733,f732])).
fof(f607,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f606])).
fof(f606,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f175])).
fof(f175,axiom,(
  ! [X38,X77,X78] : ((p(s(bool,X78)) <=> (p(s(bool,X77)) | p(s(bool,X38)))) <=> ((p(s(bool,X78)) | ~p(s(bool,X77))) & (p(s(bool,X78)) | ~p(s(bool,X38))) & (p(s(bool,X77)) | p(s(bool,X38)) | ~p(s(bool,X78)))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.sat.sat_25_prove_ltb')).
fof(f12350,plain,(
  ( ! [X32] : (~p(s(bool,X32)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f12345])).
fof(f12345,plain,(
  ( ! [X32] : (~p(s(bool,X32)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X32))) )),
  inference(resolution,[],[f5675,f1500])).
fof(f1500,plain,(
  ( ! [X140,X139] : (sP7('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X139,X140) | ~p(s(bool,X139))) )),
  inference(resolution,[],[f917,f1098])).
fof(f1098,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f845])).
fof(f12327,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f12323])).
fof(f12323,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X34))) )),
  inference(resolution,[],[f5674,f1498])).
fof(f12322,plain,(
  ( ! [X33,X32] : (~p(s(bool,X32)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X33))) )),
  inference(resolution,[],[f5674,f1500])).
fof(f12236,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X35))) )),
  inference(resolution,[],[f5661,f1484])).
fof(f1484,plain,(
  ( ! [X107,X108] : (sP5('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X107,X108) | p(s(bool,X108))) )),
  inference(resolution,[],[f917,f1080])).
fof(f12240,plain,(
  ( ! [X33] : (~p(s(bool,X33)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f12235])).
fof(f12235,plain,(
  ( ! [X33] : (~p(s(bool,X33)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X33))) )),
  inference(resolution,[],[f5661,f1486])).
fof(f1486,plain,(
  ( ! [X111,X112] : (sP5('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),X111,X112) | ~p(s(bool,X111))) )),
  inference(resolution,[],[f917,f1081])).
fof(f1081,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f837])).
fof(f12188,plain,(
  ( ! [X36] : (p(s(bool,X36)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f12184])).
fof(f12184,plain,(
  ( ! [X36] : (p(s(bool,X36)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X36))) )),
  inference(resolution,[],[f5659,f1484])).
fof(f12183,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f5659,f1486])).
fof(f12148,plain,(
  ( ! [X35] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f12142,f1111])).
fof(f12142,plain,(
  ( ! [X35,X34] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~sP8(X34,X35,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X35))) )),
  inference(resolution,[],[f5490,f1504])).
fof(f5490,plain,(
  ( ! [X341,X340,X342] : (sP9(X340,X341,X342) | sP12(X340,sK62) | ~sP8(X341,X342,X340)) )),
  inference(resolution,[],[f5230,f1255])).
fof(f12147,plain,(
  ( ! [X32] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X32))) )),
  inference(subsumption_resolution,[],[f12141,f1112])).
fof(f12141,plain,(
  ( ! [X33,X32] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~sP8(X32,X33,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X32))) )),
  inference(resolution,[],[f5490,f1505])).
fof(f12101,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~p(s(bool,X29))) )),
  inference(resolution,[],[f5462,f1533])).
fof(f5462,plain,(
  ( ! [X288,X287,X286] : (~sP11(X286,X287,X288) | p(s(bool,X287)) | sP12(X286,sK62)) )),
  inference(resolution,[],[f5230,f1136])).
fof(f7291,plain,(
  sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | $spl48),
  inference(cnf_transformation,[],[f7291_D])).
fof(f7291_D,plain,(
  sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f12082,plain,(
  ( ! [X25] : (p(s(bool,X25)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5461,f3755])).
fof(f5461,plain,(
  ( ! [X285,X283,X284] : (~sP11(X284,X283,X285) | p(s(bool,X284)) | sP12(X283,sK62)) )),
  inference(resolution,[],[f5230,f1136])).
fof(f7289,plain,(
  sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | $spl46),
  inference(cnf_transformation,[],[f7289_D])).
fof(f7289_D,plain,(
  sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f12081,plain,(
  ( ! [X24] : (p(s(bool,X24)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5461,f3583])).
fof(f12079,plain,(
  ( ! [X22] : (p(s(bool,X22)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5461,f1535])).
fof(f12059,plain,(
  ( ! [X17,X18] : (p(s(bool,X17)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f5460,f1534])).
fof(f5460,plain,(
  ( ! [X280,X282,X281] : (~sP11(X282,X281,X280) | p(s(bool,X281)) | sP12(X280,sK62)) )),
  inference(resolution,[],[f5230,f1135])).
fof(f12048,plain,(
  ( ! [X25] : (p(s(bool,X25)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5459,f3755])).
fof(f5459,plain,(
  ( ! [X277,X279,X278] : (~sP11(X279,X277,X278) | p(s(bool,X278)) | sP12(X277,sK62)) )),
  inference(resolution,[],[f5230,f1135])).
fof(f12047,plain,(
  ( ! [X24] : (p(s(bool,X24)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5459,f3583])).
fof(f12045,plain,(
  ( ! [X22] : (p(s(bool,X22)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5459,f1535])).
fof(f12008,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5454,f3748])).
fof(f5454,plain,(
  ( ! [X263,X262,X264] : (~sP10(X262,X263,X264) | p(s(bool,X263)) | sP12(X262,sK62)) )),
  inference(resolution,[],[f5230,f1126])).
fof(f12007,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5454,f3576])).
fof(f12005,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5454,f1528])).
fof(f11985,plain,(
  ( ! [X24,X23] : (p(s(bool,X23)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~p(s(bool,X24))) )),
  inference(resolution,[],[f5453,f1520])).
fof(f5453,plain,(
  ( ! [X261,X259,X260] : (~sP10(X260,X259,X261) | p(s(bool,X260)) | sP12(X259,sK62)) )),
  inference(resolution,[],[f5230,f1126])).
fof(f10710,plain,(
  ( ! [X39] : (sP12(X39,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(sK62,X39)) ) | $spl104),
  inference(cnf_transformation,[],[f10710_D])).
fof(f10710_D,plain,(
  ( ! [X39] : (sP12(X39,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(sK62,X39)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f11963,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | sP12(X36,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(sK62,X36)) )),
  inference(resolution,[],[f1340,f8348])).
fof(f10709,plain,(
  ( ! [X33] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X33) | sP12(sK62,X33)) ) | $spl102),
  inference(cnf_transformation,[],[f10709_D])).
fof(f10709_D,plain,(
  ( ! [X33] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X33) | sP12(sK62,X33)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f11961,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X30) | sP12(sK62,X30)) )),
  inference(resolution,[],[f1340,f8096])).
fof(f10708,plain,(
  ( ! [X27] : (sP12(X27,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(X27,sK62)) ) | $spl100),
  inference(cnf_transformation,[],[f10708_D])).
fof(f10708_D,plain,(
  ( ! [X27] : (sP12(X27,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(X27,sK62)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f11959,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | sP12(X24,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(X24,sK62)) )),
  inference(resolution,[],[f1340,f7663])).
fof(f10707,plain,(
  ( ! [X21] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X21) | sP12(X21,sK62)) ) | $spl98),
  inference(cnf_transformation,[],[f10707_D])).
fof(f10707_D,plain,(
  ( ! [X21] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X21) | sP12(X21,sK62)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f11957,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X18) | sP12(X18,sK62)) )),
  inference(resolution,[],[f1340,f7401])).
fof(f11939,plain,(
  ( ! [X17,X18] : (p(s(bool,X17)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f5452,f1521])).
fof(f5452,plain,(
  ( ! [X257,X256,X258] : (~sP10(X257,X258,X256) | p(s(bool,X257)) | sP12(X256,sK62)) )),
  inference(resolution,[],[f5230,f1125])).
fof(f11931,plain,(
  ( ! [X31] : (p(s(bool,X31)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5451,f3748])).
fof(f5451,plain,(
  ( ! [X255,X253,X254] : (~sP10(X253,X255,X254) | p(s(bool,X254)) | sP12(X253,sK62)) )),
  inference(resolution,[],[f5230,f1125])).
fof(f11930,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5451,f3576])).
fof(f11928,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5451,f1528])).
fof(f11906,plain,(
  ( ! [X28,X29] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X28)) | p(s(bool,X29))) )),
  inference(resolution,[],[f5447,f1504])).
fof(f5447,plain,(
  ( ! [X243,X241,X242] : (sP9(X241,X242,X243) | sP12(X241,sK62) | p(s(bool,X242))) )),
  inference(resolution,[],[f5230,f1114])).
fof(f11910,plain,(
  ( ! [X27] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X27))) )),
  inference(duplicate_literal_removal,[],[f11905])).
fof(f11905,plain,(
  ( ! [X27] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X27)) | p(s(bool,X27))) )),
  inference(resolution,[],[f5447,f1505])).
fof(f11882,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | sP12(X39,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(sK62,X39)) )),
  inference(resolution,[],[f1339,f8349])).
fof(f11880,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X33) | sP12(sK62,X33)) )),
  inference(resolution,[],[f1339,f8097])).
fof(f11878,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | sP12(X27,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(X27,sK62)) )),
  inference(resolution,[],[f1339,f7664])).
fof(f11876,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X21) | sP12(X21,sK62)) )),
  inference(resolution,[],[f1339,f7402])).
fof(f11841,plain,(
  ( ! [X29] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X29))) )),
  inference(duplicate_literal_removal,[],[f11837])).
fof(f11837,plain,(
  ( ! [X29] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X29)) | p(s(bool,X29))) )),
  inference(resolution,[],[f5445,f1504])).
fof(f5445,plain,(
  ( ! [X237,X235,X236] : (sP9(X235,X236,X237) | sP12(X235,sK62) | p(s(bool,X237))) )),
  inference(resolution,[],[f5230,f1113])).
fof(f11836,plain,(
  ( ! [X28,X27] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X27)) | p(s(bool,X28))) )),
  inference(resolution,[],[f5445,f1505])).
fof(f11737,plain,(
  ( ! [X33,X34] : (~p(s(bool,X33)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~p(s(bool,X34))) )),
  inference(resolution,[],[f5422,f1498])).
fof(f11741,plain,(
  ( ! [X32] : (~p(s(bool,X32)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f11736])).
fof(f11736,plain,(
  ( ! [X32] : (~p(s(bool,X32)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~p(s(bool,X32))) )),
  inference(resolution,[],[f5422,f1500])).
fof(f11718,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f11714])).
fof(f11714,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~p(s(bool,X34))) )),
  inference(resolution,[],[f5421,f1498])).
fof(f11713,plain,(
  ( ! [X33,X32] : (~p(s(bool,X32)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~p(s(bool,X33))) )),
  inference(resolution,[],[f5421,f1500])).
fof(f11685,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,X22)) | sP12(X23,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(sK62,X23)) )),
  inference(resolution,[],[f1323,f8335])).
fof(f11684,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X20) | sP12(sK62,X20)) )),
  inference(resolution,[],[f1323,f8083])).
fof(f11683,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | sP12(X17,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(X17,sK62)) )),
  inference(resolution,[],[f1323,f7650])).
fof(f11682,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X14) | sP12(X14,sK62)) )),
  inference(resolution,[],[f1323,f7388])).
fof(f11646,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X32))) )),
  inference(resolution,[],[f5408,f1484])).
fof(f11650,plain,(
  ( ! [X30] : (~p(s(bool,X30)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f11645])).
fof(f11645,plain,(
  ( ! [X30] : (~p(s(bool,X30)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~p(s(bool,X30))) )),
  inference(resolution,[],[f5408,f1486])).
fof(f11577,plain,(
  ( ! [X34] : (p(s(bool,X34)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f11573])).
fof(f11573,plain,(
  ( ! [X34] : (p(s(bool,X34)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X34))) )),
  inference(resolution,[],[f5406,f1484])).
fof(f11572,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | ~p(s(bool,X33))) )),
  inference(resolution,[],[f5406,f1486])).
fof(f11456,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | sP12(X36,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(sK62,X36)) )),
  inference(resolution,[],[f1310,f8318])).
fof(f11454,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X30) | sP12(sK62,X30)) )),
  inference(resolution,[],[f1310,f8066])).
fof(f11452,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | sP12(X24,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(X24,sK62)) )),
  inference(resolution,[],[f1310,f7633])).
fof(f11450,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X18) | sP12(X18,sK62)) )),
  inference(resolution,[],[f1310,f7371])).
fof(f11411,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,X37)) | p(s(bool,X38)) | sP12(X39,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(sK62,X39)) )),
  inference(resolution,[],[f1309,f8319])).
fof(f11409,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X33) | sP12(sK62,X33)) )),
  inference(resolution,[],[f1309,f8067])).
fof(f11407,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | sP12(X27,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(X27,sK62)) )),
  inference(resolution,[],[f1309,f7634])).
fof(f11405,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | p(s(bool,X20)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X21) | sP12(X21,sK62)) )),
  inference(resolution,[],[f1309,f7372])).
fof(f11384,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | sP12(X36,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(sK62,X36)) )),
  inference(resolution,[],[f1279,f8289])).
fof(f11382,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X30) | sP12(sK62,X30)) )),
  inference(resolution,[],[f1279,f8037])).
fof(f11380,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | sP12(X24,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(X24,sK62)) )),
  inference(resolution,[],[f1279,f7604])).
fof(f11378,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X18) | sP12(X18,sK62)) )),
  inference(resolution,[],[f1279,f7342])).
fof(f11359,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | sP12(X39,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(sK62,X39)) )),
  inference(resolution,[],[f1278,f8290])).
fof(f11357,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X33) | sP12(sK62,X33)) )),
  inference(resolution,[],[f1278,f8038])).
fof(f11355,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | sP12(X27,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(X27,sK62)) )),
  inference(resolution,[],[f1278,f7605])).
fof(f11353,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X21) | sP12(X21,sK62)) )),
  inference(resolution,[],[f1278,f7343])).
fof(f11329,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f11301,f1112])).
fof(f11301,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X62,X63,X64) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f8349,f3758])).
fof(f11328,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f11300,f1112])).
fof(f11300,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X59,X60,X61) | ~p(s(bool,X59))) )),
  inference(resolution,[],[f8349,f3586])).
fof(f11327,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f11298,f1112])).
fof(f11298,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X53,X54,X55) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f8349,f1538])).
fof(f11271,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | sP12(X36,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(sK62,X36)) )),
  inference(resolution,[],[f1276,f8286])).
fof(f11269,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X30) | sP12(sK62,X30)) )),
  inference(resolution,[],[f1276,f8034])).
fof(f11267,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | sP12(X24,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(X24,sK62)) )),
  inference(resolution,[],[f1276,f7601])).
fof(f11265,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X18) | sP12(X18,sK62)) )),
  inference(resolution,[],[f1276,f7339])).
fof(f11241,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f11213,f1111])).
fof(f11213,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X62,X63,X64) | ~p(s(bool,X63))) )),
  inference(resolution,[],[f8348,f3758])).
fof(f11240,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f11212,f1111])).
fof(f11212,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X59,X60,X61) | ~p(s(bool,X60))) )),
  inference(resolution,[],[f8348,f3586])).
fof(f11239,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f11210,f1111])).
fof(f11210,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X53,X54,X55) | ~p(s(bool,X54))) )),
  inference(resolution,[],[f8348,f1538])).
fof(f11181,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f11153,f1095])).
fof(f11153,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X62,X63,X64) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f8335,f3758])).
fof(f11180,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f11152,f1095])).
fof(f11152,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X59,X60,X61) | ~p(s(bool,X59))) )),
  inference(resolution,[],[f8335,f3586])).
fof(f11179,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f11150,f1095])).
fof(f11150,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X53,X54,X55) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f8335,f1538])).
fof(f11121,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f11093,f1078])).
fof(f11093,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X62,X63,X64) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f8319,f3758])).
fof(f11120,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f11092,f1078])).
fof(f11092,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X59,X60,X61) | ~p(s(bool,X59))) )),
  inference(resolution,[],[f8319,f3586])).
fof(f11119,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f11090,f1078])).
fof(f11090,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X53,X54,X55) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f8319,f1538])).
fof(f11060,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f11032,f1076])).
fof(f11032,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X62,X63,X64) | ~p(s(bool,X64))) )),
  inference(resolution,[],[f8318,f3758])).
fof(f11059,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f11031,f1076])).
fof(f11031,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X59,X60,X61) | ~p(s(bool,X61))) )),
  inference(resolution,[],[f8318,f3586])).
fof(f11058,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f11029,f1076])).
fof(f11029,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X53,X54,X55) | ~p(s(bool,X55))) )),
  inference(resolution,[],[f8318,f1538])).
fof(f10841,plain,(
  $false | $spl10),
  inference(resolution,[],[f1423,f916])).
fof(f10843,plain,(
  $false | $spl10),
  inference(resolution,[],[f1423,f1051])).
fof(f1051,plain,(
  ( ! [X0] : (p(s(bool,'const.operator.MONOID_2'(s(fun('type.list.list'(X0),fun('type.list.list'(X0),'type.list.list'(X0))),'const.list.APPEND_0'),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f601])).
fof(f601,plain,(
  ! [X0] : p(s(bool,'const.operator.MONOID_2'(s(fun('type.list.list'(X0),fun('type.list.list'(X0),'type.list.list'(X0))),'const.list.APPEND_0'),s('type.list.list'(X0),'const.list.NIL_0'))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : p(s(bool,'const.operator.MONOID_2'(s(fun('type.list.list'(X6),fun('type.list.list'(X6),'type.list.list'(X6))),'const.list.APPEND_0'),s('type.list.list'(X6),'const.list.NIL_0'))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.rich_list.MONOID_APPEND_NIL')).
fof(f10844,plain,(
  $false | $spl10),
  inference(resolution,[],[f1423,f917])).
fof(f10882,plain,(
  $false | ($spl0 | $spl10)),
  inference(resolution,[],[f1423,f1417])).
fof(f10884,plain,(
  $false | $spl10),
  inference(resolution,[],[f1423,f1432])).
fof(f10886,plain,(
  $false | $spl10),
  inference(resolution,[],[f1423,f1433])).
fof(f10698,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | sP12(X39,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(sK62,X39)) )),
  inference(resolution,[],[f1275,f8287])).
fof(f10696,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X33) | sP12(sK62,X33)) )),
  inference(resolution,[],[f1275,f8035])).
fof(f10694,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | sP12(X27,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | sP12(X27,sK62)) )),
  inference(resolution,[],[f1275,f7602])).
fof(f10692,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X21) | sP12(X21,sK62)) )),
  inference(resolution,[],[f1275,f7340])).
fof(f10681,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(sK62,X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) )),
  inference(subsumption_resolution,[],[f10660,f5721])).
fof(f10660,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,X9) | sP12(sK62,X8) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) )),
  inference(resolution,[],[f8290,f938])).
fof(f10666,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f10639,f947])).
fof(f10639,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X62,X63,X64) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f8290,f3758])).
fof(f10665,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f10638,f947])).
fof(f10638,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X59,X60,X61) | ~p(s(bool,X59))) )),
  inference(resolution,[],[f8290,f3586])).
fof(f10664,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f10636,f947])).
fof(f10636,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X53,X54,X55) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f8290,f1538])).
fof(f10605,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f10578,f946])).
fof(f10578,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X62,X63,X64) | ~p(s(bool,X63))) )),
  inference(resolution,[],[f8289,f3758])).
fof(f10604,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f10577,f946])).
fof(f10577,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X59,X60,X61) | ~p(s(bool,X60))) )),
  inference(resolution,[],[f8289,f3586])).
fof(f10603,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f10575,f946])).
fof(f10575,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X53,X54,X55) | ~p(s(bool,X54))) )),
  inference(resolution,[],[f8289,f1538])).
fof(f10559,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(sK62,X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) )),
  inference(subsumption_resolution,[],[f10538,f5721])).
fof(f10538,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,X13) | sP12(sK62,X12) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) )),
  inference(resolution,[],[f8287,f939])).
fof(f10543,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f10516,f943])).
fof(f10516,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X62,X63,X64) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f8287,f3758])).
fof(f10542,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f10515,f943])).
fof(f10515,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X59,X60,X61) | ~p(s(bool,X59))) )),
  inference(resolution,[],[f8287,f3586])).
fof(f10541,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f10513,f943])).
fof(f10513,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X53,X54,X55) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f8287,f1538])).
fof(f10483,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f10456,f942])).
fof(f10456,plain,(
  ( ! [X64,X62,X63] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X62,X63,X64) | ~p(s(bool,X63))) )),
  inference(resolution,[],[f8286,f3758])).
fof(f10482,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f10455,f942])).
fof(f10455,plain,(
  ( ! [X61,X59,X60] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X59,X60,X61) | ~p(s(bool,X60))) )),
  inference(resolution,[],[f8286,f3586])).
fof(f10481,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f10453,f942])).
fof(f10453,plain,(
  ( ! [X54,X55,X53] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X53,X54,X55) | ~p(s(bool,X54))) )),
  inference(resolution,[],[f8286,f1538])).
fof(f10423,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f10389,f1112])).
fof(f10389,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X44,X45,X46) | ~p(s(bool,X44))) )),
  inference(resolution,[],[f8097,f3759])).
fof(f10422,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f10388,f1112])).
fof(f10388,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X41,X42,X43) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f8097,f3587])).
fof(f10421,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f10386,f1112])).
fof(f10386,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X35,X36,X37) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f8097,f1539])).
fof(f10349,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f10315,f1111])).
fof(f10315,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X44,X45,X46) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f8096,f3759])).
fof(f10348,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f10314,f1111])).
fof(f10314,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X41,X42,X43) | ~p(s(bool,X42))) )),
  inference(resolution,[],[f8096,f3587])).
fof(f10347,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f10312,f1111])).
fof(f10312,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP8(X35,X36,X37) | ~p(s(bool,X36))) )),
  inference(resolution,[],[f8096,f1539])).
fof(f10289,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f10255,f1095])).
fof(f10255,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X44,X45,X46) | ~p(s(bool,X44))) )),
  inference(resolution,[],[f8083,f3759])).
fof(f10288,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f10254,f1095])).
fof(f10254,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X41,X42,X43) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f8083,f3587])).
fof(f10287,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f10252,f1095])).
fof(f10252,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP6(X35,X36,X37) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f8083,f1539])).
fof(f10229,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f10195,f1078])).
fof(f10195,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X44,X45,X46) | ~p(s(bool,X44))) )),
  inference(resolution,[],[f8067,f3759])).
fof(f10228,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f10194,f1078])).
fof(f10194,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X41,X42,X43) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f8067,f3587])).
fof(f10227,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f10192,f1078])).
fof(f10192,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X35,X36,X37) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f8067,f1539])).
fof(f10168,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f10134,f1076])).
fof(f10134,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X44,X45,X46) | ~p(s(bool,X46))) )),
  inference(resolution,[],[f8066,f3759])).
fof(f10167,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f10133,f1076])).
fof(f10133,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X41,X42,X43) | ~p(s(bool,X43))) )),
  inference(resolution,[],[f8066,f3587])).
fof(f10166,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f10131,f1076])).
fof(f10131,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP4(X35,X36,X37) | ~p(s(bool,X37))) )),
  inference(resolution,[],[f8066,f1539])).
fof(f10120,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(sK62,X9) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) )),
  inference(subsumption_resolution,[],[f10099,f5720])).
fof(f5720,plain,(
  ( ! [X302,X301] : (~sP12(X302,X301) | p(s(bool,X302)) | sP12(sK62,X301)) )),
  inference(resolution,[],[f5231,f1165])).
fof(f10099,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,X9) | sP12(sK62,X9) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) )),
  inference(resolution,[],[f8038,f938])).
fof(f10105,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f10072,f947])).
fof(f10072,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X44,X45,X46) | ~p(s(bool,X44))) )),
  inference(resolution,[],[f8038,f3759])).
fof(f10104,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f10071,f947])).
fof(f10071,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X41,X42,X43) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f8038,f3587])).
fof(f10103,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f10069,f947])).
fof(f10069,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X35,X36,X37) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f8038,f1539])).
fof(f10031,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f9998,f946])).
fof(f9998,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X44,X45,X46) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f8037,f3759])).
fof(f10030,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f9997,f946])).
fof(f9997,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X41,X42,X43) | ~p(s(bool,X42))) )),
  inference(resolution,[],[f8037,f3587])).
fof(f10029,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f9995,f946])).
fof(f9995,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP1(X35,X36,X37) | ~p(s(bool,X36))) )),
  inference(resolution,[],[f8037,f1539])).
fof(f9984,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(sK62,X13) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) )),
  inference(subsumption_resolution,[],[f9963,f5720])).
fof(f9963,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,X13) | sP12(sK62,X13) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) )),
  inference(resolution,[],[f8035,f939])).
fof(f9968,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f9935,f943])).
fof(f9935,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X44,X45,X46) | ~p(s(bool,X44))) )),
  inference(resolution,[],[f8035,f3759])).
fof(f9967,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f9934,f943])).
fof(f9934,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X41,X42,X43) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f8035,f3587])).
fof(f9966,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f9932,f943])).
fof(f9932,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X35,X36,X37) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f8035,f1539])).
fof(f9908,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f9875,f942])).
fof(f9875,plain,(
  ( ! [X45,X46,X44] : (sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X44,X45,X46) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f8034,f3759])).
fof(f9907,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f9874,f942])).
fof(f9874,plain,(
  ( ! [X43,X41,X42] : (sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X41,X42,X43) | ~p(s(bool,X42))) )),
  inference(resolution,[],[f8034,f3587])).
fof(f9906,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f9872,f942])).
fof(f9872,plain,(
  ( ! [X37,X35,X36] : (sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | sP2(X35,X36,X37) | ~p(s(bool,X36))) )),
  inference(resolution,[],[f8034,f1539])).
fof(f9846,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f9818,f1112])).
fof(f9818,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X62,X63,X64) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f7664,f3758])).
fof(f9845,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f9817,f1112])).
fof(f9817,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X59,X60,X61) | ~p(s(bool,X59))) )),
  inference(resolution,[],[f7664,f3586])).
fof(f9844,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f9815,f1112])).
fof(f9815,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X53,X54,X55) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f7664,f1538])).
fof(f9782,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f9754,f1111])).
fof(f9754,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X62,X63,X64) | ~p(s(bool,X63))) )),
  inference(resolution,[],[f7663,f3758])).
fof(f9781,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f9753,f1111])).
fof(f9753,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X59,X60,X61) | ~p(s(bool,X60))) )),
  inference(resolution,[],[f7663,f3586])).
fof(f9780,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f9751,f1111])).
fof(f9751,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X53,X54,X55) | ~p(s(bool,X54))) )),
  inference(resolution,[],[f7663,f1538])).
fof(f9720,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f9692,f1095])).
fof(f9692,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X62,X63,X64) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f7650,f3758])).
fof(f9719,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f9691,f1095])).
fof(f9691,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X59,X60,X61) | ~p(s(bool,X59))) )),
  inference(resolution,[],[f7650,f3586])).
fof(f9718,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f9689,f1095])).
fof(f9689,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X53,X54,X55) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f7650,f1538])).
fof(f9658,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f9630,f1078])).
fof(f9630,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X62,X63,X64) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f7634,f3758])).
fof(f9657,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f9629,f1078])).
fof(f9629,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X59,X60,X61) | ~p(s(bool,X59))) )),
  inference(resolution,[],[f7634,f3586])).
fof(f9656,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f9627,f1078])).
fof(f9627,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X53,X54,X55) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f7634,f1538])).
fof(f9595,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f9567,f1076])).
fof(f9567,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X62,X63,X64) | ~p(s(bool,X64))) )),
  inference(resolution,[],[f7633,f3758])).
fof(f9594,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f9566,f1076])).
fof(f9566,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X59,X60,X61) | ~p(s(bool,X61))) )),
  inference(resolution,[],[f7633,f3586])).
fof(f9593,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f9564,f1076])).
fof(f9564,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X53,X54,X55) | ~p(s(bool,X55))) )),
  inference(resolution,[],[f7633,f1538])).
fof(f9547,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,sK62) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) )),
  inference(subsumption_resolution,[],[f9524,f5468])).
fof(f9524,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,X9) | sP12(X8,sK62) | sP3(X10,X9,X11) | p(s(bool,X9)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) )),
  inference(resolution,[],[f7605,f938])).
fof(f9530,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f9503,f947])).
fof(f9503,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X62,X63,X64) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f7605,f3758])).
fof(f9529,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f9502,f947])).
fof(f9502,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X59,X60,X61) | ~p(s(bool,X59))) )),
  inference(resolution,[],[f7605,f3586])).
fof(f9528,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f9500,f947])).
fof(f9500,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X53,X54,X55) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f7605,f1538])).
fof(f9467,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f9440,f946])).
fof(f9440,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X62,X63,X64) | ~p(s(bool,X63))) )),
  inference(resolution,[],[f7604,f3758])).
fof(f9466,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f9439,f946])).
fof(f9439,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X59,X60,X61) | ~p(s(bool,X60))) )),
  inference(resolution,[],[f7604,f3586])).
fof(f9465,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f9437,f946])).
fof(f9437,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X53,X54,X55) | ~p(s(bool,X54))) )),
  inference(resolution,[],[f7604,f1538])).
fof(f9421,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,sK62) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) )),
  inference(subsumption_resolution,[],[f9398,f5468])).
fof(f9398,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,X13) | sP12(X12,sK62) | sP3(X13,X14,X15) | p(s(bool,X13)) | ~sP1(X14,X15,X13) | ~sP0(X15,X14,X13)) )),
  inference(resolution,[],[f7602,f939])).
fof(f9403,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f9376,f943])).
fof(f9376,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X62,X63,X64) | ~p(s(bool,X62))) )),
  inference(resolution,[],[f7602,f3758])).
fof(f9402,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f9375,f943])).
fof(f9375,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X59,X60,X61) | ~p(s(bool,X59))) )),
  inference(resolution,[],[f7602,f3586])).
fof(f9401,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f9373,f943])).
fof(f9373,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X53,X54,X55) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f7602,f1538])).
fof(f9341,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X62,X63,X64)) )),
  inference(subsumption_resolution,[],[f9314,f942])).
fof(f9314,plain,(
  ( ! [X64,X62,X63] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X62,X63,X64) | ~p(s(bool,X63))) )),
  inference(resolution,[],[f7601,f3758])).
fof(f9340,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X59,X60,X61)) )),
  inference(subsumption_resolution,[],[f9313,f942])).
fof(f9313,plain,(
  ( ! [X61,X59,X60] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X59,X60,X61) | ~p(s(bool,X60))) )),
  inference(resolution,[],[f7601,f3586])).
fof(f9339,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X53,X54,X55)) )),
  inference(subsumption_resolution,[],[f9311,f942])).
fof(f9311,plain,(
  ( ! [X54,X55,X53] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X53,X54,X55) | ~p(s(bool,X54))) )),
  inference(resolution,[],[f7601,f1538])).
fof(f9268,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f9234,f1112])).
fof(f9234,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X44,X45,X46) | ~p(s(bool,X44))) )),
  inference(resolution,[],[f7402,f3759])).
fof(f9267,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f9233,f1112])).
fof(f9233,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X41,X42,X43) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f7402,f3587])).
fof(f9266,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f9231,f1112])).
fof(f9231,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X35,X36,X37) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f7402,f1539])).
fof(f9204,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f9169,f1111])).
fof(f9169,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X44,X45,X46) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f7401,f3759])).
fof(f9203,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f9168,f1111])).
fof(f9168,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X41,X42,X43) | ~p(s(bool,X42))) )),
  inference(resolution,[],[f7401,f3587])).
fof(f9201,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f9166,f1111])).
fof(f9166,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP8(X35,X36,X37) | ~p(s(bool,X36))) )),
  inference(resolution,[],[f7401,f1539])).
fof(f9141,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f9106,f1095])).
fof(f9106,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X44,X45,X46) | ~p(s(bool,X44))) )),
  inference(resolution,[],[f7388,f3759])).
fof(f9140,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f9105,f1095])).
fof(f9105,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X41,X42,X43) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f7388,f3587])).
fof(f9138,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f9103,f1095])).
fof(f9103,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP6(X35,X36,X37) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f7388,f1539])).
fof(f9078,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f9044,f1078])).
fof(f9044,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X44,X45,X46) | ~p(s(bool,X44))) )),
  inference(resolution,[],[f7372,f3759])).
fof(f9077,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f9043,f1078])).
fof(f9043,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X41,X42,X43) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f7372,f3587])).
fof(f9076,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f9041,f1078])).
fof(f9041,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X35,X36,X37) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f7372,f1539])).
fof(f9015,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f8980,f1076])).
fof(f8980,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X44,X45,X46) | ~p(s(bool,X46))) )),
  inference(resolution,[],[f7371,f3759])).
fof(f9014,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f8979,f1076])).
fof(f8979,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X41,X42,X43) | ~p(s(bool,X43))) )),
  inference(resolution,[],[f7371,f3587])).
fof(f9012,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f8977,f1076])).
fof(f8977,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP4(X35,X36,X37) | ~p(s(bool,X37))) )),
  inference(resolution,[],[f7371,f1539])).
fof(f8956,plain,(
  $false | $spl84),
  inference(subsumption_resolution,[],[f8955,f1247])).
fof(f8955,plain,(
  p(s(bool,'F_0')) | $spl84),
  inference(forward_demodulation,[],[f8941,f5240])).
fof(f8941,plain,(
  p(s(bool,sK62)) | $spl84),
  inference(resolution,[],[f8796,f5229])).
fof(f8940,plain,(
  $false | ($spl79 | $spl84)),
  inference(resolution,[],[f8796,f8676])).
fof(f8923,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X9,sK62) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) )),
  inference(subsumption_resolution,[],[f8900,f5467])).
fof(f8900,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,X9) | sP12(X9,sK62) | sP3(X10,X8,X11) | p(s(bool,X8)) | ~sP2(X10,X11,X8) | ~sP0(X11,X8,X10)) )),
  inference(resolution,[],[f7343,f938])).
fof(f8906,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f8873,f947])).
fof(f8873,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X44,X45,X46) | ~p(s(bool,X44))) )),
  inference(resolution,[],[f7343,f3759])).
fof(f8905,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f8872,f947])).
fof(f8872,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X41,X42,X43) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f7343,f3587])).
fof(f8904,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f8870,f947])).
fof(f8870,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X35,X36,X37) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f7343,f1539])).
fof(f8843,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f8809,f946])).
fof(f8809,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X44,X45,X46) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f7342,f3759])).
fof(f8842,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f8808,f946])).
fof(f8808,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X41,X42,X43) | ~p(s(bool,X42))) )),
  inference(resolution,[],[f7342,f3587])).
fof(f8840,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f8806,f946])).
fof(f8806,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP1(X35,X36,X37) | ~p(s(bool,X36))) )),
  inference(resolution,[],[f7342,f1539])).
fof(f8794,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X13,sK62) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) )),
  inference(subsumption_resolution,[],[f8771,f5467])).
fof(f8771,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,X13) | sP12(X13,sK62) | sP3(X12,X14,X15) | p(s(bool,X12)) | ~sP1(X14,X15,X12) | ~sP0(X15,X14,X12)) )),
  inference(resolution,[],[f7340,f939])).
fof(f8776,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f8743,f943])).
fof(f8743,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X44,X45,X46) | ~p(s(bool,X44))) )),
  inference(resolution,[],[f7340,f3759])).
fof(f8775,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f8742,f943])).
fof(f8742,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X41,X42,X43) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f7340,f3587])).
fof(f8774,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f8740,f943])).
fof(f8740,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X35,X36,X37) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f7340,f1539])).
fof(f8722,plain,(
  $false | $spl78),
  inference(subsumption_resolution,[],[f8721,f1247])).
fof(f8721,plain,(
  p(s(bool,'F_0')) | $spl78),
  inference(forward_demodulation,[],[f8720,f5240])).
fof(f8720,plain,(
  p(s(bool,sK62)) | $spl78),
  inference(subsumption_resolution,[],[f8696,f5228])).
fof(f8696,plain,(
  p(s(bool,sK62)) | sP12(sK62,sK62) | $spl78),
  inference(resolution,[],[f8677,f5721])).
fof(f8719,plain,(
  $false | $spl78),
  inference(subsumption_resolution,[],[f8718,f1247])).
fof(f8718,plain,(
  p(s(bool,'F_0')) | $spl78),
  inference(forward_demodulation,[],[f8717,f5240])).
fof(f8717,plain,(
  p(s(bool,sK62)) | $spl78),
  inference(subsumption_resolution,[],[f8695,f5228])).
fof(f8695,plain,(
  p(s(bool,sK62)) | sP12(sK62,sK62) | $spl78),
  inference(resolution,[],[f8677,f5720])).
fof(f8716,plain,(
  $false | $spl78),
  inference(subsumption_resolution,[],[f8715,f1247])).
fof(f8715,plain,(
  p(s(bool,'F_0')) | $spl78),
  inference(forward_demodulation,[],[f8714,f5240])).
fof(f8714,plain,(
  p(s(bool,sK62)) | $spl78),
  inference(subsumption_resolution,[],[f8694,f5228])).
fof(f8694,plain,(
  p(s(bool,sK62)) | sP12(sK62,sK62) | $spl78),
  inference(resolution,[],[f8677,f5468])).
fof(f8713,plain,(
  $false | $spl78),
  inference(subsumption_resolution,[],[f8712,f1247])).
fof(f8712,plain,(
  p(s(bool,'F_0')) | $spl78),
  inference(forward_demodulation,[],[f8711,f5240])).
fof(f8711,plain,(
  p(s(bool,sK62)) | $spl78),
  inference(subsumption_resolution,[],[f8693,f5228])).
fof(f8693,plain,(
  p(s(bool,sK62)) | sP12(sK62,sK62) | $spl78),
  inference(resolution,[],[f8677,f5467])).
fof(f8710,plain,(
  $false | $spl78),
  inference(subsumption_resolution,[],[f8709,f1247])).
fof(f8709,plain,(
  p(s(bool,'F_0')) | $spl78),
  inference(forward_demodulation,[],[f8697,f5240])).
fof(f8697,plain,(
  p(s(bool,sK62)) | $spl78),
  inference(duplicate_literal_removal,[],[f8692])).
fof(f8692,plain,(
  p(s(bool,sK62)) | p(s(bool,sK62)) | $spl78),
  inference(resolution,[],[f8677,f1174])).
fof(f8706,plain,(
  $false | $spl78),
  inference(subsumption_resolution,[],[f8705,f1247])).
fof(f8705,plain,(
  p(s(bool,'F_0')) | $spl78),
  inference(forward_demodulation,[],[f8700,f5240])).
fof(f8700,plain,(
  p(s(bool,sK62)) | $spl78),
  inference(duplicate_literal_removal,[],[f8689])).
fof(f8689,plain,(
  p(s(bool,sK62)) | p(s(bool,sK62)) | $spl78),
  inference(resolution,[],[f8677,f1165])).
fof(f8704,plain,(
  $false | $spl78),
  inference(subsumption_resolution,[],[f8703,f1247])).
fof(f8703,plain,(
  p(s(bool,'F_0')) | $spl78),
  inference(forward_demodulation,[],[f8688,f5240])).
fof(f8688,plain,(
  p(s(bool,sK62)) | $spl78),
  inference(resolution,[],[f8677,f5228])).
fof(f8702,plain,(
  $false | $spl78),
  inference(subsumption_resolution,[],[f8701,f1247])).
fof(f8701,plain,(
  p(s(bool,'F_0')) | $spl78),
  inference(forward_demodulation,[],[f8687,f5240])).
fof(f8687,plain,(
  p(s(bool,sK62)) | $spl78),
  inference(resolution,[],[f8677,f5229])).
fof(f8677,plain,(
  sP12(sK62,sK62) | $spl78),
  inference(cnf_transformation,[],[f8677_D])).
fof(f8677_D,plain,(
  sP12(sK62,sK62) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f8678,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | p(s(bool,X21)) | sP2(X22,X20,X23)) ) | $spl80),
  inference(cnf_transformation,[],[f8678_D])).
fof(f8678_D,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | p(s(bool,X21)) | sP2(X22,X20,X23)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f8652,plain,(
  ( ! [X30,X28,X31,X29] : (sP12(X28,X29) | sP2(X30,X28,X31) | p(s(bool,X29)) | sP12(sK62,sK62)) )),
  inference(resolution,[],[f7339,f5720])).
fof(f8650,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,X21) | sP2(X22,X20,X23) | p(s(bool,X21)) | sP12(sK62,sK62)) )),
  inference(resolution,[],[f7339,f5467])).
fof(f8669,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X44,X45,X46)) )),
  inference(subsumption_resolution,[],[f8637,f942])).
fof(f8637,plain,(
  ( ! [X45,X46,X44] : (sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X44,X45,X46) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f7339,f3759])).
fof(f8668,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X41,X42,X43)) )),
  inference(subsumption_resolution,[],[f8636,f942])).
fof(f8636,plain,(
  ( ! [X43,X41,X42] : (sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X41,X42,X43) | ~p(s(bool,X42))) )),
  inference(resolution,[],[f7339,f3587])).
fof(f8666,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X35,X36,X37)) )),
  inference(subsumption_resolution,[],[f8634,f942])).
fof(f8634,plain,(
  ( ! [X37,X35,X36] : (sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | sP2(X35,X36,X37) | ~p(s(bool,X36))) )),
  inference(resolution,[],[f7339,f1539])).
fof(f8622,plain,(
  ( ! [X21] : (p(s(bool,X21)) | sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5865,f3767])).
fof(f8621,plain,(
  ( ! [X20] : (p(s(bool,X20)) | sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5865,f3595])).
fof(f8619,plain,(
  ( ! [X18] : (p(s(bool,X18)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5865,f1547])).
fof(f8598,plain,(
  ( ! [X16] : (p(s(bool,X16)) | sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5864,f3766])).
fof(f8597,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5864,f3594])).
fof(f8595,plain,(
  ( ! [X13] : (p(s(bool,X13)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f5864,f1546])).
fof(f5615,plain,(
  ( ! [X344,X345] : (p(s(bool,X344)) | ~sP13(X344,X345) | p(s(bool,X345))) ) | $spl32),
  inference(cnf_transformation,[],[f5615_D])).
fof(f5615_D,plain,(
  ( ! [X344,X345] : (p(s(bool,X344)) | ~sP13(X344,X345) | p(s(bool,X345))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f8521,plain,(
  ( ! [X468] : (sP12(X468,'F_0') | sP12(sK62,X468)) ) | $spl74),
  inference(cnf_transformation,[],[f8521_D])).
fof(f8521_D,plain,(
  ( ! [X468] : (sP12(X468,'F_0') | sP12(sK62,X468)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f8400,plain,(
  ( ! [X476,X474,X475] : (sP12(X474,'F_0') | sP12(sK62,X474) | p(s(bool,X475)) | p(s(bool,X476)) | ~sP13(X475,X476)) )),
  inference(resolution,[],[f5723,f1181])).
fof(f8398,plain,(
  ( ! [X469,X468,X470] : (sP12(X468,'F_0') | sP12(sK62,X468) | p(s(bool,X469)) | p(s(bool,X470)) | ~sP13(X469,X470)) )),
  inference(resolution,[],[f5723,f1185])).
fof(f8515,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP12(sK62,X0)) ) | $spl72),
  inference(cnf_transformation,[],[f8515_D])).
fof(f8515_D,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP12(sK62,X0)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f8394,plain,(
  ( ! [X455,X457,X454,X456] : (sP12(X454,X455) | sP12(sK62,X454) | p(s(bool,X456)) | ~p(s(bool,X457)) | s(bool,X455) != s(bool,X457)) )),
  inference(resolution,[],[f5723,f1252])).
fof(f8281,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP12(sK62,X35) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f5723,f926])).
fof(f8279,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP12(sK62,X25) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) )),
  inference(resolution,[],[f5723,f925])).
fof(f8277,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP12(sK62,X15) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f5723,f924])).
fof(f8276,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP12(sK62,X10) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f5723,f922])).
fof(f8274,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP12(sK62,X0) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f5723,f920])).
fof(f8270,plain,(
  ( ! [X468] : (sP12('F_0',X468) | sP12(sK62,X468)) ) | $spl70),
  inference(cnf_transformation,[],[f8270_D])).
fof(f8270_D,plain,(
  ( ! [X468] : (sP12('F_0',X468) | sP12(sK62,X468)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f8148,plain,(
  ( ! [X476,X474,X475] : (sP12('F_0',X474) | sP12(sK62,X474) | p(s(bool,X475)) | p(s(bool,X476)) | ~sP13(X475,X476)) )),
  inference(resolution,[],[f5722,f1181])).
fof(f8146,plain,(
  ( ! [X469,X468,X470] : (sP12('F_0',X468) | sP12(sK62,X468) | p(s(bool,X469)) | p(s(bool,X470)) | ~sP13(X469,X470)) )),
  inference(resolution,[],[f5722,f1185])).
fof(f8264,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP12(sK62,X1)) ) | $spl68),
  inference(cnf_transformation,[],[f8264_D])).
fof(f8264_D,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP12(sK62,X1)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f8142,plain,(
  ( ! [X455,X457,X454,X456] : (sP12(X454,X455) | sP12(sK62,X455) | p(s(bool,X456)) | ~p(s(bool,X457)) | s(bool,X454) != s(bool,X457)) )),
  inference(resolution,[],[f5722,f1252])).
fof(f8029,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP12(sK62,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f5722,f926])).
fof(f8027,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP12(sK62,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) )),
  inference(resolution,[],[f5722,f925])).
fof(f8025,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP12(sK62,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) )),
  inference(resolution,[],[f5722,f924])).
fof(f8024,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP12(sK62,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f5722,f922])).
fof(f8022,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP12(sK62,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) )),
  inference(resolution,[],[f5722,f920])).
fof(f8007,plain,(
  ( ! [X18] : (p(s(bool,X18)) | sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f8004])).
fof(f8004,plain,(
  ( ! [X18] : (p(s(bool,X18)) | sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X18))) )),
  inference(resolution,[],[f5721,f3760])).
fof(f8008,plain,(
  ( ! [X17] : (p(s(bool,X17)) | sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f8003])).
fof(f8003,plain,(
  ( ! [X17] : (p(s(bool,X17)) | sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X17))) )),
  inference(resolution,[],[f5721,f3588])).
fof(f7977,plain,(
  sP12(sK62,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | $spl62),
  inference(cnf_transformation,[],[f7977_D])).
fof(f7977_D,plain,(
  sP12(sK62,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f8009,plain,(
  ( ! [X16] : (p(s(bool,X16)) | sP12(sK62,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f8001])).
fof(f8001,plain,(
  ( ! [X16] : (p(s(bool,X16)) | sP12(sK62,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X16))) ) | $spl0),
  inference(resolution,[],[f5721,f3247])).
fof(f8010,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f8000])).
fof(f8000,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X15))) )),
  inference(resolution,[],[f5721,f1540])).
fof(f7973,plain,(
  sP12(sK62,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | $spl58),
  inference(cnf_transformation,[],[f7973_D])).
fof(f7973_D,plain,(
  sP12(sK62,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f8011,plain,(
  ( ! [X14] : (p(s(bool,X14)) | sP12(sK62,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) )),
  inference(duplicate_literal_removal,[],[f7998])).
fof(f7998,plain,(
  ( ! [X14] : (p(s(bool,X14)) | sP12(sK62,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X14))) )),
  inference(resolution,[],[f5721,f1369])).
fof(f1369,plain,(
  ( ! [X221] : (sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),X221) | p(s(bool,X221))) )),
  inference(resolution,[],[f916,f1167])).
fof(f7965,plain,(
  ( ! [X12] : (p(s(bool,X12)) | sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f7954])).
fof(f7954,plain,(
  ( ! [X12] : (p(s(bool,X12)) | sP12(sK62,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X12))) )),
  inference(resolution,[],[f5720,f3761])).
fof(f7966,plain,(
  ( ! [X11] : (p(s(bool,X11)) | sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f7953])).
fof(f7953,plain,(
  ( ! [X11] : (p(s(bool,X11)) | sP12(sK62,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X11))) )),
  inference(resolution,[],[f5720,f3589])).
fof(f7967,plain,(
  ( ! [X10] : (p(s(bool,X10)) | sP12(sK62,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f7952])).
fof(f7952,plain,(
  ( ! [X10] : (p(s(bool,X10)) | sP12(sK62,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X10))) ) | $spl0),
  inference(resolution,[],[f5720,f3248])).
fof(f7968,plain,(
  ( ! [X9] : (p(s(bool,X9)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(duplicate_literal_removal,[],[f7951])).
fof(f7951,plain,(
  ( ! [X9] : (p(s(bool,X9)) | sP12(sK62,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | p(s(bool,X9))) )),
  inference(resolution,[],[f5720,f1541])).
fof(f7969,plain,(
  ( ! [X8] : (p(s(bool,X8)) | sP12(sK62,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) )),
  inference(duplicate_literal_removal,[],[f7950])).
fof(f7950,plain,(
  ( ! [X8] : (p(s(bool,X8)) | sP12(sK62,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X8))) )),
  inference(resolution,[],[f5720,f1370])).
fof(f1370,plain,(
  ( ! [X222] : (sP12(X222,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X222))) )),
  inference(resolution,[],[f916,f1170])).
fof(f7943,plain,(
  ( ! [X21] : (p(s(bool,X21)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5612,f3767])).
fof(f7942,plain,(
  ( ! [X20] : (p(s(bool,X20)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5612,f3595])).
fof(f7940,plain,(
  ( ! [X18] : (p(s(bool,X18)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5612,f1547])).
fof(f7919,plain,(
  ( ! [X16] : (p(s(bool,X16)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5611,f3766])).
fof(f7918,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5611,f3594])).
fof(f7916,plain,(
  ( ! [X13] : (p(s(bool,X13)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(resolution,[],[f5611,f1546])).
fof(f7837,plain,(
  ( ! [X468] : (sP12(X468,'F_0') | sP12(X468,sK62)) ) | $spl56),
  inference(cnf_transformation,[],[f7837_D])).
fof(f7837_D,plain,(
  ( ! [X468] : (sP12(X468,'F_0') | sP12(X468,sK62)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f7715,plain,(
  ( ! [X476,X474,X475] : (sP12(X474,'F_0') | sP12(X474,sK62) | p(s(bool,X475)) | p(s(bool,X476)) | ~sP13(X475,X476)) )),
  inference(resolution,[],[f5470,f1181])).
fof(f7713,plain,(
  ( ! [X469,X468,X470] : (sP12(X468,'F_0') | sP12(X468,sK62) | p(s(bool,X469)) | p(s(bool,X470)) | ~sP13(X469,X470)) )),
  inference(resolution,[],[f5470,f1185])).
fof(f7831,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP12(X0,sK62)) ) | $spl54),
  inference(cnf_transformation,[],[f7831_D])).
fof(f7831_D,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X3)) | sP12(X0,sK62)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f7709,plain,(
  ( ! [X455,X457,X454,X456] : (sP12(X454,X455) | sP12(X454,sK62) | p(s(bool,X456)) | ~p(s(bool,X457)) | s(bool,X455) != s(bool,X457)) )),
  inference(resolution,[],[f5470,f1252])).
fof(f7596,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X35,sK62) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f5470,f926])).
fof(f7594,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP12(X25,sK62) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28)) )),
  inference(resolution,[],[f5470,f925])).
fof(f7592,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP12(X15,sK62) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f5470,f924])).
fof(f7591,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP12(X10,sK62) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f5470,f922])).
fof(f7589,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP12(X0,sK62) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f5470,f920])).
fof(f7574,plain,(
  ( ! [X468] : (sP12('F_0',X468) | sP12(X468,sK62)) ) | $spl52),
  inference(cnf_transformation,[],[f7574_D])).
fof(f7574_D,plain,(
  ( ! [X468] : (sP12('F_0',X468) | sP12(X468,sK62)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f7453,plain,(
  ( ! [X476,X474,X475] : (sP12('F_0',X474) | sP12(X474,sK62) | p(s(bool,X475)) | p(s(bool,X476)) | ~sP13(X475,X476)) )),
  inference(resolution,[],[f5469,f1181])).
fof(f7451,plain,(
  ( ! [X469,X468,X470] : (sP12('F_0',X468) | sP12(X468,sK62) | p(s(bool,X469)) | p(s(bool,X470)) | ~sP13(X469,X470)) )),
  inference(resolution,[],[f5469,f1185])).
fof(f7568,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP12(X1,sK62)) ) | $spl50),
  inference(cnf_transformation,[],[f7568_D])).
fof(f7568_D,plain,(
  ( ! [X0,X3,X1] : (sP12(X0,X1) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | sP12(X1,sK62)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f7447,plain,(
  ( ! [X455,X457,X454,X456] : (sP12(X454,X455) | sP12(X455,sK62) | p(s(bool,X456)) | ~p(s(bool,X457)) | s(bool,X454) != s(bool,X457)) )),
  inference(resolution,[],[f5469,f1252])).
fof(f7334,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP12(X35,X36) | sP12(X36,sK62) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f5469,f926])).
fof(f7332,plain,(
  ( ! [X28,X26,X29,X27,X25] : (sP12(X25,X26) | sP12(X26,sK62) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) )),
  inference(resolution,[],[f5469,f925])).
fof(f7330,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP12(X15,X16) | sP12(X16,sK62) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) )),
  inference(resolution,[],[f5469,f924])).
fof(f7329,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP12(X10,X11) | sP12(X11,sK62) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f5469,f922])).
fof(f7327,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1) | sP12(X1,sK62) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) )),
  inference(resolution,[],[f5469,f920])).
fof(f7314,plain,(
  ( ! [X17] : (p(s(bool,X17)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f7311])).
fof(f7311,plain,(
  ( ! [X17] : (p(s(bool,X17)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X17))) )),
  inference(resolution,[],[f5468,f3760])).
fof(f7315,plain,(
  ( ! [X16] : (p(s(bool,X16)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f7310])).
fof(f7310,plain,(
  ( ! [X16] : (p(s(bool,X16)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X16))) )),
  inference(resolution,[],[f5468,f3588])).
fof(f7287,plain,(
  sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62) | $spl44),
  inference(cnf_transformation,[],[f7287_D])).
fof(f7287_D,plain,(
  sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f7316,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f7309])).
fof(f7309,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62) | p(s(bool,X15))) ) | $spl0),
  inference(resolution,[],[f5468,f3247])).
fof(f7317,plain,(
  ( ! [X14] : (p(s(bool,X14)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f7308])).
fof(f7308,plain,(
  ( ! [X14] : (p(s(bool,X14)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X14))) )),
  inference(resolution,[],[f5468,f1540])).
fof(f7283,plain,(
  sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62) | $spl40),
  inference(cnf_transformation,[],[f7283_D])).
fof(f7283_D,plain,(
  sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f7318,plain,(
  ( ! [X13] : (p(s(bool,X13)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62)) )),
  inference(duplicate_literal_removal,[],[f7307])).
fof(f7307,plain,(
  ( ! [X13] : (p(s(bool,X13)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62) | p(s(bool,X13))) )),
  inference(resolution,[],[f5468,f1369])).
fof(f7275,plain,(
  ( ! [X12] : (p(s(bool,X12)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f7267])).
fof(f7267,plain,(
  ( ! [X12] : (p(s(bool,X12)) | sP12('const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X12))) )),
  inference(resolution,[],[f5467,f3761])).
fof(f7276,plain,(
  ( ! [X11] : (p(s(bool,X11)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f7266])).
fof(f7266,plain,(
  ( ! [X11] : (p(s(bool,X11)) | sP12('const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X11))) )),
  inference(resolution,[],[f5467,f3589])).
fof(f7277,plain,(
  ( ! [X10] : (p(s(bool,X10)) | sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f7265])).
fof(f7265,plain,(
  ( ! [X10] : (p(s(bool,X10)) | sP12('const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62) | p(s(bool,X10))) ) | $spl0),
  inference(resolution,[],[f5467,f3248])).
fof(f7278,plain,(
  ( ! [X9] : (p(s(bool,X9)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62)) )),
  inference(duplicate_literal_removal,[],[f7264])).
fof(f7264,plain,(
  ( ! [X9] : (p(s(bool,X9)) | sP12('const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)),sK62) | p(s(bool,X9))) )),
  inference(resolution,[],[f5467,f1541])).
fof(f7279,plain,(
  ( ! [X8] : (p(s(bool,X8)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62)) )),
  inference(duplicate_literal_removal,[],[f7263])).
fof(f7263,plain,(
  ( ! [X8] : (p(s(bool,X8)) | sP12('const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)),sK62) | p(s(bool,X8))) )),
  inference(resolution,[],[f5467,f1370])).
fof(f5748,plain,(
  ( ! [X349,X348] : (sP12(sK62,'F_0') | p(s(bool,X348)) | p(s(bool,X349)) | ~sP13(X348,X349)) )),
  inference(resolution,[],[f5231,f1181])).
fof(f5746,plain,(
  ( ! [X344,X345] : (sP12(sK62,'F_0') | p(s(bool,X344)) | p(s(bool,X345)) | ~sP13(X344,X345)) )),
  inference(resolution,[],[f5231,f1185])).
fof(f5862,plain,(
  ( ! [X2,X0] : (sP12(sK62,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl36),
  inference(cnf_transformation,[],[f5862_D])).
fof(f5862_D,plain,(
  ( ! [X2,X0] : (sP12(sK62,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5741,plain,(
  ( ! [X335,X334,X336] : (sP12(sK62,X334) | p(s(bool,X335)) | ~p(s(bool,X336)) | s(bool,X334) != s(bool,X336)) )),
  inference(resolution,[],[f5231,f1252])).
fof(f5628,plain,(
  ( ! [X30,X28,X31,X29] : (sP12(sK62,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f5231,f926])).
fof(f5626,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(sK62,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f5231,f925])).
fof(f5624,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(sK62,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f5231,f924])).
fof(f5623,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(sK62,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f5231,f922])).
fof(f5621,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(sK62,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f5231,f920])).
fof(f5495,plain,(
  ( ! [X349,X348] : (sP12('F_0',sK62) | p(s(bool,X348)) | p(s(bool,X349)) | ~sP13(X348,X349)) )),
  inference(resolution,[],[f5230,f1181])).
fof(f5493,plain,(
  ( ! [X344,X345] : (sP12('F_0',sK62) | p(s(bool,X344)) | p(s(bool,X345)) | ~sP13(X344,X345)) )),
  inference(resolution,[],[f5230,f1185])).
fof(f5609,plain,(
  ( ! [X2,X0] : (sP12(X0,sK62) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl30),
  inference(cnf_transformation,[],[f5609_D])).
fof(f5609_D,plain,(
  ( ! [X2,X0] : (sP12(X0,sK62) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f5488,plain,(
  ( ! [X335,X334,X336] : (sP12(X334,sK62) | p(s(bool,X335)) | ~p(s(bool,X336)) | s(bool,X334) != s(bool,X336)) )),
  inference(resolution,[],[f5230,f1252])).
fof(f5375,plain,(
  ( ! [X30,X28,X31,X29] : (sP12(X28,sK62) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f5230,f926])).
fof(f5373,plain,(
  ( ! [X23,X21,X22,X20] : (sP12(X20,sK62) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f5230,f925])).
fof(f5371,plain,(
  ( ! [X14,X12,X15,X13] : (sP12(X12,sK62) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f5230,f924])).
fof(f5370,plain,(
  ( ! [X10,X8,X11,X9] : (sP12(X8,sK62) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f5230,f922])).
fof(f5368,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(X0,sK62) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f5230,f920])).
fof(f5279,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,sK62) != s(bool,X1)) ) | $spl28),
  inference(cnf_transformation,[],[f5279_D])).
fof(f5279_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,sK62) != s(bool,X1)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5249,plain,(
  ( ! [X215,X214] : (p(s(bool,X214)) | ~p(s(bool,X215)) | s(bool,sK62) != s(bool,X215)) )),
  inference(resolution,[],[f5116,f1252])).
fof(f5136,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK62) != s(bool,X22) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f5116,f926])).
fof(f5134,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,sK62) != s(bool,X16)) )),
  inference(resolution,[],[f5116,f925])).
fof(f5132,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK62) != s(bool,X10)) )),
  inference(resolution,[],[f5116,f924])).
fof(f5131,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK62) != s(bool,X7) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f5116,f922])).
fof(f5129,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK62) != s(bool,X1)) )),
  inference(resolution,[],[f5116,f920])).
fof(f5110,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f5110_D])).
fof(f5110_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4989,plain,(
  ( ! [X335,X334,X336] : (s(bool,X334) = s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X335)) | ~p(s(bool,X336)) | s(bool,X334) != s(bool,X336)) )),
  inference(resolution,[],[f1288,f1252])).
fof(f4876,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f1288,f926])).
fof(f4874,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f1288,f925])).
fof(f4872,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1288,f924])).
fof(f4871,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f1288,f922])).
fof(f4869,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1288,f920])).
fof(f3808,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl24),
  inference(cnf_transformation,[],[f3808_D])).
fof(f3808_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3775,plain,(
  ( ! [X225,X224] : (p(s(bool,X224)) | p(s(bool,X225)) | s(bool,X225) != s(bool,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1433,f1252])).
fof(f3662,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f1433,f926])).
fof(f3659,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1433,f925])).
fof(f3658,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1433,f924])).
fof(f3657,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f1433,f922])).
fof(f3655,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.operator.LEFT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1433,f920])).
fof(f3636,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl22),
  inference(cnf_transformation,[],[f3636_D])).
fof(f3636_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3603,plain,(
  ( ! [X225,X224] : (p(s(bool,X224)) | p(s(bool,X225)) | s(bool,X225) != s(bool,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1432,f1252])).
fof(f3490,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f1432,f926])).
fof(f3487,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1432,f925])).
fof(f3486,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1432,f924])).
fof(f3485,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f1432,f922])).
fof(f3483,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.operator.RIGHT_ID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f1432,f920])).
fof(f3296,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl20),
  inference(cnf_transformation,[],[f3296_D])).
fof(f3296_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3262,plain,(
  ( ! [X227,X228] : (p(s(bool,X227)) | p(s(bool,X228)) | s(bool,X228) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl0),
  inference(resolution,[],[f1417,f1252])).
fof(f3149,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | s(bool,X25) != s(bool,X26)) ) | $spl0),
  inference(resolution,[],[f1417,f926])).
fof(f3146,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl0),
  inference(resolution,[],[f1417,f925])).
fof(f3145,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl0),
  inference(resolution,[],[f1417,f924])).
fof(f3144,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | s(bool,X10) != s(bool,X11)) ) | $spl0),
  inference(resolution,[],[f1417,f922])).
fof(f3142,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl0),
  inference(resolution,[],[f1417,f920])).
fof(f3139,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f1431,f1418])).
fof(f1431,plain,(
  p(s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))))),
  inference(resolution,[],[f917,f1046])).
fof(f1046,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | ~p(s(bool,'const.operator.MONOID_2'(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f829])).
fof(f2970,plain,(
  $false | $spl8),
  inference(resolution,[],[f2823,f916])).
fof(f2823,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(equality_resolution,[],[f1422])).
fof(f2972,plain,(
  $false | $spl8),
  inference(resolution,[],[f2823,f1051])).
fof(f2973,plain,(
  $false | $spl8),
  inference(resolution,[],[f2823,f917])).
fof(f2692,plain,(
  $false | $spl18),
  inference(equality_resolution,[],[f2297])).
fof(f2604,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f2441,f2287])).
fof(f2287,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl6),
  inference(equality_resolution,[],[f1421])).
fof(f2441,plain,(
  ( ! [X366] : (~p(s(bool,X366))) ) | $spl6),
  inference(resolution,[],[f2287,f1243])).
fof(f2603,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f2435,f2287])).
fof(f2435,plain,(
  ( ! [X360] : (~p(s(bool,X360))) ) | $spl6),
  inference(resolution,[],[f2287,f1209])).
fof(f2443,plain,(
  $false | $spl6),
  inference(resolution,[],[f2287,f1247])).
fof(f2495,plain,(
  $false | ($spl1 | $spl6)),
  inference(resolution,[],[f2287,f1418])).
fof(f2297,plain,(
  ( ! [X53] : (s(bool,'F_0') != s(bool,X53)) ) | $spl18),
  inference(cnf_transformation,[],[f2297_D])).
fof(f2297_D,plain,(
  ( ! [X53] : (s(bool,'F_0') != s(bool,X53)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2288,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') != s(bool,X59) | p(s(bool,X58))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2190])).
fof(f2190,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') != s(bool,X59) | p(s(bool,X58)) | p(s(bool,X58))) ) | $spl6),
  inference(superposition,[],[f1421,f1235])).
fof(f2289,plain,(
  ( ! [X57,X56] : (s(bool,'F_0') != s(bool,X57) | p(s(bool,X56))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2189])).
fof(f2189,plain,(
  ( ! [X57,X56] : (s(bool,'F_0') != s(bool,X57) | p(s(bool,X56)) | p(s(bool,X56))) ) | $spl6),
  inference(superposition,[],[f1421,f1233])).
fof(f2290,plain,(
  ( ! [X52,X53] : (s(bool,'F_0') != s(bool,X53) | p(s(bool,X52))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2187])).
fof(f2187,plain,(
  ( ! [X52,X53] : (s(bool,'F_0') != s(bool,X53) | p(s(bool,X52)) | p(s(bool,X52))) ) | $spl6),
  inference(superposition,[],[f1421,f1211])).
fof(f2162,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl16),
  inference(cnf_transformation,[],[f2162_D])).
fof(f2162_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2130,plain,(
  ( ! [X216,X215] : (p(s(bool,X215)) | ~p(s(bool,X216)) | s(bool,X216) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl1),
  inference(resolution,[],[f1418,f1252])).
fof(f2017,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f1418,f926])).
fof(f2015,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl1),
  inference(resolution,[],[f1418,f925])).
fof(f2013,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl1),
  inference(resolution,[],[f1418,f924])).
fof(f2012,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl1),
  inference(resolution,[],[f1418,f922])).
fof(f2010,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl1),
  inference(resolution,[],[f1418,f920])).
fof(f1890,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1728,f1424])).
fof(f1728,plain,(
  ( ! [X366] : (~p(s(bool,X366))) ) | $spl12),
  inference(resolution,[],[f1424,f1243])).
fof(f1889,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1722,f1424])).
fof(f1722,plain,(
  ( ! [X360] : (~p(s(bool,X360))) ) | $spl12),
  inference(resolution,[],[f1424,f1209])).
fof(f1730,plain,(
  $false | $spl12),
  inference(resolution,[],[f1424,f1247])).
fof(f1589,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) | $spl14),
  inference(cnf_transformation,[],[f1589_D])).
fof(f1589_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1555,plain,(
  ( ! [X227,X228] : (p(s(bool,X227)) | p(s(bool,X228)) | s(bool,X228) != s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f917,f1252])).
fof(f1442,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | s(bool,X25) != s(bool,X26)) )),
  inference(resolution,[],[f917,f926])).
fof(f1439,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f917,f925])).
fof(f1438,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f917,f924])).
fof(f1437,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16))) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f917,f922])).
fof(f1435,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.operator.MONOID_2'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,sK16)))) )),
  inference(resolution,[],[f917,f920])).
fof(f1420,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) | $spl4),
  inference(cnf_transformation,[],[f1420_D])).
fof(f1420_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1384,plain,(
  ( ! [X231,X232] : (p(s(bool,X231)) | p(s(bool,X232)) | s(bool,X232) != s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) )),
  inference(resolution,[],[f916,f1252])).
fof(f1271,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | s(bool,X29) != s(bool,X30)) )),
  inference(resolution,[],[f916,f926])).
fof(f1268,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) )),
  inference(resolution,[],[f916,f925])).
fof(f1267,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) )),
  inference(resolution,[],[f916,f924])).
fof(f1266,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15))) | s(bool,X14) != s(bool,X15)) )),
  inference(resolution,[],[f916,f922])).
fof(f1264,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.operator.COMM_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) )),
  inference(resolution,[],[f916,f920])).
fof(f1418,plain,(
  ~p(s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) | $spl1),
  inference(cnf_transformation,[],[f1418_D])).
fof(f1418_D,plain,(
  ~p(s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1419,plain,(
  ( ! [X4,X3] : (s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X3),s('type.list.list'(sK14),X4))) = s(sK14,'const.list.FOLDL_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X3),s('type.list.list'(sK14),X4)))) ) | $spl2),
  inference(cnf_transformation,[],[f1419_D])).
fof(f1419_D,plain,(
  ( ! [X4,X3] : (s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X3),s('type.list.list'(sK14),X4))) = s(sK14,'const.list.FOLDL_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X3),s('type.list.list'(sK14),X4)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1262,plain,(
  ( ! [X4,X3] : (s(sK14,'const.list.FOLDR_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X3),s('type.list.list'(sK14),X4))) = s(sK14,'const.list.FOLDL_3'(s(fun(sK14,fun(sK14,sK14)),sK15),s(sK14,X3),s('type.list.list'(sK14),X4))) | ~p(s(bool,'const.operator.ASSOC_1'(s(fun(sK14,fun(sK14,sK14)),sK15))))) )),
  inference(resolution,[],[f916,f1033])).
fof(f1033,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.list.FOLDR_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X2))) = s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X2))) | ~p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) | ~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1))))) )),
  inference(cnf_transformation,[],[f684])).
fof(f684,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | ~p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) | s(X0,'const.list.FOLDR_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X2))) = s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X2))))),
  inference(flattening,[],[f683])).
fof(f683,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | ~p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1))))) | s(X0,'const.list.FOLDR_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X2))) = s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X2))))),
  inference(ennf_transformation,[],[f583])).
fof(f583,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) & p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1))))) => s(X0,'const.list.FOLDR_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X2))) = s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X2))))),
  inference(flattening,[],[f582])).
fof(f582,plain,(
  ! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) & p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1))))) => s(X0,'const.list.FOLDR_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X2))) = s(X0,'const.list.FOLDL_3'(s(fun(X0,fun(X0,X0)),X1),s(X0,X3),s('type.list.list'(X0),X2))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7,X10,X9] : ((p(s(bool,'const.operator.ASSOC_1'(s(fun(X6,fun(X6,X6)),X7)))) & p(s(bool,'const.operator.COMM_1'(s(fun(X6,fun(X6,X6)),X7))))) => s(X6,'const.list.FOLDR_3'(s(fun(X6,fun(X6,X6)),X7),s(X6,X9),s('type.list.list'(X6),X10))) = s(X6,'const.list.FOLDL_3'(s(fun(X6,fun(X6,X6)),X7),s(X6,X9),s('type.list.list'(X6),X10))))),
  file('thm.rich_list.COMM_MONOID_FOLDR','thm.list.FOLDL_EQ_FOLDR')).
% SZS output end Proof for thm.rich_list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 758
% Passive clauses: 12863
% Generated clauses: 32626
% Final active clauses: 745
% Final passive clauses: 11620
% Input formulas: 515
% Initial clauses: 336
% 
% Selected by SInE selection: 89
% SInE iterations: 1
% 
% Duplicate literals: 1226
% Trivial inequalities: 34
% Fw subsumption resolutions: 2707
% Fw demodulations: 982
% Bw demodulations: 2
% 
% Simple tautologies: 1641
% Equational tautologies: 7
% Forward subsumptions: 11905
% 
% Binary resolution: 21037
% Factoring: 72
% Forward superposition: 5576
% Backward superposition: 790
% Self superposition: 6
% Equality resolution: 5
% 
% Split clauses: 1538
% Split components: 3469
% SAT solver clauses: 1621
% SAT solver unit clauses: 30
% SAT solver binary clauses: 550
% 
% Sat splits: 1538
% Sat splitting refutations: 35
% 
% Memory used [KB]: 17398
% Time elapsed: 0.791 s
% ------------------------------
% ------------------------------
% Success in time 12.58 s
