Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.topology
% (106681)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 102
% Passive clauses: 2105
% Generated clauses: 17737
% Final active clauses: 92
% Final passive clauses: 1872
% Input formulas: 12
% Initial clauses: 27
% 
% Duplicate literals: 1248
% Trivial inequalities: 205
% Fw subsumption resolutions: 2450
% Fw demodulations: 943
% Bw demodulations: 6
% 
% Simple tautologies: 355
% Equational tautologies: 105
% Forward subsumptions: 7759
% Fw demodulations to eq. taut.: 3
% 
% Binary resolution: 103
% Forward superposition: 7196
% Backward superposition: 4698
% Self superposition: 381
% Equality factoring: 295
% Equality resolution: 13
% 
% Split clauses: 1003
% Split components: 2021
% SAT solver clauses: 2545
% SAT solver unit clauses: 335
% SAT solver binary clauses: 738
% 
% Sat splits: 1003
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2532
% TWLsolver calls for satisfiability: 95
% 
% Memory used [KB]: 7931
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.topology
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.topology
% SZS output start Proof for thm.topology
fof(f3622,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3621,f49,f80,f68,f98,f2412,f2422,f106,f113,f51,f82,f52,f83,f710,f711,f53,f84,f57,f85,f3571,f3573,f114,f116,f47,f78,f441,f448,f446,f369,f370,f101,f103,f48,f76,f70,f99,f104,f112,f477,f890,f899,f61,f93])).
fof(f93,plain,(
  ( ! [X0,X5,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X5)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X5),s(fun(fun(X0,bool),bool),X1)))) | ~sP0(X0,X1)) ) | $spl28),
  inference(cnf_transformation,[],[f93_D])).
fof(f93_D,plain,(
  ( ! [X0,X5,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X5)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X5),s(fun(fun(X0,bool),bool),X1)))) | ~sP0(X0,X1)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f61,plain,(
  ( ! [X0,X5,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X5)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X5),s(fun(fun(X0,bool),bool),X1)))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X6,X7] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X6)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X7)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X6),s(fun(X0,bool),X7))))))) & ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X5),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X5))))))))) & (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) | (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),sK4(X1,X0))))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),sK5(X1,X0))))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))))))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),sK6(X1,X0)),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),sK6(X1,X0)))))))) | sP0(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5,sK6])],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X6,X7] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X6)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X7)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X6),s(fun(X0,bool),X7))))))) & ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X5),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X5))))))))) & (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) | ? [X3,X4] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) | ? [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))))))) | sP0(X0,X1)))),
  inference(rectify,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))))))))) & (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) | ? [X3,X4] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) | ? [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))))))) | sP0(X0,X1)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))))))))) & ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) | ? [X3,X4] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) | ? [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f899,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) | p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))) | ~sP0(X0,'const.topology.open_1'(s('type.topology.topology'(X0),X1)))) ) | $spl156),
  inference(cnf_transformation,[],[f899_D])).
fof(f899_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) | p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))) | ~sP0(X0,'const.topology.open_1'(s('type.topology.topology'(X0),X1)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f890,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) | ~sP0(X0,'const.topology.open_1'(s('type.topology.topology'(X0),X1)))) ) | ($spl28 | $spl40)),
  inference(superposition,[],[f93,f99])).
fof(f477,plain,(
  $false | ($spl4 | $spl7 | $spl78)),
  inference(subsumption_resolution,[],[f454,f80])).
fof(f454,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl78)),
  inference(backward_demodulation,[],[f446,f82])).
fof(f112,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl46),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f104,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f83,f83])).
fof(f99,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),X2)))) ) | $spl40),
  inference(cnf_transformation,[],[f99_D])).
fof(f99_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),X2)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f70,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X14,X15] : s(bool,'const.topology.open_2'(s('type.topology.topology'(X6),X14),s(fun(X6,bool),X15))) = s(bool,happ(s(fun(fun(X6,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X6),X14))),s(fun(X6,bool),X15)))),
  file('thm.topology.TOPOLOGY|split|3','ARITY_1_2_const.topology.open')).
fof(f76,plain,(
  ~p(s(bool,happ(s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3)))))) | $spl1),
  inference(cnf_transformation,[],[f76_D])).
fof(f76_D,plain,(
  ~p(s(bool,happ(s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f48,plain,(
  ~p(s(bool,happ(s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3))))))),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),sK3),s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2)))))) & ~p(s(bool,happ(s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3])],[f31])).
fof(f31,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))),
  inference(ennf_transformation,[],[f15])).
fof(f15,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) => p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))),
  inference(flattening,[],[f14])).
fof(f14,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) => p(s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X6,bool),bool),X8),s(fun(fun(X6,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X6),X7)))))) => p(s(bool,happ(s(fun(fun(X6,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X6),X7))),s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X6,bool),bool),X8),s(fun(fun(X6,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X6),X7)))))) => p(s(bool,happ(s(fun(fun(X6,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X6),X7))),s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X8)))))))),
  file('thm.topology.TOPOLOGY|split|3','thm.topology.TOPOLOGY|split|3')).
fof(f103,plain,(
  ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3)))))) | $spl45),
  inference(cnf_transformation,[],[f103_D])).
fof(f103_D,plain,(
  ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3)))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f101,plain,(
  ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3)))))) | ($spl1 | $spl40)),
  inference(forward_demodulation,[],[f76,f99])).
fof(f370,plain,(
  ( ! [X18] : (s(bool,X18) = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3))))) | s(bool,'T_0') = s(bool,X18)) ) | $spl66),
  inference(cnf_transformation,[],[f370_D])).
fof(f370_D,plain,(
  ( ! [X18] : (s(bool,X18) = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3))))) | s(bool,'T_0') = s(bool,X18)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f369,plain,(
  ( ! [X18] : (s(bool,X18) = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3))))) | s(bool,'T_0') = s(bool,X18)) ) | ($spl4 | $spl45 | $spl46)),
  inference(subsumption_resolution,[],[f248,f80])).
fof(f248,plain,(
  ( ! [X18] : (~p(s(bool,'T_0')) | s(bool,X18) = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3))))) | s(bool,'T_0') = s(bool,X18)) ) | ($spl45 | $spl46)),
  inference(superposition,[],[f103,f112])).
fof(f446,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) | $spl78),
  inference(cnf_transformation,[],[f446_D])).
fof(f446_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f448,plain,(
  s(bool,'F_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3))))) | $spl80),
  inference(cnf_transformation,[],[f448_D])).
fof(f448_D,plain,(
  s(bool,'F_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f441,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3))))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl8 | $spl66)),
  inference(duplicate_literal_removal,[],[f396])).
fof(f396,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3))))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl8 | $spl66)),
  inference(superposition,[],[f370,f83])).
fof(f78,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),sK3),s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2)))))) | $spl2),
  inference(cnf_transformation,[],[f78_D])).
fof(f78_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),sK3),s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f47,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),sK3),s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))))))),
  inference(cnf_transformation,[],[f37])).
fof(f116,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),sK3),s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))))) | $spl50),
  inference(cnf_transformation,[],[f116_D])).
fof(f116_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),sK3),s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f114,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),sK3),s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))))) | ($spl2 | $spl7 | $spl8)),
  inference(subsumption_resolution,[],[f108,f82])).
fof(f108,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(sK1,bool),bool),sK3),s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))))) | ($spl2 | $spl8)),
  inference(superposition,[],[f78,f83])).
fof(f3573,plain,(
  ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | $spl291),
  inference(cnf_transformation,[],[f3573_D])).
fof(f3573_D,plain,(
  ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) <=> ~$spl291),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl291])])).
fof(f3571,plain,(
  ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | ($spl4 | $spl7 | $spl50 | $spl80 | $spl156)),
  inference(subsumption_resolution,[],[f3570,f82])).
fof(f3570,plain,(
  p(s(bool,'F_0')) | ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | ($spl4 | $spl50 | $spl80 | $spl156)),
  inference(forward_demodulation,[],[f3569,f448])).
fof(f3569,plain,(
  p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3)))))) | ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | ($spl4 | $spl50 | $spl156)),
  inference(subsumption_resolution,[],[f3557,f80])).
fof(f3557,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK1,bool),bool),sK3)))))) | ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | ($spl50 | $spl156)),
  inference(superposition,[],[f899,f116])).
fof(f85,plain,(
  ( ! [X0,X1] : (s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1))))) | p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))) ) | $spl12),
  inference(cnf_transformation,[],[f85_D])).
fof(f85_D,plain,(
  ( ! [X0,X1] : (s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1))))) | p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f57,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) | s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1)))))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1)))))) & (s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1))))) | p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1))))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X13] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X6,bool),bool),X13)))) <=> s(fun(fun(X6,bool),bool),X13) = s(fun(fun(X6,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X6),'const.topology.topology_1'(s(fun(fun(X6,bool),bool),X13))))))),
  file('thm.topology.TOPOLOGY|split|3','thm.topology.topology_tybij|split|1')).
fof(f84,plain,(
  ( ! [X0,X1] : (s('type.topology.topology'(X0),X1) = s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) ) | $spl10),
  inference(cnf_transformation,[],[f84_D])).
fof(f84_D,plain,(
  ( ! [X0,X1] : (s('type.topology.topology'(X0),X1) = s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f53,plain,(
  ( ! [X0,X1] : (s('type.topology.topology'(X0),X1) = s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : s('type.topology.topology'(X0),X1) = s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0] : ! [X1] : s('type.topology.topology'(X0),X1) = s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X11] : s('type.topology.topology'(X6),X11) = s('type.topology.topology'(X6),'const.topology.topology_1'(s(fun(fun(X6,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X6),X11)))))),
  file('thm.topology.TOPOLOGY|split|3','thm.topology.topology_tybij|split|0')).
fof(f711,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))))))) ) | $spl124),
  inference(cnf_transformation,[],[f711_D])).
fof(f711_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f710,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))))))) ) | ($spl10 | $spl12)),
  inference(trivial_inequality_removal,[],[f709])).
fof(f709,plain,(
  ( ! [X0,X1] : (s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))) != s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))) | p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))))))) ) | ($spl10 | $spl12)),
  inference(superposition,[],[f85,f84])).
fof(f83,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f83_D])).
fof(f83_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f52,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.topology.TOPOLOGY|split|3',aHL_BOOL_CASES)).
fof(f82,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f82_D])).
fof(f82_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f51,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.topology.TOPOLOGY|split|3',aHL_FALSITY)).
fof(f113,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl48),
  inference(cnf_transformation,[],[f113_D])).
fof(f113_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f106,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8)),
  inference(superposition,[],[f82,f83])).
fof(f2422,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.topology.istopology_1'(s(fun(fun(X2,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X2),X3)))))) ) | $spl248),
  inference(cnf_transformation,[],[f2422_D])).
fof(f2422_D,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.topology.istopology_1'(s(fun(fun(X2,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X2),X3)))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f2412,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.topology.istopology_1'(s(fun(fun(X2,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X2),X3)))))) ) | ($spl48 | $spl124)),
  inference(resolution,[],[f711,f113])).
fof(f98,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))) ) | $spl38),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f68,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) | sP0(X0,X1)) & (~sP0(X0,X1) | p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> sP0(X0,X1))),
  inference(definition_folding,[],[f33,f35])).
fof(f33,plain,(
  ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : ((p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4))))) => p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) => p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : ((p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4))))) => p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) => p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X6,bool),bool),X7)))) <=> (p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),'const.pred_set.UNIV_0')))) & ! [X11,X12] : ((p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),X11)))) & p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),X12))))) => p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),'const.topology.re_intersect_2'(s(fun(X6,bool),X11),s(fun(X6,bool),X12))))))) & ! [X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X6,bool),bool),X8),s(fun(fun(X6,bool),bool),X7)))) => p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X8)))))))))),
  file('thm.topology.TOPOLOGY|split|3','thm.topology.istopology')).
fof(f80,plain,(
  p(s(bool,'T_0')) | $spl4),
  inference(cnf_transformation,[],[f80_D])).
fof(f80_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f49,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.topology.TOPOLOGY|split|3',aHL_TRUTH)).
fof(f3621,plain,(
  $false | ($spl4 | $spl38 | $spl248 | $spl291)),
  inference(subsumption_resolution,[],[f3620,f80])).
fof(f3620,plain,(
  ~p(s(bool,'T_0')) | ($spl38 | $spl248 | $spl291)),
  inference(forward_demodulation,[],[f3584,f2422])).
fof(f3584,plain,(
  ~p(s(bool,'const.topology.istopology_1'(s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2)))))) | ($spl38 | $spl291)),
  inference(resolution,[],[f3573,f98])).
% SZS output end Proof for thm.topology
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 75
% Passive clauses: 183
% Generated clauses: 3534
% Final active clauses: 63
% Final passive clauses: 65
% Input formulas: 12
% Initial clauses: 27
% 
% Duplicate literals: 341
% Trivial inequalities: 17
% Fw subsumption resolutions: 289
% Fw demodulations: 454
% Bw demodulations: 48
% Condensations: 257
% 
% Simple tautologies: 71
% Equational tautologies: 18
% Forward subsumptions: 1653
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 51
% Forward superposition: 1239
% Backward superposition: 544
% Self superposition: 47
% Equality factoring: 32
% Equality resolution: 1
% 
% Split clauses: 23
% Split components: 68
% SAT solver clauses: 265
% SAT solver unit clauses: 32
% SAT solver binary clauses: 15
% SAT solver learnt clauses: 1
% 
% Sat splits: 23
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 264
% TWLsolver calls for satisfiability: 387
% 
% Memory used [KB]: 2686
% Time elapsed: 0.088 s
% ------------------------------
% ------------------------------
% Success in time 0.701 s
