Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (77328)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 619
% Passive clauses: 9591
% Generated clauses: 17289
% Final active clauses: 616
% Final passive clauses: 8613
% Input formulas: 50
% Initial clauses: 204
% 
% Duplicate literals: 306
% Trivial inequalities: 45
% Fw subsumption resolutions: 1439
% Fw demodulations: 165
% Bw demodulations: 3
% 
% Simple tautologies: 217
% Equational tautologies: 58
% Forward subsumptions: 5195
% 
% Binary resolution: 10686
% Factoring: 990
% Forward superposition: 1239
% Backward superposition: 2009
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 18
% 
% Split clauses: 305
% Split components: 617
% SAT solver clauses: 599
% SAT solver unit clauses: 19
% SAT solver binary clauses: 166
% 
% Sat splits: 305
% 
% TWLsolver clauses: 599
% TWLsolver calls for satisfiability: 354
% 
% Memory used [KB]: 8955
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (77352)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 511
% Passive clauses: 3187
% Generated clauses: 15162
% Final active clauses: 411
% Final passive clauses: 1525
% Input formulas: 50
% Initial clauses: 204
% 
% Duplicate literals: 225
% Trivial inequalities: 70
% Fw subsumption resolutions: 1453
% Fw demodulations: 774
% Bw demodulations: 1118
% Condensations: 343
% Global subsumptions: 18
% 
% Simple tautologies: 275
% Equational tautologies: 23
% Forward subsumptions: 5878
% Bw demodulations to eq. taut.: 123
% 
% Binary resolution: 3807
% Factoring: 296
% Forward superposition: 1609
% Backward superposition: 1780
% Self superposition: 35
% Equality factoring: 30
% Equality resolution: 19
% 
% Split clauses: 77
% Split components: 155
% SAT solver clauses: 2805
% SAT solver unit clauses: 213
% SAT solver binary clauses: 72
% SAT solver learnt clauses: 37
% SAT solver learnt literals: 15
% 
% Sat splits: 77
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 2732
% TWLsolver calls for satisfiability: 1939
% 
% Memory used [KB]: 12537
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (77375)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1751
% Passive clauses: 9905
% Generated clauses: 101810
% Final active clauses: 675
% Final passive clauses: 1651
% Input formulas: 50
% Initial clauses: 204
% 
% Duplicate literals: 5448
% Trivial inequalities: 210
% Fw subsumption resolutions: 5488
% Fw demodulations: 8696
% Bw demodulations: 4453
% Condensations: 11480
% Global subsumptions: 372
% 
% Simple tautologies: 256
% Equational tautologies: 1189
% Forward subsumptions: 57763
% Bw demodulations to eq. taut.: 321
% 
% Unit resulting resolution: 16704
% Factoring: 9
% Forward superposition: 29060
% Backward superposition: 12259
% Self superposition: 198
% Equality factoring: 162
% Equality resolution: 10
% 
% Split clauses: 425
% Split components: 863
% SAT solver clauses: 11771
% SAT solver unit clauses: 789
% SAT solver binary clauses: 1655
% 
% Sat splits: 425
% Sat splitting refutations: 642
% 
% Memory used [KB]: 59999
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (77690)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1903
% Passive clauses: 42846
% Generated clauses: 173524
% Final active clauses: 1878
% Final passive clauses: 39513
% Input formulas: 50
% Initial clauses: 202
% 
% Selected by SInE selection: 48
% SInE iterations: 1
% 
% Duplicate literals: 5243
% Trivial inequalities: 1369
% Fw subsumption resolutions: 22195
% Fw demodulations: 3590
% Bw demodulations: 8
% 
% Simple tautologies: 5597
% Equational tautologies: 380
% Forward subsumptions: 81733
% 
% Binary resolution: 126006
% Factoring: 418
% Forward superposition: 10977
% Backward superposition: 1997
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 34
% 
% Split clauses: 11843
% Split components: 27585
% SAT solver clauses: 12748
% SAT solver unit clauses: 58
% SAT solver binary clauses: 718
% 
% Sat splits: 11843
% Sat splitting refutations: 74
% 
% Memory used [KB]: 58463
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (77788)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 360
% Passive clauses: 3041
% Generated clauses: 11266
% Final active clauses: 353
% Final passive clauses: 2578
% Input formulas: 50
% Initial clauses: 204
% 
% Duplicate literals: 395
% Trivial inequalities: 48
% Fw subsumption resolutions: 911
% Fw demodulations: 167
% Bw demodulations: 32
% Global subsumptions: 10
% 
% Simple tautologies: 437
% Equational tautologies: 45
% Forward subsumptions: 5868
% 
% Unit resulting resolution: 2215
% Factoring: 14
% Forward superposition: 3063
% Backward superposition: 4036
% Self superposition: 74
% Equality factoring: 6
% Equality resolution: 7
% 
% Split clauses: 32
% Split components: 65
% SAT solver clauses: 3184
% SAT solver unit clauses: 256
% SAT solver binary clauses: 135
% 
% Sat splits: 32
% Sat splitting refutations: 99
% 
% Memory used [KB]: 13304
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (77804)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 898
% Passive clauses: 5670
% Generated clauses: 19948
% Final active clauses: 475
% Final passive clauses: 3144
% Input formulas: 50
% Initial clauses: 204
% 
% Duplicate literals: 994
% Trivial inequalities: 90
% Fw subsumption resolutions: 1270
% Bw subsumption resolutions: 320
% Fw demodulations: 1317
% Bw demodulations: 842
% Condensations: 1418
% Global subsumptions: 5
% 
% Simple tautologies: 1518
% Equational tautologies: 72
% Forward subsumptions: 7468
% Backward subsumptions: 1011
% 
% Binary resolution: 5978
% Factoring: 25
% Forward superposition: 2664
% Backward superposition: 3041
% Self superposition: 44
% Equality factoring: 12
% Equality resolution: 16
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 5775
% SAT solver unit clauses: 229
% SAT solver binary clauses: 1102
% 
% Sat splits: 3
% Sat splitting refutations: 94
% 
% Memory used [KB]: 11769
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f8708,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f274,f483,f273,f485,f272,f487,f623,f621,f620,f619,f665,f618,f617,f616,f614,f613,f612,f669,f611,f667,f610,f609,f671,f601,f673,f599,f598,f597,f677,f596,f675,f595,f594,f679,f590,f681,f575,f683,f566,f565,f564,f687,f563,f685,f562,f561,f689,f559,f558,f557,f693,f556,f691,f555,f554,f695,f553,f552,f551,f699,f550,f697,f549,f548,f701,f522,f521,f520,f519,f518,f517,f516,f515,f514,f711,f513,f709,f512,f511,f510,f707,f509,f703,f508,f507,f705,f501,f500,f498,f497,f491,f490,f663,f489,f488,f625,f626,f714,f713,f731,f730,f737,f729,f739,f723,f722,f741,f721,f743,f883,f881,f878,f877,f876,f874,f873,f872,f928,f871,f926,f870,f869,f930,f861,f932,f859,f858,f857,f936,f856,f934,f855,f854,f938,f850,f940,f835,f942,f832,f944,f831,f824,f823,f822,f950,f821,f948,f820,f819,f952,f817,f816,f815,f956,f814,f954,f813,f812,f958,f811,f810,f809,f962,f808,f960,f807,f806,f946,f780,f779,f778,f777,f776,f775,f774,f773,f772,f972,f771,f970,f770,f769,f768,f968,f767,f964,f766,f765,f966,f759,f758,f756,f755,f749,f748,f924,f747,f746,f745,f744,f884,f978,f885,f976,f886,f887,f979,f974,f1081,f1080,f1084,f1097,f1096,f1095,f1100,f1109,f1101,f1107,f1102,f1105,f1320,f1319,f1318,f1322,f1332,f1323,f1330,f1324,f1328,f1342,f1348,f1350,f1387,f1391,f1397,f1401,f1412,f1416,f1422,f1426,f1485,f1484,f1584,f1471,f1470,f1593,f1785,f1784,f1791,f1800,f1799,f1803,f1811,f1810,f1814,f2047,f2053,f2055,f2097,f2101,f2107,f2111,f2117,f2121,f2360,f2408,f2359,f2410,f2358,f2350,f2349,f2347,f2344,f2343,f2342,f2341,f2340,f2336,f2335,f2334,f2419,f2331,f2417,f2330,f2329,f2423,f2319,f2427,f2315,f2314,f2313,f2433,f2310,f2431,f2309,f2308,f2437,f2278,f2447,f2255,f2453,f2240,f2459,f2239,f2238,f2463,f2236,f2465,f2228,f2227,f2226,f2471,f2223,f2469,f2222,f2221,f2475,f2217,f2216,f2215,f2481,f2212,f2479,f2211,f2210,f2485,f2209,f2487,f2208,f2489,f2207,f2491,f2206,f2493,f2205,f2204,f2203,f2497,f2202,f2495,f2201,f2200,f2461,f2174,f2173,f2172,f2171,f2170,f2169,f2168,f2167,f2166,f2507,f2165,f2505,f2164,f2163,f2162,f2503,f2161,f2499,f2160,f2159,f2501,f2153,f2152,f2150,f2149,f2143,f2142,f2412,f2141,f2140,f2370,f2371,f2511,f2509,f2736,f2782,f2735,f2784,f2734,f2726,f2725,f2723,f2720,f2719,f2718,f2717,f2716,f2712,f2711,f2710,f2793,f2707,f2791,f2706,f2705,f2797,f2695,f2801,f2691,f2690,f2689,f2807,f2686,f2805,f2685,f2684,f2811,f2654,f2821,f2631,f2827,f2616,f2833,f2615,f2614,f2837,f2612,f2839,f2604,f2603,f2602,f2845,f2599,f2843,f2598,f2597,f2849,f2593,f2592,f2591,f2855,f2588,f2853,f2587,f2586,f2859,f2585,f2861,f2584,f2863,f2583,f2865,f2582,f2867,f2581,f2580,f2579,f2871,f2578,f2869,f2577,f2576,f2835,f2550,f2549,f2548,f2547,f2546,f2545,f2544,f2543,f2542,f2881,f2541,f2879,f2540,f2539,f2538,f2877,f2537,f2873,f2536,f2535,f2875,f2529,f2528,f2526,f2525,f2519,f2518,f2786,f2517,f2516,f2515,f2742,f2885,f2744,f2745,f2886,f2883,f2994,f2998,f3006,f3010,f3085,f3084,f3189,f3069,f3068,f3200,f3481,f3485,f3493,f3497,f3553,f3552,f3557,f3565,f3558,f3563,f3570,f3574,f3578,f3797,f3796,f3803,f3795,f3794,f3805,f3816,f3815,f3819,f3829,f3828,f3832,f4073,f4079,f4081,f4128,f4132,f4138,f4142,f4148,f4152,f4209,f4208,f4215,f4207,f4206,f4217,f4205,f4204,f4219,f4232,f4231,f4235,f4228,f4227,f4237,f4247,f4246,f4252,f4497,f4503,f4505,f4552,f4556,f4562,f4566,f4572,f4576,f4594,f4593,f4598,f4606,f4599,f4604,f4611,f4615,f4619,f4708,f4707,f4819,f4690,f4689,f4832,f4929,f4928,f5042,f4909,f4908,f5057,f5081,f5080,f5097,f5079,f5078,f5099,f5077,f5076,f5101,f5075,f5074,f5103,f5111,f5110,f5127,f5109,f5108,f5129,f5105,f5104,f5131,f5147,f5161,f5146,f5163,f5144,f5165,f5169,f5183,f5168,f5185,f5166,f5187,f5580,f5585,f5588,f5643,f5648,f5651,f6054,f6053,f6069,f6052,f6051,f6071,f6048,f6047,f6073,f6086,f6085,f6101,f6084,f6083,f6103,f6107,f6121,f6106,f6123,f6127,f6141,f6126,f6143,f6272,f6275,f6278,f6392,f6396,f6404,f6408,f6502,f6506,f6514,f6518,f6739,f6740,f6741,f6742,f6750,f6751,f6964,f6963,f6968,f6976,f6969,f6974,f6981,f6985,f6989,f6994,f6993,f6998,f7006,f6999,f7004,f7011,f7015,f7019,f7024,f7028,f7032,f7047,f7049,f7262,f7266,f7270,f7279,f7283,f7287,f7299,f7303,f7307,f7316,f7320,f7324,f7333,f7337,f7341,f7351,f7355,f7359,f7367,f7370,f7375,f7378,f7385,f7381,f7388,f7394,f7391,f7397,f7407,f7404,f7410,f7417,f7413,f7420,f7442,f7441,f7458,f7440,f7439,f7460,f7466,f7465,f7484,f7488,f7512,f7515,f7531,f7537,f7536,f7555,f8040,f8041,f8044,f8045,f8082,f8046,f8047,f8078,f8049,f8079,f8050,f8080,f8081,f8083,f8122,f8126,f8130,f8134,f8181,f8183,f8179,f8185,f8462,f8610,f8461,f8460,f8612,f8440,f8629,f8439,f8438,f8631,f8437,f8633,f8436,f8635,f8435,f8434,f8637,f8433,f8432,f8639,f8431,f8641,f8430,f8643,f8429,f8428,f8645,f8427,f8426,f8647,f8425,f8649,f8424,f8651,f8423,f8422,f8653,f8421,f8420,f8655,f8419,f8657,f8418,f8659,f8417,f8416,f8661,f8415,f8414,f8663,f8666,f8668,f8480,f8669,f8500,f8501,f8675,f8673,f8503,f8674,f8504,f8676,f8505,f8506,f8507,f8678,f8508,f8677,f8509,f8679,f8510,f8680,f8513,f8681,f8514,f8515,f8683,f8518,f8684,f8521,f8685,f8522,f8686,f8523,f8687,f8529,f8688,f8530,f8689,f8535,f8690,f8536,f8691,f8538,f8692,f8540,f8693,f8541,f8694,f8542,f8695,f8553,f6277,f8682,f8707])).
fof(f8707,plain,(
  s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)) != s(sK12,'const.pred_set.LINV_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(sK13,sK17))) | s(sK13,sK17) != s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)))) | s(sK13,sK17) = s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,'const.pred_set.LINV_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(sK13,sK17)))))),
  introduced(theory_axiom,[])).
fof(f8682,plain,(
  ( ! [X24,X25] : (sP3('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X24,X25)) ) | $spl566),
  inference(cnf_transformation,[],[f8682_D])).
fof(f8682_D,plain,(
  ( ! [X24,X25] : (sP3('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X24,X25)) ) <=> ~$spl566),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl566])])).
fof(f6277,plain,(
  ~p(s(bool,'T_0')) | $spl421),
  inference(cnf_transformation,[],[f6277_D])).
fof(f6277_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl421),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl421])])).
fof(f8553,plain,(
  ( ! [X112,X113] : (~p(s(bool,'T_0')) | sP3('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X112,X113)) ) | $spl504),
  inference(superposition,[],[f478,f8185])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP3(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f400])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X2)) | p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X2)) & ~p(s(bool,X0))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f236])).
fof(f236,plain,(
  ! [X0,X2,X1] : ((~sP3(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (~p(s(bool,X2)) & ~p(s(bool,X0))) | sP3(X0,X2,X1)))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X0,X2,X1] : ((~sP3(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (~p(s(bool,X2)) & ~p(s(bool,X0)))) | sP3(X0,X2,X1)))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X2,X1] : (sP3(X0,X2,X1) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f8695,plain,(
  ( ! [X80,X81] : (~p(s(bool,X80)) | sP11('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X80,X81)) ) | $spl592),
  inference(cnf_transformation,[],[f8695_D])).
fof(f8695_D,plain,(
  ( ! [X80,X81] : (~p(s(bool,X80)) | sP11('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X80,X81)) ) <=> ~$spl592),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl592])])).
fof(f8542,plain,(
  ( ! [X80,X81] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | sP11('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X80,X81)) ) | $spl504),
  inference(superposition,[],[f464,f8185])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f159,f179,f178])).
fof(f178,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f179,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) => p(s(bool,X20)))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (~p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.sat.sat_27_prove_ltb')).
fof(f8694,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | sP11('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X79,X78)) ) | $spl590),
  inference(cnf_transformation,[],[f8694_D])).
fof(f8694_D,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | sP11('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X79,X78)) ) <=> ~$spl590),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl590])])).
fof(f8541,plain,(
  ( ! [X78,X79] : (~p(s(bool,'T_0')) | p(s(bool,X78)) | sP11('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X79,X78)) ) | $spl504),
  inference(superposition,[],[f463,f8185])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f269])).
fof(f8693,plain,(
  ( ! [X76,X77] : (p(s(bool,X76)) | p(s(bool,X77)) | sP11(X77,X76,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl588),
  inference(cnf_transformation,[],[f8693_D])).
fof(f8693_D,plain,(
  ( ! [X76,X77] : (p(s(bool,X76)) | p(s(bool,X77)) | sP11(X77,X76,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl588),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl588])])).
fof(f8540,plain,(
  ( ! [X76,X77] : (~p(s(bool,'T_0')) | p(s(bool,X76)) | sP11(X77,X76,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X77))) ) | $spl504),
  inference(superposition,[],[f462,f8185])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP11(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f8692,plain,(
  ( ! [X72,X73] : (sP10(X72,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X73)) ) | $spl586),
  inference(cnf_transformation,[],[f8692_D])).
fof(f8692_D,plain,(
  ( ! [X72,X73] : (sP10(X72,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X73)) ) <=> ~$spl586),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl586])])).
fof(f8538,plain,(
  ( ! [X72,X73] : (~p(s(bool,'T_0')) | sP10(X72,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X73)) ) | $spl504),
  inference(superposition,[],[f460,f8185])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f266])).
fof(f266,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f178])).
fof(f8691,plain,(
  ( ! [X68,X69] : (~p(s(bool,X68)) | sP9('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X68,X69)) ) | $spl584),
  inference(cnf_transformation,[],[f8691_D])).
fof(f8691_D,plain,(
  ( ! [X68,X69] : (~p(s(bool,X68)) | sP9('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X68,X69)) ) <=> ~$spl584),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl584])])).
fof(f8536,plain,(
  ( ! [X68,X69] : (~p(s(bool,'T_0')) | ~p(s(bool,X68)) | sP9('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X68,X69)) ) | $spl504),
  inference(superposition,[],[f447,f8185])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f122,f176,f175])).
fof(f175,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f176,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f122,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) | p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.sat.sat_25_prove_ltb')).
fof(f8690,plain,(
  ( ! [X66,X67] : (~p(s(bool,X66)) | sP9('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X67,X66)) ) | $spl582),
  inference(cnf_transformation,[],[f8690_D])).
fof(f8690_D,plain,(
  ( ! [X66,X67] : (~p(s(bool,X66)) | sP9('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X67,X66)) ) <=> ~$spl582),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl582])])).
fof(f8535,plain,(
  ( ! [X66,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X66)) | sP9('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X67,X66)) ) | $spl504),
  inference(superposition,[],[f446,f8185])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f261])).
fof(f8689,plain,(
  ( ! [X57,X56] : (sP8(X56,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X57)) ) | $spl580),
  inference(cnf_transformation,[],[f8689_D])).
fof(f8689_D,plain,(
  ( ! [X57,X56] : (sP8(X56,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X57)) ) <=> ~$spl580),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl580])])).
fof(f8530,plain,(
  ( ! [X57,X56] : (~p(s(bool,'T_0')) | sP8(X56,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X57)) ) | $spl504),
  inference(superposition,[],[f443,f8185])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f258])).
fof(f258,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f175])).
fof(f8688,plain,(
  ( ! [X54,X55] : (sP8(X54,X55,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl578),
  inference(cnf_transformation,[],[f8688_D])).
fof(f8688_D,plain,(
  ( ! [X54,X55] : (sP8(X54,X55,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl578),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl578])])).
fof(f8529,plain,(
  ( ! [X54,X55] : (~p(s(bool,'T_0')) | sP8(X54,X55,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl504),
  inference(superposition,[],[f442,f8185])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f259])).
fof(f8687,plain,(
  ( ! [X43,X42] : (sP4('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X42,X43)) ) | $spl576),
  inference(cnf_transformation,[],[f8687_D])).
fof(f8687_D,plain,(
  ( ! [X43,X42] : (sP4('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X42,X43)) ) <=> ~$spl576),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl576])])).
fof(f8523,plain,(
  ( ! [X43,X42] : (~p(s(bool,'T_0')) | sP4('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X42,X43)) ) | $spl504),
  inference(superposition,[],[f429,f8185])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f8686,plain,(
  ( ! [X41,X40] : (sP4(X40,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X41)) ) | $spl574),
  inference(cnf_transformation,[],[f8686_D])).
fof(f8686_D,plain,(
  ( ! [X41,X40] : (sP4(X40,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X41)) ) <=> ~$spl574),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl574])])).
fof(f8522,plain,(
  ( ! [X41,X40] : (~p(s(bool,'T_0')) | sP4(X40,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X41)) ) | $spl504),
  inference(superposition,[],[f428,f8185])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f8685,plain,(
  ( ! [X39,X38] : (sP4(X38,X39,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl572),
  inference(cnf_transformation,[],[f8685_D])).
fof(f8685_D,plain,(
  ( ! [X39,X38] : (sP4(X38,X39,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl572),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl572])])).
fof(f8521,plain,(
  ( ! [X39,X38] : (~p(s(bool,'T_0')) | sP4(X38,X39,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl504),
  inference(superposition,[],[f427,f8185])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f8684,plain,(
  ( ! [X33,X32] : (sP5(X32,X33,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl570),
  inference(cnf_transformation,[],[f8684_D])).
fof(f8684_D,plain,(
  ( ! [X33,X32] : (sP5(X32,X33,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl570),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl570])])).
fof(f8518,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | sP5(X32,X33,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl504),
  inference(superposition,[],[f423,f8185])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f249])).
fof(f249,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X1,X0,X2] : (sP5(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f8683,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl568),
  inference(cnf_transformation,[],[f8683_D])).
fof(f8683_D,plain,(
  ( ! [X26,X27] : (sP6(X26,X27,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl568),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl568])])).
fof(f8515,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | sP6(X26,X27,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl504),
  inference(superposition,[],[f419,f8185])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f246])).
fof(f246,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f8514,plain,(
  ( ! [X24,X25] : (~p(s(bool,'T_0')) | sP3('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X24,X25)) ) | $spl504),
  inference(superposition,[],[f402,f8185])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP3(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP3(X0,X2,X1)) & (~sP3(X0,X2,X1) | (p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) | ~p(s(bool,X2))) & ~p(s(bool,X0))) | sP3(X0,X2,X1)) & (~sP3(X0,X2,X1) | ((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> sP3(X0,X2,X1))),
  inference(definition_folding,[],[f117,f168])).
fof(f117,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & p(s(bool,X2))) | p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | p(s(bool,X0)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6,X7,X16] : (((p(s(bool,X7)) & p(s(bool,X16))) | p(s(bool,X6))) <=> ((p(s(bool,X7)) | p(s(bool,X6))) & (p(s(bool,X16)) | p(s(bool,X6)))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.bool.RIGHT_OR_OVER_AND')).
fof(f8681,plain,(
  ( ! [X23,X22] : (sP3(X22,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X23) | ~p(s(bool,X23))) ) | $spl564),
  inference(cnf_transformation,[],[f8681_D])).
fof(f8681_D,plain,(
  ( ! [X23,X22] : (sP3(X22,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X23) | ~p(s(bool,X23))) ) <=> ~$spl564),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl564])])).
fof(f8513,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | sP3(X22,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X23) | ~p(s(bool,X23))) ) | $spl504),
  inference(superposition,[],[f401,f8185])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP3(X0,X2,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f239])).
fof(f8680,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | sP3(X17,X16,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl562),
  inference(cnf_transformation,[],[f8680_D])).
fof(f8680_D,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | sP3(X17,X16,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl562),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl562])])).
fof(f8510,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | sP3(X17,X16,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl504),
  inference(superposition,[],[f397,f8185])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f8679,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | ~sP1('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X15)) ) | $spl560),
  inference(cnf_transformation,[],[f8679_D])).
fof(f8679_D,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | ~sP1('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X15)) ) <=> ~$spl560),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl560])])).
fof(f8509,plain,(
  ( ! [X15] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | ~sP1('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X15)) ) | $spl504),
  inference(superposition,[],[f369,f8185])).
fof(f369,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP1(X1,X0)) & (~sP1(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP1(X1,X0))),
  inference(definition_folding,[],[f106,f164])).
fof(f164,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f106,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X21,X22] : ((p(s(bool,X22)) <=> ~p(s(bool,X21))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.sat.sat_29_prove_ltb')).
fof(f8677,plain,(
  ( ! [X11] : (p(s(bool,X11)) | sP1('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X11)) ) | $spl556),
  inference(cnf_transformation,[],[f8677_D])).
fof(f8677_D,plain,(
  ( ! [X11] : (p(s(bool,X11)) | sP1('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X11)) ) <=> ~$spl556),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl556])])).
fof(f8508,plain,(
  ( ! [X14] : (~p(s(bool,'T_0')) | p(s(bool,X14)) | sP1('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X14)) ) | $spl504),
  inference(superposition,[],[f368,f8185])).
fof(f368,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f222])).
fof(f8678,plain,(
  ( ! [X12] : (p(s(bool,X12)) | sP1(X12,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl558),
  inference(cnf_transformation,[],[f8678_D])).
fof(f8678_D,plain,(
  ( ! [X12] : (p(s(bool,X12)) | sP1(X12,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl558),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl558])])).
fof(f8507,plain,(
  ( ! [X13] : (~p(s(bool,'T_0')) | sP1(X13,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X13))) ) | $spl504),
  inference(superposition,[],[f367,f8185])).
fof(f367,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X1,X0) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f222])).
fof(f8506,plain,(
  ( ! [X12] : (~p(s(bool,'T_0')) | p(s(bool,X12)) | sP1(X12,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl504),
  inference(superposition,[],[f366,f8185])).
fof(f366,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f220])).
fof(f220,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f219])).
fof(f219,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f164])).
fof(f8505,plain,(
  ( ! [X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | sP1('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X11)) ) | $spl504),
  inference(superposition,[],[f363,f8185])).
fof(f363,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f8676,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~sP1(X10,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl554),
  inference(cnf_transformation,[],[f8676_D])).
fof(f8676_D,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~sP1(X10,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl554),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl554])])).
fof(f8504,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | ~sP1(X10,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl504),
  inference(superposition,[],[f362,f8185])).
fof(f362,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f8674,plain,(
  ( ! [X7] : (sP0('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X7)) ) | $spl552),
  inference(cnf_transformation,[],[f8674_D])).
fof(f8674_D,plain,(
  ( ! [X7] : (sP0('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X7)) ) <=> ~$spl552),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl552])])).
fof(f8503,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | sP0('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X9)) ) | $spl504),
  inference(superposition,[],[f353,f8185])).
fof(f353,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f149,f162])).
fof(f162,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f149,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X7,X6] : ((~(p(s(bool,X6)) | p(s(bool,X7))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X7)) => p(s(bool,'F_0')))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.sat.sat_9_prove_ltb')).
fof(f8673,plain,(
  ( ! [X6] : (sP0(X6,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl550),
  inference(cnf_transformation,[],[f8673_D])).
fof(f8673_D,plain,(
  ( ! [X6] : (sP0(X6,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl550),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl550])])).
fof(f8675,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | sP0(X8,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl504),
  inference(global_subsumption,[],[f274,f273,f272,f8185,f8462,f8461,f8460,f8440,f8439,f8438,f8480,f8500,f8501,f8502])).
fof(f8502,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | p(s(bool,'F_0')) | sP0(X8,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl504),
  inference(superposition,[],[f352,f8185])).
fof(f352,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f216])).
fof(f8501,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | sP0('const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),X7)) ) | $spl504),
  inference(superposition,[],[f348,f8185])).
fof(f348,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f162])).
fof(f8500,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | sP0(X6,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl504),
  inference(superposition,[],[f347,f8185])).
fof(f347,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f214])).
fof(f8669,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(sK12,X0),X1),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)))),s(fun(X0,bool),X2)))) | ~sP2(sK12,X0,X1,sK15,X2)) ) | $spl548),
  inference(cnf_transformation,[],[f8669_D])).
fof(f8669_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(sK12,X0),X1),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)))),s(fun(X0,bool),X2)))) | ~sP2(sK12,X0,X1,sK15,X2)) ) <=> ~$spl548),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl548])])).
fof(f8480,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(sK12,X0),X1),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)))),s(fun(X0,bool),X2)))) | ~sP2(sK12,X0,X1,sK15,X2)) ) | $spl504),
  inference(superposition,[],[f383,f8185])).
fof(f383,plain,(
  ( ! [X4,X2,X0,X10,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4)))) | ~sP2(X0,X1,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP2(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK22(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK22(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK23(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,sK24(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,sK23(X4,X3,X2,X1,X0)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK24(X4,X3,X2,X1,X0)))) & s(X0,sK23(X4,X3,X2,X1,X0)) != s(X0,sK24(X4,X3,X2,X1,X0))) | sP2(X0,X1,X2,X3,X4)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22,sK23,sK24])],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP2(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))) & s(X0,X5) != s(X0,X6)) | sP2(X0,X1,X2,X3,X4)))),
  inference(rectify,[],[f231])).
fof(f231,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP2(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6)) | sP2(X1,X0,X2,X3,X4)))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP2(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6))) | sP2(X1,X0,X2,X3,X4)))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X1,X0,X2,X3,X4] : (sP2(X1,X0,X2,X3,X4) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f8668,plain,(
  s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)) = s(sK12,'const.pred_set.LINV_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(sK13,sK17))) | $spl546),
  inference(cnf_transformation,[],[f8668_D])).
fof(f8668_D,plain,(
  s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)) = s(sK12,'const.pred_set.LINV_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(sK13,sK17))) <=> ~$spl546),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl546])])).
fof(f8666,plain,(
  s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)) = s(sK12,'const.pred_set.LINV_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(sK13,sK17))) | ~p(s(bool,'T_0')) | ($spl120 | $spl418 | $spl504)),
  inference(forward_demodulation,[],[f8479,f6275])).
fof(f8479,plain,(
  ~p(s(bool,'T_0')) | s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)) = s(sK12,'const.pred_set.LINV_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)))))) | ($spl120 | $spl504)),
  inference(superposition,[],[f2743,f8185])).
fof(f2743,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15)))) | s(sK12,X0) = s(sK12,'const.pred_set.LINV_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0)))))) ) | $spl120),
  inference(resolution,[],[f978,f372])).
fof(f372,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | s(X1,X5) = s(X1,'const.pred_set.LINV_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | s(X1,X5) = s(X1,'const.pred_set.LINV_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) => ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) => s(X1,X5) = s(X1,'const.pred_set.LINV_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) => ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) => s(X1,X5) = s(X1,'const.pred_set.LINV_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X7,X6] : ! [X8,X9,X10] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X10)))) => ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X9)))) => s(X6,X11) = s(X6,'const.pred_set.LINV_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(X7,happ(s(fun(X6,X7),X8),s(X6,X11)))))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.pred_set.LINV_DEF')).
fof(f8663,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl544),
  inference(cnf_transformation,[],[f8663_D])).
fof(f8663_D,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl544])])).
fof(f8414,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl504)),
  inference(unit_resulting_resolution,[],[f276,f8185,f1539])).
fof(f1539,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,X46) | sP7(X47,X46,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ~p(s(bool,X47))) ) | $spl56),
  inference(superposition,[],[f431,f713])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP7(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP7(X2,X1,X0)) & (~sP7(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP7(X2,X1,X0))),
  inference(definition_folding,[],[f120,f173,f172,f171,f170])).
fof(f173,plain,(
  ! [X2,X1,X0] : (sP7(X2,X1,X0) <=> (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f120,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X22)) | p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.sat.sat_2|49|_prove_ltb')).
fof(f276,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.BIJ_LINV_INV','thm.bool.TRUTH')).
fof(f8415,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl504)),
  inference(unit_resulting_resolution,[],[f275,f8185,f1539])).
fof(f275,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.BIJ_LINV_INV',aHL_TRUTH)).
fof(f8661,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl542),
  inference(cnf_transformation,[],[f8661_D])).
fof(f8661_D,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl542),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl542])])).
fof(f8416,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl56 | $spl504)),
  inference(unit_resulting_resolution,[],[f276,f8185,f1540])).
fof(f1540,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X48) | sP7(X49,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X48) | ~p(s(bool,X49))) ) | $spl56),
  inference(superposition,[],[f431,f713])).
fof(f8417,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl56 | $spl504)),
  inference(unit_resulting_resolution,[],[f275,f8185,f1540])).
fof(f8659,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl541),
  inference(cnf_transformation,[],[f8659_D])).
fof(f8659_D,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl541),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl541])])).
fof(f8418,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl504)),
  inference(unit_resulting_resolution,[],[f281,f8185,f1541])).
fof(f1541,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X50) | p(s(bool,X51)) | ~sP7(X51,X50,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl56),
  inference(superposition,[],[f433,f713])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~sP7(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f253])).
fof(f281,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.BIJ_LINV_INV',aHL_FALSITY)).
fof(f8657,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl539),
  inference(cnf_transformation,[],[f8657_D])).
fof(f8657_D,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl539),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl539])])).
fof(f8419,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl56 | $spl504)),
  inference(unit_resulting_resolution,[],[f281,f8185,f1542])).
fof(f1542,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,X52) | p(s(bool,X53)) | ~sP7(X53,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X52)) ) | $spl56),
  inference(superposition,[],[f433,f713])).
fof(f8655,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl536),
  inference(cnf_transformation,[],[f8655_D])).
fof(f8655_D,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl536),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl536])])).
fof(f8420,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl504)),
  inference(unit_resulting_resolution,[],[f276,f8185,f3142])).
fof(f3142,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') != s(bool,X43) | sP7(X44,X43,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ~p(s(bool,X44))) ) | $spl116),
  inference(superposition,[],[f431,f974])).
fof(f8421,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl504)),
  inference(unit_resulting_resolution,[],[f275,f8185,f3142])).
fof(f8653,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl534),
  inference(cnf_transformation,[],[f8653_D])).
fof(f8653_D,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl534),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl534])])).
fof(f8422,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl116 | $spl504)),
  inference(unit_resulting_resolution,[],[f276,f8185,f3143])).
fof(f3143,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,X45) | sP7(X46,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X45) | ~p(s(bool,X46))) ) | $spl116),
  inference(superposition,[],[f431,f974])).
fof(f8423,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl116 | $spl504)),
  inference(unit_resulting_resolution,[],[f275,f8185,f3143])).
fof(f8651,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl533),
  inference(cnf_transformation,[],[f8651_D])).
fof(f8651_D,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl533),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl533])])).
fof(f8424,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl504)),
  inference(unit_resulting_resolution,[],[f281,f8185,f3144])).
fof(f3144,plain,(
  ( ! [X47,X48] : (s(bool,'T_0') != s(bool,X47) | p(s(bool,X48)) | ~sP7(X48,X47,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl116),
  inference(superposition,[],[f433,f974])).
fof(f8649,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl531),
  inference(cnf_transformation,[],[f8649_D])).
fof(f8649_D,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl531),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl531])])).
fof(f8425,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl116 | $spl504)),
  inference(unit_resulting_resolution,[],[f281,f8185,f3145])).
fof(f3145,plain,(
  ( ! [X50,X49] : (s(bool,'T_0') != s(bool,X49) | p(s(bool,X50)) | ~sP7(X50,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X49)) ) | $spl116),
  inference(superposition,[],[f433,f974])).
fof(f8647,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl528),
  inference(cnf_transformation,[],[f8647_D])).
fof(f8647_D,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl528),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl528])])).
fof(f8426,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl230 | $spl504)),
  inference(unit_resulting_resolution,[],[f276,f8185,f4770])).
fof(f4770,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,X46) | sP7(X47,X46,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ~p(s(bool,X47))) ) | $spl230),
  inference(superposition,[],[f431,f2509])).
fof(f8427,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl230 | $spl504)),
  inference(unit_resulting_resolution,[],[f275,f8185,f4770])).
fof(f8645,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl526),
  inference(cnf_transformation,[],[f8645_D])).
fof(f8645_D,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl526),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl526])])).
fof(f8428,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl230 | $spl504)),
  inference(unit_resulting_resolution,[],[f276,f8185,f4771])).
fof(f4771,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X48) | sP7(X49,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X48) | ~p(s(bool,X49))) ) | $spl230),
  inference(superposition,[],[f431,f2509])).
fof(f8429,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl230 | $spl504)),
  inference(unit_resulting_resolution,[],[f275,f8185,f4771])).
fof(f8643,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl525),
  inference(cnf_transformation,[],[f8643_D])).
fof(f8643_D,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl525),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl525])])).
fof(f8430,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl230 | $spl504)),
  inference(unit_resulting_resolution,[],[f281,f8185,f4772])).
fof(f4772,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X50) | p(s(bool,X51)) | ~sP7(X51,X50,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl230),
  inference(superposition,[],[f433,f2509])).
fof(f8641,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl523),
  inference(cnf_transformation,[],[f8641_D])).
fof(f8641_D,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl523),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl523])])).
fof(f8431,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl230 | $spl504)),
  inference(unit_resulting_resolution,[],[f281,f8185,f4773])).
fof(f4773,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,X52) | p(s(bool,X53)) | ~sP7(X53,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X52)) ) | $spl230),
  inference(superposition,[],[f433,f2509])).
fof(f8639,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl520),
  inference(cnf_transformation,[],[f8639_D])).
fof(f8639_D,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl520),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl520])])).
fof(f8432,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl298 | $spl504)),
  inference(unit_resulting_resolution,[],[f276,f8185,f4991])).
fof(f4991,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,X44) | sP7(X45,X44,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ~p(s(bool,X45))) ) | $spl298),
  inference(superposition,[],[f431,f2883])).
fof(f8433,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl298 | $spl504)),
  inference(unit_resulting_resolution,[],[f275,f8185,f4991])).
fof(f8637,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl518),
  inference(cnf_transformation,[],[f8637_D])).
fof(f8637_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl518),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl518])])).
fof(f8434,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl298 | $spl504)),
  inference(unit_resulting_resolution,[],[f276,f8185,f4992])).
fof(f4992,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,X46) | sP7(X47,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X46) | ~p(s(bool,X47))) ) | $spl298),
  inference(superposition,[],[f431,f2883])).
fof(f8435,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl298 | $spl504)),
  inference(unit_resulting_resolution,[],[f275,f8185,f4992])).
fof(f8635,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl517),
  inference(cnf_transformation,[],[f8635_D])).
fof(f8635_D,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl517),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl517])])).
fof(f8436,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl298 | $spl504)),
  inference(unit_resulting_resolution,[],[f281,f8185,f4993])).
fof(f4993,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X48) | p(s(bool,X49)) | ~sP7(X49,X48,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl298),
  inference(superposition,[],[f433,f2883])).
fof(f8633,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl515),
  inference(cnf_transformation,[],[f8633_D])).
fof(f8633_D,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl515),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl515])])).
fof(f8437,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl298 | $spl504)),
  inference(unit_resulting_resolution,[],[f281,f8185,f4994])).
fof(f4994,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X50) | p(s(bool,X51)) | ~sP7(X51,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X50)) ) | $spl298),
  inference(superposition,[],[f433,f2883])).
fof(f8631,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | $spl512),
  inference(cnf_transformation,[],[f8631_D])).
fof(f8631_D,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f8438,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | $spl504),
  inference(unit_resulting_resolution,[],[f276,f8185,f431])).
fof(f8439,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | $spl504),
  inference(unit_resulting_resolution,[],[f275,f8185,f431])).
fof(f8629,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | $spl511),
  inference(cnf_transformation,[],[f8629_D])).
fof(f8629_D,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') <=> ~$spl511),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl511])])).
fof(f8440,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)),'T_0') | $spl504),
  inference(unit_resulting_resolution,[],[f281,f8185,f433])).
fof(f8612,plain,(
  sP7('T_0','T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl508),
  inference(cnf_transformation,[],[f8612_D])).
fof(f8612_D,plain,(
  sP7('T_0','T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl508),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl508])])).
fof(f8460,plain,(
  sP7('T_0','T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl504),
  inference(unit_resulting_resolution,[],[f276,f8185,f431])).
fof(f8461,plain,(
  sP7('T_0','T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl504),
  inference(unit_resulting_resolution,[],[f275,f8185,f431])).
fof(f8610,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl507),
  inference(cnf_transformation,[],[f8610_D])).
fof(f8610_D,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl507),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl507])])).
fof(f8462,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl504),
  inference(unit_resulting_resolution,[],[f281,f8185,f433])).
fof(f8185,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl504),
  inference(cnf_transformation,[],[f8185_D])).
fof(f8185_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f8179,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl56 | $spl118)),
  inference(superposition,[],[f6010,f713])).
fof(f6010,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X4),s(fun(sK13,bool),sK16)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK19(X4,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl118),
  inference(resolution,[],[f2367,f290])).
fof(f290,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.BIJ_LINV_INV','thm.bool.EQ_CLAUSES|split|0')).
fof(f2367,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK19(X1,sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,X1),s(fun(sK13,bool),sK16))))) ) | $spl118),
  inference(resolution,[],[f976,f377])).
fof(f377,plain,(
  ( ! [X4,X2,X0,X8,X3,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK19(X8,X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK19(X8,X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK19(X8,X4,X3,X2,X1,X0)))))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK20(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK20(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X1,sK21(X4,X3,X2,X1,X0)),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,sK21(X4,X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20,sK21])],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) & s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))))))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X10)))) <=> (! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X9)))) => p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,X11))),s(fun(X7,bool),X10))))) & ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X7,X11),s(fun(X7,bool),X10)))) => ? [X15] : (p(s(bool,'const.bool.IN_2'(s(X6,X15),s(fun(X6,bool),X9)))) & s(X7,X11) = s(X7,happ(s(fun(X6,X7),X8),s(X6,X15)))))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.pred_set.SURJ_DEF')).
fof(f8183,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK19(happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12))),sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl502),
  inference(cnf_transformation,[],[f8183_D])).
fof(f8183_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK19(happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12))),sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f8181,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK19(happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12))),sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl56 | $spl118 | $spl418)),
  inference(forward_demodulation,[],[f8178,f713])).
fof(f8178,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK19(happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12))),sK16,sK15,sK14,sK13,sK12)),s(fun(sK12,bool),sK15))) | ($spl118 | $spl418)),
  inference(superposition,[],[f6010,f6275])).
fof(f8134,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl500),
  inference(cnf_transformation,[],[f8134_D])).
fof(f8134_D,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','F_0') <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f8130,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl498),
  inference(resolution,[],[f8122,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP5(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP4(X2,X1,X0) & sP5(X1,X2,X0) & sP6(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X2,X1,X0) | ~sP5(X1,X2,X0) | ~sP6(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f243])).
fof(f243,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP7(X2,X1,X0)))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP7(X2,X1,X0)))),
  inference(nnf_transformation,[],[f173])).
fof(f8126,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl498),
  inference(unit_resulting_resolution,[],[f8122,f412])).
fof(f8122,plain,(
  sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl498),
  inference(cnf_transformation,[],[f8122_D])).
fof(f8122_D,plain,(
  sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f8083,plain,(
  sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f281,f281,f569,f6768])).
fof(f6768,plain,(
  ( ! [X0,X1] : (~sP4(X1,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0) | p(s(bool,X0)) | sP7(X0,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X1) | p(s(bool,X1))) ) | $spl2),
  inference(resolution,[],[f1067,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f250])).
fof(f1067,plain,(
  ( ! [X6,X7] : (~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X7,X6) | sP7(X6,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X7) | p(s(bool,X6)) | ~sP4(X7,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X6)) ) | $spl2),
  inference(resolution,[],[f560,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X2,X1) | p(s(bool,X0)) | sP7(X0,X1,X2) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f560,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f419])).
fof(f569,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f428])).
fof(f8081,plain,(
  ( ! [X169] : (p(s(bool,X169))) ) | $spl494),
  inference(cnf_transformation,[],[f8081_D])).
fof(f8081_D,plain,(
  ( ! [X169] : (p(s(bool,X169))) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f8080,plain,(
  ( ! [X154,X155] : (s(bool,X155) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X154))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X154))) | p(s(bool,happ(s(fun(bool,bool),X154),s(bool,'F_0')))) | ~p(s(bool,X155))) ) | $spl492),
  inference(cnf_transformation,[],[f8080_D])).
fof(f8080_D,plain,(
  ( ! [X154,X155] : (s(bool,X155) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X154))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X154))) | p(s(bool,happ(s(fun(bool,bool),X154),s(bool,'F_0')))) | ~p(s(bool,X155))) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f8050,plain,(
  ( ! [X189,X190,X188] : (s(bool,X189) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X188))) | ~p(s(bool,X189)) | p(s(bool,happ(s(fun(bool,bool),X188),s(bool,'F_0')))) | p(s(bool,X190)) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X188)))) ) | $spl430),
  inference(superposition,[],[f481,f6740])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X11,X18,X15,X19] : ((s(bool,X11) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X15) = s(bool,X19))) => ((p(s(bool,X11)) => p(s(bool,X15))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.bool.IMP_CONG')).
fof(f8079,plain,(
  ( ! [X151,X150] : (s(bool,X151) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X150))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X150))) | p(s(bool,X151)) | ~p(s(bool,happ(s(fun(bool,bool),X150),s(bool,'F_0'))))) ) | $spl490),
  inference(cnf_transformation,[],[f8079_D])).
fof(f8079_D,plain,(
  ( ! [X151,X150] : (s(bool,X151) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X150))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X150))) | p(s(bool,X151)) | ~p(s(bool,happ(s(fun(bool,bool),X150),s(bool,'F_0'))))) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f8049,plain,(
  ( ! [X187,X185,X186] : (s(bool,X186) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X185))) | ~p(s(bool,happ(s(fun(bool,bool),X185),s(bool,'F_0')))) | p(s(bool,X186)) | p(s(bool,X187)) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X185)))) ) | $spl430),
  inference(superposition,[],[f481,f6740])).
fof(f8078,plain,(
  ( ! [X152,X153] : (~p(s(bool,X152)) | s(bool,X152) != s(bool,X153)) ) | $spl488),
  inference(cnf_transformation,[],[f8078_D])).
fof(f8078_D,plain,(
  ( ! [X152,X153] : (~p(s(bool,X152)) | s(bool,X152) != s(bool,X153)) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f8047,plain,(
  ( ! [X180,X178,X181,X179] : (s(bool,X179) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X178))) | ~p(s(bool,X180)) | ~p(s(bool,happ(s(fun(bool,bool),X178),s(bool,'F_0')))) | p(s(bool,X179)) | s(bool,X180) != s(bool,X181) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X178)))) ) | $spl430),
  inference(superposition,[],[f475,f6740])).
fof(f475,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f271])).
fof(f8046,plain,(
  ( ! [X177,X175,X176,X174] : (s(bool,X175) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X174))) | ~p(s(bool,X176)) | ~p(s(bool,X175)) | p(s(bool,happ(s(fun(bool,bool),X174),s(bool,'F_0')))) | s(bool,X176) != s(bool,X177) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X174)))) ) | $spl430),
  inference(superposition,[],[f475,f6740])).
fof(f8082,plain,(
  ( ! [X168] : (~p(s(bool,X168))) ) | $spl496),
  inference(cnf_transformation,[],[f8082_D])).
fof(f8082_D,plain,(
  ( ! [X168] : (~p(s(bool,X168))) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f8045,plain,(
  ( ! [X173,X171,X172,X170] : (s(bool,X171) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X170))) | ~p(s(bool,X171)) | ~p(s(bool,X172)) | p(s(bool,happ(s(fun(bool,bool),X170),s(bool,'F_0')))) | p(s(bool,X173)) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X170)))) ) | $spl430),
  inference(superposition,[],[f474,f6740])).
fof(f474,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f271])).
fof(f8044,plain,(
  ( ! [X167,X169,X166,X168] : (s(bool,X167) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X166))) | ~p(s(bool,happ(s(fun(bool,bool),X166),s(bool,'F_0')))) | ~p(s(bool,X168)) | p(s(bool,X167)) | p(s(bool,X169)) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X166)))) ) | $spl430),
  inference(superposition,[],[f474,f6740])).
fof(f8041,plain,(
  ( ! [X156,X154,X157,X155] : (s(bool,X155) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X154))) | ~p(s(bool,X156)) | ~p(s(bool,X155)) | p(s(bool,happ(s(fun(bool,bool),X154),s(bool,'F_0')))) | s(bool,X156) != s(bool,X157) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X154)))) ) | $spl430),
  inference(superposition,[],[f471,f6740])).
fof(f471,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f271])).
fof(f8040,plain,(
  ( ! [X152,X151,X153,X150] : (s(bool,X151) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X150))) | ~p(s(bool,X152)) | ~p(s(bool,happ(s(fun(bool,bool),X150),s(bool,'F_0')))) | p(s(bool,X151)) | s(bool,X152) != s(bool,X153) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X150)))) ) | $spl430),
  inference(superposition,[],[f471,f6740])).
fof(f7555,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl486),
  inference(cnf_transformation,[],[f7555_D])).
fof(f7555_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f7536,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl298),
  inference(unit_resulting_resolution,[],[f276,f2883,f4991])).
fof(f7537,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl298),
  inference(unit_resulting_resolution,[],[f275,f2883,f4991])).
fof(f7531,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl485),
  inference(cnf_transformation,[],[f7531_D])).
fof(f7531_D,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl485),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl485])])).
fof(f7515,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl230 | $spl298)),
  inference(unit_resulting_resolution,[],[f281,f2883,f4773])).
fof(f7512,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl483),
  inference(cnf_transformation,[],[f7512_D])).
fof(f7512_D,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl483),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl483])])).
fof(f7488,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl230 | $spl298)),
  inference(unit_resulting_resolution,[],[f281,f2883,f4772])).
fof(f7484,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl480),
  inference(cnf_transformation,[],[f7484_D])).
fof(f7484_D,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f7465,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl230 | $spl298)),
  inference(unit_resulting_resolution,[],[f276,f2883,f4771])).
fof(f7466,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl230 | $spl298)),
  inference(unit_resulting_resolution,[],[f275,f2883,f4771])).
fof(f7460,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl478),
  inference(cnf_transformation,[],[f7460_D])).
fof(f7460_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f7439,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl230 | $spl298)),
  inference(unit_resulting_resolution,[],[f276,f2883,f4770])).
fof(f7440,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl230 | $spl298)),
  inference(unit_resulting_resolution,[],[f275,f2883,f4770])).
fof(f7458,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl476),
  inference(cnf_transformation,[],[f7458_D])).
fof(f7458_D,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f7441,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl230),
  inference(unit_resulting_resolution,[],[f276,f2509,f4770])).
fof(f7442,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl230),
  inference(unit_resulting_resolution,[],[f275,f2509,f4770])).
fof(f7420,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl475),
  inference(cnf_transformation,[],[f7420_D])).
fof(f7420_D,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl475),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl475])])).
fof(f7413,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | ($spl2 | $spl4 | $spl395)),
  inference(unit_resulting_resolution,[],[f281,f828,f5187,f1067])).
fof(f828,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f429])).
fof(f7417,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | ($spl2 | $spl4 | $spl395)),
  inference(unit_resulting_resolution,[],[f281,f828,f560,f5187,f417])).
fof(f7410,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl473),
  inference(cnf_transformation,[],[f7410_D])).
fof(f7410_D,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl473),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl473])])).
fof(f7404,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | ($spl2 | $spl120 | $spl393)),
  inference(unit_resulting_resolution,[],[f281,f2610,f5185,f1067])).
fof(f2610,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,X1)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f429])).
fof(f7407,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | ($spl2 | $spl120 | $spl393)),
  inference(unit_resulting_resolution,[],[f281,f2610,f560,f5185,f417])).
fof(f7397,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl471),
  inference(cnf_transformation,[],[f7397_D])).
fof(f7397_D,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl471),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl471])])).
fof(f7391,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | ($spl2 | $spl118 | $spl391)),
  inference(unit_resulting_resolution,[],[f281,f2234,f5183,f1067])).
fof(f2234,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,X1)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f429])).
fof(f7394,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | ($spl2 | $spl118 | $spl391)),
  inference(unit_resulting_resolution,[],[f281,f2234,f560,f5183,f417])).
fof(f7388,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl469),
  inference(cnf_transformation,[],[f7388_D])).
fof(f7388_D,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl469),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl469])])).
fof(f7381,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | ($spl2 | $spl4 | $spl389)),
  inference(unit_resulting_resolution,[],[f281,f570,f5165,f1789])).
fof(f1789,plain,(
  ( ! [X6,X7] : (~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X7,X6) | sP7(X6,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X7) | p(s(bool,X6)) | ~sP4(X7,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X6)) ) | $spl4),
  inference(resolution,[],[f818,f417])).
fof(f818,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f419])).
fof(f570,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f429])).
fof(f7385,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | ($spl2 | $spl4 | $spl389)),
  inference(unit_resulting_resolution,[],[f281,f570,f818,f5165,f417])).
fof(f7378,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl467),
  inference(cnf_transformation,[],[f7378_D])).
fof(f7378_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl467),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl467])])).
fof(f7375,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | ($spl2 | $spl120 | $spl387)),
  inference(unit_resulting_resolution,[],[f281,f570,f2596,f5163,f417])).
fof(f2596,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f419])).
fof(f7370,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl465),
  inference(cnf_transformation,[],[f7370_D])).
fof(f7370_D,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl465),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl465])])).
fof(f7367,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | ($spl2 | $spl118 | $spl385)),
  inference(unit_resulting_resolution,[],[f281,f570,f2220,f5161,f417])).
fof(f2220,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f419])).
fof(f7359,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl462),
  inference(cnf_transformation,[],[f7359_D])).
fof(f7359_D,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f7355,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl382),
  inference(resolution,[],[f5131,f412])).
fof(f7351,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl382),
  inference(unit_resulting_resolution,[],[f5131,f412])).
fof(f7341,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl460),
  inference(cnf_transformation,[],[f7341_D])).
fof(f7341_D,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f7337,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl380),
  inference(resolution,[],[f5129,f412])).
fof(f7333,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl380),
  inference(unit_resulting_resolution,[],[f5129,f412])).
fof(f7324,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl458),
  inference(cnf_transformation,[],[f7324_D])).
fof(f7324_D,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f7320,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl378),
  inference(resolution,[],[f5127,f412])).
fof(f7316,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl378),
  inference(unit_resulting_resolution,[],[f5127,f412])).
fof(f7307,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl456),
  inference(cnf_transformation,[],[f7307_D])).
fof(f7307_D,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f7303,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl376),
  inference(resolution,[],[f5103,f412])).
fof(f7299,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl376),
  inference(unit_resulting_resolution,[],[f5103,f412])).
fof(f7287,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl454),
  inference(cnf_transformation,[],[f7287_D])).
fof(f7287_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f7283,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl372),
  inference(resolution,[],[f5099,f412])).
fof(f7279,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f5099,f412])).
fof(f7270,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl452),
  inference(cnf_transformation,[],[f7270_D])).
fof(f7270_D,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f7266,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl370),
  inference(resolution,[],[f5097,f412])).
fof(f7262,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl370),
  inference(unit_resulting_resolution,[],[f5097,f412])).
fof(f7049,plain,(
  s(sK13,sK17) = s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12))),sK16,sK15,sK14,sK13,sK12)))) | $spl450),
  inference(cnf_transformation,[],[f7049_D])).
fof(f7049_D,plain,(
  s(sK13,sK17) = s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12))),sK16,sK15,sK14,sK13,sK12)))) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f7047,plain,(
  s(sK13,sK17) = s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12))),sK16,sK15,sK14,sK13,sK12)))) | ~p(s(bool,'T_0')) | ($spl56 | $spl118 | $spl418)),
  inference(forward_demodulation,[],[f7046,f6275])).
fof(f7046,plain,(
  ~p(s(bool,'T_0')) | s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)))) = s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12))),sK16,sK15,sK14,sK13,sK12)))) | ($spl56 | $spl118 | $spl418)),
  inference(forward_demodulation,[],[f7034,f713])).
fof(f7034,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) | s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)))) = s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12))),sK16,sK15,sK14,sK13,sK12)))) | ($spl118 | $spl418)),
  inference(superposition,[],[f2368,f6275])).
fof(f2368,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X2),s(fun(sK13,bool),sK16)))) | s(sK13,X2) = s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(X2,sK16,sK15,sK14,sK13,sK12))))) ) | $spl118),
  inference(resolution,[],[f976,f378])).
fof(f378,plain,(
  ( ! [X4,X2,X0,X8,X3,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK19(X8,X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f7032,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl448),
  inference(cnf_transformation,[],[f7032_D])).
fof(f7032_D,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f7028,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl374),
  inference(resolution,[],[f5101,f412])).
fof(f7024,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl374),
  inference(unit_resulting_resolution,[],[f5101,f412])).
fof(f7019,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl446),
  inference(cnf_transformation,[],[f7019_D])).
fof(f7019_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f7015,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl368),
  inference(resolution,[],[f5057,f412])).
fof(f7011,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl368),
  inference(unit_resulting_resolution,[],[f5057,f412])).
fof(f7004,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl442),
  inference(cnf_transformation,[],[f7004_D])).
fof(f7004_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f6999,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl366),
  inference(resolution,[],[f5042,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f244])).
fof(f7006,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl444),
  inference(cnf_transformation,[],[f7006_D])).
fof(f7006_D,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f6998,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl366),
  inference(resolution,[],[f5042,f412])).
fof(f6993,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl366),
  inference(unit_resulting_resolution,[],[f5042,f412])).
fof(f6994,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl366),
  inference(unit_resulting_resolution,[],[f5042,f413])).
fof(f6989,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl440),
  inference(cnf_transformation,[],[f6989_D])).
fof(f6989_D,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f6985,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl364),
  inference(resolution,[],[f4832,f412])).
fof(f6981,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl364),
  inference(unit_resulting_resolution,[],[f4832,f412])).
fof(f6974,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl436),
  inference(cnf_transformation,[],[f6974_D])).
fof(f6974_D,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f6969,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl362),
  inference(resolution,[],[f4819,f413])).
fof(f6976,plain,(
  sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl438),
  inference(cnf_transformation,[],[f6976_D])).
fof(f6976_D,plain,(
  sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f6968,plain,(
  sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl362),
  inference(resolution,[],[f4819,f412])).
fof(f6963,plain,(
  sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f4819,f412])).
fof(f6964,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f4819,f413])).
fof(f6751,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4)))) ) | $spl434),
  inference(cnf_transformation,[],[f6751_D])).
fof(f6751_D,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4)))) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f6750,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | ~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))))) ) | ($spl2 | $spl56)),
  inference(forward_demodulation,[],[f6747,f713])).
fof(f6747,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | ~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))))) ) | $spl2),
  inference(resolution,[],[f3387,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f244])).
fof(f3387,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),happ(s(fun(bool,bool),X0),s(bool,'F_0')),'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) ) | $spl2),
  inference(superposition,[],[f572,f282])).
fof(f282,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.BIJ_LINV_INV',aHL_BOOL_CASES)).
fof(f572,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f332,f485,f431])).
fof(f332,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X6,bool),X5))) = s(bool,happ(s(fun(X6,bool),X5),s(X6,'const.min.@_1'(s(fun(X6,bool),X5)))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.bool.EXISTS_DEF')).
fof(f6742,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))))) ) | $spl432),
  inference(cnf_transformation,[],[f6742_D])).
fof(f6742_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f6741,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4)))) | p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0'))))) ) | ($spl2 | $spl56)),
  inference(forward_demodulation,[],[f6736,f713])).
fof(f6736,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4)))) | p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))))) ) | $spl2),
  inference(resolution,[],[f3375,f414])).
fof(f3375,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'|dollar|exists_1'(s(fun(bool,bool),X0)),happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) ) | $spl2),
  inference(superposition,[],[f571,f282])).
fof(f571,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f332,f485,f431])).
fof(f6740,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) ) | $spl430),
  inference(cnf_transformation,[],[f6740_D])).
fof(f6740_D,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f6739,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) ) | ($spl2 | $spl56)),
  inference(forward_demodulation,[],[f6732,f713])).
fof(f6732,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) ) | $spl2),
  inference(resolution,[],[f3375,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f253])).
fof(f6518,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl429),
  inference(cnf_transformation,[],[f6518_D])).
fof(f6518_D,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl429),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl429])])).
fof(f6514,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl271),
  inference(unit_resulting_resolution,[],[f2853,f413])).
fof(f6506,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl427),
  inference(cnf_transformation,[],[f6506_D])).
fof(f6506_D,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl427),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl427])])).
fof(f6502,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl265),
  inference(unit_resulting_resolution,[],[f2843,f412])).
fof(f6408,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl425),
  inference(cnf_transformation,[],[f6408_D])).
fof(f6408_D,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl425),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl425])])).
fof(f6404,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl203),
  inference(unit_resulting_resolution,[],[f2479,f413])).
fof(f6396,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl423),
  inference(cnf_transformation,[],[f6396_D])).
fof(f6396_D,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl423),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl423])])).
fof(f6392,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl197),
  inference(unit_resulting_resolution,[],[f2469,f412])).
fof(f6278,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) | p(s(bool,'T_0'))),
  introduced(theory_axiom,[])).
fof(f6275,plain,(
  s(sK13,sK17) = s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)))) | $spl418),
  inference(cnf_transformation,[],[f6275_D])).
fof(f6275_D,plain,(
  s(sK13,sK17) = s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)))) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f6272,plain,(
  ~p(s(bool,'T_0')) | s(sK13,sK17) = s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK19(sK17,sK16,sK15,sK14,sK13,sK12)))) | ($spl56 | $spl118)),
  inference(superposition,[],[f2368,f713])).
fof(f6143,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl417),
  inference(cnf_transformation,[],[f6143_D])).
fof(f6143_D,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl417),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl417])])).
fof(f6126,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl298)),
  inference(unit_resulting_resolution,[],[f281,f2883,f3145])).
fof(f6141,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl415),
  inference(cnf_transformation,[],[f6141_D])).
fof(f6141_D,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl415),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl415])])).
fof(f6127,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl230)),
  inference(unit_resulting_resolution,[],[f281,f2509,f3145])).
fof(f6123,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl413),
  inference(cnf_transformation,[],[f6123_D])).
fof(f6123_D,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl413),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl413])])).
fof(f6106,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl298)),
  inference(unit_resulting_resolution,[],[f281,f2883,f3144])).
fof(f6121,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl411),
  inference(cnf_transformation,[],[f6121_D])).
fof(f6121_D,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl411),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl411])])).
fof(f6107,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl230)),
  inference(unit_resulting_resolution,[],[f281,f2509,f3144])).
fof(f6103,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl408),
  inference(cnf_transformation,[],[f6103_D])).
fof(f6103_D,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f6083,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl298)),
  inference(unit_resulting_resolution,[],[f276,f2883,f3143])).
fof(f6084,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl298)),
  inference(unit_resulting_resolution,[],[f275,f2883,f3143])).
fof(f6101,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl406),
  inference(cnf_transformation,[],[f6101_D])).
fof(f6101_D,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f6085,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl230)),
  inference(unit_resulting_resolution,[],[f276,f2509,f3143])).
fof(f6086,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl230)),
  inference(unit_resulting_resolution,[],[f275,f2509,f3143])).
fof(f6073,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl404),
  inference(cnf_transformation,[],[f6073_D])).
fof(f6073_D,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f6047,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl116),
  inference(unit_resulting_resolution,[],[f276,f974,f3142])).
fof(f6048,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl116),
  inference(unit_resulting_resolution,[],[f275,f974,f3142])).
fof(f6071,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl402),
  inference(cnf_transformation,[],[f6071_D])).
fof(f6071_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f6051,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl298)),
  inference(unit_resulting_resolution,[],[f276,f2883,f3142])).
fof(f6052,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl298)),
  inference(unit_resulting_resolution,[],[f275,f2883,f3142])).
fof(f6069,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl400),
  inference(cnf_transformation,[],[f6069_D])).
fof(f6069_D,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f6053,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl230)),
  inference(unit_resulting_resolution,[],[f276,f2509,f3142])).
fof(f6054,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl116 | $spl230)),
  inference(unit_resulting_resolution,[],[f275,f2509,f3142])).
fof(f5651,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl398),
  inference(cnf_transformation,[],[f5651_D])).
fof(f5651_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f5648,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl262),
  inference(resolution,[],[f2839,f413])).
fof(f5643,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl262),
  inference(unit_resulting_resolution,[],[f2839,f413])).
fof(f5588,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl396),
  inference(cnf_transformation,[],[f5588_D])).
fof(f5588_D,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f5585,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl194),
  inference(resolution,[],[f2465,f413])).
fof(f5580,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl194),
  inference(unit_resulting_resolution,[],[f2465,f413])).
fof(f5187,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl395),
  inference(cnf_transformation,[],[f5187_D])).
fof(f5187_D,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl395),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl395])])).
fof(f5166,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl116)),
  inference(unit_resulting_resolution,[],[f281,f974,f1542])).
fof(f5185,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl393),
  inference(cnf_transformation,[],[f5185_D])).
fof(f5185_D,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl393),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl393])])).
fof(f5168,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl298)),
  inference(unit_resulting_resolution,[],[f281,f2883,f1542])).
fof(f5183,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl391),
  inference(cnf_transformation,[],[f5183_D])).
fof(f5183_D,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl391),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl391])])).
fof(f5169,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl230)),
  inference(unit_resulting_resolution,[],[f281,f2509,f1542])).
fof(f5165,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl389),
  inference(cnf_transformation,[],[f5165_D])).
fof(f5165_D,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl389),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl389])])).
fof(f5144,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl116)),
  inference(unit_resulting_resolution,[],[f281,f974,f1541])).
fof(f5163,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl387),
  inference(cnf_transformation,[],[f5163_D])).
fof(f5163_D,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl387),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl387])])).
fof(f5146,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl298)),
  inference(unit_resulting_resolution,[],[f281,f2883,f1541])).
fof(f5161,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl385),
  inference(cnf_transformation,[],[f5161_D])).
fof(f5161_D,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl385),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl385])])).
fof(f5147,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl230)),
  inference(unit_resulting_resolution,[],[f281,f2509,f1541])).
fof(f5131,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl382),
  inference(cnf_transformation,[],[f5131_D])).
fof(f5131_D,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f5104,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl116)),
  inference(unit_resulting_resolution,[],[f276,f974,f1540])).
fof(f5105,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl116)),
  inference(unit_resulting_resolution,[],[f275,f974,f1540])).
fof(f5129,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl380),
  inference(cnf_transformation,[],[f5129_D])).
fof(f5129_D,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f5108,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl298)),
  inference(unit_resulting_resolution,[],[f276,f2883,f1540])).
fof(f5109,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl298)),
  inference(unit_resulting_resolution,[],[f275,f2883,f1540])).
fof(f5127,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl378),
  inference(cnf_transformation,[],[f5127_D])).
fof(f5127_D,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f5110,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl230)),
  inference(unit_resulting_resolution,[],[f276,f2509,f1540])).
fof(f5111,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl230)),
  inference(unit_resulting_resolution,[],[f275,f2509,f1540])).
fof(f5103,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl376),
  inference(cnf_transformation,[],[f5103_D])).
fof(f5103_D,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f5074,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl116)),
  inference(unit_resulting_resolution,[],[f276,f974,f1539])).
fof(f5075,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl116)),
  inference(unit_resulting_resolution,[],[f275,f974,f1539])).
fof(f5101,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl374),
  inference(cnf_transformation,[],[f5101_D])).
fof(f5101_D,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f5076,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl56),
  inference(unit_resulting_resolution,[],[f276,f713,f1539])).
fof(f5077,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl56),
  inference(unit_resulting_resolution,[],[f275,f713,f1539])).
fof(f5099,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl372),
  inference(cnf_transformation,[],[f5099_D])).
fof(f5099_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f5078,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl298)),
  inference(unit_resulting_resolution,[],[f276,f2883,f1539])).
fof(f5079,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl298)),
  inference(unit_resulting_resolution,[],[f275,f2883,f1539])).
fof(f5097,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl370),
  inference(cnf_transformation,[],[f5097_D])).
fof(f5097_D,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f5080,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl230)),
  inference(unit_resulting_resolution,[],[f276,f2509,f1539])).
fof(f5081,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl230)),
  inference(unit_resulting_resolution,[],[f275,f2509,f1539])).
fof(f5057,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl368),
  inference(cnf_transformation,[],[f5057_D])).
fof(f5057_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f4908,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl298),
  inference(unit_resulting_resolution,[],[f276,f2883,f431])).
fof(f4909,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl298),
  inference(unit_resulting_resolution,[],[f275,f2883,f431])).
fof(f5042,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl366),
  inference(cnf_transformation,[],[f5042_D])).
fof(f5042_D,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f4928,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl298),
  inference(unit_resulting_resolution,[],[f276,f2883,f431])).
fof(f4929,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl298),
  inference(unit_resulting_resolution,[],[f275,f2883,f431])).
fof(f4832,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl364),
  inference(cnf_transformation,[],[f4832_D])).
fof(f4832_D,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f4689,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl230),
  inference(unit_resulting_resolution,[],[f276,f2509,f431])).
fof(f4690,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl230),
  inference(unit_resulting_resolution,[],[f275,f2509,f431])).
fof(f4819,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl362),
  inference(cnf_transformation,[],[f4819_D])).
fof(f4819_D,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f4707,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl230),
  inference(unit_resulting_resolution,[],[f276,f2509,f431])).
fof(f4708,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl230),
  inference(unit_resulting_resolution,[],[f275,f2509,f431])).
fof(f4619,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl360),
  inference(cnf_transformation,[],[f4619_D])).
fof(f4619_D,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f4615,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl308),
  inference(resolution,[],[f3200,f412])).
fof(f4611,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl308),
  inference(unit_resulting_resolution,[],[f3200,f412])).
fof(f4604,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl356),
  inference(cnf_transformation,[],[f4604_D])).
fof(f4604_D,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f4599,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl306),
  inference(resolution,[],[f3189,f413])).
fof(f4606,plain,(
  sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl358),
  inference(cnf_transformation,[],[f4606_D])).
fof(f4606_D,plain,(
  sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f4598,plain,(
  sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl306),
  inference(resolution,[],[f3189,f412])).
fof(f4593,plain,(
  sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl306),
  inference(unit_resulting_resolution,[],[f3189,f412])).
fof(f4594,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl306),
  inference(unit_resulting_resolution,[],[f3189,f413])).
fof(f4576,plain,(
  ~sP7('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl355),
  inference(cnf_transformation,[],[f4576_D])).
fof(f4576_D,plain,(
  ~sP7('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl355),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl355])])).
fof(f4572,plain,(
  ~sP7('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl273),
  inference(unit_resulting_resolution,[],[f2855,f413])).
fof(f4566,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl353),
  inference(cnf_transformation,[],[f4566_D])).
fof(f4566_D,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl353),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl353])])).
fof(f4562,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl269),
  inference(unit_resulting_resolution,[],[f2849,f412])).
fof(f4556,plain,(
  ~sP7('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl351),
  inference(cnf_transformation,[],[f4556_D])).
fof(f4556_D,plain,(
  ~sP7('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl351),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl351])])).
fof(f4552,plain,(
  ~sP7('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl267),
  inference(unit_resulting_resolution,[],[f2845,f412])).
fof(f4505,plain,(
  sP10('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl348),
  inference(cnf_transformation,[],[f4505_D])).
fof(f4505_D,plain,(
  sP10('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f4503,plain,(
  sP10('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl244),
  inference(resolution,[],[f2801,f453])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f263])).
fof(f263,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f179])).
fof(f4497,plain,(
  sP10('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl244),
  inference(unit_resulting_resolution,[],[f2801,f453])).
fof(f4252,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl346),
  inference(cnf_transformation,[],[f4252_D])).
fof(f4252_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f4246,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl120),
  inference(unit_resulting_resolution,[],[f276,f2596,f2607,f2608,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X2,X1) | ~p(s(bool,X2)) | sP7(X0,X1,X2) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f2608,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f427])).
fof(f2607,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f423])).
fof(f4247,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl120),
  inference(unit_resulting_resolution,[],[f275,f2596,f2607,f2608,f415])).
fof(f4237,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl344),
  inference(cnf_transformation,[],[f4237_D])).
fof(f4237_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f4227,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl4 | $spl120)),
  inference(unit_resulting_resolution,[],[f276,f827,f818,f2607,f415])).
fof(f827,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f428])).
fof(f4228,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl4 | $spl120)),
  inference(unit_resulting_resolution,[],[f275,f827,f818,f2607,f415])).
fof(f4235,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl342),
  inference(cnf_transformation,[],[f4235_D])).
fof(f4235_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f4231,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl118 | $spl120)),
  inference(unit_resulting_resolution,[],[f276,f2233,f2220,f2607,f415])).
fof(f2233,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X1)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f428])).
fof(f4232,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl118 | $spl120)),
  inference(unit_resulting_resolution,[],[f275,f2233,f2220,f2607,f415])).
fof(f4219,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl340),
  inference(cnf_transformation,[],[f4219_D])).
fof(f4219_D,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f4204,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl2 | $spl120)),
  inference(unit_resulting_resolution,[],[f276,f568,f567,f2596,f415])).
fof(f567,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f423])).
fof(f568,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f427])).
fof(f4205,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl2 | $spl120)),
  inference(unit_resulting_resolution,[],[f275,f568,f567,f2596,f415])).
fof(f4217,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl338),
  inference(cnf_transformation,[],[f4217_D])).
fof(f4217_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f4206,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl4 | $spl120)),
  inference(unit_resulting_resolution,[],[f276,f826,f825,f2596,f415])).
fof(f825,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f423])).
fof(f826,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f427])).
fof(f4207,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl4 | $spl120)),
  inference(unit_resulting_resolution,[],[f275,f826,f825,f2596,f415])).
fof(f4215,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl336),
  inference(cnf_transformation,[],[f4215_D])).
fof(f4215_D,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f4208,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl118 | $spl120)),
  inference(unit_resulting_resolution,[],[f276,f2232,f2231,f2596,f415])).
fof(f2231,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f423])).
fof(f2232,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f427])).
fof(f4209,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl118 | $spl120)),
  inference(unit_resulting_resolution,[],[f275,f2232,f2231,f2596,f415])).
fof(f4152,plain,(
  ~sP7('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl335),
  inference(cnf_transformation,[],[f4152_D])).
fof(f4152_D,plain,(
  ~sP7('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl335),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl335])])).
fof(f4148,plain,(
  ~sP7('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl205),
  inference(unit_resulting_resolution,[],[f2481,f413])).
fof(f4142,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl333),
  inference(cnf_transformation,[],[f4142_D])).
fof(f4142_D,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl333),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl333])])).
fof(f4138,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl201),
  inference(unit_resulting_resolution,[],[f2475,f412])).
fof(f4132,plain,(
  ~sP7('F_0','T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl331),
  inference(cnf_transformation,[],[f4132_D])).
fof(f4132_D,plain,(
  ~sP7('F_0','T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl331),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl331])])).
fof(f4128,plain,(
  ~sP7('F_0','T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl199),
  inference(unit_resulting_resolution,[],[f2471,f412])).
fof(f4081,plain,(
  sP10('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl328),
  inference(cnf_transformation,[],[f4081_D])).
fof(f4081_D,plain,(
  sP10('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f4079,plain,(
  sP10('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl176),
  inference(resolution,[],[f2427,f453])).
fof(f4073,plain,(
  sP10('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl176),
  inference(unit_resulting_resolution,[],[f2427,f453])).
fof(f3832,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl326),
  inference(cnf_transformation,[],[f3832_D])).
fof(f3832_D,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f3828,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl118),
  inference(unit_resulting_resolution,[],[f276,f2220,f2231,f2232,f415])).
fof(f3829,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl118),
  inference(unit_resulting_resolution,[],[f275,f2220,f2231,f2232,f415])).
fof(f3819,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl324),
  inference(cnf_transformation,[],[f3819_D])).
fof(f3819_D,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f3815,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl4 | $spl118)),
  inference(unit_resulting_resolution,[],[f276,f827,f818,f2231,f415])).
fof(f3816,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl4 | $spl118)),
  inference(unit_resulting_resolution,[],[f275,f827,f818,f2231,f415])).
fof(f3805,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl322),
  inference(cnf_transformation,[],[f3805_D])).
fof(f3805_D,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f3794,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl2 | $spl118)),
  inference(unit_resulting_resolution,[],[f276,f568,f567,f2220,f415])).
fof(f3795,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl2 | $spl118)),
  inference(unit_resulting_resolution,[],[f275,f568,f567,f2220,f415])).
fof(f3803,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl320),
  inference(cnf_transformation,[],[f3803_D])).
fof(f3803_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f3796,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl4 | $spl118)),
  inference(unit_resulting_resolution,[],[f276,f826,f825,f2220,f415])).
fof(f3797,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl4 | $spl118)),
  inference(unit_resulting_resolution,[],[f275,f826,f825,f2220,f415])).
fof(f3578,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl318),
  inference(cnf_transformation,[],[f3578_D])).
fof(f3578_D,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','T_0') <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f3574,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl148),
  inference(resolution,[],[f1593,f412])).
fof(f3570,plain,(
  sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','T_0') | $spl148),
  inference(unit_resulting_resolution,[],[f1593,f412])).
fof(f3563,plain,(
  sP6('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl314),
  inference(cnf_transformation,[],[f3563_D])).
fof(f3563_D,plain,(
  sP6('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f3558,plain,(
  sP6('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl146),
  inference(resolution,[],[f1584,f413])).
fof(f3565,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl316),
  inference(cnf_transformation,[],[f3565_D])).
fof(f3565_D,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f3557,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl146),
  inference(resolution,[],[f1584,f412])).
fof(f3552,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f1584,f412])).
fof(f3553,plain,(
  sP6('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f1584,f413])).
fof(f3497,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl313),
  inference(cnf_transformation,[],[f3497_D])).
fof(f3497_D,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl313),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl313])])).
fof(f3493,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl97),
  inference(unit_resulting_resolution,[],[f954,f413])).
fof(f3485,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl311),
  inference(cnf_transformation,[],[f3485_D])).
fof(f3485_D,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl311),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl311])])).
fof(f3481,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl91),
  inference(unit_resulting_resolution,[],[f948,f412])).
fof(f3200,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl308),
  inference(cnf_transformation,[],[f3200_D])).
fof(f3200_D,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f3068,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl116),
  inference(unit_resulting_resolution,[],[f276,f974,f431])).
fof(f3069,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl116),
  inference(unit_resulting_resolution,[],[f275,f974,f431])).
fof(f3189,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl306),
  inference(cnf_transformation,[],[f3189_D])).
fof(f3189_D,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f3084,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl116),
  inference(unit_resulting_resolution,[],[f276,f974,f431])).
fof(f3085,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl116),
  inference(unit_resulting_resolution,[],[f275,f974,f431])).
fof(f3010,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl305),
  inference(cnf_transformation,[],[f3010_D])).
fof(f3010_D,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl305),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl305])])).
fof(f3006,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl35),
  inference(unit_resulting_resolution,[],[f691,f413])).
fof(f2998,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl303),
  inference(cnf_transformation,[],[f2998_D])).
fof(f2998_D,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl303),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl303])])).
fof(f2994,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl29),
  inference(unit_resulting_resolution,[],[f685,f412])).
fof(f2883,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl298),
  inference(cnf_transformation,[],[f2883_D])).
fof(f2883_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f2886,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(global_subsumption,[],[f274,f273,f272,f978,f2734,f2727,f2723,f2720,f2719,f2718,f2713,f2712,f2711,f2710,f2707,f2706,f2705,f2704,f2703,f2700,f2698,f2695,f2694,f2691,f2690,f2689,f2686,f2685,f2684,f2660,f2655,f2654,f2653,f2652,f2647,f2645,f2644,f2639,f2637,f2636,f2631,f2613,f2611,f2610,f2609,f2608,f2607,f2604,f2603,f2602,f2599,f2598,f2597,f2596,f2593,f2592,f2591,f2588,f2587,f2586,f2581,f2580,f2579,f2578,f2577,f2576,f2574,f2573,f2570,f2567,f2561,f2550,f2549,f2548,f2547,f2546,f2545,f2544,f2543,f2542,f2541,f2540,f2539,f2538,f2537,f2536,f2535,f2534,f2533,f2529,f2528,f2526,f2525,f2519,f2518,f2517,f2516,f2515,f2742,f2743,f2744,f2745,f2750,f2755,f2756,f2757,f2760,f2761,f2764,f2773,f2774,f2776,f2777,f2778,f2780])).
fof(f2780,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(superposition,[],[f978,f282])).
fof(f2778,plain,(
  ( ! [X47,X46] : (sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X46,X47) | ~p(s(bool,X46))) ) | $spl120),
  inference(resolution,[],[f978,f464])).
fof(f2777,plain,(
  ( ! [X45,X44] : (sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X45,X44) | p(s(bool,X44))) ) | $spl120),
  inference(resolution,[],[f978,f463])).
fof(f2776,plain,(
  ( ! [X43,X42] : (sP11(X43,X42,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | p(s(bool,X42)) | p(s(bool,X43))) ) | $spl120),
  inference(resolution,[],[f978,f462])).
fof(f2774,plain,(
  ( ! [X39,X38] : (sP9('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X38,X39) | ~p(s(bool,X38))) ) | $spl120),
  inference(resolution,[],[f978,f447])).
fof(f2773,plain,(
  ( ! [X37,X36] : (sP9('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X37,X36) | ~p(s(bool,X36))) ) | $spl120),
  inference(resolution,[],[f978,f446])).
fof(f2764,plain,(
  ( ! [X19,X18] : (sP3(X18,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X19) | ~p(s(bool,X19))) ) | $spl120),
  inference(resolution,[],[f978,f401])).
fof(f2761,plain,(
  ( ! [X12,X13] : (sP3(X13,X12,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ~p(s(bool,X12))) ) | $spl120),
  inference(resolution,[],[f978,f397])).
fof(f2760,plain,(
  ( ! [X11] : (~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X11) | ~p(s(bool,X11))) ) | $spl120),
  inference(resolution,[],[f978,f369])).
fof(f2757,plain,(
  ( ! [X8] : (sP1(X8,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | p(s(bool,X8))) ) | $spl120),
  inference(resolution,[],[f978,f366])).
fof(f2756,plain,(
  ( ! [X7] : (sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X7) | p(s(bool,X7))) ) | $spl120),
  inference(resolution,[],[f978,f363])).
fof(f2755,plain,(
  ( ! [X6] : (~sP1(X6,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ~p(s(bool,X6))) ) | $spl120),
  inference(resolution,[],[f978,f362])).
fof(f2750,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(resolution,[],[f978,f341])).
fof(f341,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.pred_set.BIJ_LINV_INV','thm.bool.IMP_ANTISYM_AX')).
fof(f2533,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f352])).
fof(f2534,plain,(
  ( ! [X0] : (sP0('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f353])).
fof(f2561,plain,(
  ( ! [X0] : (sP3('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f237])).
fof(f2567,plain,(
  ( ! [X0] : (sP3('T_0',X0,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f2570,plain,(
  ( ! [X0] : (sP3(X0,'T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f401])).
fof(f2573,plain,(
  ( ! [X0] : (sP3(X0,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0')) ) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f401])).
fof(f2574,plain,(
  ( ! [X0] : (sP3(X0,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f978,f401])).
fof(f2609,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X1)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f428])).
fof(f2611,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f332,f978,f431])).
fof(f2613,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f332,f978,f431])).
fof(f2636,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f442])).
fof(f2637,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X1)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f443])).
fof(f2639,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f446])).
fof(f2644,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,'T_0')) ) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f446])).
fof(f2645,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f978,f446])).
fof(f2647,plain,(
  ( ! [X0] : (sP9('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f447])).
fof(f2652,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0',X0)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f447])).
fof(f2653,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f978,f447])).
fof(f2655,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f449])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f261])).
fof(f2660,plain,(
  ( ! [X0] : (~sP9('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f450])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f261])).
fof(f2694,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X1)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f460])).
fof(f2698,plain,(
  ( ! [X0] : (sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,'F_0')) ) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f463])).
fof(f2700,plain,(
  ( ! [X0] : (sP11('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f464])).
fof(f2703,plain,(
  ( ! [X0] : (sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0',X0)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f464])).
fof(f2704,plain,(
  ( ! [X0] : (sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f978,f464])).
fof(f2713,plain,(
  ( ! [X0] : (~sP11('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f467])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f2727,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,X1)) ) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f478])).
fof(f2745,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(resolution,[],[f978,f292])).
fof(f292,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.BIJ_LINV_INV','thm.bool.EQ_CLAUSES|split|1')).
fof(f2744,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(resolution,[],[f978,f290])).
fof(f2885,plain,(
  sP2(sK12,sK13,sK14,sK15,sK16) | $spl300),
  inference(cnf_transformation,[],[f2885_D])).
fof(f2885_D,plain,(
  sP2(sK12,sK13,sK14,sK15,sK16) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f2742,plain,(
  sP2(sK12,sK13,sK14,sK15,sK16) | $spl120),
  inference(resolution,[],[f978,f393])).
fof(f393,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | sP2(X1,X0,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | sP2(X1,X0,X2,X3,X4)) & (~sP2(X1,X0,X2,X3,X4) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> sP2(X1,X0,X2,X3,X4))),
  inference(definition_folding,[],[f156,f166])).
fof(f156,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) | (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X7,X6] : ! [X8,X9,X10] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X10)))) <=> (! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X9)))) => p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,X11))),s(fun(X7,bool),X10))))) & ! [X11,X15] : ((p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X9)))) & p(s(bool,'const.bool.IN_2'(s(X6,X15),s(fun(X6,bool),X9))))) => (s(X7,happ(s(fun(X6,X7),X8),s(X6,X11))) = s(X7,happ(s(fun(X6,X7),X8),s(X6,X15))) => s(X6,X11) = s(X6,X15)))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.pred_set.INJ_DEF')).
fof(f2515,plain,(
  sP2(sK12,sK13,sK14,sK15,sK16) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f393])).
fof(f2516,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f290])).
fof(f2517,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f292])).
fof(f2786,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl237),
  inference(cnf_transformation,[],[f2786_D])).
fof(f2786_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f2518,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f293])).
fof(f293,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.BIJ_LINV_INV','thm.bool.EQ_CLAUSES|split|2')).
fof(f2519,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f295])).
fof(f295,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.BIJ_LINV_INV','thm.bool.EQ_CLAUSES|split|3')).
fof(f2525,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f276,f978,f341])).
fof(f2526,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f341])).
fof(f2528,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f276,f978,f341])).
fof(f2529,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f341])).
fof(f2875,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl291),
  inference(cnf_transformation,[],[f2875_D])).
fof(f2875_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl291),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl291])])).
fof(f2535,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl120),
  inference(unit_resulting_resolution,[],[f276,f978,f362])).
fof(f2536,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f362])).
fof(f2873,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl289),
  inference(cnf_transformation,[],[f2873_D])).
fof(f2873_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl289),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl289])])).
fof(f2537,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f978,f362])).
fof(f2877,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl293),
  inference(cnf_transformation,[],[f2877_D])).
fof(f2877_D,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl293),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl293])])).
fof(f2538,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f276,f978,f362])).
fof(f2539,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f362])).
fof(f2540,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f978,f362])).
fof(f2879,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl294),
  inference(cnf_transformation,[],[f2879_D])).
fof(f2879_D,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f2541,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f363])).
fof(f2881,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl296),
  inference(cnf_transformation,[],[f2881_D])).
fof(f2881_D,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f2542,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f366])).
fof(f2543,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f367])).
fof(f2544,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f368])).
fof(f2545,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f276,f978,f369])).
fof(f2546,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f369])).
fof(f2547,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f978,f369])).
fof(f2548,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl120),
  inference(unit_resulting_resolution,[],[f276,f978,f369])).
fof(f2549,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl120),
  inference(unit_resulting_resolution,[],[f275,f978,f369])).
fof(f2550,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f978,f369])).
fof(f2835,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl259),
  inference(cnf_transformation,[],[f2835_D])).
fof(f2835_D,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl259),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl259])])).
fof(f2576,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f276,f281,f978,f414])).
fof(f2577,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f275,f281,f978,f414])).
fof(f2869,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl285),
  inference(cnf_transformation,[],[f2869_D])).
fof(f2869_D,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl285),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl285])])).
fof(f2578,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f978,f281,f978,f414])).
fof(f2871,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl287),
  inference(cnf_transformation,[],[f2871_D])).
fof(f2871_D,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl287),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl287])])).
fof(f2579,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f276,f281,f978,f414])).
fof(f2580,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f275,f281,f978,f414])).
fof(f2581,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f978,f281,f978,f414])).
fof(f2867,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl282),
  inference(cnf_transformation,[],[f2867_D])).
fof(f2867_D,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f2582,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl2 | $spl120)),
  inference(unit_resulting_resolution,[],[f569,f567,f560,f978,f415])).
fof(f2865,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl280),
  inference(cnf_transformation,[],[f2865_D])).
fof(f2865_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f2583,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl2 | $spl4 | $spl120)),
  inference(unit_resulting_resolution,[],[f569,f825,f560,f978,f415])).
fof(f2863,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl278),
  inference(cnf_transformation,[],[f2863_D])).
fof(f2863_D,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f2584,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl2 | $spl4 | $spl120)),
  inference(unit_resulting_resolution,[],[f827,f567,f818,f978,f415])).
fof(f2861,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl276),
  inference(cnf_transformation,[],[f2861_D])).
fof(f2861_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f2585,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl4 | $spl120)),
  inference(unit_resulting_resolution,[],[f827,f825,f818,f978,f415])).
fof(f2859,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl275),
  inference(cnf_transformation,[],[f2859_D])).
fof(f2859_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl275),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl275])])).
fof(f2586,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f276,f281,f978,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f247])).
fof(f2587,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f275,f281,f978,f418])).
fof(f2853,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl271),
  inference(cnf_transformation,[],[f2853_D])).
fof(f2853_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl271),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl271])])).
fof(f2588,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f978,f281,f978,f418])).
fof(f2855,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl273),
  inference(cnf_transformation,[],[f2855_D])).
fof(f2855_D,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl273),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl273])])).
fof(f2591,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f276,f281,f978,f418])).
fof(f2592,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f275,f281,f978,f418])).
fof(f2593,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f978,f281,f978,f418])).
fof(f2849,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl269),
  inference(cnf_transformation,[],[f2849_D])).
fof(f2849_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl269),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl269])])).
fof(f2597,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f276,f281,f978,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f250])).
fof(f2598,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f275,f281,f978,f422])).
fof(f2843,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl265),
  inference(cnf_transformation,[],[f2843_D])).
fof(f2843_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl265),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl265])])).
fof(f2599,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f978,f281,f978,f422])).
fof(f2845,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl267),
  inference(cnf_transformation,[],[f2845_D])).
fof(f2845_D,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl267),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl267])])).
fof(f2602,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f276,f281,f978,f422])).
fof(f2603,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f275,f281,f978,f422])).
fof(f2604,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f978,f281,f978,f422])).
fof(f2839,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl262),
  inference(cnf_transformation,[],[f2839_D])).
fof(f2839_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f2612,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl120)),
  inference(unit_resulting_resolution,[],[f713,f978,f431])).
fof(f2837,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl260),
  inference(cnf_transformation,[],[f2837_D])).
fof(f2837_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f2614,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | ($spl56 | $spl120)),
  inference(unit_resulting_resolution,[],[f713,f978,f431])).
fof(f2615,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | ($spl9 | $spl120)),
  inference(unit_resulting_resolution,[],[f665,f978,f432])).
fof(f2833,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl257),
  inference(cnf_transformation,[],[f2833_D])).
fof(f2833_D,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f2616,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | ($spl9 | $spl120)),
  inference(unit_resulting_resolution,[],[f665,f978,f432])).
fof(f2827,plain,(
  ~sP8('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl255),
  inference(cnf_transformation,[],[f2827_D])).
fof(f2827_D,plain,(
  ~sP8('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f2631,plain,(
  ~sP8('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f281,f281,f978,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f259])).
fof(f2821,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl253),
  inference(cnf_transformation,[],[f2821_D])).
fof(f2821_D,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') <=> ~$spl253),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl253])])).
fof(f2654,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl120),
  inference(unit_resulting_resolution,[],[f281,f281,f978,f448])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f261])).
fof(f2811,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl251),
  inference(cnf_transformation,[],[f2811_D])).
fof(f2811_D,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl251),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl251])])).
fof(f2684,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl120),
  inference(unit_resulting_resolution,[],[f281,f276,f978,f458])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f267])).
fof(f2685,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl120),
  inference(unit_resulting_resolution,[],[f281,f275,f978,f458])).
fof(f2805,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl247),
  inference(cnf_transformation,[],[f2805_D])).
fof(f2805_D,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl247),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl247])])).
fof(f2686,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f978,f458])).
fof(f2807,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl249),
  inference(cnf_transformation,[],[f2807_D])).
fof(f2807_D,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl249),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl249])])).
fof(f2689,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f276,f281,f978,f458])).
fof(f2690,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f275,f281,f978,f458])).
fof(f2691,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f978,f281,f978,f458])).
fof(f2801,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl244),
  inference(cnf_transformation,[],[f2801_D])).
fof(f2801_D,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f2695,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f281,f978,f462])).
fof(f2797,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl243),
  inference(cnf_transformation,[],[f2797_D])).
fof(f2797_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f2705,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f276,f978,f465])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f269])).
fof(f2706,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f275,f978,f465])).
fof(f2791,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl239),
  inference(cnf_transformation,[],[f2791_D])).
fof(f2791_D,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl239),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl239])])).
fof(f2707,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f978,f465])).
fof(f2793,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl241),
  inference(cnf_transformation,[],[f2793_D])).
fof(f2793_D,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl241),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl241])])).
fof(f2710,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl120),
  inference(unit_resulting_resolution,[],[f281,f276,f978,f465])).
fof(f2711,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl120),
  inference(unit_resulting_resolution,[],[f281,f275,f978,f465])).
fof(f2712,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f978,f978,f465])).
fof(f2716,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl120)),
  inference(unit_resulting_resolution,[],[f281,f276,f713,f978,f471])).
fof(f2717,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl120)),
  inference(unit_resulting_resolution,[],[f281,f275,f713,f978,f471])).
fof(f2718,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f281,f276,f978,f474])).
fof(f2719,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f281,f275,f978,f474])).
fof(f2720,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f281,f978,f978,f474])).
fof(f2723,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f281,f978,f978,f474])).
fof(f2725,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl120)),
  inference(unit_resulting_resolution,[],[f281,f276,f713,f978,f475])).
fof(f2726,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl120)),
  inference(unit_resulting_resolution,[],[f281,f275,f713,f978,f475])).
fof(f2734,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl120),
  inference(unit_resulting_resolution,[],[f281,f281,f978,f481])).
fof(f2784,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl235),
  inference(cnf_transformation,[],[f2784_D])).
fof(f2784_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f2735,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl2 | $spl120)),
  inference(unit_resulting_resolution,[],[f978,f636])).
fof(f636,plain,(
  ( ! [X8] : (~sP1(X8,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ~p(s(bool,X8))) ) | $spl2),
  inference(resolution,[],[f485,f362])).
fof(f2782,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl233),
  inference(cnf_transformation,[],[f2782_D])).
fof(f2782_D,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f2736,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl2 | $spl120)),
  inference(unit_resulting_resolution,[],[f978,f641])).
fof(f641,plain,(
  ( ! [X13] : (~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X13) | ~p(s(bool,X13))) ) | $spl2),
  inference(resolution,[],[f485,f369])).
fof(f2509,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl230),
  inference(cnf_transformation,[],[f2509_D])).
fof(f2509_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f2511,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(global_subsumption,[],[f274,f273,f272,f976,f2358,f2351,f2347,f2344,f2343,f2342,f2337,f2336,f2335,f2334,f2331,f2330,f2329,f2328,f2327,f2324,f2322,f2319,f2318,f2315,f2314,f2313,f2310,f2309,f2308,f2284,f2279,f2278,f2277,f2276,f2271,f2269,f2268,f2263,f2261,f2260,f2255,f2237,f2235,f2234,f2233,f2232,f2231,f2228,f2227,f2226,f2223,f2222,f2221,f2220,f2217,f2216,f2215,f2212,f2211,f2210,f2205,f2204,f2203,f2202,f2201,f2200,f2198,f2197,f2194,f2191,f2185,f2174,f2173,f2172,f2171,f2170,f2169,f2168,f2167,f2166,f2165,f2164,f2163,f2162,f2161,f2160,f2159,f2158,f2157,f2153,f2152,f2150,f2149,f2143,f2142,f2141,f2140,f2366,f2367,f2368,f2370,f2371,f2376,f2381,f2382,f2383,f2386,f2387,f2390,f2399,f2400,f2402,f2403,f2404,f2406])).
fof(f2406,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(superposition,[],[f976,f282])).
fof(f2404,plain,(
  ( ! [X48,X49] : (sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X48,X49) | ~p(s(bool,X48))) ) | $spl118),
  inference(resolution,[],[f976,f464])).
fof(f2403,plain,(
  ( ! [X47,X46] : (sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X47,X46) | p(s(bool,X46))) ) | $spl118),
  inference(resolution,[],[f976,f463])).
fof(f2402,plain,(
  ( ! [X45,X44] : (sP11(X45,X44,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | p(s(bool,X44)) | p(s(bool,X45))) ) | $spl118),
  inference(resolution,[],[f976,f462])).
fof(f2400,plain,(
  ( ! [X41,X40] : (sP9('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X40,X41) | ~p(s(bool,X40))) ) | $spl118),
  inference(resolution,[],[f976,f447])).
fof(f2399,plain,(
  ( ! [X39,X38] : (sP9('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X39,X38) | ~p(s(bool,X38))) ) | $spl118),
  inference(resolution,[],[f976,f446])).
fof(f2390,plain,(
  ( ! [X21,X20] : (sP3(X20,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X21) | ~p(s(bool,X21))) ) | $spl118),
  inference(resolution,[],[f976,f401])).
fof(f2387,plain,(
  ( ! [X14,X15] : (sP3(X15,X14,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ~p(s(bool,X14))) ) | $spl118),
  inference(resolution,[],[f976,f397])).
fof(f2386,plain,(
  ( ! [X13] : (~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X13) | ~p(s(bool,X13))) ) | $spl118),
  inference(resolution,[],[f976,f369])).
fof(f2383,plain,(
  ( ! [X10] : (sP1(X10,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | p(s(bool,X10))) ) | $spl118),
  inference(resolution,[],[f976,f366])).
fof(f2382,plain,(
  ( ! [X9] : (sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X9) | p(s(bool,X9))) ) | $spl118),
  inference(resolution,[],[f976,f363])).
fof(f2381,plain,(
  ( ! [X8] : (~sP1(X8,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ~p(s(bool,X8))) ) | $spl118),
  inference(resolution,[],[f976,f362])).
fof(f2376,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(resolution,[],[f976,f341])).
fof(f2366,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0))),s(fun(sK13,bool),sK16)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl118),
  inference(resolution,[],[f976,f376])).
fof(f376,plain,(
  ( ! [X4,X2,X0,X10,X3,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f2157,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f352])).
fof(f2158,plain,(
  ( ! [X0] : (sP0('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f353])).
fof(f2185,plain,(
  ( ! [X0] : (sP3('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f398])).
fof(f2191,plain,(
  ( ! [X0] : (sP3('T_0',X0,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f399])).
fof(f2194,plain,(
  ( ! [X0] : (sP3(X0,'T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f401])).
fof(f2197,plain,(
  ( ! [X0] : (sP3(X0,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0')) ) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f401])).
fof(f2198,plain,(
  ( ! [X0] : (sP3(X0,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f976,f401])).
fof(f2235,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f332,f976,f431])).
fof(f2237,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f332,f976,f431])).
fof(f2260,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f442])).
fof(f2261,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X1)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f443])).
fof(f2263,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f446])).
fof(f2268,plain,(
  ( ! [X0] : (sP9('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,'T_0')) ) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f446])).
fof(f2269,plain,(
  ( ! [X0] : (sP9('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f976,f446])).
fof(f2271,plain,(
  ( ! [X0] : (sP9('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f447])).
fof(f2276,plain,(
  ( ! [X0] : (sP9('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0',X0)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f447])).
fof(f2277,plain,(
  ( ! [X0] : (sP9('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f976,f447])).
fof(f2279,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f449])).
fof(f2284,plain,(
  ( ! [X0] : (~sP9('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f450])).
fof(f2318,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X1)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f460])).
fof(f2322,plain,(
  ( ! [X0] : (sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,'F_0')) ) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f463])).
fof(f2324,plain,(
  ( ! [X0] : (sP11('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f464])).
fof(f2327,plain,(
  ( ! [X0] : (sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0',X0)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f464])).
fof(f2328,plain,(
  ( ! [X0] : (sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f976,f464])).
fof(f2337,plain,(
  ( ! [X0] : (~sP11('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f467])).
fof(f2351,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,X1)) ) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f478])).
fof(f2371,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(resolution,[],[f976,f292])).
fof(f2370,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(resolution,[],[f976,f290])).
fof(f2140,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f290])).
fof(f2141,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f292])).
fof(f2412,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl169),
  inference(cnf_transformation,[],[f2412_D])).
fof(f2412_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f2142,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f293])).
fof(f2143,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f295])).
fof(f2149,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f276,f976,f341])).
fof(f2150,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f341])).
fof(f2152,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f276,f976,f341])).
fof(f2153,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f341])).
fof(f2501,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl223),
  inference(cnf_transformation,[],[f2501_D])).
fof(f2501_D,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f2159,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl118),
  inference(unit_resulting_resolution,[],[f276,f976,f362])).
fof(f2160,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f362])).
fof(f2499,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl221),
  inference(cnf_transformation,[],[f2499_D])).
fof(f2499_D,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl221),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl221])])).
fof(f2161,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f976,f362])).
fof(f2503,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl225),
  inference(cnf_transformation,[],[f2503_D])).
fof(f2503_D,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f2162,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f276,f976,f362])).
fof(f2163,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f362])).
fof(f2164,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f976,f362])).
fof(f2505,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl226),
  inference(cnf_transformation,[],[f2505_D])).
fof(f2505_D,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f2165,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f363])).
fof(f2507,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl228),
  inference(cnf_transformation,[],[f2507_D])).
fof(f2507_D,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f2166,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f366])).
fof(f2167,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f367])).
fof(f2168,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f368])).
fof(f2169,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f276,f976,f369])).
fof(f2170,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f369])).
fof(f2171,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f976,f369])).
fof(f2172,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl118),
  inference(unit_resulting_resolution,[],[f276,f976,f369])).
fof(f2173,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl118),
  inference(unit_resulting_resolution,[],[f275,f976,f369])).
fof(f2174,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f976,f369])).
fof(f2461,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl191),
  inference(cnf_transformation,[],[f2461_D])).
fof(f2461_D,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl191),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl191])])).
fof(f2200,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f276,f281,f976,f414])).
fof(f2201,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f275,f281,f976,f414])).
fof(f2495,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl217),
  inference(cnf_transformation,[],[f2495_D])).
fof(f2495_D,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f2202,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f976,f281,f976,f414])).
fof(f2497,plain,(
  ~sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl219),
  inference(cnf_transformation,[],[f2497_D])).
fof(f2497_D,plain,(
  ~sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl219),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl219])])).
fof(f2203,plain,(
  ~sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f276,f281,f976,f414])).
fof(f2204,plain,(
  ~sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f275,f281,f976,f414])).
fof(f2205,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f976,f281,f976,f414])).
fof(f2493,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl214),
  inference(cnf_transformation,[],[f2493_D])).
fof(f2493_D,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f2206,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl2 | $spl118)),
  inference(unit_resulting_resolution,[],[f569,f567,f560,f976,f415])).
fof(f2491,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl212),
  inference(cnf_transformation,[],[f2491_D])).
fof(f2491_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f2207,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl2 | $spl4 | $spl118)),
  inference(unit_resulting_resolution,[],[f569,f825,f560,f976,f415])).
fof(f2489,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl210),
  inference(cnf_transformation,[],[f2489_D])).
fof(f2489_D,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f2208,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl2 | $spl4 | $spl118)),
  inference(unit_resulting_resolution,[],[f827,f567,f818,f976,f415])).
fof(f2487,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl208),
  inference(cnf_transformation,[],[f2487_D])).
fof(f2487_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f2209,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl4 | $spl118)),
  inference(unit_resulting_resolution,[],[f827,f825,f818,f976,f415])).
fof(f2485,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl207),
  inference(cnf_transformation,[],[f2485_D])).
fof(f2485_D,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl207),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl207])])).
fof(f2210,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f276,f281,f976,f418])).
fof(f2211,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f275,f281,f976,f418])).
fof(f2479,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl203),
  inference(cnf_transformation,[],[f2479_D])).
fof(f2479_D,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f2212,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f976,f281,f976,f418])).
fof(f2481,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl205),
  inference(cnf_transformation,[],[f2481_D])).
fof(f2481_D,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl205),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl205])])).
fof(f2215,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f276,f281,f976,f418])).
fof(f2216,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f275,f281,f976,f418])).
fof(f2217,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f976,f281,f976,f418])).
fof(f2475,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl201),
  inference(cnf_transformation,[],[f2475_D])).
fof(f2475_D,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f2221,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f276,f281,f976,f422])).
fof(f2222,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f275,f281,f976,f422])).
fof(f2469,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl197),
  inference(cnf_transformation,[],[f2469_D])).
fof(f2469_D,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f2223,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f976,f281,f976,f422])).
fof(f2471,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl199),
  inference(cnf_transformation,[],[f2471_D])).
fof(f2471_D,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f2226,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f276,f281,f976,f422])).
fof(f2227,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f275,f281,f976,f422])).
fof(f2228,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f976,f281,f976,f422])).
fof(f2465,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl194),
  inference(cnf_transformation,[],[f2465_D])).
fof(f2465_D,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f2236,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl56 | $spl118)),
  inference(unit_resulting_resolution,[],[f713,f976,f431])).
fof(f2463,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl192),
  inference(cnf_transformation,[],[f2463_D])).
fof(f2463_D,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f2238,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | ($spl56 | $spl118)),
  inference(unit_resulting_resolution,[],[f713,f976,f431])).
fof(f2239,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | ($spl9 | $spl118)),
  inference(unit_resulting_resolution,[],[f665,f976,f432])).
fof(f2459,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl189),
  inference(cnf_transformation,[],[f2459_D])).
fof(f2459_D,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f2240,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | ($spl9 | $spl118)),
  inference(unit_resulting_resolution,[],[f665,f976,f432])).
fof(f2453,plain,(
  ~sP8('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl187),
  inference(cnf_transformation,[],[f2453_D])).
fof(f2453_D,plain,(
  ~sP8('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f2255,plain,(
  ~sP8('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f281,f281,f976,f441])).
fof(f2447,plain,(
  ~sP9('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl185),
  inference(cnf_transformation,[],[f2447_D])).
fof(f2447_D,plain,(
  ~sP9('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f2278,plain,(
  ~sP9('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f281,f281,f976,f448])).
fof(f2437,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl183),
  inference(cnf_transformation,[],[f2437_D])).
fof(f2437_D,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f2308,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl118),
  inference(unit_resulting_resolution,[],[f281,f276,f976,f458])).
fof(f2309,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl118),
  inference(unit_resulting_resolution,[],[f281,f275,f976,f458])).
fof(f2431,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl179),
  inference(cnf_transformation,[],[f2431_D])).
fof(f2431_D,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f2310,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f976,f458])).
fof(f2433,plain,(
  ~sP10('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl181),
  inference(cnf_transformation,[],[f2433_D])).
fof(f2433_D,plain,(
  ~sP10('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f2313,plain,(
  ~sP10('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f276,f281,f976,f458])).
fof(f2314,plain,(
  ~sP10('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f275,f281,f976,f458])).
fof(f2315,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f976,f281,f976,f458])).
fof(f2427,plain,(
  sP11('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl176),
  inference(cnf_transformation,[],[f2427_D])).
fof(f2427_D,plain,(
  sP11('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f2319,plain,(
  sP11('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f281,f976,f462])).
fof(f2423,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl175),
  inference(cnf_transformation,[],[f2423_D])).
fof(f2423_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f2329,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f276,f976,f465])).
fof(f2330,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f275,f976,f465])).
fof(f2417,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl171),
  inference(cnf_transformation,[],[f2417_D])).
fof(f2417_D,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f2331,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f976,f465])).
fof(f2419,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl173),
  inference(cnf_transformation,[],[f2419_D])).
fof(f2419_D,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f2334,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl118),
  inference(unit_resulting_resolution,[],[f281,f276,f976,f465])).
fof(f2335,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl118),
  inference(unit_resulting_resolution,[],[f281,f275,f976,f465])).
fof(f2336,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f976,f976,f465])).
fof(f2340,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl118)),
  inference(unit_resulting_resolution,[],[f281,f276,f713,f976,f471])).
fof(f2341,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl118)),
  inference(unit_resulting_resolution,[],[f281,f275,f713,f976,f471])).
fof(f2342,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f281,f276,f976,f474])).
fof(f2343,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f281,f275,f976,f474])).
fof(f2344,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f281,f976,f976,f474])).
fof(f2347,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f281,f976,f976,f474])).
fof(f2349,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl118)),
  inference(unit_resulting_resolution,[],[f281,f276,f713,f976,f475])).
fof(f2350,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl56 | $spl118)),
  inference(unit_resulting_resolution,[],[f281,f275,f713,f976,f475])).
fof(f2358,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl118),
  inference(unit_resulting_resolution,[],[f281,f281,f976,f481])).
fof(f2410,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl167),
  inference(cnf_transformation,[],[f2410_D])).
fof(f2410_D,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl167])])).
fof(f2359,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ($spl2 | $spl118)),
  inference(unit_resulting_resolution,[],[f976,f636])).
fof(f2408,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl165),
  inference(cnf_transformation,[],[f2408_D])).
fof(f2408_D,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl165),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl165])])).
fof(f2360,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ($spl2 | $spl118)),
  inference(unit_resulting_resolution,[],[f976,f641])).
fof(f2121,plain,(
  ~sP7('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl163),
  inference(cnf_transformation,[],[f2121_D])).
fof(f2121_D,plain,(
  ~sP7('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f2117,plain,(
  ~sP7('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl99),
  inference(unit_resulting_resolution,[],[f956,f413])).
fof(f2111,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl161),
  inference(cnf_transformation,[],[f2111_D])).
fof(f2111_D,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f2107,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl95),
  inference(unit_resulting_resolution,[],[f952,f412])).
fof(f2101,plain,(
  ~sP7('F_0','T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl159),
  inference(cnf_transformation,[],[f2101_D])).
fof(f2101_D,plain,(
  ~sP7('F_0','T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f2097,plain,(
  ~sP7('F_0','T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl93),
  inference(unit_resulting_resolution,[],[f950,f412])).
fof(f2055,plain,(
  sP10('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl156),
  inference(cnf_transformation,[],[f2055_D])).
fof(f2055_D,plain,(
  sP10('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f2053,plain,(
  sP10('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl74),
  inference(resolution,[],[f932,f453])).
fof(f2047,plain,(
  sP10('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl74),
  inference(unit_resulting_resolution,[],[f932,f453])).
fof(f1814,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl154),
  inference(cnf_transformation,[],[f1814_D])).
fof(f1814_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f1810,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f276,f818,f825,f826,f415])).
fof(f1811,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f275,f818,f825,f826,f415])).
fof(f1803,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl152),
  inference(cnf_transformation,[],[f1803_D])).
fof(f1803_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f1799,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f276,f569,f560,f825,f415])).
fof(f1800,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f275,f569,f560,f825,f415])).
fof(f1791,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl150),
  inference(cnf_transformation,[],[f1791_D])).
fof(f1791_D,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f1784,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f276,f568,f567,f818,f415])).
fof(f1785,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f275,f568,f567,f818,f415])).
fof(f1593,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl148),
  inference(cnf_transformation,[],[f1593_D])).
fof(f1593_D,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f1470,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl56),
  inference(unit_resulting_resolution,[],[f276,f713,f431])).
fof(f1471,plain,(
  sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl56),
  inference(unit_resulting_resolution,[],[f275,f713,f431])).
fof(f1584,plain,(
  sP7('T_0','T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl146),
  inference(cnf_transformation,[],[f1584_D])).
fof(f1584_D,plain,(
  sP7('T_0','T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f1484,plain,(
  sP7('T_0','T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl56),
  inference(unit_resulting_resolution,[],[f276,f713,f431])).
fof(f1485,plain,(
  sP7('T_0','T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl56),
  inference(unit_resulting_resolution,[],[f275,f713,f431])).
fof(f1426,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl145),
  inference(cnf_transformation,[],[f1426_D])).
fof(f1426_D,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f1422,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl39),
  inference(unit_resulting_resolution,[],[f695,f413])).
fof(f1416,plain,(
  ~sP7('T_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl143),
  inference(cnf_transformation,[],[f1416_D])).
fof(f1416_D,plain,(
  ~sP7('T_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f1412,plain,(
  ~sP7('T_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl37),
  inference(unit_resulting_resolution,[],[f693,f413])).
fof(f1401,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl141),
  inference(cnf_transformation,[],[f1401_D])).
fof(f1401_D,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f1397,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl33),
  inference(unit_resulting_resolution,[],[f689,f412])).
fof(f1391,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl139),
  inference(cnf_transformation,[],[f1391_D])).
fof(f1391_D,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f1387,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl31),
  inference(unit_resulting_resolution,[],[f687,f412])).
fof(f1350,plain,(
  sP10('F_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl136),
  inference(cnf_transformation,[],[f1350_D])).
fof(f1350_D,plain,(
  sP10('F_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f1348,plain,(
  sP10('F_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl16),
  inference(resolution,[],[f673,f453])).
fof(f1342,plain,(
  sP10('F_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl16),
  inference(unit_resulting_resolution,[],[f673,f453])).
fof(f1328,plain,(
  sP6('F_0','T_0','F_0') | $spl130),
  inference(cnf_transformation,[],[f1328_D])).
fof(f1328_D,plain,(
  sP6('F_0','T_0','F_0') <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f1324,plain,(
  sP6('F_0','T_0','F_0') | $spl64),
  inference(resolution,[],[f743,f413])).
fof(f1330,plain,(
  sP5('F_0','T_0','F_0') | $spl132),
  inference(cnf_transformation,[],[f1330_D])).
fof(f1330_D,plain,(
  sP5('F_0','T_0','F_0') <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f1323,plain,(
  sP5('F_0','T_0','F_0') | $spl64),
  inference(resolution,[],[f743,f412])).
fof(f1332,plain,(
  sP4('T_0','F_0','F_0') | $spl134),
  inference(cnf_transformation,[],[f1332_D])).
fof(f1332_D,plain,(
  sP4('T_0','F_0','F_0') <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f1322,plain,(
  sP4('T_0','F_0','F_0') | $spl64),
  inference(resolution,[],[f743,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f1318,plain,(
  sP4('T_0','F_0','F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f743,f411])).
fof(f1319,plain,(
  sP5('F_0','T_0','F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f743,f412])).
fof(f1320,plain,(
  sP6('F_0','T_0','F_0') | $spl64),
  inference(unit_resulting_resolution,[],[f743,f413])).
fof(f1105,plain,(
  sP6('F_0','F_0','T_0') | $spl124),
  inference(cnf_transformation,[],[f1105_D])).
fof(f1105_D,plain,(
  sP6('F_0','F_0','T_0') <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f1102,plain,(
  sP6('F_0','F_0','T_0') | $spl60),
  inference(resolution,[],[f739,f413])).
fof(f1107,plain,(
  sP5('T_0','F_0','F_0') | $spl126),
  inference(cnf_transformation,[],[f1107_D])).
fof(f1107_D,plain,(
  sP5('T_0','F_0','F_0') <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f1101,plain,(
  sP5('T_0','F_0','F_0') | $spl60),
  inference(resolution,[],[f739,f412])).
fof(f1109,plain,(
  sP4('F_0','T_0','F_0') | $spl128),
  inference(cnf_transformation,[],[f1109_D])).
fof(f1109_D,plain,(
  sP4('F_0','T_0','F_0') <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f1100,plain,(
  sP4('F_0','T_0','F_0') | $spl60),
  inference(resolution,[],[f739,f411])).
fof(f1095,plain,(
  sP4('F_0','T_0','F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f739,f411])).
fof(f1096,plain,(
  sP5('T_0','F_0','F_0') | $spl60),
  inference(unit_resulting_resolution,[],[f739,f412])).
fof(f1097,plain,(
  sP6('F_0','F_0','T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f739,f413])).
fof(f1084,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl122),
  inference(cnf_transformation,[],[f1084_D])).
fof(f1084_D,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f1080,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f276,f560,f567,f568,f415])).
fof(f1081,plain,(
  sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f275,f560,f567,f568,f415])).
fof(f974,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl116),
  inference(cnf_transformation,[],[f974_D])).
fof(f974_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f979,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(global_subsumption,[],[f274,f273,f272,f487,f883,f882,f881,f878,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f867,f864,f862,f861,f860,f859,f858,f857,f856,f855,f854,f852,f851,f850,f849,f848,f845,f843,f842,f839,f837,f836,f835,f830,f829,f828,f827,f826,f825,f824,f823,f822,f821,f820,f819,f818,f817,f816,f815,f814,f813,f812,f811,f810,f809,f808,f807,f806,f804,f803,f800,f797,f791,f780,f779,f778,f777,f776,f775,f774,f773,f772,f771,f770,f769,f768,f767,f766,f765,f764,f763,f759,f758,f756,f755,f749,f748,f747,f746,f745,f744,f884,f885,f886,f887,f892,f897,f898,f899,f902,f903,f906,f915,f916,f918,f919,f920,f922])).
fof(f922,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(superposition,[],[f487,f282])).
fof(f920,plain,(
  ( ! [X45,X46] : (sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X45,X46) | ~p(s(bool,X45))) ) | $spl4),
  inference(resolution,[],[f487,f464])).
fof(f919,plain,(
  ( ! [X43,X44] : (sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X44,X43) | p(s(bool,X43))) ) | $spl4),
  inference(resolution,[],[f487,f463])).
fof(f918,plain,(
  ( ! [X41,X42] : (sP11(X42,X41,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | p(s(bool,X41)) | p(s(bool,X42))) ) | $spl4),
  inference(resolution,[],[f487,f462])).
fof(f916,plain,(
  ( ! [X37,X38] : (sP9('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X37,X38) | ~p(s(bool,X37))) ) | $spl4),
  inference(resolution,[],[f487,f447])).
fof(f915,plain,(
  ( ! [X35,X36] : (sP9('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X36,X35) | ~p(s(bool,X35))) ) | $spl4),
  inference(resolution,[],[f487,f446])).
fof(f906,plain,(
  ( ! [X17,X18] : (sP3(X17,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X18) | ~p(s(bool,X18))) ) | $spl4),
  inference(resolution,[],[f487,f401])).
fof(f903,plain,(
  ( ! [X12,X11] : (sP3(X12,X11,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ~p(s(bool,X11))) ) | $spl4),
  inference(resolution,[],[f487,f397])).
fof(f902,plain,(
  ( ! [X10] : (~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X10) | ~p(s(bool,X10))) ) | $spl4),
  inference(resolution,[],[f487,f369])).
fof(f899,plain,(
  ( ! [X7] : (sP1(X7,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | p(s(bool,X7))) ) | $spl4),
  inference(resolution,[],[f487,f366])).
fof(f898,plain,(
  ( ! [X6] : (sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X6) | p(s(bool,X6))) ) | $spl4),
  inference(resolution,[],[f487,f363])).
fof(f897,plain,(
  ( ! [X5] : (~sP1(X5,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | ~p(s(bool,X5))) ) | $spl4),
  inference(resolution,[],[f487,f362])).
fof(f892,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(resolution,[],[f487,f341])).
fof(f763,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f352])).
fof(f764,plain,(
  ( ! [X0] : (sP0('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f353])).
fof(f791,plain,(
  ( ! [X0] : (sP3('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f398])).
fof(f797,plain,(
  ( ! [X0] : (sP3('T_0',X0,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f399])).
fof(f800,plain,(
  ( ! [X0] : (sP3(X0,'T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f401])).
fof(f803,plain,(
  ( ! [X0] : (sP3(X0,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f401])).
fof(f804,plain,(
  ( ! [X0] : (sP3(X0,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f487,f401])).
fof(f829,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f332,f487,f431])).
fof(f830,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f332,f487,f431])).
fof(f836,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f442])).
fof(f837,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f443])).
fof(f839,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f446])).
fof(f842,plain,(
  ( ! [X0] : (sP9('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f446])).
fof(f843,plain,(
  ( ! [X0] : (sP9('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f487,f446])).
fof(f845,plain,(
  ( ! [X0] : (sP9('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f447])).
fof(f848,plain,(
  ( ! [X0] : (sP9('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f447])).
fof(f849,plain,(
  ( ! [X0] : (sP9('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f487,f447])).
fof(f851,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f449])).
fof(f852,plain,(
  ( ! [X0] : (~sP9('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f450])).
fof(f860,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f460])).
fof(f862,plain,(
  ( ! [X0] : (sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f463])).
fof(f864,plain,(
  ( ! [X0] : (sP11('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f464])).
fof(f867,plain,(
  ( ! [X0] : (sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f464])).
fof(f868,plain,(
  ( ! [X0] : (sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f487,f464])).
fof(f875,plain,(
  ( ! [X0] : (~sP11('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f467])).
fof(f882,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f478])).
fof(f887,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(resolution,[],[f487,f292])).
fof(f886,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(resolution,[],[f487,f290])).
fof(f976,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) | $spl118),
  inference(cnf_transformation,[],[f976_D])).
fof(f976_D,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f885,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) | $spl4),
  inference(resolution,[],[f487,f374])).
fof(f374,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X10)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X10)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X10))))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.pred_set.BIJ_DEF')).
fof(f978,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) | $spl120),
  inference(cnf_transformation,[],[f978_D])).
fof(f978_D,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f884,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) | $spl4),
  inference(resolution,[],[f487,f373])).
fof(f373,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f744,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f373])).
fof(f745,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f374])).
fof(f746,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f290])).
fof(f747,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f292])).
fof(f924,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl67),
  inference(cnf_transformation,[],[f924_D])).
fof(f924_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f748,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f293])).
fof(f749,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f295])).
fof(f755,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f276,f487,f341])).
fof(f756,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f341])).
fof(f758,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f276,f487,f341])).
fof(f759,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f341])).
fof(f966,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl109),
  inference(cnf_transformation,[],[f966_D])).
fof(f966_D,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f765,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f276,f487,f362])).
fof(f766,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f362])).
fof(f964,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl107),
  inference(cnf_transformation,[],[f964_D])).
fof(f964_D,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f767,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f487,f362])).
fof(f968,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl111),
  inference(cnf_transformation,[],[f968_D])).
fof(f968_D,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f768,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f276,f487,f362])).
fof(f769,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f362])).
fof(f770,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f487,f362])).
fof(f970,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl112),
  inference(cnf_transformation,[],[f970_D])).
fof(f970_D,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f771,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f363])).
fof(f972,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl114),
  inference(cnf_transformation,[],[f972_D])).
fof(f972_D,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f772,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f366])).
fof(f773,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f367])).
fof(f774,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f368])).
fof(f775,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f276,f487,f369])).
fof(f776,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f369])).
fof(f777,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f487,f369])).
fof(f778,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f276,f487,f369])).
fof(f779,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f275,f487,f369])).
fof(f780,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f487,f369])).
fof(f946,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl89),
  inference(cnf_transformation,[],[f946_D])).
fof(f946_D,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f806,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f276,f281,f487,f414])).
fof(f807,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f275,f281,f487,f414])).
fof(f960,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl103),
  inference(cnf_transformation,[],[f960_D])).
fof(f960_D,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f808,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f487,f281,f487,f414])).
fof(f962,plain,(
  ~sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl105),
  inference(cnf_transformation,[],[f962_D])).
fof(f962_D,plain,(
  ~sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f809,plain,(
  ~sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f276,f281,f487,f414])).
fof(f810,plain,(
  ~sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f275,f281,f487,f414])).
fof(f811,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f487,f281,f487,f414])).
fof(f958,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl101),
  inference(cnf_transformation,[],[f958_D])).
fof(f958_D,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f812,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f276,f281,f487,f418])).
fof(f813,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f275,f281,f487,f418])).
fof(f954,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl97),
  inference(cnf_transformation,[],[f954_D])).
fof(f954_D,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f814,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f487,f281,f487,f418])).
fof(f956,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl99),
  inference(cnf_transformation,[],[f956_D])).
fof(f956_D,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f815,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f276,f281,f487,f418])).
fof(f816,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f275,f281,f487,f418])).
fof(f817,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f487,f281,f487,f418])).
fof(f952,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl95),
  inference(cnf_transformation,[],[f952_D])).
fof(f952_D,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f819,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f276,f281,f487,f422])).
fof(f820,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f275,f281,f487,f422])).
fof(f948,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl91),
  inference(cnf_transformation,[],[f948_D])).
fof(f948_D,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f821,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f487,f281,f487,f422])).
fof(f950,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl93),
  inference(cnf_transformation,[],[f950_D])).
fof(f950_D,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f822,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f276,f281,f487,f422])).
fof(f823,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f275,f281,f487,f422])).
fof(f824,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f487,f281,f487,f422])).
fof(f831,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'T_0','F_0') | ($spl4 | $spl9)),
  inference(unit_resulting_resolution,[],[f665,f487,f432])).
fof(f944,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl87),
  inference(cnf_transformation,[],[f944_D])).
fof(f944_D,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f832,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | ($spl4 | $spl9)),
  inference(unit_resulting_resolution,[],[f665,f487,f432])).
fof(f942,plain,(
  ~sP8('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl85),
  inference(cnf_transformation,[],[f942_D])).
fof(f942_D,plain,(
  ~sP8('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f835,plain,(
  ~sP8('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f281,f281,f487,f441])).
fof(f940,plain,(
  ~sP9('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl83),
  inference(cnf_transformation,[],[f940_D])).
fof(f940_D,plain,(
  ~sP9('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f850,plain,(
  ~sP9('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f281,f281,f487,f448])).
fof(f938,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl81),
  inference(cnf_transformation,[],[f938_D])).
fof(f938_D,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f854,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f281,f276,f487,f458])).
fof(f855,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f281,f275,f487,f458])).
fof(f934,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl77),
  inference(cnf_transformation,[],[f934_D])).
fof(f934_D,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f856,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f487,f458])).
fof(f936,plain,(
  ~sP10('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl79),
  inference(cnf_transformation,[],[f936_D])).
fof(f936_D,plain,(
  ~sP10('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f857,plain,(
  ~sP10('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f276,f281,f487,f458])).
fof(f858,plain,(
  ~sP10('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f275,f281,f487,f458])).
fof(f859,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f487,f281,f487,f458])).
fof(f932,plain,(
  sP11('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl74),
  inference(cnf_transformation,[],[f932_D])).
fof(f932_D,plain,(
  sP11('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f861,plain,(
  sP11('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f281,f487,f462])).
fof(f930,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl73),
  inference(cnf_transformation,[],[f930_D])).
fof(f930_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f869,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f276,f487,f465])).
fof(f870,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f275,f487,f465])).
fof(f926,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl69),
  inference(cnf_transformation,[],[f926_D])).
fof(f926_D,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f871,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f487,f465])).
fof(f928,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl71),
  inference(cnf_transformation,[],[f928_D])).
fof(f928_D,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f872,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f281,f276,f487,f465])).
fof(f873,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f281,f275,f487,f465])).
fof(f874,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)),'F_0','const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f487,f487,f465])).
fof(f876,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f281,f276,f487,f474])).
fof(f877,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f281,f275,f487,f474])).
fof(f878,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f281,f487,f487,f474])).
fof(f881,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f281,f487,f487,f474])).
fof(f883,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))) | $spl4),
  inference(unit_resulting_resolution,[],[f281,f281,f487,f481])).
fof(f743,plain,(
  sP7('F_0','F_0','T_0') | $spl64),
  inference(cnf_transformation,[],[f743_D])).
fof(f743_D,plain,(
  sP7('F_0','F_0','T_0') <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f721,plain,(
  sP7('F_0','F_0','T_0') | $spl9),
  inference(unit_resulting_resolution,[],[f281,f665,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f253])).
fof(f741,plain,(
  ~sP7('T_0','F_0','T_0') | $spl63),
  inference(cnf_transformation,[],[f741_D])).
fof(f741_D,plain,(
  ~sP7('T_0','F_0','T_0') <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f722,plain,(
  ~sP7('T_0','F_0','T_0') | $spl9),
  inference(unit_resulting_resolution,[],[f276,f665,f432])).
fof(f723,plain,(
  ~sP7('T_0','F_0','T_0') | $spl9),
  inference(unit_resulting_resolution,[],[f275,f665,f432])).
fof(f739,plain,(
  sP7('F_0','T_0','F_0') | $spl60),
  inference(cnf_transformation,[],[f739_D])).
fof(f739_D,plain,(
  sP7('F_0','T_0','F_0') <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f729,plain,(
  sP7('F_0','T_0','F_0') | $spl9),
  inference(unit_resulting_resolution,[],[f281,f665,f430])).
fof(f737,plain,(
  ~sP7('T_0','T_0','F_0') | $spl59),
  inference(cnf_transformation,[],[f737_D])).
fof(f737_D,plain,(
  ~sP7('T_0','T_0','F_0') <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f730,plain,(
  ~sP7('T_0','T_0','F_0') | $spl9),
  inference(unit_resulting_resolution,[],[f276,f665,f432])).
fof(f731,plain,(
  ~sP7('T_0','T_0','F_0') | $spl9),
  inference(unit_resulting_resolution,[],[f275,f665,f432])).
fof(f713,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl56),
  inference(cnf_transformation,[],[f713_D])).
fof(f713_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f714,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(global_subsumption,[],[f274,f273,f272,f485,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f604,f602,f601,f600,f599,f598,f597,f596,f595,f594,f592,f591,f590,f589,f588,f585,f583,f582,f579,f577,f576,f575,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f546,f545,f542,f539,f533,f522,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f507,f506,f505,f501,f500,f498,f497,f491,f490,f489,f488,f624,f625,f626,f631,f636,f637,f638,f641,f642,f645,f654,f655,f657,f658,f659,f661])).
fof(f661,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(superposition,[],[f485,f282])).
fof(f659,plain,(
  ( ! [X48,X49] : (sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X48,X49) | ~p(s(bool,X48))) ) | $spl2),
  inference(resolution,[],[f485,f464])).
fof(f658,plain,(
  ( ! [X47,X46] : (sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X47,X46) | p(s(bool,X46))) ) | $spl2),
  inference(resolution,[],[f485,f463])).
fof(f657,plain,(
  ( ! [X45,X44] : (sP11(X45,X44,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | p(s(bool,X44)) | p(s(bool,X45))) ) | $spl2),
  inference(resolution,[],[f485,f462])).
fof(f655,plain,(
  ( ! [X41,X40] : (sP9('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X40,X41) | ~p(s(bool,X40))) ) | $spl2),
  inference(resolution,[],[f485,f447])).
fof(f654,plain,(
  ( ! [X39,X38] : (sP9('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X39,X38) | ~p(s(bool,X38))) ) | $spl2),
  inference(resolution,[],[f485,f446])).
fof(f645,plain,(
  ( ! [X21,X20] : (sP3(X20,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X21) | ~p(s(bool,X21))) ) | $spl2),
  inference(resolution,[],[f485,f401])).
fof(f642,plain,(
  ( ! [X14,X15] : (sP3(X15,X14,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | ~p(s(bool,X14))) ) | $spl2),
  inference(resolution,[],[f485,f397])).
fof(f638,plain,(
  ( ! [X10] : (sP1(X10,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | p(s(bool,X10))) ) | $spl2),
  inference(resolution,[],[f485,f366])).
fof(f637,plain,(
  ( ! [X9] : (sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X9) | p(s(bool,X9))) ) | $spl2),
  inference(resolution,[],[f485,f363])).
fof(f631,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(resolution,[],[f485,f341])).
fof(f624,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK17))),s(fun(X0,bool),X2)))) | ~sP2(sK13,X0,X1,sK16,X2)) ) | $spl2),
  inference(resolution,[],[f485,f383])).
fof(f505,plain,(
  ( ! [X0] : (sP0(X0,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f352])).
fof(f506,plain,(
  ( ! [X0] : (sP0('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f353])).
fof(f533,plain,(
  ( ! [X0] : (sP3('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f398])).
fof(f539,plain,(
  ( ! [X0] : (sP3('T_0',X0,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f399])).
fof(f542,plain,(
  ( ! [X0] : (sP3(X0,'T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f401])).
fof(f545,plain,(
  ( ! [X0] : (sP3(X0,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f401])).
fof(f546,plain,(
  ( ! [X0] : (sP3(X0,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f485,f401])).
fof(f576,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f442])).
fof(f577,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f443])).
fof(f579,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f446])).
fof(f582,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f446])).
fof(f583,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f485,f446])).
fof(f585,plain,(
  ( ! [X0] : (sP9('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f447])).
fof(f588,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f447])).
fof(f589,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f485,f447])).
fof(f591,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f449])).
fof(f592,plain,(
  ( ! [X0] : (~sP9('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f450])).
fof(f600,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f460])).
fof(f602,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f463])).
fof(f604,plain,(
  ( ! [X0] : (sP11('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f464])).
fof(f607,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f464])).
fof(f608,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f485,f464])).
fof(f615,plain,(
  ( ! [X0] : (~sP11('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f467])).
fof(f622,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f478])).
fof(f626,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(resolution,[],[f485,f292])).
fof(f625,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(resolution,[],[f485,f290])).
fof(f488,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f290])).
fof(f489,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f292])).
fof(f663,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl7),
  inference(cnf_transformation,[],[f663_D])).
fof(f663_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f490,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f293])).
fof(f491,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f295])).
fof(f497,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f276,f485,f341])).
fof(f498,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f341])).
fof(f500,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f276,f485,f341])).
fof(f501,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f341])).
fof(f705,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl49),
  inference(cnf_transformation,[],[f705_D])).
fof(f705_D,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f507,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f276,f485,f362])).
fof(f508,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f362])).
fof(f703,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl47),
  inference(cnf_transformation,[],[f703_D])).
fof(f703_D,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f509,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f485,f362])).
fof(f707,plain,(
  ~sP1('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl51),
  inference(cnf_transformation,[],[f707_D])).
fof(f707_D,plain,(
  ~sP1('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f510,plain,(
  ~sP1('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f276,f485,f362])).
fof(f511,plain,(
  ~sP1('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f362])).
fof(f512,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f485,f362])).
fof(f709,plain,(
  sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl52),
  inference(cnf_transformation,[],[f709_D])).
fof(f709_D,plain,(
  sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f513,plain,(
  sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f363])).
fof(f711,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl54),
  inference(cnf_transformation,[],[f711_D])).
fof(f711_D,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f514,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f366])).
fof(f515,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f367])).
fof(f516,plain,(
  sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f368])).
fof(f517,plain,(
  ~sP1('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f276,f485,f369])).
fof(f518,plain,(
  ~sP1('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f369])).
fof(f519,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f485,f369])).
fof(f520,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f276,f485,f369])).
fof(f521,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f275,f485,f369])).
fof(f522,plain,(
  ~sP1('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f485,f369])).
fof(f701,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl45),
  inference(cnf_transformation,[],[f701_D])).
fof(f701_D,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f548,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f276,f281,f485,f414])).
fof(f549,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f275,f281,f485,f414])).
fof(f697,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl41),
  inference(cnf_transformation,[],[f697_D])).
fof(f697_D,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f550,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f485,f281,f485,f414])).
fof(f699,plain,(
  ~sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl43),
  inference(cnf_transformation,[],[f699_D])).
fof(f699_D,plain,(
  ~sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f551,plain,(
  ~sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f276,f281,f485,f414])).
fof(f552,plain,(
  ~sP7('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f275,f281,f485,f414])).
fof(f553,plain,(
  ~sP7('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f485,f281,f485,f414])).
fof(f695,plain,(
  ~sP6('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl39),
  inference(cnf_transformation,[],[f695_D])).
fof(f695_D,plain,(
  ~sP6('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f554,plain,(
  ~sP6('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f276,f281,f485,f418])).
fof(f555,plain,(
  ~sP6('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f275,f281,f485,f418])).
fof(f691,plain,(
  ~sP6('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl35),
  inference(cnf_transformation,[],[f691_D])).
fof(f691_D,plain,(
  ~sP6('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f556,plain,(
  ~sP6('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f485,f281,f485,f418])).
fof(f693,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl37),
  inference(cnf_transformation,[],[f693_D])).
fof(f693_D,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f557,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f276,f281,f485,f418])).
fof(f558,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f275,f281,f485,f418])).
fof(f559,plain,(
  ~sP6('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f485,f281,f485,f418])).
fof(f689,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl33),
  inference(cnf_transformation,[],[f689_D])).
fof(f689_D,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f561,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f276,f281,f485,f422])).
fof(f562,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f275,f281,f485,f422])).
fof(f685,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl29),
  inference(cnf_transformation,[],[f685_D])).
fof(f685_D,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f563,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f485,f281,f485,f422])).
fof(f687,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl31),
  inference(cnf_transformation,[],[f687_D])).
fof(f687_D,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f564,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f276,f281,f485,f422])).
fof(f565,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f275,f281,f485,f422])).
fof(f566,plain,(
  ~sP5('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f485,f281,f485,f422])).
fof(f683,plain,(
  ~sP8('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl27),
  inference(cnf_transformation,[],[f683_D])).
fof(f683_D,plain,(
  ~sP8('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','F_0') <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f575,plain,(
  ~sP8('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f281,f281,f485,f441])).
fof(f681,plain,(
  ~sP9('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl25),
  inference(cnf_transformation,[],[f681_D])).
fof(f681_D,plain,(
  ~sP9('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','F_0') <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f590,plain,(
  ~sP9('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f281,f281,f485,f448])).
fof(f679,plain,(
  ~sP10('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl23),
  inference(cnf_transformation,[],[f679_D])).
fof(f679_D,plain,(
  ~sP10('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f594,plain,(
  ~sP10('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f281,f276,f485,f458])).
fof(f595,plain,(
  ~sP10('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f281,f275,f485,f458])).
fof(f675,plain,(
  ~sP10('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl19),
  inference(cnf_transformation,[],[f675_D])).
fof(f675_D,plain,(
  ~sP10('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f596,plain,(
  ~sP10('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f485,f458])).
fof(f677,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl21),
  inference(cnf_transformation,[],[f677_D])).
fof(f677_D,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f597,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f276,f281,f485,f458])).
fof(f598,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f275,f281,f485,f458])).
fof(f599,plain,(
  ~sP10('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f485,f281,f485,f458])).
fof(f673,plain,(
  sP11('F_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl16),
  inference(cnf_transformation,[],[f673_D])).
fof(f673_D,plain,(
  sP11('F_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f601,plain,(
  sP11('F_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f281,f485,f462])).
fof(f671,plain,(
  ~sP11('T_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl15),
  inference(cnf_transformation,[],[f671_D])).
fof(f671_D,plain,(
  ~sP11('T_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f609,plain,(
  ~sP11('T_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f276,f485,f465])).
fof(f610,plain,(
  ~sP11('T_0','F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f275,f485,f465])).
fof(f667,plain,(
  ~sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl11),
  inference(cnf_transformation,[],[f667_D])).
fof(f667_D,plain,(
  ~sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f611,plain,(
  ~sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f485,f465])).
fof(f669,plain,(
  ~sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl13),
  inference(cnf_transformation,[],[f669_D])).
fof(f669_D,plain,(
  ~sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','T_0') <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f612,plain,(
  ~sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f281,f276,f485,f465])).
fof(f613,plain,(
  ~sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f281,f275,f485,f465])).
fof(f614,plain,(
  ~sP11('const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)),'F_0','const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f485,f485,f465])).
fof(f616,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f281,f276,f485,f474])).
fof(f617,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f281,f275,f485,f474])).
fof(f618,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f281,f485,f485,f474])).
fof(f665,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f665_D])).
fof(f665_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f619,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f281,f281,f276,f485,f474])).
fof(f620,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f281,f281,f275,f485,f474])).
fof(f621,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f281,f485,f485,f474])).
fof(f623,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))) | $spl2),
  inference(unit_resulting_resolution,[],[f281,f281,f485,f481])).
fof(f487,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) | $spl4),
  inference(cnf_transformation,[],[f487_D])).
fof(f487_D,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f272,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16))))),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(fun(sK13,bool),sK16)))) & (p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) & s(sK13,sK17) != s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,'const.pred_set.LINV_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(sK13,sK17))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16,sK17])],[f126])).
fof(f126,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,'const.pred_set.LINV_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(X1,X5)))))))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) => ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,'const.pred_set.LINV_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(X1,X5)))))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) => ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,'const.pred_set.LINV_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(X1,X5)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X10)))) => ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X7,X11),s(fun(X7,bool),X10)))) => s(X7,X11) = s(X7,happ(s(fun(X6,X7),X8),s(X6,'const.pred_set.LINV_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(X7,X11)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X10)))) => ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X7,X11),s(fun(X7,bool),X10)))) => s(X7,X11) = s(X7,happ(s(fun(X6,X7),X8),s(X6,'const.pred_set.LINV_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(X7,X11)))))))),
  file('thm.pred_set.BIJ_LINV_INV','thm.pred_set.BIJ_LINV_INV')).
fof(f485,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) | $spl2),
  inference(cnf_transformation,[],[f485_D])).
fof(f485_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f273,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK17),s(fun(sK13,bool),sK16))))),
  inference(cnf_transformation,[],[f181])).
fof(f483,plain,(
  s(sK13,sK17) != s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,'const.pred_set.LINV_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(sK13,sK17))))) | $spl1),
  inference(cnf_transformation,[],[f483_D])).
fof(f483_D,plain,(
  s(sK13,sK17) != s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,'const.pred_set.LINV_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(sK13,sK17))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f274,plain,(
  s(sK13,sK17) != s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,'const.pred_set.LINV_3'(s(fun(sK12,sK13),sK14),s(fun(sK12,bool),sK15),s(sK13,sK17)))))),
  inference(cnf_transformation,[],[f181])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 705
% Passive clauses: 886
% Generated clauses: 7914
% Final active clauses: 700
% Final passive clauses: 322
% Input formulas: 50
% Initial clauses: 204
% 
% Duplicate literals: 203
% Trivial inequalities: 90
% Fw demodulations: 1582
% Bw demodulations: 8
% Global subsumptions: 9
% 
% Simple tautologies: 458
% Equational tautologies: 19
% Forward subsumptions: 4164
% Backward subsumptions: 1
% 
% Binary resolution: 2453
% Unit resulting resolution: 2550
% Forward superposition: 207
% Backward superposition: 549
% Equality resolution: 16
% 
% Split clauses: 44
% Split components: 90
% SAT solver clauses: 2029
% SAT solver unit clauses: 6
% SAT solver binary clauses: 1184
% 
% Sat splits: 44
% 
% Memory used [KB]: 6652
% Time elapsed: 0.208 s
% ------------------------------
% ------------------------------
% Success in time 18.261 s
