Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.combin
% (89190)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 108
% Passive clauses: 3062
% Generated clauses: 19512
% Final active clauses: 96
% Final passive clauses: 2809
% Input formulas: 18
% Initial clauses: 31
% 
% Duplicate literals: 929
% Trivial inequalities: 163
% Fw subsumption resolutions: 1614
% Fw demodulations: 2753
% Bw demodulations: 9
% 
% Simple tautologies: 217
% Equational tautologies: 32
% Forward subsumptions: 10083
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 12
% Forward superposition: 7395
% Backward superposition: 6326
% Self superposition: 218
% Equality factoring: 68
% Equality resolution: 18
% 
% Split clauses: 356
% Split components: 721
% SAT solver clauses: 797
% SAT solver unit clauses: 19
% SAT solver binary clauses: 115
% 
% Sat splits: 356
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 796
% TWLsolver calls for satisfiability: 124
% 
% Memory used [KB]: 15223
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.combin
% (89198)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 1292
% Generated clauses: 13833
% Final active clauses: 85
% Final passive clauses: 1178
% Input formulas: 18
% Initial clauses: 31
% 
% Duplicate literals: 696
% Trivial inequalities: 66
% Fw subsumption resolutions: 507
% Fw demodulations: 1548
% Bw demodulations: 25
% Condensations: 1337
% Global subsumptions: 4
% 
% Simple tautologies: 126
% Equational tautologies: 30
% Forward subsumptions: 6245
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 7
% Forward superposition: 5190
% Backward superposition: 2906
% Self superposition: 155
% Equality factoring: 49
% Equality resolution: 13
% 
% Split clauses: 53
% Split components: 114
% SAT solver clauses: 1996
% SAT solver unit clauses: 38
% SAT solver binary clauses: 36
% SAT solver learnt clauses: 8
% 
% Sat splits: 53
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1980
% TWLsolver calls for satisfiability: 136
% 
% Memory used [KB]: 17014
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.combin
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.combin
% SZS output start Proof for thm.combin
fof(f1771,plain,(
  $false),
  inference(subsumption_resolution,[],[f1770,f782])).
fof(f782,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X2) = s(X0,'const.combin.o_3'(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(fun(X3,X1),X4),s(X3,X5)))) )),
  inference(superposition,[],[f92,f96])).
fof(f96,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f88,f87])).
fof(f87,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X7] : ! [X12,X14] : s(X7,X12) = s(X7,'const.combin.K_2'(s(X7,X12),s(X6,X14)))),
  file('thm.combin.K_o_THM|split|0','thm.combin.K_THM')).
fof(f88,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X7,X6] : ! [X15,X16] : s(X7,'const.combin.K_2'(s(X7,X15),s(X6,X16))) = s(X7,happ(s(fun(X6,X7),'const.combin.K_1'(s(X7,X15))),s(X6,X16)))),
  file('thm.combin.K_o_THM|split|0','ARITY_1_2_const.combin.K')).
fof(f92,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X7,X8] : ! [X9,X13,X12] : s(X6,'const.combin.o_3'(s(fun(X7,X6),X9),s(fun(X8,X7),X13),s(X8,X12))) = s(X6,happ(s(fun(X7,X6),X9),s(X7,happ(s(fun(X8,X7),X13),s(X8,X12)))))),
  file('thm.combin.K_o_THM|split|0','thm.combin.o_THM')).
fof(f1770,plain,(
  s(sK2,sK4) != s(sK2,'const.combin.o_3'(s(fun(sK0,sK2),'const.combin.K_1'(s(sK2,sK4))),s(fun(sK1,sK0),sK3),s(sK1,sK6('const.combin.K_1'(s(sK2,sK4)),'const.combin.o_2'(s(fun(sK0,sK2),'const.combin.K_1'(s(sK2,sK4))),s(fun(sK1,sK0),sK3)),sK1,sK2))))),
  inference(superposition,[],[f1555,f93])).
fof(f93,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X8,X7] : ! [X15,X16,X17] : s(X6,'const.combin.o_3'(s(fun(X8,X6),X15),s(fun(X7,X8),X16),s(X7,X17))) = s(X6,happ(s(fun(X7,X6),'const.combin.o_2'(s(fun(X8,X6),X15),s(fun(X7,X8),X16))),s(X7,X17)))),
  file('thm.combin.K_o_THM|split|0','ARITY_2_3_const.combin.o')).
fof(f1555,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.combin.o_2'(s(fun(sK0,sK2),'const.combin.K_1'(s(sK2,sK4))),s(fun(sK1,sK0),sK3))),s(sK1,sK6('const.combin.K_1'(s(sK2,sK4)),'const.combin.o_2'(s(fun(sK0,sK2),'const.combin.K_1'(s(sK2,sK4))),s(fun(sK1,sK0),sK3)),sK1,sK2))))),
  inference(forward_demodulation,[],[f1513,f96])).
fof(f1513,plain,(
  s(sK2,happ(s(fun(sK1,sK2),'const.combin.K_1'(s(sK2,sK4))),s(sK1,sK6('const.combin.K_1'(s(sK2,sK4)),'const.combin.o_2'(s(fun(sK0,sK2),'const.combin.K_1'(s(sK2,sK4))),s(fun(sK1,sK0),sK3)),sK1,sK2)))) != s(sK2,happ(s(fun(sK1,sK2),'const.combin.o_2'(s(fun(sK0,sK2),'const.combin.K_1'(s(sK2,sK4))),s(fun(sK1,sK0),sK3))),s(sK1,sK6('const.combin.K_1'(s(sK2,sK4)),'const.combin.o_2'(s(fun(sK0,sK2),'const.combin.K_1'(s(sK2,sK4))),s(fun(sK1,sK0),sK3)),sK1,sK2))))),
  inference(unit_resulting_resolution,[],[f63,f91])).
fof(f91,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK6(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK6(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK6(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK6(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : ! [X9,X13] : (s(fun(X7,X6),X9) = s(fun(X7,X6),X13) <=> ! [X12] : s(X6,happ(s(fun(X7,X6),X9),s(X7,X12))) = s(X6,happ(s(fun(X7,X6),X13),s(X7,X12))))),
  file('thm.combin.K_o_THM|split|0','thm.bool.FUN_EQ_THM')).
fof(f63,plain,(
  s(fun(sK1,sK2),'const.combin.K_1'(s(sK2,sK4))) != s(fun(sK1,sK2),'const.combin.o_2'(s(fun(sK0,sK2),'const.combin.K_1'(s(sK2,sK4))),s(fun(sK1,sK0),sK3)))),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  s(fun(sK1,sK2),'const.combin.K_1'(s(sK2,sK4))) != s(fun(sK1,sK2),'const.combin.o_2'(s(fun(sK0,sK2),'const.combin.K_1'(s(sK2,sK4))),s(fun(sK1,sK0),sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f43])).
fof(f43,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X1,X2),'const.combin.K_1'(s(X2,X4))) != s(fun(X1,X2),'const.combin.o_2'(s(fun(X0,X2),'const.combin.K_1'(s(X2,X4))),s(fun(X1,X0),X3)))),
  inference(ennf_transformation,[],[f21])).
fof(f21,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X1,X2),'const.combin.K_1'(s(X2,X4))) = s(fun(X1,X2),'const.combin.o_2'(s(fun(X0,X2),'const.combin.K_1'(s(X2,X4))),s(fun(X1,X0),X3)))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ~! [X0,X1,X2] : ! [X3,X4] : s(fun(X1,X2),'const.combin.K_1'(s(X2,X4))) = s(fun(X1,X2),'const.combin.o_2'(s(fun(X0,X2),'const.combin.K_1'(s(X2,X4))),s(fun(X1,X0),X3)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10] : s(fun(X7,X8),'const.combin.K_1'(s(X8,X10))) = s(fun(X7,X8),'const.combin.o_2'(s(fun(X6,X8),'const.combin.K_1'(s(X8,X10))),s(fun(X7,X6),X9)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10] : s(fun(X7,X8),'const.combin.K_1'(s(X8,X10))) = s(fun(X7,X8),'const.combin.o_2'(s(fun(X6,X8),'const.combin.K_1'(s(X8,X10))),s(fun(X7,X6),X9)))),
  file('thm.combin.K_o_THM|split|0','thm.combin.K_o_THM|split|0')).
% SZS output end Proof for thm.combin
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 26
% Passive clauses: 146
% Generated clauses: 1666
% Final active clauses: 26
% Final passive clauses: 120
% Input formulas: 18
% Initial clauses: 31
% 
% Duplicate literals: 145
% Trivial inequalities: 18
% Fw subsumption resolutions: 22
% Fw demodulations: 97
% Condensations: 281
% 
% Simple tautologies: 43
% Equational tautologies: 27
% Forward subsumptions: 874
% Fw demodulations to eq. taut.: 10
% 
% Unit resulting resolution: 36
% Forward superposition: 509
% Backward superposition: 404
% Self superposition: 78
% Equality factoring: 41
% Equality resolution: 4
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 156
% SAT solver unit clauses: 16
% SAT solver binary clauses: 52
% 
% Sat splits: 4
% 
% Memory used [KB]: 2302
% Time elapsed: 0.100 s
% ------------------------------
% ------------------------------
% Success in time 1.44 s
