Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.rat
% (103368)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 201
% Passive clauses: 2350
% Generated clauses: 28973
% Final active clauses: 170
% Final passive clauses: 1772
% Input formulas: 30
% Initial clauses: 71
% 
% Duplicate literals: 1276
% Trivial inequalities: 279
% Fw subsumption resolutions: 4201
% Fw demodulations: 4499
% Bw demodulations: 29
% 
% Simple tautologies: 453
% Equational tautologies: 50
% Forward subsumptions: 14348
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 943
% Factoring: 2
% Forward superposition: 9650
% Backward superposition: 7574
% Self superposition: 162
% Equality factoring: 103
% Equality resolution: 22
% 
% Split clauses: 449
% Split components: 912
% SAT solver clauses: 1482
% SAT solver unit clauses: 27
% SAT solver binary clauses: 276
% 
% Sat splits: 449
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 1477
% TWLsolver calls for satisfiability: 202
% 
% Memory used [KB]: 12281
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (103388)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 237
% Passive clauses: 995
% Generated clauses: 24913
% Final active clauses: 140
% Final passive clauses: 439
% Input formulas: 30
% Initial clauses: 71
% 
% Duplicate literals: 897
% Trivial inequalities: 478
% Fw subsumption resolutions: 3033
% Fw demodulations: 3702
% Bw demodulations: 371
% Condensations: 1263
% Global subsumptions: 9
% 
% Simple tautologies: 433
% Equational tautologies: 52
% Forward subsumptions: 12383
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 589
% Factoring: 8
% Forward superposition: 5939
% Backward superposition: 6637
% Self superposition: 163
% Equality factoring: 106
% Equality resolution: 40
% 
% Split clauses: 278
% Split components: 563
% SAT solver clauses: 1711
% SAT solver unit clauses: 79
% SAT solver binary clauses: 45
% SAT solver learnt clauses: 19
% SAT solver learnt literals: 2
% 
% Sat splits: 278
% Sat splitting refutations: 32
% 
% TWLsolver clauses: 1674
% TWLsolver calls for satisfiability: 473
% 
% Memory used [KB]: 12281
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f4318,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f277,f281,f279,f854,f856,f931,f978,f972,f858,f983,f719,f857,f982,f1001,f1250,f1252,f980,f2413,f2414,f2421,f2418,f2415,f2416,f2424,f2417,f2426,f2419,f2423,f2469,f2475,f2477,f2479,f2481,f2482,f2484,f2486,f2488,f2489,f2490,f2560,f2562,f2565,f2566,f2568,f2574,f2589,f2590,f4298,f4310,f4311,f4313,f4316,f4317])).
fof(f4317,plain,(
  $false | ($spl0 | $spl7)),
  inference(global_subsumption,[],[f103,f105,f108,f109,f112,f111,f176,f156,f153,f173,f216,f218,f233,f236,f267,f277,f287,f371,f574,f594,f618,f620,f623,f636,f661,f692,f701,f776,f868,f867,f860,f883,f937,f972,f719,f104,f279,f1031,f1053,f1216,f1383,f1462,f1853,f1854,f1573,f2103,f2104,f2143,f2146,f2158,f2148,f2171,f2392,f2382,f2381,f2406,f2413,f2414,f2415,f2416,f2417,f2380,f1861,f1859,f210,f180,f1867,f1866,f1865,f1864,f1863,f1862,f1860,f1858,f1857,f985,f213,f207,f204,f201,f198,f195,f192,f189,f186,f183,f2106,f1868,f1395,f1360,f442,f393,f354,f215,f214,f177,f151,f158,f157,f1856,f1463,f1574,f2391,f2322,f2321,f2157,f470,f469,f308,f116,f115,f107,f4176,f4309,f858])).
fof(f4309,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(subsumption_resolution,[],[f4308,f105])).
fof(f4308,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f4307,f279])).
fof(f4307,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(subsumption_resolution,[],[f4306,f105])).
fof(f4306,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f3957,f287])).
fof(f3957,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f183,f574])).
fof(f4176,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(subsumption_resolution,[],[f4149,f105])).
fof(f4149,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f236,f2381])).
fof(f107,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF',aHL_FALSITY)).
fof(f115,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF','thm.bool.EQ_CLAUSES|split|2')).
fof(f116,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f73])).
fof(f308,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X3))) | s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X3),s('type.integer.int',X2)))) )),
  inference(superposition,[],[f151,f109])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,X2,bool))) )),
  inference(superposition,[],[f173,f109])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f173,f109])).
fof(f2157,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2143,f116])).
fof(f2321,plain,(
  ( ! [X23,X21,X22] : (s(bool,X23) = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X22),s('type.integer.int',X21))) | ~p(s(bool,X23)) | s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X22),s('type.integer.int',X21)))) )),
  inference(forward_demodulation,[],[f2320,f151])).
fof(f2320,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X23)) | s(bool,X23) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X21),s('type.integer.int',X22))) | s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X22),s('type.integer.int',X21)))) )),
  inference(subsumption_resolution,[],[f2299,f105])).
fof(f2299,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | s(bool,X23) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X21),s('type.integer.int',X22))) | s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X22),s('type.integer.int',X21)))) )),
  inference(superposition,[],[f156,f308])).
fof(f2322,plain,(
  ( ! [X35,X36,X34] : (s(bool,X36) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X34),s('type.integer.int',X35))) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X35),s('type.integer.int',X34)))) )),
  inference(subsumption_resolution,[],[f2302,f105])).
fof(f2302,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | s(bool,X36) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X34),s('type.integer.int',X35))) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X35),s('type.integer.int',X34)))) )),
  inference(superposition,[],[f236,f308])).
fof(f2391,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2148,f116])).
fof(f1574,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,'T_0') != s(bool,f(s(X122,sK2(X123,X124,bool,X122)))) | s(fun(X122,bool),X123) = s(fun(X122,bool),X124) | s(bool,X125) = s(bool,g(s(X122,sK2(X123,X124,bool,X122)))) | s(bool,'T_0') = s(bool,X125)) )),
  inference(superposition,[],[f173,f233])).
fof(f1463,plain,(
  ( ! [X138,X136,X135,X137] : (s(bool,X138) != s(bool,f(s(X135,sK2(X136,X137,bool,X135)))) | s(fun(X135,bool),X136) = s(fun(X135,bool),X137) | s(bool,'T_0') = s(bool,X138) | s(bool,'T_0') = s(bool,g(s(X135,sK2(X136,X137,bool,X135))))) )),
  inference(superposition,[],[f173,f233])).
fof(f1856,plain,(
  ( ! [X47,X48,X46] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X46))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X47))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X47))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X46))))) | s(bool,X48) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X46),s('type.frac.frac',X47))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f1548,f105])).
fof(f1548,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X46))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X47))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X47))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X46))))) | s(bool,X48) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X46),s('type.frac.frac',X47))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f157,f233])).
fof(f157,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1)))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1)))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) & (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1))))))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1)))) <=> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X17,X18] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X17),s('type.frac.frac',X18)))) <=> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X17))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X18))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X18))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X17))))))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF','thm.rat.rat_equiv_def')).
fof(f158,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f151,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X12,X13] : s(bool,'const.integer.int_gt_2'(s('type.integer.int',X12),s('type.integer.int',X13))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X13),s('type.integer.int',X12)))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF','thm.integer.int_gt')).
fof(f177,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f152,f151])).
fof(f152,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X12,X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int',X13)))))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF','thm.integer.INT_LT_IMP_LE')).
fof(f214,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f172,f151])).
fof(f172,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X12,X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X13),s('type.integer.int',X12)))))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF','thm.integer.INT_NOT_LT')).
fof(f215,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f171,f151])).
fof(f171,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f354,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(subsumption_resolution,[],[f353,f107])).
fof(f353,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(superposition,[],[f177,f109])).
fof(f393,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int',X3),s('type.integer.int',X2))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) )),
  inference(forward_demodulation,[],[f381,f218])).
fof(f381,plain,(
  ( ! [X2,X3] : (s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1))) != s(bool,'const.integer.int_gt_2'(s('type.integer.int',X3),s('type.integer.int',X2))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) )),
  inference(superposition,[],[f371,f151])).
fof(f442,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(subsumption_resolution,[],[f441,f107])).
fof(f441,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(superposition,[],[f215,f109])).
fof(f1360,plain,(
  ( ! [X30,X31,X29] : (s(bool,X31) = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X30),s('type.integer.int',X29))) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X29),s('type.integer.int',X30)))) )),
  inference(superposition,[],[f233,f151])).
fof(f1395,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X19),s('type.integer.int',X18))) | s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X18),s('type.integer.int',X19))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f233,f151])).
fof(f1868,plain,(
  ( ! [X111,X109,X110] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X109),s('type.integer.int',X110)))) | s(bool,X111) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X109),s('type.integer.int',X110))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(subsumption_resolution,[],[f1571,f105])).
fof(f1571,plain,(
  ( ! [X111,X109,X110] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X109),s('type.integer.int',X110)))) | s(bool,X111) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X109),s('type.integer.int',X110))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f214,f233])).
fof(f2106,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X5),s('type.integer.int',X4))) | s(bool,X6) = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X4),s('type.integer.int',X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2096,f105])).
fof(f2096,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X5),s('type.integer.int',X4))) | s(bool,X6) = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X4),s('type.integer.int',X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f354,f233])).
fof(f183,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f182,f151])).
fof(f182,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f181,f151])).
fof(f181,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f163,f151])).
fof(f163,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) & (((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) & (((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) <=> ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X15,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',X16)))))) <=> ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X15)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X16),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF','thm.integer.INT_MUL_SIGN_CASES|split|0')).
fof(f186,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) )),
  inference(forward_demodulation,[],[f185,f151])).
fof(f185,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(forward_demodulation,[],[f184,f151])).
fof(f184,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(forward_demodulation,[],[f162,f151])).
fof(f162,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(cnf_transformation,[],[f98])).
fof(f189,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) )),
  inference(forward_demodulation,[],[f188,f151])).
fof(f188,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(forward_demodulation,[],[f187,f151])).
fof(f187,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(forward_demodulation,[],[f161,f151])).
fof(f161,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(cnf_transformation,[],[f98])).
fof(f192,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f191,f151])).
fof(f191,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(forward_demodulation,[],[f190,f151])).
fof(f190,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(forward_demodulation,[],[f160,f151])).
fof(f160,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(cnf_transformation,[],[f98])).
fof(f195,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f194,f151])).
fof(f194,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(forward_demodulation,[],[f193,f151])).
fof(f193,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(forward_demodulation,[],[f159,f151])).
fof(f159,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(cnf_transformation,[],[f98])).
fof(f198,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(forward_demodulation,[],[f197,f151])).
fof(f197,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f196,f151])).
fof(f196,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f170,f151])).
fof(f170,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))))) & (((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))))) & (((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X15,X16] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int',X16))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X15)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X16),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))))))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF','thm.integer.INT_MUL_SIGN_CASES|split|1')).
fof(f201,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) )),
  inference(forward_demodulation,[],[f200,f151])).
fof(f200,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f199,f151])).
fof(f199,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f169,f151])).
fof(f169,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f204,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f203,f151])).
fof(f203,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f202,f151])).
fof(f202,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f168,f151])).
fof(f168,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f207,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f206,f151])).
fof(f206,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f205,f151])).
fof(f205,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f167,f151])).
fof(f167,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f213,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f212,f151])).
fof(f212,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f211,f151])).
fof(f211,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f165,f151])).
fof(f165,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f985,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f715,f107])).
fof(f715,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f183,f109])).
fof(f1857,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X61),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X61),s('type.integer.int',X60)))))) | s(bool,X62) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X60))) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1554,f105])).
fof(f1554,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X61),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X61),s('type.integer.int',X60)))))) | s(bool,X62) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X60))) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f201,f233])).
fof(f1858,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X64),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X63),s('type.integer.int',X64)))))) | s(bool,X65) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X63))) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f1555,f105])).
fof(f1555,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X64),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X63),s('type.integer.int',X64)))))) | s(bool,X65) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X63))) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f198,f233])).
fof(f1860,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X69),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X69)))) | s(bool,X71) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X69),s('type.integer.int',X70))))) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f1557,f105])).
fof(f1557,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X69),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X69)))) | s(bool,X71) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X69),s('type.integer.int',X70))))) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f213,f233])).
fof(f1862,plain,(
  ( ! [X76,X77,X75] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X75),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X76),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X77) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X75),s('type.integer.int',X76))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f1559,f105])).
fof(f1559,plain,(
  ( ! [X76,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X75),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X76),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X77) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X75),s('type.integer.int',X76))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f207,f233])).
fof(f1863,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X79)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X79),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X80) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X78),s('type.integer.int',X79))))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f1560,f105])).
fof(f1560,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X79)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X79),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X80) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X78),s('type.integer.int',X79))))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f204,f233])).
fof(f1864,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X84),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X84)))) | s(bool,X86) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X84),s('type.integer.int',X85))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1563,f105])).
fof(f1563,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X84),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X84)))) | s(bool,X86) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X84),s('type.integer.int',X85))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f195,f233])).
fof(f1865,plain,(
  ( ! [X88,X87,X89] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X88),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X87)))) | s(bool,X89) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X87),s('type.integer.int',X88))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f1564,f105])).
fof(f1564,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X88),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X87)))) | s(bool,X89) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X87),s('type.integer.int',X88))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f192,f233])).
fof(f1866,plain,(
  ( ! [X92,X90,X91] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X90),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X91)))) | s(bool,X92) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X90),s('type.integer.int',X91))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f1565,f105])).
fof(f1565,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X90),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X91)))) | s(bool,X92) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X90),s('type.integer.int',X91))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f189,f233])).
fof(f1867,plain,(
  ( ! [X94,X95,X93] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X94),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X94)))) | s(bool,X95) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X93),s('type.integer.int',X94))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f1566,f105])).
fof(f1566,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X94),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X94)))) | s(bool,X95) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X93),s('type.integer.int',X94))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f186,f233])).
fof(f180,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f179,f151])).
fof(f179,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f178,f151])).
fof(f178,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f164,f151])).
fof(f164,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f98])).
fof(f210,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(forward_demodulation,[],[f209,f151])).
fof(f209,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f208,f151])).
fof(f208,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f166,f151])).
fof(f166,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f1859,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X67)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X66))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X68) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X66))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f1556,f105])).
fof(f1556,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X67)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X66))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X68) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X66))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f180,f233])).
fof(f1861,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X73)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X72)))) | s(bool,X74) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X73))))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1558,f105])).
fof(f1558,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X73)))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X72)))) | s(bool,X74) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X72),s('type.integer.int',X73))))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f210,f233])).
fof(f2380,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1)))))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2148,f2148,f210])).
fof(f2406,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f2394,f105])).
fof(f2394,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f2148,f233])).
fof(f2381,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2148,f442])).
fof(f2382,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1)))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2148,f215])).
fof(f2392,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2148,f111])).
fof(f2171,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f2160,f105])).
fof(f2160,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f2143,f233])).
fof(f2148,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1)))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2143,f2143,f210])).
fof(f2158,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2143,f111])).
fof(f2146,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(subsumption_resolution,[],[f2126,f105])).
fof(f2126,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f236,f2104])).
fof(f2143,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1)))))) | $spl0),
  inference(subsumption_resolution,[],[f2114,f105])).
fof(f2114,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1)))))) | $spl0),
  inference(superposition,[],[f214,f2104])).
fof(f2104,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | $spl0),
  inference(subsumption_resolution,[],[f2091,f105])).
fof(f2091,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | $spl0),
  inference(superposition,[],[f354,f279])).
fof(f2103,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) )),
  inference(subsumption_resolution,[],[f2090,f105])).
fof(f2090,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) )),
  inference(superposition,[],[f354,f287])).
fof(f1573,plain,(
  ( ! [X121,X118,X120,X119] : (s(X120,f(s(bool,'T_0'))) != s(X120,g(s(bool,'T_0'))) | s(fun(bool,X120),X118) = s(fun(bool,X120),X119) | s(bool,X121) = s(bool,sK2(X118,X119,X120,bool)) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f173,f233])).
fof(f1854,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1536,f105])).
fof(f1536,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f236,f233])).
fof(f1853,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1535,f105])).
fof(f1535,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f156,f233])).
fof(f1462,plain,(
  ( ! [X134,X132,X133,X131] : (s(X133,f(s(bool,X134))) != s(X133,g(s(bool,X134))) | s(fun(bool,X133),X131) = s(fun(bool,X133),X132) | s(bool,'T_0') = s(bool,X134) | s(bool,'T_0') = s(bool,sK2(X131,X132,X133,bool))) )),
  inference(superposition,[],[f173,f233])).
fof(f1383,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f233,f233])).
fof(f1216,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(subsumption_resolution,[],[f1206,f105])).
fof(f1206,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f236,f1031])).
fof(f1053,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(subsumption_resolution,[],[f1044,f105])).
fof(f1044,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,X3) != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f236,f279])).
fof(f1031,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | $spl0),
  inference(unit_resulting_resolution,[],[f279,f393])).
fof(f104,plain,(
  s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1)))) & s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f64])).
fof(f64,plain,(
  ? [X0,X1] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1)))) & s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ~! [X0,X1] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1)))) => s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X6),s('type.frac.frac',X7)))) => s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X7))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X6),s('type.frac.frac',X7)))) => s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X7))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF','thm.rat.RAT_EQUIV_NMR_GTZ_IFF')).
fof(f937,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl7),
  inference(subsumption_resolution,[],[f916,f107])).
fof(f916,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl7),
  inference(superposition,[],[f267,f867])).
fof(f883,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl7),
  inference(unit_resulting_resolution,[],[f860,f112])).
fof(f860,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl7),
  inference(unit_resulting_resolution,[],[f858,f215])).
fof(f867,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl7),
  inference(unit_resulting_resolution,[],[f858,f116])).
fof(f868,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl7),
  inference(unit_resulting_resolution,[],[f858,f111])).
fof(f776,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X12),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f766,f105])).
fof(f766,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X12),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f236,f701])).
fof(f701,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X0)))) )),
  inference(unit_resulting_resolution,[],[f623,f112])).
fof(f692,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X12))))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f681,f105])).
fof(f681,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X12))))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f236,f594])).
fof(f661,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f655,f105])).
fof(f655,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,X3) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f236,f636])).
fof(f636,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK0)))),
  inference(unit_resulting_resolution,[],[f620,f112])).
fof(f623,plain,(
  ( ! [X0] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X0))))) )),
  inference(equality_resolution,[],[f158])).
fof(f620,plain,(
  p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK0))))),
  inference(unit_resulting_resolution,[],[f574,f158])).
fof(f618,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f608,f105])).
fof(f608,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f236,f287])).
fof(f594,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) )),
  inference(unit_resulting_resolution,[],[f287,f393])).
fof(f574,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))),
  inference(subsumption_resolution,[],[f571,f105])).
fof(f571,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))),
  inference(superposition,[],[f157,f218])).
fof(f371,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f364,f105])).
fof(f364,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,X4) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f236,f218])).
fof(f287,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f176,f112])).
fof(f267,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f263,f107])).
fof(f263,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f153,f109])).
fof(f236,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f115,f109])).
fof(f233,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f109,f109])).
fof(f218,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1)))),
  inference(unit_resulting_resolution,[],[f103,f112])).
fof(f216,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f107,f111])).
fof(f173,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF',aHL_EXT)).
fof(f153,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF','thm.bool.IMP_ANTISYM_AX')).
fof(f156,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f68])).
fof(f176,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(backward_demodulation,[],[f151,f108])).
fof(f111,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF','thm.bool.EQ_CLAUSES|split|0')).
fof(f112,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f71])).
fof(f109,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF',aHL_BOOL_CASES)).
fof(f108,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X14] : p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X14))))))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF','thm.frac.FRAC_DNMPOS')).
fof(f105,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.RAT_EQUIV_NMR_GTZ_IFF',aHL_TRUTH)).
fof(f103,plain,(
  p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1))))),
  inference(cnf_transformation,[],[f70])).
fof(f4316,plain,(
  $false | ($spl0 | $spl11)),
  inference(subsumption_resolution,[],[f4315,f105])).
fof(f4315,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl11)),
  inference(forward_demodulation,[],[f982,f279])).
fof(f4313,plain,(
  $false | ($spl0 | $spl7 | $spl13)),
  inference(global_subsumption,[],[f1275,f103,f105,f108,f109,f112,f111,f176,f156,f153,f173,f216,f218,f233,f236,f267,f277,f287,f371,f574,f594,f618,f620,f623,f636,f661,f692,f701,f776,f858,f868,f867,f860,f883,f937,f972,f719,f104,f279,f1031,f1053,f1216,f1383,f1462,f1853,f1854,f1573,f2103,f2104,f2143,f2146,f2158,f2148,f2171,f2392,f2382,f2381,f2406,f2413,f2414,f2415,f2416,f2417,f2380,f1861,f1859,f210,f180,f1867,f1866,f1865,f1864,f1863,f1862,f1860,f1858,f1857,f985,f213,f207,f204,f201,f198,f195,f192,f189,f186,f183,f2106,f1868,f1395,f1360,f442,f393,f354,f215,f214,f177,f151,f158,f157,f1856,f1463,f1574,f2391,f2322,f2321,f2157,f470,f469,f308,f116,f115,f107,f4176,f4309])).
fof(f1275,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))))) | $spl13),
  inference(unit_resulting_resolution,[],[f1252,f215])).
fof(f4311,plain,(
  $false | ($spl0 | $spl2 | $spl7 | $spl13)),
  inference(global_subsumption,[],[f3805,f103,f105,f108,f109,f112,f111,f176,f156,f153,f173,f216,f218,f233,f236,f267,f277,f287,f371,f574,f594,f618,f620,f623,f636,f661,f692,f701,f776,f858,f868,f867,f860,f883,f937,f972,f719,f104,f279,f1031,f1053,f1216,f1383,f1462,f1853,f1854,f1573,f2103,f2104,f2143,f2146,f2158,f2148,f2171,f2392,f2382,f2381,f2406,f2413,f2414,f2415,f2416,f2417,f2380,f1861,f1859,f210,f180,f1867,f1866,f1865,f1864,f1863,f1862,f1860,f1858,f1857,f985,f213,f207,f204,f201,f198,f195,f192,f189,f186,f183,f2106,f1868,f1395,f1360,f442,f393,f354,f215,f214,f177,f151,f158,f157,f1856,f1463,f1574,f2391,f2322,f2321,f2157,f470,f469,f308,f116,f115,f107,f4176,f4309])).
fof(f3805,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1284,f2190])).
fof(f2190,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int',X3))) | s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) ) | $spl2),
  inference(subsumption_resolution,[],[f2189,f107])).
fof(f2189,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int',X3))) | s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) ) | $spl2),
  inference(forward_demodulation,[],[f2179,f280])).
fof(f280,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(cnf_transformation,[],[f280_D])).
fof(f280_D,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2179,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int',X3))) | s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) ) | $spl2),
  inference(superposition,[],[f442,f1081])).
fof(f1081,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl2),
  inference(subsumption_resolution,[],[f1068,f107])).
fof(f1068,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl2),
  inference(superposition,[],[f267,f280])).
fof(f1284,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))))) | $spl13),
  inference(unit_resulting_resolution,[],[f1252,f111])).
fof(f4310,plain,(
  $false | ($spl0 | $spl7)),
  inference(global_subsumption,[],[f103,f105,f108,f109,f112,f111,f176,f156,f153,f173,f216,f218,f233,f236,f267,f277,f287,f371,f574,f594,f618,f620,f623,f636,f661,f692,f701,f776,f858,f868,f867,f860,f883,f937,f972,f719,f104,f279,f1031,f1053,f1216,f1383,f1462,f1853,f1854,f1573,f2103,f2104,f2143,f2146,f2158,f2148,f2171,f2392,f2382,f2381,f2406,f2413,f2414,f2415,f2416,f2417,f2380,f1861,f1859,f210,f180,f1867,f1866,f1865,f1864,f1863,f1862,f1860,f1858,f1857,f985,f213,f207,f204,f201,f198,f195,f192,f189,f186,f183,f2106,f1868,f1395,f1360,f442,f393,f354,f215,f214,f177,f151,f158,f157,f1856,f1463,f1574,f2391,f2322,f2321,f2157,f470,f469,f308,f116,f115,f107,f4176,f4309])).
fof(f4298,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f4297,f105])).
fof(f4297,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f4264,f2103])).
fof(f4264,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))) | ($spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f4258,f214])).
fof(f4258,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f4257,f107])).
fof(f4257,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f4256,f280])).
fof(f4256,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))) | $spl6),
  inference(subsumption_resolution,[],[f4242,f105])).
fof(f4242,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))) | $spl6),
  inference(superposition,[],[f189,f1266])).
fof(f1266,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(unit_resulting_resolution,[],[f857,f112])).
fof(f2590,plain,(
  $false | ($spl0 | $spl19 | $spl22)),
  inference(subsumption_resolution,[],[f2550,f105])).
fof(f2550,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl19 | $spl22)),
  inference(backward_demodulation,[],[f2491,f2380])).
fof(f2491,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f2467,f2424])).
fof(f2467,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl19),
  inference(trivial_inequality_removal,[],[f2461])).
fof(f2461,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl19),
  inference(superposition,[],[f2421,f233])).
fof(f2589,plain,(
  $false | ($spl13 | $spl19 | $spl22)),
  inference(subsumption_resolution,[],[f2549,f105])).
fof(f2549,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl19 | $spl22)),
  inference(backward_demodulation,[],[f2491,f1274])).
fof(f1274,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))))))) | $spl13),
  inference(unit_resulting_resolution,[],[f1252,f1252,f210])).
fof(f2574,plain,(
  $false | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f2573,f105])).
fof(f2573,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl22)),
  inference(forward_demodulation,[],[f2572,f2491])).
fof(f2572,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) ) | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f2517,f105])).
fof(f2517,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) ) | ($spl19 | $spl22)),
  inference(backward_demodulation,[],[f2491,f214])).
fof(f2568,plain,(
  $false | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f2567,f105])).
fof(f2567,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl22)),
  inference(forward_demodulation,[],[f2509,f2491])).
fof(f2509,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl19 | $spl22)),
  inference(backward_demodulation,[],[f2491,f2454])).
fof(f2454,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl19),
  inference(unit_resulting_resolution,[],[f105,f2421,f156])).
fof(f2566,plain,(
  $false | ($spl19 | $spl22 | $spl25)),
  inference(subsumption_resolution,[],[f2508,f2491])).
fof(f2508,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl19 | $spl22 | $spl25)),
  inference(backward_demodulation,[],[f2491,f2426])).
fof(f2565,plain,(
  $false | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f2506,f2491])).
fof(f2506,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl19 | $spl22)),
  inference(backward_demodulation,[],[f2491,f2421])).
fof(f2562,plain,(
  $false | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f2561,f2491])).
fof(f2561,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f2493,f111])).
fof(f2493,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl19 | $spl22)),
  inference(backward_demodulation,[],[f2491,f115])).
fof(f2560,plain,(
  $false | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f2492,f105])).
fof(f2492,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl22)),
  inference(backward_demodulation,[],[f2491,f107])).
fof(f2490,plain,(
  $false | ($spl19 | $spl25)),
  inference(subsumption_resolution,[],[f2427,f2426])).
fof(f2427,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl19),
  inference(unit_resulting_resolution,[],[f2421,f109])).
fof(f2489,plain,(
  $false | ($spl19 | $spl25)),
  inference(subsumption_resolution,[],[f2429,f2426])).
fof(f2429,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl19),
  inference(unit_resulting_resolution,[],[f216,f2421,f233])).
fof(f2488,plain,(
  $false | ($spl2 | $spl19 | $spl25)),
  inference(subsumption_resolution,[],[f2487,f2426])).
fof(f2487,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl2 | $spl19)),
  inference(forward_demodulation,[],[f2430,f280])).
fof(f2430,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl19)),
  inference(unit_resulting_resolution,[],[f1075,f2421,f233])).
fof(f1075,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(subsumption_resolution,[],[f1060,f107])).
fof(f1060,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(superposition,[],[f111,f280])).
fof(f2486,plain,(
  $false | ($spl5 | $spl19 | $spl25)),
  inference(subsumption_resolution,[],[f2485,f2426])).
fof(f2485,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl5 | $spl19)),
  inference(forward_demodulation,[],[f2431,f1011])).
fof(f1011,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f855,f116])).
fof(f855,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))) | $spl5),
  inference(cnf_transformation,[],[f855_D])).
fof(f855_D,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f2431,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) | ($spl5 | $spl19)),
  inference(unit_resulting_resolution,[],[f1012,f2421,f233])).
fof(f1012,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f855,f111])).
fof(f2484,plain,(
  $false | ($spl0 | $spl19 | $spl25)),
  inference(subsumption_resolution,[],[f2483,f2426])).
fof(f2483,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl0 | $spl19)),
  inference(forward_demodulation,[],[f2432,f2157])).
fof(f2432,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | ($spl0 | $spl19)),
  inference(unit_resulting_resolution,[],[f2158,f2421,f233])).
fof(f2482,plain,(
  $false | ($spl19 | $spl25)),
  inference(subsumption_resolution,[],[f2435,f2426])).
fof(f2435,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl19),
  inference(unit_resulting_resolution,[],[f216,f2421,f233])).
fof(f2481,plain,(
  $false | ($spl2 | $spl19 | $spl25)),
  inference(subsumption_resolution,[],[f2480,f2426])).
fof(f2480,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl2 | $spl19)),
  inference(forward_demodulation,[],[f2436,f280])).
fof(f2436,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl19)),
  inference(unit_resulting_resolution,[],[f1075,f2421,f233])).
fof(f2479,plain,(
  $false | ($spl5 | $spl19 | $spl25)),
  inference(subsumption_resolution,[],[f2478,f2426])).
fof(f2478,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl5 | $spl19)),
  inference(forward_demodulation,[],[f2437,f1011])).
fof(f2437,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) | ($spl5 | $spl19)),
  inference(unit_resulting_resolution,[],[f1012,f2421,f233])).
fof(f2477,plain,(
  $false | ($spl0 | $spl19 | $spl25)),
  inference(subsumption_resolution,[],[f2476,f2426])).
fof(f2476,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl0 | $spl19)),
  inference(forward_demodulation,[],[f2438,f2157])).
fof(f2438,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))) | ($spl0 | $spl19)),
  inference(unit_resulting_resolution,[],[f2158,f2421,f233])).
fof(f2475,plain,(
  $false | ($spl19 | $spl25)),
  inference(subsumption_resolution,[],[f2442,f2426])).
fof(f2442,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl19),
  inference(unit_resulting_resolution,[],[f107,f2421,f267])).
fof(f2469,plain,(
  $false | ($spl2 | $spl19 | $spl25)),
  inference(subsumption_resolution,[],[f2468,f2426])).
fof(f2468,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl2 | $spl19)),
  inference(forward_demodulation,[],[f2449,f280])).
fof(f2449,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl19)),
  inference(unit_resulting_resolution,[],[f2421,f1081])).
fof(f2423,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl20),
  inference(cnf_transformation,[],[f2423_D])).
fof(f2423_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2419,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool))) ) | $spl16),
  inference(cnf_transformation,[],[f2419_D])).
fof(f2419_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2426,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(cnf_transformation,[],[f2426_D])).
fof(f2426_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f2417,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK2(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f469,f109])).
fof(f2424,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl22),
  inference(cnf_transformation,[],[f2424_D])).
fof(f2424_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2416,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK2(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f469,f233])).
fof(f2415,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK2(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f469,f233])).
fof(f2418,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(cnf_transformation,[],[f2418_D])).
fof(f2418_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2421,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl19),
  inference(cnf_transformation,[],[f2421_D])).
fof(f2421_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f2414,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK2(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f469,f233])).
fof(f2413,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f469,f233])).
fof(f980,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1)))))) | $spl8),
  inference(cnf_transformation,[],[f980_D])).
fof(f980_D,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1)))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1252,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))))) | $spl13),
  inference(cnf_transformation,[],[f1252_D])).
fof(f1252_D,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f1250,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1)))))) | $spl5),
  inference(subsumption_resolution,[],[f1249,f107])).
fof(f1249,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1)))))) | $spl5),
  inference(forward_demodulation,[],[f1245,f1011])).
fof(f1245,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1)))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))))),
  inference(superposition,[],[f210,f574])).
fof(f1001,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1000,f105])).
fof(f1000,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f990,f287])).
fof(f990,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl4),
  inference(unit_resulting_resolution,[],[f953,f177])).
fof(f953,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))) | $spl4),
  inference(unit_resulting_resolution,[],[f856,f214])).
fof(f982,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl11),
  inference(cnf_transformation,[],[f982_D])).
fof(f982_D,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f857,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl6),
  inference(cnf_transformation,[],[f857_D])).
fof(f857_D,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f719,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(subsumption_resolution,[],[f718,f105])).
fof(f718,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f713,f287])).
fof(f713,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f183,f574])).
fof(f983,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl10),
  inference(cnf_transformation,[],[f983_D])).
fof(f983_D,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f858,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl7),
  inference(cnf_transformation,[],[f858_D])).
fof(f858_D,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f972,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))))))),
  inference(superposition,[],[f195,f574])).
fof(f978,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1)))))) | $spl3),
  inference(subsumption_resolution,[],[f977,f107])).
fof(f977,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1)))))) | $spl3),
  inference(forward_demodulation,[],[f972,f302])).
fof(f302,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f301,f109])).
fof(f301,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(backward_demodulation,[],[f293,f104])).
fof(f293,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f281,f109])).
fof(f931,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f930,f105])).
fof(f930,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl7)),
  inference(forward_demodulation,[],[f929,f293])).
fof(f929,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl7),
  inference(subsumption_resolution,[],[f928,f105])).
fof(f928,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl7),
  inference(forward_demodulation,[],[f927,f287])).
fof(f927,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl7),
  inference(subsumption_resolution,[],[f906,f107])).
fof(f906,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl7),
  inference(superposition,[],[f183,f867])).
fof(f856,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))) | $spl4),
  inference(cnf_transformation,[],[f856_D])).
fof(f856_D,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f854,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))) | $spl3),
  inference(subsumption_resolution,[],[f853,f107])).
fof(f853,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))) | $spl3),
  inference(forward_demodulation,[],[f848,f302])).
fof(f848,plain,(
  ~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))),
  inference(superposition,[],[f189,f574])).
fof(f279,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(cnf_transformation,[],[f279_D])).
fof(f279_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f281,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(cnf_transformation,[],[f281_D])).
fof(f281_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f277,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f104,f109])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 123
% Passive clauses: 237
% Generated clauses: 4269
% Final active clauses: 60
% Final passive clauses: 53
% Input formulas: 30
% Initial clauses: 71
% 
% Duplicate literals: 217
% Trivial inequalities: 59
% Fw subsumption resolutions: 323
% Fw demodulations: 721
% Bw demodulations: 87
% Condensations: 241
% Global subsumptions: 4
% 
% Simple tautologies: 114
% Equational tautologies: 112
% Forward subsumptions: 2203
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 538
% Factoring: 1
% Forward superposition: 779
% Backward superposition: 1056
% Self superposition: 30
% Equality factoring: 34
% Equality resolution: 15
% 
% Split clauses: 11
% Split components: 28
% SAT solver clauses: 292
% SAT solver unit clauses: 21
% SAT solver binary clauses: 121
% 
% Sat splits: 11
% Sat splitting refutations: 27
% 
% Memory used [KB]: 2302
% Time elapsed: 0.089 s
% ------------------------------
% ------------------------------
% Success in time 1.406 s
