Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% (76395)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 949
% Passive clauses: 6973
% Generated clauses: 30361
% Final active clauses: 895
% Final passive clauses: 5567
% Input formulas: 146
% Initial clauses: 308
% 
% Duplicate literals: 245
% Trivial inequalities: 60
% Fw subsumption resolutions: 1904
% Fw demodulations: 2119
% Bw demodulations: 52
% 
% Simple tautologies: 216
% Equational tautologies: 49
% Forward subsumptions: 18740
% Fw demodulations to eq. taut.: 36
% 
% Binary resolution: 7627
% Factoring: 1364
% Forward superposition: 8894
% Backward superposition: 7646
% Self superposition: 46
% Equality factoring: 2
% Equality resolution: 17
% 
% Split clauses: 227
% Split components: 454
% SAT solver clauses: 479
% SAT solver unit clauses: 15
% SAT solver binary clauses: 68
% 
% Sat splits: 227
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 479
% TWLsolver calls for satisfiability: 420
% 
% Memory used [KB]: 17398
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (76423)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 453
% Passive clauses: 2228
% Generated clauses: 23121
% Final active clauses: 425
% Final passive clauses: 1519
% Input formulas: 146
% Initial clauses: 308
% 
% Duplicate literals: 78
% Trivial inequalities: 141
% Fw subsumption resolutions: 659
% Fw demodulations: 4285
% Bw demodulations: 279
% Condensations: 28
% Global subsumptions: 19
% 
% Simple tautologies: 204
% Equational tautologies: 11
% Forward subsumptions: 12572
% Fw demodulations to eq. taut.: 34
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 1422
% Factoring: 79
% Forward superposition: 8028
% Backward superposition: 5411
% Self superposition: 77
% Equality factoring: 2
% Equality resolution: 47
% 
% Split clauses: 5
% Split components: 10
% SAT solver clauses: 3191
% SAT solver unit clauses: 378
% SAT solver binary clauses: 97
% SAT solver learnt clauses: 38
% SAT solver learnt literals: 12
% 
% Sat splits: 5
% 
% TWLsolver clauses: 3115
% TWLsolver calls for satisfiability: 1830
% 
% Memory used [KB]: 20596
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (76440)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2685
% Passive clauses: 18579
% Generated clauses: 319277
% Final active clauses: 2309
% Final passive clauses: 5723
% Input formulas: 146
% Initial clauses: 308
% Discarded non-redundant clauses: 48434
% 
% Duplicate literals: 206
% Trivial inequalities: 225
% Fw subsumption resolutions: 2580
% Fw demodulations: 54652
% Bw demodulations: 1086
% Condensations: 254
% Global subsumptions: 1
% 
% Simple tautologies: 127
% Equational tautologies: 112
% Forward subsumptions: 239086
% Fw demodulations to eq. taut.: 655
% Bw demodulations to eq. taut.: 19
% 
% Unit resulting resolution: 34693
% Factoring: 3
% Forward superposition: 163746
% Backward superposition: 60717
% Self superposition: 151
% Equality factoring: 30
% Equality resolution: 72
% 
% Split clauses: 43
% Split components: 86
% SAT solver clauses: 18724
% SAT solver unit clauses: 11185
% SAT solver binary clauses: 5664
% 
% Sat splits: 43
% Sat splitting refutations: 56
% 
% Memory used [KB]: 161191
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f24051,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f757,f762,f759,f710,f711,f712,f713,f714,f715,f769,f773,f771,f740,f742,f743,f776,f770,f809,f811,f895,f972,f973,f976,f977,f978,f979,f1145,f1148,f761,f1093,f1094,f1095,f1096,f1097,f1098,f1154,f1155,f1177,f1413,f1414,f1415,f1416,f1417,f1418,f1179,f1489,f1490,f1491,f1492,f1493,f1494,f1512,f1541,f1542,f1543,f1544,f1545,f1546,f1566,f1595,f1596,f1597,f1598,f1599,f1600,f1619,f1648,f1649,f1650,f1651,f1652,f1653,f1673,f1730,f1731,f1740,f1741,f1743,f1744,f1802,f1803,f1804,f1805,f1806,f1807,f1826,f1910,f1911,f1912,f1913,f1914,f1962,f1963,f1964,f1965,f1966,f1967,f1985,f2015,f2016,f2017,f2018,f2019,f2020,f2038,f767,f2129,f2130,f2139,f2140,f2142,f2143,f2333,f2334,f2336,f2337,f2357,f2423,f2424,f2429,f2430,f2574,f2575,f2577,f2578,f2598,f2629,f2630,f2632,f2633,f2653,f768,f2719,f2720,f2437,f2725,f2726,f2438,f2763,f2764,f2787,f3200,f3201,f3415,f3536,f3537,f3561,f3592,f3593,f3616,f3649,f3650,f3651,f3652,f3653,f3654,f765,f3708,f3709,f3825,f3946,f3948,f4153,f4154,f4298,f4299,f4327,f4359,f763,f4360,f4388,f4466,f4586,f4587,f4624,f4649,f4682,f4707,f4741,f4769,f4803,f4826,f4964,f4965,f5021,f5022,f5064,f5126,f5149,f766,f5287,f5288,f5020,f5343,f5344,f3947,f5466,f5471,f5468,f5426,f5427,f5503,f5641,f5646,f5643,f6983,f6984,f6985,f6986,f6987,f6988,f1150,f7281,f7282,f7283,f7449,f7448,f7440,f7426,f7425,f7424,f7529,f7530,f7535,f7536,f7539,f7540,f7852,f7854,f7912,f7914,f7919,f7921,f7920,f8544,f8547,f8546,f8494,f8495,f8577,f9131,f9132,f9347,f9351,f9349,f9357,f9360,f9359,f9831,f9832,f9833,f9834,f9835,f9836,f5092,f3700,f9950,f9951,f9995,f9952,f9953,f1767,f10003,f10004,f10116,f10018,f10019,f10118,f10021,f10022,f10117,f10126,f10130,f10168,f10170,f10175,f10177,f10176,f10296,f10298,f10198,f10119,f10199,f1772,f10304,f10306,f10344,f10345,f10169,f10350,f10352,f10351,f10470,f11083,f11084,f9180,f11253,f11256,f11156,f10120,f3945,f11262,f11264,f11302,f11303,f11308,f11311,f11309,f11430,f11432,f11438,f11440,f11478,f11479,f11254,f11484,f11486,f11485,f11545,f11546,f11547,f11548,f11615,f11616,f11657,f11737,f11738,f11743,f11744,f12024,f12028,f12026,f12123,f12170,f12171,f12127,f7326,f12128,f7329,f12129,f7328,f12130,f12027,f12629,f12630,f1147,f12954,f12958,f12960,f12961,f12962,f12964,f12956,f12981,f12989,f12987,f13046,f13051,f12909,f12910,f13079,f13081,f13083,f13082,f13233,f13234,f13235,f13421,f13423,f13649,f13653,f13650,f13664,f13666,f13665,f13671,f13276,f13672,f13279,f13673,f13278,f13779,f13781,f13780,f13786,f13788,f13787,f14125,f14127,f14126,f14134,f14136,f14135,f14257,f14259,f14258,f14261,f14264,f14262,f14268,f14270,f14269,f14396,f14399,f5502,f14397,f14416,f14424,f10128,f14441,f14445,f14443,f14467,f14418,f14469,f5470,f14516,f14518,f7913,f14523,f14525,f14524,f14623,f14624,f14626,f14627,f14697,f14700,f14699,f13048,f14860,f14861,f14863,f14864,f14908,f14909,f14910,f15099,f15102,f15289,f15290,f15292,f15293,f15330,f3701,f15381,f15382,f15383,f12959,f15384,f12957,f15385,f12963,f15386,f12965,f15501,f15509,f12983,f15525,f15531,f15529,f15527,f15553,f15505,f15503,f5645,f15584,f15588,f9350,f12984,f15586,f15603,f15606,f15609,f15612,f15430,f15431,f9996,f15613,f15457,f15478,f15476,f15614,f15616,f11310,f15615,f15621,f14949,f15622,f14952,f15623,f14951,f16170,f16619,f16620,f16621,f16622,f16623,f16624,f16662,f16750,f16751,f16752,f16753,f16754,f16755,f17745,f17746,f11755,f17751,f17752,f11756,f17964,f17965,f17978,f17979,f17980,f18898,f18899,f18900,f18925,f18926,f18982,f18983,f19127,f12985,f19080,f19081,f14685,f19199,f19200,f19239,f19215,f19216,f19217,f19218,f19219,f19220,f19252,f19283,f19284,f19323,f19299,f19300,f19301,f19302,f19303,f19304,f19336,f19420,f19421,f19494,f19528,f19529,f20242,f20325,f20326,f14895,f20449,f20450,f19493,f20555,f20556,f20607,f20638,f20639,f20690,f20654,f11263,f20655,f11439,f20656,f20657,f20658,f10129,f20659,f10305,f20806,f20807,f14686,f20934,f18955,f20935,f18958,f20936,f18957,f20961,f20962,f18969,f21043,f21044,f21084,f21123,f21124,f21163,f21202,f21203,f21243,f21283,f21284,f21323,f21406,f21407,f21445,f21485,f21486,f21524,f21537,f21538,f21684,f21682,f21619,f21620,f21694,f21731,f21732,f21783,f21794,f21798,f21797,f21795,f21813,f21814,f21854,f21877,f21878,f21929,f21893,f21942,f21894,f21943,f21895,f21896,f21897,f21945,f21898,f16792,f21944,f22141,f22200,f22142,f22203,f22143,f22202,f22166,f22167,f22168,f22169,f22170,f22171,f22215,f22254,f772,f810,f22297,f22298,f22299,f22300,f22301,f22302,f22342,f22373,f22374,f22375,f22376,f22377,f22378,f22417,f22448,f22449,f22450,f22451,f22452,f22453,f22493,f22525,f22526,f22527,f22528,f22529,f22530,f22569,f22601,f22602,f22603,f22604,f22605,f22606,f22644,f22676,f22214,f22677,f22216,f22678,f22679,f22680,f22218,f22681,f22719,f22217,f23146,f23148,f21683,f23215,f3698,f23216,f23250,f3699,f23289,f23290,f23291,f23292,f23293,f23294,f23331,f23364,f23365,f23366,f23367,f23368,f23369,f23406,f23441,f23442,f23443,f23444,f23445,f23446,f23484,f23517,f23518,f23519,f23520,f23521,f23522,f23560,f23595,f23596,f23597,f23598,f23599,f23600,f23638,f23672,f23673,f23674,f23675,f23676,f23677,f23713,f23747,f23748,f23749,f23750,f23751,f23752,f23788,f23822,f16661,f23823,f16663,f23824,f23825,f23826,f16665,f23827,f23864,f16664,f24026,f24030,f1766,f24028,f24039,f24042,f24045,f24048,f24050])).
fof(f24050,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f24049,f672])).
fof(f672,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_LE_REDUCE|split|1',aHL_FALSITY)).
fof(f24049,plain,(
  p(s(bool,'F_0')) | $spl80),
  inference(forward_demodulation,[],[f23952,f1242])).
fof(f1242,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))),
  inference(resolution,[],[f1130,f646])).
fof(f646,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f394])).
fof(f394,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X10))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.IMP_F_EQ_F')).
fof(f1130,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))))),
  inference(forward_demodulation,[],[f991,f487])).
fof(f487,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X6,X8] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X8))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.integer.INT_LT_CALCULATE|split|0')).
fof(f991,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))),
  inference(resolution,[],[f745,f463])).
fof(f463,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X7,X15] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X15)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X15)))) | s('type.integer.int',X7) = s('type.integer.int',X15)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.integer.INT_LE_LT')).
fof(f745,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))),
  inference(trivial_inequality_removal,[],[f737])).
fof(f737,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))),
  inference(superposition,[],[f679,f607])).
fof(f607,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f379])).
fof(f379,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f679,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(forward_demodulation,[],[f406,f458])).
fof(f458,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.arithmetic.NUMERAL_DEF')).
fof(f406,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK10)))))))),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK10)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10])],[f270])).
fof(f270,plain,(
  ? [X0] : s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ~! [X0] : s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.integer.INT_LE_REDUCE|split|1')).
fof(f23952,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | $spl80),
  inference(equality_resolution,[],[f3947])).
fof(f24048,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f24047,f672])).
fof(f24047,plain,(
  p(s(bool,'F_0')) | $spl80),
  inference(forward_demodulation,[],[f24046,f1242])).
fof(f24046,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | $spl80),
  inference(forward_demodulation,[],[f23953,f491])).
fof(f491,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.arithmetic.MULT|split|0')).
fof(f23953,plain,(
  ( ! [X91] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X91))),s('type.num.num',sK10))))) ) | $spl80),
  inference(trivial_inequality_removal,[],[f23951])).
fof(f23951,plain,(
  ( ! [X91] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X91))),s('type.num.num',sK10))))) ) | $spl80),
  inference(superposition,[],[f3947,f491])).
fof(f24045,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f24044,f672])).
fof(f24044,plain,(
  p(s(bool,'F_0')) | $spl80),
  inference(forward_demodulation,[],[f24043,f1242])).
fof(f24043,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | $spl80),
  inference(forward_demodulation,[],[f23954,f491])).
fof(f23954,plain,(
  ( ! [X90] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X90))),s('type.num.num',sK10))))) ) | $spl80),
  inference(trivial_inequality_removal,[],[f23950])).
fof(f23950,plain,(
  ( ! [X90] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X90))),s('type.num.num',sK10))))) ) | $spl80),
  inference(superposition,[],[f3947,f494])).
fof(f494,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.arithmetic.MULT_CLAUSES|split|0')).
fof(f24042,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f24041,f672])).
fof(f24041,plain,(
  p(s(bool,'F_0')) | $spl80),
  inference(forward_demodulation,[],[f24040,f1242])).
fof(f24040,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | $spl80),
  inference(forward_demodulation,[],[f23955,f490])).
fof(f490,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_distrib|split|4')).
fof(f23955,plain,(
  ( ! [X89] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10))))) ) | $spl80),
  inference(trivial_inequality_removal,[],[f23949])).
fof(f23949,plain,(
  ( ! [X89] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10))))) ) | $spl80),
  inference(superposition,[],[f3947,f490])).
fof(f24039,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f24038,f672])).
fof(f24038,plain,(
  p(s(bool,'F_0')) | $spl80),
  inference(forward_demodulation,[],[f24037,f1242])).
fof(f24037,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | $spl80),
  inference(forward_demodulation,[],[f23956,f490])).
fof(f23956,plain,(
  ( ! [X88] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X88),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10))))) ) | $spl80),
  inference(trivial_inequality_removal,[],[f23948])).
fof(f23948,plain,(
  ( ! [X88] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X88),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10))))) ) | $spl80),
  inference(superposition,[],[f3947,f493])).
fof(f493,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.arithmetic.MULT_|48|')).
fof(f24028,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) | $spl398),
  inference(cnf_transformation,[],[f24028_D])).
fof(f24028_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f1766,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl40),
  inference(cnf_transformation,[],[f1766_D])).
fof(f1766_D,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f24030,plain,(
  ( ! [X82] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X82) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10))))) ) | $spl80),
  inference(duplicate_literal_removal,[],[f24029])).
fof(f24029,plain,(
  ( ! [X82] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X82) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X82) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10))))) ) | $spl80),
  inference(forward_demodulation,[],[f23943,f458])).
fof(f23943,plain,(
  ( ! [X82] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X82) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X82)))) ) | $spl80),
  inference(superposition,[],[f3947,f445])).
fof(f445,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X6] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X6))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_distrib|split|17')).
fof(f24026,plain,(
  ( ! [X81] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X81) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10))))) ) | $spl80),
  inference(duplicate_literal_removal,[],[f24025])).
fof(f24025,plain,(
  ( ! [X81] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X81) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X81) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10))))) ) | $spl80),
  inference(forward_demodulation,[],[f23942,f458])).
fof(f23942,plain,(
  ( ! [X81] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X81) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) ) | $spl80),
  inference(superposition,[],[f3947,f447])).
fof(f447,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X6] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X6))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_distrib|split|18')).
fof(f16664,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl282),
  inference(cnf_transformation,[],[f16664_D])).
fof(f16664_D,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f23864,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) | $spl396),
  inference(cnf_transformation,[],[f23864_D])).
fof(f23864_D,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f23827,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,X26,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16425,f482])).
fof(f482,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X6,X8] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X8)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.integer.integer_244_prove_ltb|split|0')).
fof(f16425,plain,(
  ( ! [X78,X79] : (sP6(X78,X79,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) ) | $spl258),
  inference(resolution,[],[f15478,f528])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f341])).
fof(f341,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f16665,plain,(
  ( ! [X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) | $spl284),
  inference(cnf_transformation,[],[f16665_D])).
fof(f16665_D,plain,(
  ( ! [X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f23826,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl258),
  inference(superposition,[],[f16425,f481])).
fof(f481,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f327])).
fof(f23825,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl258),
  inference(superposition,[],[f16425,f479])).
fof(f479,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))))),
  inference(flattening,[],[f324])).
fof(f324,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X6,X8] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X8)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.integer.integer_244_prove_ltb|split|1')).
fof(f23824,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,X17,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16425,f478])).
fof(f478,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f325])).
fof(f16663,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl280),
  inference(cnf_transformation,[],[f16663_D])).
fof(f16663_D,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f23823,plain,(
  ( ! [X14,X15,X13] : (sP6(X13,X14,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl258),
  inference(superposition,[],[f16425,f477])).
fof(f477,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))))),
  inference(flattening,[],[f322])).
fof(f322,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) <=> (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X6,X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X8)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.integer.INT_LT_CALCULATE|split|2')).
fof(f16661,plain,(
  ( ! [X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) | $spl276),
  inference(cnf_transformation,[],[f16661_D])).
fof(f16661_D,plain,(
  ( ! [X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f23822,plain,(
  ( ! [X12,X10,X11] : (sP6(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl258),
  inference(superposition,[],[f16425,f476])).
fof(f476,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f323])).
fof(f23788,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X11)) ) | $spl394),
  inference(cnf_transformation,[],[f23788_D])).
fof(f23788_D,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X11)) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f23752,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16409,f482])).
fof(f16409,plain,(
  ( ! [X47,X46] : (sP4(X46,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X47)) ) | $spl258),
  inference(resolution,[],[f15478,f512])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f334])).
fof(f334,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f333])).
fof(f333,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f23751,plain,(
  ( ! [X24,X23,X22] : (sP4(X22,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl258),
  inference(superposition,[],[f16409,f481])).
fof(f23750,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl258),
  inference(superposition,[],[f16409,f479])).
fof(f23749,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16409,f478])).
fof(f23748,plain,(
  ( ! [X14,X15,X13] : (sP4(X13,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl258),
  inference(superposition,[],[f16409,f477])).
fof(f23747,plain,(
  ( ! [X12,X10,X11] : (sP4(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl258),
  inference(superposition,[],[f16409,f476])).
fof(f23713,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) | $spl392),
  inference(cnf_transformation,[],[f23713_D])).
fof(f23713_D,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f23677,plain,(
  ( ! [X26,X27,X25] : (sP4(X25,X26,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16408,f482])).
fof(f16408,plain,(
  ( ! [X45,X44] : (sP4(X44,X45,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) ) | $spl258),
  inference(resolution,[],[f15478,f511])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f335])).
fof(f23676,plain,(
  ( ! [X24,X23,X22] : (sP4(X22,X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl258),
  inference(superposition,[],[f16408,f481])).
fof(f23675,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl258),
  inference(superposition,[],[f16408,f479])).
fof(f23674,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16408,f478])).
fof(f23673,plain,(
  ( ! [X14,X15,X13] : (sP4(X13,X14,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl258),
  inference(superposition,[],[f16408,f477])).
fof(f23672,plain,(
  ( ! [X12,X10,X11] : (sP4(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl258),
  inference(superposition,[],[f16408,f476])).
fof(f23638,plain,(
  ( ! [X10,X11] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X10,X11)) ) | $spl390),
  inference(cnf_transformation,[],[f23638_D])).
fof(f23638_D,plain,(
  ( ! [X10,X11] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X10,X11)) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f23600,plain,(
  ( ! [X26,X27,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X25,X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16396,f482])).
fof(f16396,plain,(
  ( ! [X24,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X24,X25)) ) | $spl258),
  inference(resolution,[],[f15478,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f23599,plain,(
  ( ! [X24,X23,X22] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X22,X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl258),
  inference(superposition,[],[f16396,f481])).
fof(f23598,plain,(
  ( ! [X21,X19,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X19,X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl258),
  inference(superposition,[],[f16396,f479])).
fof(f23597,plain,(
  ( ! [X17,X18,X16] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X16,X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16396,f478])).
fof(f23596,plain,(
  ( ! [X14,X15,X13] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X13,X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl258),
  inference(superposition,[],[f16396,f477])).
fof(f23595,plain,(
  ( ! [X12,X10,X11] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X10,X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl258),
  inference(superposition,[],[f16396,f476])).
fof(f23560,plain,(
  ( ! [X10,X11] : (sP0(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X11)) ) | $spl388),
  inference(cnf_transformation,[],[f23560_D])).
fof(f23560_D,plain,(
  ( ! [X10,X11] : (sP0(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X11)) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f23522,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16395,f482])).
fof(f16395,plain,(
  ( ! [X23,X22] : (sP0(X22,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X23)) ) | $spl258),
  inference(resolution,[],[f15478,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f312])).
fof(f23521,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl258),
  inference(superposition,[],[f16395,f481])).
fof(f23520,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl258),
  inference(superposition,[],[f16395,f479])).
fof(f23519,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16395,f478])).
fof(f23518,plain,(
  ( ! [X14,X15,X13] : (sP0(X13,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl258),
  inference(superposition,[],[f16395,f477])).
fof(f23517,plain,(
  ( ! [X12,X10,X11] : (sP0(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl258),
  inference(superposition,[],[f16395,f476])).
fof(f23484,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) | $spl386),
  inference(cnf_transformation,[],[f23484_D])).
fof(f23484_D,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f23446,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,X26,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16394,f482])).
fof(f16394,plain,(
  ( ! [X21,X20] : (sP0(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) ) | $spl258),
  inference(resolution,[],[f15478,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f312])).
fof(f23445,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl258),
  inference(superposition,[],[f16394,f481])).
fof(f23444,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl258),
  inference(superposition,[],[f16394,f479])).
fof(f23443,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,X17,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16394,f478])).
fof(f23442,plain,(
  ( ! [X14,X15,X13] : (sP0(X13,X14,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl258),
  inference(superposition,[],[f16394,f477])).
fof(f23441,plain,(
  ( ! [X12,X10,X11] : (sP0(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl258),
  inference(superposition,[],[f16394,f476])).
fof(f23406,plain,(
  ( ! [X10,X11] : (sP1(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) | $spl384),
  inference(cnf_transformation,[],[f23406_D])).
fof(f23406_D,plain,(
  ( ! [X10,X11] : (sP1(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f23369,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16393,f482])).
fof(f16393,plain,(
  ( ! [X19,X18] : (sP1(X18,X19,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) ) | $spl258),
  inference(resolution,[],[f15478,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f310])).
fof(f310,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f309])).
fof(f309,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f23368,plain,(
  ( ! [X24,X23,X22] : (sP1(X22,X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl258),
  inference(superposition,[],[f16393,f481])).
fof(f23367,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl258),
  inference(superposition,[],[f16393,f479])).
fof(f23366,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,X17,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16393,f478])).
fof(f23365,plain,(
  ( ! [X14,X15,X13] : (sP1(X13,X14,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl258),
  inference(superposition,[],[f16393,f477])).
fof(f23364,plain,(
  ( ! [X12,X10,X11] : (sP1(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl258),
  inference(superposition,[],[f16393,f476])).
fof(f23331,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) | $spl382),
  inference(cnf_transformation,[],[f23331_D])).
fof(f23331_D,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f23294,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,X26,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16390,f482])).
fof(f16390,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) ) | $spl258),
  inference(resolution,[],[f15478,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f306])).
fof(f306,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f23293,plain,(
  ( ! [X24,X23,X22] : (sP2(X22,X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl258),
  inference(superposition,[],[f16390,f481])).
fof(f23292,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl258),
  inference(superposition,[],[f16390,f479])).
fof(f23291,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16390,f478])).
fof(f23290,plain,(
  ( ! [X14,X15,X13] : (sP2(X13,X14,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl258),
  inference(superposition,[],[f16390,f477])).
fof(f23289,plain,(
  ( ! [X12,X10,X11] : (sP2(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl258),
  inference(superposition,[],[f16390,f476])).
fof(f3699,plain,(
  ( ! [X5] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl72),
  inference(cnf_transformation,[],[f3699_D])).
fof(f3699_D,plain,(
  ( ! [X5] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f23250,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | $spl381),
  inference(cnf_transformation,[],[f23250_D])).
fof(f23250_D,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl381),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl381])])).
fof(f23216,plain,(
  ( ! [X5] : (s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl329),
  inference(superposition,[],[f21683,f477])).
fof(f3698,plain,(
  ( ! [X4] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl70),
  inference(cnf_transformation,[],[f3698_D])).
fof(f3698_D,plain,(
  ( ! [X4] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f23215,plain,(
  ( ! [X4] : (s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl329),
  inference(superposition,[],[f21683,f476])).
fof(f21683,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl329),
  inference(cnf_transformation,[],[f21683_D])).
fof(f21683_D,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl329),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl329])])).
fof(f23148,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl379),
  inference(cnf_transformation,[],[f23148_D])).
fof(f23148_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl379),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl379])])).
fof(f23146,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl3),
  inference(forward_demodulation,[],[f23145,f458])).
fof(f23145,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl3),
  inference(forward_demodulation,[],[f23144,f458])).
fof(f23144,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl3),
  inference(forward_demodulation,[],[f23049,f458])).
fof(f23049,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl3),
  inference(superposition,[],[f1174,f474])).
fof(f474,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f11])).
fof(f11,axiom,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.arithmetic.TWO')).
fof(f1174,plain,(
  ( ! [X0] : (s('type.num.num',sK10) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | $spl3),
  inference(forward_demodulation,[],[f1152,f458])).
fof(f1152,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK10))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) ) | $spl3),
  inference(superposition,[],[f761,f431])).
fof(f431,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X6,X8] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))) <=> s('type.num.num',X6) = s('type.num.num',X8))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_distrib|split|19')).
fof(f22217,plain,(
  ( ! [X18] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl362),
  inference(cnf_transformation,[],[f22217_D])).
fof(f22217_D,plain,(
  ( ! [X18] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f22719,plain,(
  ( ! [X10,X11] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl376),
  inference(cnf_transformation,[],[f22719_D])).
fof(f22719_D,plain,(
  ( ! [X10,X11] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f22681,plain,(
  ( ! [X26,X27,X25] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16244,f482])).
fof(f16244,plain,(
  ( ! [X78,X77] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')),X77,X78)) ) | $spl257),
  inference(resolution,[],[f15475,f530])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f342])).
fof(f15475,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | $spl257),
  inference(cnf_transformation,[],[f15475_D])).
fof(f15475_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f22218,plain,(
  ( ! [X21] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl364),
  inference(cnf_transformation,[],[f22218_D])).
fof(f22218_D,plain,(
  ( ! [X21] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f22680,plain,(
  ( ! [X24,X23,X22] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl257),
  inference(superposition,[],[f16244,f481])).
fof(f22679,plain,(
  ( ! [X21,X19,X20] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl257),
  inference(superposition,[],[f16244,f479])).
fof(f22678,plain,(
  ( ! [X17,X18,X16] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16244,f478])).
fof(f22216,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))))) ) | $spl360),
  inference(cnf_transformation,[],[f22216_D])).
fof(f22216_D,plain,(
  ( ! [X15] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f22677,plain,(
  ( ! [X14,X15,X13] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))))) ) | $spl257),
  inference(superposition,[],[f16244,f477])).
fof(f22214,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl356),
  inference(cnf_transformation,[],[f22214_D])).
fof(f22214_D,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f22676,plain,(
  ( ! [X12,X10,X11] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl257),
  inference(superposition,[],[f16244,f476])).
fof(f22644,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl374),
  inference(cnf_transformation,[],[f22644_D])).
fof(f22644_D,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f22606,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16243,f482])).
fof(f16243,plain,(
  ( ! [X76,X75] : (sP6(X75,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')),X76)) ) | $spl257),
  inference(resolution,[],[f15475,f529])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f342])).
fof(f22605,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl257),
  inference(superposition,[],[f16243,f481])).
fof(f22604,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl257),
  inference(superposition,[],[f16243,f479])).
fof(f22603,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16243,f478])).
fof(f22602,plain,(
  ( ! [X14,X15,X13] : (sP6(X13,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))))) ) | $spl257),
  inference(superposition,[],[f16243,f477])).
fof(f22601,plain,(
  ( ! [X12,X10,X11] : (sP6(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl257),
  inference(superposition,[],[f16243,f476])).
fof(f22569,plain,(
  ( ! [X10,X11] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl372),
  inference(cnf_transformation,[],[f22569_D])).
fof(f22569_D,plain,(
  ( ! [X10,X11] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f22530,plain,(
  ( ! [X26,X27,X25] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16230,f482])).
fof(f16230,plain,(
  ( ! [X50,X49] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')),X49,X50)) ) | $spl257),
  inference(resolution,[],[f15475,f513])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f335])).
fof(f22529,plain,(
  ( ! [X24,X23,X22] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl257),
  inference(superposition,[],[f16230,f481])).
fof(f22528,plain,(
  ( ! [X21,X19,X20] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl257),
  inference(superposition,[],[f16230,f479])).
fof(f22527,plain,(
  ( ! [X17,X18,X16] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16230,f478])).
fof(f22526,plain,(
  ( ! [X14,X15,X13] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))))) ) | $spl257),
  inference(superposition,[],[f16230,f477])).
fof(f22525,plain,(
  ( ! [X12,X10,X11] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl257),
  inference(superposition,[],[f16230,f476])).
fof(f22493,plain,(
  ( ! [X10,X11] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl370),
  inference(cnf_transformation,[],[f22493_D])).
fof(f22493_D,plain,(
  ( ! [X10,X11] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f22453,plain,(
  ( ! [X26,X27,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16212,f482])).
fof(f16212,plain,(
  ( ! [X17,X18] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')),X17,X18)) ) | $spl257),
  inference(resolution,[],[f15475,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f310])).
fof(f22452,plain,(
  ( ! [X24,X23,X22] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl257),
  inference(superposition,[],[f16212,f481])).
fof(f22451,plain,(
  ( ! [X21,X19,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl257),
  inference(superposition,[],[f16212,f479])).
fof(f22450,plain,(
  ( ! [X17,X18,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16212,f478])).
fof(f22449,plain,(
  ( ! [X14,X15,X13] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))))) ) | $spl257),
  inference(superposition,[],[f16212,f477])).
fof(f22448,plain,(
  ( ! [X12,X10,X11] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl257),
  inference(superposition,[],[f16212,f476])).
fof(f22417,plain,(
  ( ! [X10,X11] : (sP1(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl368),
  inference(cnf_transformation,[],[f22417_D])).
fof(f22417_D,plain,(
  ( ! [X10,X11] : (sP1(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f22378,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16211,f482])).
fof(f16211,plain,(
  ( ! [X15,X16] : (sP1(X15,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')),X16)) ) | $spl257),
  inference(resolution,[],[f15475,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f310])).
fof(f22377,plain,(
  ( ! [X24,X23,X22] : (sP1(X22,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl257),
  inference(superposition,[],[f16211,f481])).
fof(f22376,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl257),
  inference(superposition,[],[f16211,f479])).
fof(f22375,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16211,f478])).
fof(f22374,plain,(
  ( ! [X14,X15,X13] : (sP1(X13,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))))) ) | $spl257),
  inference(superposition,[],[f16211,f477])).
fof(f22373,plain,(
  ( ! [X12,X10,X11] : (sP1(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl257),
  inference(superposition,[],[f16211,f476])).
fof(f22342,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl366),
  inference(cnf_transformation,[],[f22342_D])).
fof(f22342_D,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f22302,plain,(
  ( ! [X26,X27,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16209,f482])).
fof(f16209,plain,(
  ( ! [X12,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')),X11,X12)) ) | $spl257),
  inference(resolution,[],[f15475,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f307])).
fof(f22301,plain,(
  ( ! [X24,X23,X22] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl257),
  inference(superposition,[],[f16209,f481])).
fof(f22300,plain,(
  ( ! [X21,X19,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl257),
  inference(superposition,[],[f16209,f479])).
fof(f22299,plain,(
  ( ! [X17,X18,X16] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16209,f478])).
fof(f22298,plain,(
  ( ! [X14,X15,X13] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))))) ) | $spl257),
  inference(superposition,[],[f16209,f477])).
fof(f22297,plain,(
  ( ! [X12,X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl257),
  inference(superposition,[],[f16209,f476])).
fof(f810,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl20),
  inference(cnf_transformation,[],[f810_D])).
fof(f810_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f772,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) | $spl16),
  inference(cnf_transformation,[],[f772_D])).
fof(f772_D,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f22254,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl358),
  inference(subsumption_resolution,[],[f22225,f496])).
fof(f496,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.arithmetic.ZERO_LESS_EQ')).
fof(f22225,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | $spl358),
  inference(resolution,[],[f22215,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f307])).
fof(f22215,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl358),
  inference(cnf_transformation,[],[f22215_D])).
fof(f22215_D,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f22171,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16208,f482])).
fof(f16208,plain,(
  ( ! [X10,X9] : (sP2(X9,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')),X10)) ) | $spl257),
  inference(resolution,[],[f15475,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f307])).
fof(f22170,plain,(
  ( ! [X24,X23,X22] : (sP2(X22,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) ) | $spl257),
  inference(superposition,[],[f16208,f481])).
fof(f22169,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) ) | $spl257),
  inference(superposition,[],[f16208,f479])).
fof(f22168,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10)))))))) ) | $spl257),
  inference(superposition,[],[f16208,f478])).
fof(f22167,plain,(
  ( ! [X14,X15,X13] : (sP2(X13,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))))) ) | $spl257),
  inference(superposition,[],[f16208,f477])).
fof(f22166,plain,(
  ( ! [X12,X10,X11] : (sP2(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl257),
  inference(superposition,[],[f16208,f476])).
fof(f22202,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0'))) | $spl353),
  inference(cnf_transformation,[],[f22202_D])).
fof(f22202_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0'))) <=> ~$spl353),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl353])])).
fof(f22143,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) | $spl257),
  inference(resolution,[],[f16208,f6146])).
fof(f6146,plain,(
  ( ! [X14,X15] : (~sP2(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f6145,f672])).
fof(f6145,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0')) )),
  inference(subsumption_resolution,[],[f6144,f6016])).
fof(f6016,plain,(
  ( ! [X4,X5] : (sP0(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f6011,f423])).
fof(f6011,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP0(X4,'F_0',X5)) )),
  inference(resolution,[],[f784,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (sP0(X2,X1,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f303])).
fof(f303,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f784,plain,(
  ( ! [X8,X7] : (sP3(X8,'F_0',X7) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8))) )),
  inference(superposition,[],[f774,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP3(X2,X1,X0)) & (~sP3(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP3(X2,X1,X0))),
  inference(definition_folding,[],[f150,f289,f288,f287,f286])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f144])).
fof(f144,axiom,(
  ! [X18,X19,X9] : ((p(s(bool,X9)) <=> s(bool,X18) = s(bool,X19)) <=> ((p(s(bool,X9)) | p(s(bool,X19)) | p(s(bool,X18))) & (p(s(bool,X9)) | ~p(s(bool,X18)) | ~p(s(bool,X19))) & (p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,X9))) & (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X9)))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.sat.sat_2|49|_prove_ltb')).
fof(f774,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f739,f679])).
fof(f739,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(superposition,[],[f679,f608])).
fof(f608,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.INT_LE_REDUCE|split|1',aHL_BOOL_CASES)).
fof(f6144,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0') | ~sP0(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f6137,f421])).
fof(f6137,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0') | ~sP1('F_0',X15,X14) | ~sP0(X15,'F_0',X14)) )),
  inference(resolution,[],[f786,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f304])).
fof(f786,plain,(
  ( ! [X12,X11] : (~sP3(X12,'F_0',X11) | ~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X11)) )),
  inference(superposition,[],[f774,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f313])).
fof(f22203,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) ) | $spl354),
  inference(cnf_transformation,[],[f22203_D])).
fof(f22203_D,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f22142,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) ) | $spl257),
  inference(resolution,[],[f16208,f6024])).
fof(f6024,plain,(
  ( ! [X4,X5] : (~sP2(X5,X4,'F_0') | sP3(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f6023,f672])).
fof(f6023,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f6020,f421])).
fof(f6020,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0') | ~sP1('F_0',X4,X5)) )),
  inference(resolution,[],[f6016,f412])).
fof(f22200,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0'))) | $spl350),
  inference(cnf_transformation,[],[f22200_D])).
fof(f22200_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0'))) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f22141,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) ) | $spl257),
  inference(resolution,[],[f16208,f6100])).
fof(f6100,plain,(
  ( ! [X2,X3] : (~sP2(X2,X3,'F_0') | ~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f6093])).
fof(f6093,plain,(
  ( ! [X2,X3] : (~sP2(X2,X3,'F_0') | ~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f6081,f428])).
fof(f6081,plain,(
  ( ! [X2,X3] : (sP3(X2,'F_0',X3) | ~sP2(X2,X3,'F_0') | ~p(s(bool,X2))) )),
  inference(resolution,[],[f6055,f423])).
fof(f6055,plain,(
  ( ! [X4,X5] : (~sP0(X5,'F_0',X4) | ~sP2(X4,X5,'F_0') | sP3(X4,'F_0',X5)) )),
  inference(subsumption_resolution,[],[f6053,f672])).
fof(f6053,plain,(
  ( ! [X4,X5] : (sP3(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP2(X4,X5,'F_0') | ~sP0(X5,'F_0',X4)) )),
  inference(resolution,[],[f6050,f412])).
fof(f6050,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5)) )),
  inference(subsumption_resolution,[],[f6044,f419])).
fof(f6044,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f6041,f417])).
fof(f6041,plain,(
  ( ! [X6,X7] : (~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f6037,f420])).
fof(f6037,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) )),
  inference(resolution,[],[f6031,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,X2,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f304])).
fof(f6031,plain,(
  ( ! [X2,X3] : (sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f6030,f605])).
fof(f605,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f378])).
fof(f378,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f6030,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f6026,f6016])).
fof(f6026,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0') | ~sP0(X2,'F_0',X3)) )),
  inference(resolution,[],[f6017,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f304])).
fof(f6017,plain,(
  ( ! [X6,X7] : (sP1('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6012,f419])).
fof(f6012,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP1('F_0',X6,X7)) )),
  inference(resolution,[],[f784,f408])).
fof(f21944,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl346),
  inference(cnf_transformation,[],[f21944_D])).
fof(f21944_D,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f16792,plain,(
  ( ! [X6] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X6)) ) | $spl286),
  inference(cnf_transformation,[],[f16792_D])).
fof(f16792_D,plain,(
  ( ! [X6] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X6)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f21898,plain,(
  ( ! [X17,X16] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X16) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl199),
  inference(superposition,[],[f21606,f482])).
fof(f21606,plain,(
  ( ! [X103] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)),X103)) ) | $spl199),
  inference(resolution,[],[f14670,f631])).
fof(f631,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f388])).
fof(f388,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(flattening,[],[f387])).
fof(f387,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(nnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1] : (sP9(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f14670,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | $spl199),
  inference(forward_demodulation,[],[f14529,f500])).
fof(f500,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_suc|split|1')).
fof(f14529,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | $spl199),
  inference(resolution,[],[f12987,f618])).
fof(f618,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X8,X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X6)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X8)))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.arithmetic.NOT_LEQ')).
fof(f21945,plain,(
  ( ! [X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) | $spl348),
  inference(cnf_transformation,[],[f21945_D])).
fof(f21945_D,plain,(
  ( ! [X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f21897,plain,(
  ( ! [X14,X15] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X14) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15)))))) ) | $spl199),
  inference(superposition,[],[f21606,f481])).
fof(f21896,plain,(
  ( ! [X12,X13] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X12) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) | $spl199),
  inference(superposition,[],[f21606,f479])).
fof(f21895,plain,(
  ( ! [X10,X11] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X10) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl199),
  inference(superposition,[],[f21606,f478])).
fof(f21943,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl344),
  inference(cnf_transformation,[],[f21943_D])).
fof(f21943_D,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f21894,plain,(
  ( ! [X8,X9] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X8) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl199),
  inference(superposition,[],[f21606,f477])).
fof(f21942,plain,(
  ( ! [X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) | $spl342),
  inference(cnf_transformation,[],[f21942_D])).
fof(f21942_D,plain,(
  ( ! [X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f21893,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X6) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) | $spl199),
  inference(superposition,[],[f21606,f476])).
fof(f21929,plain,(
  ( ! [X10] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X10)) ) | $spl340),
  inference(cnf_transformation,[],[f21929_D])).
fof(f21929_D,plain,(
  ( ! [X10] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X10)) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f21878,plain,(
  ( ! [X12,X13] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f21606,f477])).
fof(f21877,plain,(
  ( ! [X10,X11] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X10) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))))) ) | $spl199),
  inference(superposition,[],[f21606,f476])).
fof(f21854,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | $spl339),
  inference(cnf_transformation,[],[f21854_D])).
fof(f21854_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl339),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl339])])).
fof(f21814,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl337),
  inference(superposition,[],[f21797,f477])).
fof(f21813,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl337),
  inference(superposition,[],[f21797,f476])).
fof(f21795,plain,(
  ( ! [X4] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) ) | $spl334),
  inference(cnf_transformation,[],[f21795_D])).
fof(f21795_D,plain,(
  ( ! [X4] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f21797,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl337),
  inference(cnf_transformation,[],[f21797_D])).
fof(f21797_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl337),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl337])])).
fof(f21798,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | sP9(X5,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) ) | $spl199),
  inference(forward_demodulation,[],[f21746,f458])).
fof(f21746,plain,(
  ( ! [X5] : (sP9(X5,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl199),
  inference(superposition,[],[f21605,f447])).
fof(f21605,plain,(
  ( ! [X102] : (sP9(X102,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) ) | $spl199),
  inference(resolution,[],[f14670,f630])).
fof(f630,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f388])).
fof(f21794,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) ) | $spl199),
  inference(forward_demodulation,[],[f21745,f458])).
fof(f21745,plain,(
  ( ! [X4] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl199),
  inference(superposition,[],[f21605,f445])).
fof(f21783,plain,(
  ( ! [X10] : (sP9(X10,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) ) | $spl332),
  inference(cnf_transformation,[],[f21783_D])).
fof(f21783_D,plain,(
  ( ! [X10] : (sP9(X10,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f21732,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f21605,f477])).
fof(f21731,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))))) ) | $spl199),
  inference(superposition,[],[f21605,f476])).
fof(f21694,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | $spl330),
  inference(cnf_transformation,[],[f21694_D])).
fof(f21694_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f21620,plain,(
  ( ! [X5] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f14670,f477])).
fof(f21619,plain,(
  ( ! [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl199),
  inference(superposition,[],[f14670,f476])).
fof(f21682,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl327),
  inference(cnf_transformation,[],[f21682_D])).
fof(f21682_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl327),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl327])])).
fof(f21684,plain,(
  s('type.num.num',sK10) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl328),
  inference(cnf_transformation,[],[f21684_D])).
fof(f21684_D,plain,(
  s('type.num.num',sK10) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f21538,plain,(
  s('type.num.num',sK10) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl199),
  inference(resolution,[],[f14670,f642])).
fof(f642,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f393])).
fof(f393,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(flattening,[],[f392])).
fof(f392,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X8,X6] : (s('type.num.num',X6) = s('type.num.num',X8) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X6)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X8))))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.arithmetic.EQ_LESS_EQ')).
fof(f21537,plain,(
  s('type.num.num',sK10) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl199),
  inference(resolution,[],[f14670,f642])).
fof(f21524,plain,(
  ( ! [X17,X16] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17)) ) | $spl324),
  inference(cnf_transformation,[],[f21524_D])).
fof(f21524_D,plain,(
  ( ! [X17,X16] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17)) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f21486,plain,(
  ( ! [X21,X19,X20] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X19,X20) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f14571,f477])).
fof(f14571,plain,(
  ( ! [X78,X77] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X77,X78)) ) | $spl199),
  inference(resolution,[],[f12987,f530])).
fof(f21485,plain,(
  ( ! [X17,X18,X16] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))))))) ) | $spl199),
  inference(superposition,[],[f14571,f476])).
fof(f21445,plain,(
  ( ! [X17,X16] : (sP6(X16,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X17)) ) | $spl322),
  inference(cnf_transformation,[],[f21445_D])).
fof(f21445_D,plain,(
  ( ! [X17,X16] : (sP6(X16,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X17)) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f21407,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X20) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f14570,f477])).
fof(f14570,plain,(
  ( ! [X76,X75] : (sP6(X75,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X76)) ) | $spl199),
  inference(resolution,[],[f12987,f529])).
fof(f21406,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X17) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))))))) ) | $spl199),
  inference(superposition,[],[f14570,f476])).
fof(f21323,plain,(
  ( ! [X17,X16] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17)) ) | $spl320),
  inference(cnf_transformation,[],[f21323_D])).
fof(f21323_D,plain,(
  ( ! [X17,X16] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17)) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f21284,plain,(
  ( ! [X21,X19,X20] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X19,X20) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f14557,f477])).
fof(f14557,plain,(
  ( ! [X50,X49] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X49,X50)) ) | $spl199),
  inference(resolution,[],[f12987,f513])).
fof(f21283,plain,(
  ( ! [X17,X18,X16] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))))))) ) | $spl199),
  inference(superposition,[],[f14557,f476])).
fof(f21243,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17)) ) | $spl318),
  inference(cnf_transformation,[],[f21243_D])).
fof(f21243_D,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17)) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f21203,plain,(
  ( ! [X21,X19,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X19,X20) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f14539,f477])).
fof(f14539,plain,(
  ( ! [X17,X18] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X17,X18)) ) | $spl199),
  inference(resolution,[],[f12987,f421])).
fof(f21202,plain,(
  ( ! [X17,X18,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))))))) ) | $spl199),
  inference(superposition,[],[f14539,f476])).
fof(f21163,plain,(
  ( ! [X17,X16] : (sP1(X16,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X17)) ) | $spl316),
  inference(cnf_transformation,[],[f21163_D])).
fof(f21163_D,plain,(
  ( ! [X17,X16] : (sP1(X16,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X17)) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f21124,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X20) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f14538,f477])).
fof(f14538,plain,(
  ( ! [X15,X16] : (sP1(X15,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X16)) ) | $spl199),
  inference(resolution,[],[f12987,f420])).
fof(f21123,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X17) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))))))) ) | $spl199),
  inference(superposition,[],[f14538,f476])).
fof(f21084,plain,(
  ( ! [X17,X16] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17)) ) | $spl314),
  inference(cnf_transformation,[],[f21084_D])).
fof(f21084_D,plain,(
  ( ! [X17,X16] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f21044,plain,(
  ( ! [X21,X19,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X19,X20) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f14536,f477])).
fof(f14536,plain,(
  ( ! [X12,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X11,X12)) ) | $spl199),
  inference(resolution,[],[f12987,f417])).
fof(f21043,plain,(
  ( ! [X17,X18,X16] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,X17) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))))))) ) | $spl199),
  inference(superposition,[],[f14536,f476])).
fof(f18969,plain,(
  ( ! [X17,X16] : (sP2(X16,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X17)) ) | $spl296),
  inference(cnf_transformation,[],[f18969_D])).
fof(f18969_D,plain,(
  ( ! [X17,X16] : (sP2(X16,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X17)) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f20962,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X20) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f14535,f477])).
fof(f14535,plain,(
  ( ! [X10,X9] : (sP2(X9,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X10)) ) | $spl199),
  inference(resolution,[],[f12987,f416])).
fof(f20961,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X17) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))))))) ) | $spl199),
  inference(superposition,[],[f14535,f476])).
fof(f18957,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl293),
  inference(cnf_transformation,[],[f18957_D])).
fof(f18957_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl293),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl293])])).
fof(f20936,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,X10))) ) | $spl199),
  inference(resolution,[],[f14535,f6146])).
fof(f18958,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl294),
  inference(cnf_transformation,[],[f18958_D])).
fof(f18958_D,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f20935,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl199),
  inference(resolution,[],[f14535,f6024])).
fof(f18955,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl290),
  inference(cnf_transformation,[],[f18955_D])).
fof(f18955_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f20934,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl199),
  inference(resolution,[],[f14535,f6100])).
fof(f14686,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | $spl243),
  inference(cnf_transformation,[],[f14686_D])).
fof(f14686_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f20807,plain,(
  ( ! [X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f12987,f477])).
fof(f20806,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl199),
  inference(superposition,[],[f12987,f476])).
fof(f10305,plain,(
  ( ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl156),
  inference(cnf_transformation,[],[f10305_D])).
fof(f10305_D,plain,(
  ( ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f20659,plain,(
  ( ! [X17,X16] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X16) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl196),
  inference(superposition,[],[f20436,f482])).
fof(f20436,plain,(
  ( ! [X103] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)),X103)) ) | $spl196),
  inference(resolution,[],[f12984,f631])).
fof(f10129,plain,(
  ( ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76)))))) ) | $spl150),
  inference(cnf_transformation,[],[f10129_D])).
fof(f10129_D,plain,(
  ( ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76)))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f20658,plain,(
  ( ! [X14,X15] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X14) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15)))))) ) | $spl196),
  inference(superposition,[],[f20436,f481])).
fof(f20657,plain,(
  ( ! [X12,X13] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X12) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) | $spl196),
  inference(superposition,[],[f20436,f479])).
fof(f20656,plain,(
  ( ! [X10,X11] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X10) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl196),
  inference(superposition,[],[f20436,f478])).
fof(f11439,plain,(
  ( ! [X76] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl168),
  inference(cnf_transformation,[],[f11439_D])).
fof(f11439_D,plain,(
  ( ! [X76] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f20655,plain,(
  ( ! [X8,X9] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X8) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl196),
  inference(superposition,[],[f20436,f477])).
fof(f11263,plain,(
  ( ! [X76] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))))))) ) | $spl162),
  inference(cnf_transformation,[],[f11263_D])).
fof(f11263_D,plain,(
  ( ! [X76] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f20654,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X6) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) | $spl196),
  inference(superposition,[],[f20436,f476])).
fof(f20690,plain,(
  ( ! [X10] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X10)) ) | $spl312),
  inference(cnf_transformation,[],[f20690_D])).
fof(f20690_D,plain,(
  ( ! [X10] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X10)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f20639,plain,(
  ( ! [X12,X13] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl196),
  inference(superposition,[],[f20436,f477])).
fof(f20638,plain,(
  ( ! [X10,X11] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X10) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))))) ) | $spl196),
  inference(superposition,[],[f20436,f476])).
fof(f20607,plain,(
  ( ! [X10] : (sP9(X10,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) ) | $spl310),
  inference(cnf_transformation,[],[f20607_D])).
fof(f20607_D,plain,(
  ( ! [X10] : (sP9(X10,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f20556,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl196),
  inference(superposition,[],[f20435,f477])).
fof(f20435,plain,(
  ( ! [X102] : (sP9(X102,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) ) | $spl196),
  inference(resolution,[],[f12984,f630])).
fof(f20555,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))))) ) | $spl196),
  inference(superposition,[],[f20435,f476])).
fof(f19493,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | $spl308),
  inference(cnf_transformation,[],[f19493_D])).
fof(f19493_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f20450,plain,(
  ( ! [X5] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl196),
  inference(superposition,[],[f12984,f477])).
fof(f20449,plain,(
  ( ! [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl196),
  inference(superposition,[],[f12984,f476])).
fof(f14895,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | $spl247),
  inference(cnf_transformation,[],[f14895_D])).
fof(f14895_D,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl247),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl247])])).
fof(f20326,plain,(
  ( ! [X5] : (s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f14526,f477])).
fof(f14526,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl199),
  inference(resolution,[],[f12987,f640])).
fof(f640,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f393])).
fof(f20325,plain,(
  ( ! [X4] : (s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl199),
  inference(superposition,[],[f14526,f476])).
fof(f20242,plain,(
  $false | ($spl197 | $spl201)),
  inference(subsumption_resolution,[],[f20241,f12985])).
fof(f20241,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | $spl201),
  inference(forward_demodulation,[],[f20240,f461])).
fof(f461,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(cnf_transformation,[],[f75])).
fof(f75,axiom,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_suc|split|0')).
fof(f20240,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | $spl201),
  inference(subsumption_resolution,[],[f20035,f660])).
fof(f660,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f403])).
fof(f403,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f402])).
fof(f402,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) & p(s(bool,X10))) <=> p(s(bool,'F_0')))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.AND_CLAUSES|split|2')).
fof(f20035,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | $spl201),
  inference(superposition,[],[f618,f15200])).
fof(f15200,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl201),
  inference(resolution,[],[f13048,f646])).
fof(f19529,plain,(
  ( ! [X5] : (s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl197),
  inference(superposition,[],[f19339,f477])).
fof(f19339,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl197),
  inference(resolution,[],[f12985,f640])).
fof(f19528,plain,(
  ( ! [X4] : (s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl197),
  inference(superposition,[],[f19339,f476])).
fof(f19494,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | $spl309),
  inference(cnf_transformation,[],[f19494_D])).
fof(f19494_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) <=> ~$spl309),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl309])])).
fof(f19421,plain,(
  ( ! [X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl197),
  inference(superposition,[],[f12985,f477])).
fof(f19420,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl197),
  inference(superposition,[],[f12985,f476])).
fof(f19336,plain,(
  ( ! [X6] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')),X6)) ) | $spl306),
  inference(cnf_transformation,[],[f19336_D])).
fof(f19336_D,plain,(
  ( ! [X6] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')),X6)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f19304,plain,(
  ( ! [X17,X16] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')),X16) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl198),
  inference(superposition,[],[f19051,f482])).
fof(f19051,plain,(
  ( ! [X103] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X103)) ) | $spl198),
  inference(resolution,[],[f12986,f631])).
fof(f12986,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl198),
  inference(cnf_transformation,[],[f12986_D])).
fof(f12986_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f19303,plain,(
  ( ! [X14,X15] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')),X14) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15)))))) ) | $spl198),
  inference(superposition,[],[f19051,f481])).
fof(f19302,plain,(
  ( ! [X12,X13] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')),X12) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) | $spl198),
  inference(superposition,[],[f19051,f479])).
fof(f19301,plain,(
  ( ! [X10,X11] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')),X10) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl198),
  inference(superposition,[],[f19051,f478])).
fof(f19300,plain,(
  ( ! [X8,X9] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl198),
  inference(superposition,[],[f19051,f477])).
fof(f19299,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')),X6) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) | $spl198),
  inference(superposition,[],[f19051,f476])).
fof(f19323,plain,(
  ( ! [X10] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X10)) ) | $spl304),
  inference(cnf_transformation,[],[f19323_D])).
fof(f19323_D,plain,(
  ( ! [X10] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X10)) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f19284,plain,(
  ( ! [X12,X13] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl198),
  inference(superposition,[],[f19051,f477])).
fof(f19283,plain,(
  ( ! [X10,X11] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X10) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))))) ) | $spl198),
  inference(superposition,[],[f19051,f476])).
fof(f19252,plain,(
  ( ! [X6] : (sP9(X6,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) ) | $spl302),
  inference(cnf_transformation,[],[f19252_D])).
fof(f19252_D,plain,(
  ( ! [X6] : (sP9(X6,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f19220,plain,(
  ( ! [X17,X16] : (sP9(X16,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl198),
  inference(superposition,[],[f19050,f482])).
fof(f19050,plain,(
  ( ! [X102] : (sP9(X102,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl198),
  inference(resolution,[],[f12986,f630])).
fof(f19219,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15)))))) ) | $spl198),
  inference(superposition,[],[f19050,f481])).
fof(f19218,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) | $spl198),
  inference(superposition,[],[f19050,f479])).
fof(f19217,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl198),
  inference(superposition,[],[f19050,f478])).
fof(f19216,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl198),
  inference(superposition,[],[f19050,f477])).
fof(f19215,plain,(
  ( ! [X6,X7] : (sP9(X6,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) | $spl198),
  inference(superposition,[],[f19050,f476])).
fof(f19239,plain,(
  ( ! [X10] : (sP9(X10,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl300),
  inference(cnf_transformation,[],[f19239_D])).
fof(f19239_D,plain,(
  ( ! [X10] : (sP9(X10,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f19200,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl198),
  inference(superposition,[],[f19050,f477])).
fof(f19199,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))))) ) | $spl198),
  inference(superposition,[],[f19050,f476])).
fof(f14685,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | $spl242),
  inference(cnf_transformation,[],[f14685_D])).
fof(f14685_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f19081,plain,(
  ( ! [X5] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl198),
  inference(superposition,[],[f12986,f477])).
fof(f19080,plain,(
  ( ! [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl198),
  inference(superposition,[],[f12986,f476])).
fof(f12985,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | $spl197),
  inference(cnf_transformation,[],[f12985_D])).
fof(f12985_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f19127,plain,(
  s('type.num.num',sK10) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl298),
  inference(cnf_transformation,[],[f19127_D])).
fof(f19127_D,plain,(
  s('type.num.num',sK10) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f18983,plain,(
  s('type.num.num',sK10) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | $spl198),
  inference(resolution,[],[f12986,f642])).
fof(f18982,plain,(
  s('type.num.num',sK10) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | $spl198),
  inference(resolution,[],[f12986,f642])).
fof(f18926,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X20) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f14535,f477])).
fof(f18925,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X17) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))))))) ) | $spl199),
  inference(superposition,[],[f14535,f476])).
fof(f18900,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,X10))) ) | $spl199),
  inference(resolution,[],[f14535,f6146])).
fof(f18899,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl199),
  inference(resolution,[],[f14535,f6024])).
fof(f18898,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl199),
  inference(resolution,[],[f14535,f6100])).
fof(f17980,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) ) | $spl140),
  inference(subsumption_resolution,[],[f17971,f485])).
fof(f485,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X8,X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))) <=> s('type.num.num',X6) = s('type.num.num',X8))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.integer.INT_INJ')).
fof(f17971,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) ) | $spl140),
  inference(superposition,[],[f10117,f480])).
fof(f480,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f325])).
fof(f17979,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl140),
  inference(subsumption_resolution,[],[f17970,f485])).
fof(f17970,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl140),
  inference(superposition,[],[f10117,f483])).
fof(f483,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f327])).
fof(f17978,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) | $spl288),
  inference(cnf_transformation,[],[f17978_D])).
fof(f17978_D,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f17965,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl140),
  inference(superposition,[],[f10117,f480])).
fof(f17964,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl140),
  inference(superposition,[],[f10117,f483])).
fof(f11756,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl176),
  inference(cnf_transformation,[],[f11756_D])).
fof(f11756_D,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f17752,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) ) | $spl76),
  inference(superposition,[],[f3701,f480])).
fof(f17751,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl76),
  inference(superposition,[],[f3701,f483])).
fof(f11755,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) | $spl174),
  inference(cnf_transformation,[],[f11755_D])).
fof(f11755_D,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f17746,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl76),
  inference(superposition,[],[f3701,f480])).
fof(f17745,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl76),
  inference(superposition,[],[f3701,f483])).
fof(f16755,plain,(
  ( ! [X17,X16] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X16) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16448,f482])).
fof(f16448,plain,(
  ( ! [X103] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)),X103)) ) | $spl258),
  inference(resolution,[],[f15478,f631])).
fof(f16754,plain,(
  ( ! [X14,X15] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X14) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15)))))) ) | $spl258),
  inference(superposition,[],[f16448,f481])).
fof(f16753,plain,(
  ( ! [X12,X13] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X12) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) | $spl258),
  inference(superposition,[],[f16448,f479])).
fof(f16752,plain,(
  ( ! [X10,X11] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X10) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16448,f478])).
fof(f16751,plain,(
  ( ! [X8,X9] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X8) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl258),
  inference(superposition,[],[f16448,f477])).
fof(f16750,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X6) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) | $spl258),
  inference(superposition,[],[f16448,f476])).
fof(f16662,plain,(
  ( ! [X6] : (sP9(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) | $spl278),
  inference(cnf_transformation,[],[f16662_D])).
fof(f16662_D,plain,(
  ( ! [X6] : (sP9(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f16624,plain,(
  ( ! [X17,X16] : (sP9(X16,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16447,f482])).
fof(f16447,plain,(
  ( ! [X102] : (sP9(X102,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) ) | $spl258),
  inference(resolution,[],[f15478,f630])).
fof(f16623,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15)))))) ) | $spl258),
  inference(superposition,[],[f16447,f481])).
fof(f16622,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) | $spl258),
  inference(superposition,[],[f16447,f479])).
fof(f16621,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl258),
  inference(superposition,[],[f16447,f478])).
fof(f16620,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl258),
  inference(superposition,[],[f16447,f477])).
fof(f16619,plain,(
  ( ! [X6,X7] : (sP9(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) | $spl258),
  inference(superposition,[],[f16447,f476])).
fof(f16170,plain,(
  $false | $spl256),
  inference(subsumption_resolution,[],[f16023,f496])).
fof(f16023,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | $spl256),
  inference(resolution,[],[f15476,f619])).
fof(f619,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f382])).
fof(f14951,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl251),
  inference(cnf_transformation,[],[f14951_D])).
fof(f14951_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl251),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl251])])).
fof(f15623,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,X10))) ) | $spl201),
  inference(resolution,[],[f14713,f6146])).
fof(f14713,plain,(
  ( ! [X10,X9] : (sP2(X9,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')),X10)) ) | $spl201),
  inference(resolution,[],[f13048,f416])).
fof(f14952,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl252),
  inference(cnf_transformation,[],[f14952_D])).
fof(f14952_D,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f15622,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl201),
  inference(resolution,[],[f14713,f6024])).
fof(f14949,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl248),
  inference(cnf_transformation,[],[f14949_D])).
fof(f14949_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f15621,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl201),
  inference(resolution,[],[f14713,f6100])).
fof(f15615,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num',sK10)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num',X12))))) ) | $spl274),
  inference(cnf_transformation,[],[f15615_D])).
fof(f15615_D,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num',sK10)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num',X12))))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f11310,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,X18)) ) | $spl166),
  inference(cnf_transformation,[],[f11310_D])).
fof(f11310_D,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,X18)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f15616,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))))))) )),
  inference(subsumption_resolution,[],[f15443,f604])).
fof(f604,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f378])).
fof(f15443,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14)))))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f718,f437])).
fof(f437,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1))))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X12,X13] : (s(bool,X12) = s(bool,X13) <=> ((p(s(bool,X12)) & p(s(bool,X13))) | (~p(s(bool,X12)) & ~p(s(bool,X13)))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.EQ_EXPAND')).
fof(f718,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num',X10))))) )),
  inference(superposition,[],[f679,f620])).
fof(f620,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f384])).
fof(f384,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(flattening,[],[f383])).
fof(f383,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) = s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0,X1] : (~s('type.num.num',X0) = s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X8,X6] : (~s('type.num.num',X6) = s('type.num.num',X8) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X6)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X8))))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.arithmetic.NOT_NUM_EQ')).
fof(f15614,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(subsumption_resolution,[],[f15442,f604])).
fof(f15442,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num',X12)))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(superposition,[],[f718,f437])).
fof(f15476,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | $spl256),
  inference(cnf_transformation,[],[f15476_D])).
fof(f15476_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f15478,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | $spl258),
  inference(cnf_transformation,[],[f15478_D])).
fof(f15478_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f15457,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | $spl6),
  inference(trivial_inequality_removal,[],[f15437])).
fof(f15437,plain,(
  s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | $spl6),
  inference(superposition,[],[f718,f764])).
fof(f764,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(cnf_transformation,[],[f764_D])).
fof(f764_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f15613,plain,(
  ( ! [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num',X4))))) ) | $spl272),
  inference(cnf_transformation,[],[f15613_D])).
fof(f15613_D,plain,(
  ( ! [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num',X4))))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f9996,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X16))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X16)) ) | $spl136),
  inference(cnf_transformation,[],[f9996_D])).
fof(f9996_D,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X16))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X16)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f15431,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num',X6)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(superposition,[],[f718,f480])).
fof(f15430,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num',X4)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) )),
  inference(superposition,[],[f718,f483])).
fof(f15612,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | $spl6),
  inference(forward_demodulation,[],[f15611,f491])).
fof(f15611,plain,(
  ( ! [X91] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X91))))))) ) | $spl6),
  inference(forward_demodulation,[],[f15610,f491])).
fof(f15610,plain,(
  ( ! [X91] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X91))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X91))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f15427,f764])).
fof(f15427,plain,(
  ( ! [X91] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X91))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X91))))))) )),
  inference(superposition,[],[f718,f491])).
fof(f15609,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | $spl6),
  inference(forward_demodulation,[],[f15608,f491])).
fof(f15608,plain,(
  ( ! [X90] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X90))))))) ) | $spl6),
  inference(forward_demodulation,[],[f15607,f491])).
fof(f15607,plain,(
  ( ! [X90] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X90))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X90))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f15426,f764])).
fof(f15426,plain,(
  ( ! [X90] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X90))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X90))))))) )),
  inference(superposition,[],[f718,f494])).
fof(f15606,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | $spl6),
  inference(forward_demodulation,[],[f15605,f490])).
fof(f15605,plain,(
  ( ! [X89] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0'))))))) ) | $spl6),
  inference(forward_demodulation,[],[f15604,f490])).
fof(f15604,plain,(
  ( ! [X89] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0'))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0'))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f15425,f764])).
fof(f15425,plain,(
  ( ! [X89] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0'))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X89),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f718,f490])).
fof(f15603,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | $spl6),
  inference(forward_demodulation,[],[f15602,f490])).
fof(f15602,plain,(
  ( ! [X88] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X88),s('type.num.num','const.num.|48|_0'))))))) ) | $spl6),
  inference(forward_demodulation,[],[f15601,f490])).
fof(f15601,plain,(
  ( ! [X88] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X88),s('type.num.num','const.num.|48|_0'))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X88),s('type.num.num','const.num.|48|_0'))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f15424,f764])).
fof(f15424,plain,(
  ( ! [X88] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X88),s('type.num.num','const.num.|48|_0'))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X88),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f718,f493])).
fof(f15586,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl270),
  inference(cnf_transformation,[],[f15586_D])).
fof(f15586_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f12984,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | $spl196),
  inference(cnf_transformation,[],[f12984_D])).
fof(f12984_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f9350,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4)))))) ) | $spl130),
  inference(cnf_transformation,[],[f9350_D])).
fof(f9350_D,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4)))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f15588,plain,(
  ( ! [X82] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X82) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X82))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f15587,f458])).
fof(f15587,plain,(
  ( ! [X82] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X82))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X82)))) )),
  inference(forward_demodulation,[],[f15419,f461])).
fof(f15419,plain,(
  ( ! [X82] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X82))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X82)))) )),
  inference(superposition,[],[f718,f445])).
fof(f15584,plain,(
  ( ! [X81] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X81) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X81))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f15583,f458])).
fof(f15583,plain,(
  ( ! [X81] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X81))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) )),
  inference(forward_demodulation,[],[f15418,f461])).
fof(f15418,plain,(
  ( ! [X81] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X81))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) )),
  inference(superposition,[],[f718,f447])).
fof(f5645,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl111),
  inference(cnf_transformation,[],[f5645_D])).
fof(f5645_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f15503,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | $spl260),
  inference(cnf_transformation,[],[f15503_D])).
fof(f15503_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f15505,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl262),
  inference(cnf_transformation,[],[f15505_D])).
fof(f15505_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f15553,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f15552,f461])).
fof(f15552,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(forward_demodulation,[],[f15551,f500])).
fof(f15551,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(forward_demodulation,[],[f15403,f461])).
fof(f15403,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(superposition,[],[f718,f461])).
fof(f15527,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl265),
  inference(cnf_transformation,[],[f15527_D])).
fof(f15527_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl265),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl265])])).
fof(f15529,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | $spl266),
  inference(cnf_transformation,[],[f15529_D])).
fof(f15529_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f15531,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl268),
  inference(cnf_transformation,[],[f15531_D])).
fof(f15531_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f15525,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f15524,f500])).
fof(f15524,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f15523,f458])).
fof(f15523,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f15522,f453])).
fof(f453,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_distrib|split|14')).
fof(f15522,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f15521,f461])).
fof(f15521,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f15520,f499])).
fof(f499,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_suc|split|2')).
fof(f15520,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f15519,f500])).
fof(f15519,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f15518,f458])).
fof(f15518,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f15517,f453])).
fof(f15517,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f15516,f453])).
fof(f15516,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f15395,f458])).
fof(f15395,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(superposition,[],[f718,f474])).
fof(f12983,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl195),
  inference(cnf_transformation,[],[f12983_D])).
fof(f12983_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl195),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl195])])).
fof(f15509,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f15508,f458])).
fof(f15508,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f15507,f500])).
fof(f15507,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f15506,f458])).
fof(f15506,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f15393,f453])).
fof(f15393,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(superposition,[],[f718,f459])).
fof(f459,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f10])).
fof(f10,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.arithmetic.ONE')).
fof(f15501,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f15500,f458])).
fof(f15500,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f15499,f500])).
fof(f15499,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f15498,f458])).
fof(f15498,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f15392,f453])).
fof(f15392,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(superposition,[],[f718,f460])).
fof(f460,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f91])).
fof(f91,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_distrib|split|13')).
fof(f12965,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X41))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))))))) ) | $spl192),
  inference(cnf_transformation,[],[f12965_D])).
fof(f12965_D,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X41))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f15386,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X41))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))))))) )),
  inference(superposition,[],[f718,f477])).
fof(f12963,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X39))))))) ) | $spl190),
  inference(cnf_transformation,[],[f12963_D])).
fof(f12963_D,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X39))))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f15385,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X39))))))) )),
  inference(superposition,[],[f718,f476])).
fof(f12957,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X31)))))) ) | $spl186),
  inference(cnf_transformation,[],[f12957_D])).
fof(f12957_D,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X31)))))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f15384,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X36))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X36))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X36))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X37)))))) )),
  inference(superposition,[],[f718,f479])).
fof(f12959,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X33))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32)))))) ) | $spl188),
  inference(cnf_transformation,[],[f12959_D])).
fof(f12959_D,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X33))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32)))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f15383,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X34))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X34))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X35))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X34)))))) )),
  inference(superposition,[],[f718,f478])).
fof(f15382,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X33))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32)))))) )),
  inference(superposition,[],[f718,f482])).
fof(f15381,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK10)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK10))),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X31)))))) )),
  inference(superposition,[],[f718,f481])).
fof(f3701,plain,(
  ( ! [X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))) ) | $spl76),
  inference(cnf_transformation,[],[f3701_D])).
fof(f3701_D,plain,(
  ( ! [X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f15330,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | $spl255),
  inference(cnf_transformation,[],[f15330_D])).
fof(f15330_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f15293,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) ) | $spl245),
  inference(superposition,[],[f14699,f481])).
fof(f15292,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))) ) | $spl245),
  inference(superposition,[],[f14699,f479])).
fof(f15290,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl245),
  inference(superposition,[],[f14699,f477])).
fof(f15289,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl245),
  inference(superposition,[],[f14699,f476])).
fof(f15102,plain,(
  $false | ($spl29 | $spl200)),
  inference(subsumption_resolution,[],[f15101,f615])).
fof(f615,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_LE_REDUCE|split|1',aHL_TRUTH)).
fof(f15101,plain,(
  ~p(s(bool,'T_0')) | ($spl29 | $spl200)),
  inference(forward_demodulation,[],[f15100,f609])).
fof(f609,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X6)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_lte|split|0')).
fof(f15100,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) | ($spl29 | $spl200)),
  inference(subsumption_resolution,[],[f14970,f1371])).
fof(f1371,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',sK10) | $spl29),
  inference(trivial_inequality_removal,[],[f1346])).
fof(f1346,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',sK10) | $spl29),
  inference(superposition,[],[f1177,f446])).
fof(f446,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f318])).
fof(f14970,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',sK10) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) | $spl200),
  inference(resolution,[],[f13047,f642])).
fof(f13047,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl200),
  inference(cnf_transformation,[],[f13047_D])).
fof(f13047_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f15099,plain,(
  $false | ($spl29 | $spl200)),
  inference(subsumption_resolution,[],[f15098,f615])).
fof(f15098,plain,(
  ~p(s(bool,'T_0')) | ($spl29 | $spl200)),
  inference(forward_demodulation,[],[f15097,f609])).
fof(f15097,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) | ($spl29 | $spl200)),
  inference(subsumption_resolution,[],[f14969,f1371])).
fof(f14969,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',sK10) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) | $spl200),
  inference(resolution,[],[f13047,f642])).
fof(f14910,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,X10))) ) | $spl201),
  inference(resolution,[],[f14713,f6146])).
fof(f14909,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl201),
  inference(resolution,[],[f14713,f6024])).
fof(f14908,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl201),
  inference(resolution,[],[f14713,f6100])).
fof(f14864,plain,(
  ( ! [X8] : (s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) ) | $spl199),
  inference(superposition,[],[f14526,f481])).
fof(f14863,plain,(
  ( ! [X7] : (s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))) ) | $spl199),
  inference(superposition,[],[f14526,f479])).
fof(f14861,plain,(
  ( ! [X5] : (s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f14526,f477])).
fof(f14860,plain,(
  ( ! [X4] : (s('type.num.num',sK10) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl199),
  inference(superposition,[],[f14526,f476])).
fof(f13048,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl201),
  inference(cnf_transformation,[],[f13048_D])).
fof(f13048_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f14699,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl245),
  inference(cnf_transformation,[],[f14699_D])).
fof(f14699_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl245),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl245])])).
fof(f14700,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl199),
  inference(forward_demodulation,[],[f14638,f458])).
fof(f14638,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl199),
  inference(superposition,[],[f12987,f447])).
fof(f14697,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl199),
  inference(forward_demodulation,[],[f14637,f458])).
fof(f14637,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl199),
  inference(superposition,[],[f12987,f445])).
fof(f14627,plain,(
  ( ! [X8] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) ) | $spl199),
  inference(superposition,[],[f12987,f481])).
fof(f14626,plain,(
  ( ! [X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))) ) | $spl199),
  inference(superposition,[],[f12987,f479])).
fof(f14624,plain,(
  ( ! [X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl199),
  inference(superposition,[],[f12987,f477])).
fof(f14623,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl199),
  inference(superposition,[],[f12987,f476])).
fof(f14524,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X12))))) ) | $spl240),
  inference(cnf_transformation,[],[f14524_D])).
fof(f14524_D,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X12))))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f14525,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f14370,f604])).
fof(f14370,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f696,f437])).
fof(f696,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X22))))) )),
  inference(superposition,[],[f679,f620])).
fof(f14523,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f14369,f604])).
fof(f14369,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X12)))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f696,f437])).
fof(f7913,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) | $spl118),
  inference(cnf_transformation,[],[f7913_D])).
fof(f7913_D,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f14518,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(subsumption_resolution,[],[f14517,f622])).
fof(f622,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f384])).
fof(f14517,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(subsumption_resolution,[],[f14337,f622])).
fof(f14337,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(superposition,[],[f696,f480])).
fof(f14516,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) )),
  inference(subsumption_resolution,[],[f14515,f622])).
fof(f14515,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X4)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) )),
  inference(subsumption_resolution,[],[f14336,f622])).
fof(f14336,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X4)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) )),
  inference(superposition,[],[f696,f483])).
fof(f5470,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | $spl103),
  inference(cnf_transformation,[],[f5470_D])).
fof(f5470_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f14469,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.|48|_0')))) | $spl238),
  inference(cnf_transformation,[],[f14469_D])).
fof(f14469_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f14418,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl232),
  inference(cnf_transformation,[],[f14418_D])).
fof(f14418_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f14467,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(forward_demodulation,[],[f14466,f461])).
fof(f14466,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(forward_demodulation,[],[f14465,f500])).
fof(f14465,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(forward_demodulation,[],[f14309,f461])).
fof(f14309,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(superposition,[],[f696,f461])).
fof(f14443,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | $spl235),
  inference(cnf_transformation,[],[f14443_D])).
fof(f14443_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f14445,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl236),
  inference(cnf_transformation,[],[f14445_D])).
fof(f14445_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f14441,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(forward_demodulation,[],[f14440,f500])).
fof(f14440,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(forward_demodulation,[],[f14439,f458])).
fof(f14439,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(forward_demodulation,[],[f14438,f453])).
fof(f14438,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(subsumption_resolution,[],[f14437,f649])).
fof(f649,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f396])).
fof(f396,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) | p(s(bool,X10))) <=> p(s(bool,X10)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.OR_CLAUSES|split|2')).
fof(f14437,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f14436,f670])).
fof(f670,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f133])).
fof(f133,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_lte|split|1')).
fof(f14436,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f14435,f499])).
fof(f14435,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f14434,f500])).
fof(f14434,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f14433,f471])).
fof(f471,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X6] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_distrib|split|27')).
fof(f14433,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f14432,f453])).
fof(f14432,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f14431,f453])).
fof(f14431,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(forward_demodulation,[],[f14301,f458])).
fof(f14301,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(superposition,[],[f696,f474])).
fof(f10128,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | $spl149),
  inference(cnf_transformation,[],[f10128_D])).
fof(f10128_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f14424,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(forward_demodulation,[],[f14423,f458])).
fof(f14423,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(subsumption_resolution,[],[f14422,f649])).
fof(f14422,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f14421,f671])).
fof(f671,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_lte|split|2')).
fof(f14421,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f14420,f500])).
fof(f14420,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f14419,f471])).
fof(f14419,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f14299,f453])).
fof(f14299,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(superposition,[],[f696,f459])).
fof(f14416,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(forward_demodulation,[],[f14415,f458])).
fof(f14415,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(subsumption_resolution,[],[f14414,f649])).
fof(f14414,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f14413,f671])).
fof(f14413,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f14412,f500])).
fof(f14412,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f14411,f471])).
fof(f14411,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f14298,f453])).
fof(f14298,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(superposition,[],[f696,f460])).
fof(f14397,plain,(
  ( ! [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X4))))) ) | $spl230),
  inference(cnf_transformation,[],[f14397_D])).
fof(f14397_D,plain,(
  ( ! [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X4))))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f5502,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | $spl105),
  inference(cnf_transformation,[],[f5502_D])).
fof(f5502_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f14399,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X5))))) )),
  inference(subsumption_resolution,[],[f14398,f620])).
fof(f14398,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X5))))) )),
  inference(forward_demodulation,[],[f14274,f458])).
fof(f14274,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X5)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5)))) )),
  inference(superposition,[],[f696,f447])).
fof(f14396,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X4))))) )),
  inference(subsumption_resolution,[],[f14395,f620])).
fof(f14395,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X4))))) )),
  inference(forward_demodulation,[],[f14273,f458])).
fof(f14273,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X4)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4)))) )),
  inference(superposition,[],[f696,f445])).
fof(f14269,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X12))))) ) | $spl228),
  inference(cnf_transformation,[],[f14269_D])).
fof(f14269_D,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X12))))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f14270,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14))))) )),
  inference(subsumption_resolution,[],[f14234,f604])).
fof(f14234,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14)))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f727,f437])).
fof(f727,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X7)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X7))))) )),
  inference(superposition,[],[f679,f462])).
fof(f462,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f321])).
fof(f14268,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12))))) )),
  inference(subsumption_resolution,[],[f14233,f604])).
fof(f14233,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int',X12)))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12))))) )),
  inference(superposition,[],[f727,f437])).
fof(f14262,plain,(
  ( ! [X16] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X16)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X16)) ) | $spl226),
  inference(cnf_transformation,[],[f14262_D])).
fof(f14262_D,plain,(
  ( ! [X16] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X16)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X16)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f14264,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X19) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X18)) )),
  inference(subsumption_resolution,[],[f14263,f672])).
fof(f14263,plain,(
  ( ! [X19,X18] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X19) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X18)) )),
  inference(forward_demodulation,[],[f14224,f488])).
fof(f488,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X6,X8] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.integer.INT_LT_CALCULATE|split|3')).
fof(f14224,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X19) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X18)) )),
  inference(superposition,[],[f727,f480])).
fof(f14261,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X16)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X16) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X17)) )),
  inference(subsumption_resolution,[],[f14260,f672])).
fof(f14260,plain,(
  ( ! [X17,X16] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X16)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X16) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X17)) )),
  inference(forward_demodulation,[],[f14223,f488])).
fof(f14223,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X16)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X16)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X16) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X17)) )),
  inference(superposition,[],[f727,f483])).
fof(f14258,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK10),s('type.num.num',X12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl224),
  inference(cnf_transformation,[],[f14258_D])).
fof(f14258_D,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK10),s('type.num.num',X12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f14259,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK10),s('type.num.num',X14)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X15)) )),
  inference(forward_demodulation,[],[f14222,f487])).
fof(f14222,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X15)) )),
  inference(superposition,[],[f727,f480])).
fof(f14257,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK10),s('type.num.num',X12)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) )),
  inference(forward_demodulation,[],[f14221,f487])).
fof(f14221,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) )),
  inference(superposition,[],[f727,f483])).
fof(f14135,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl222),
  inference(cnf_transformation,[],[f14135_D])).
fof(f14135_D,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f14136,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14))))) )),
  inference(subsumption_resolution,[],[f14102,f604])).
fof(f14102,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14)))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f726,f437])).
fof(f726,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f679,f462])).
fof(f14134,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12))))) )),
  inference(subsumption_resolution,[],[f14101,f604])).
fof(f14101,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12))))) )),
  inference(superposition,[],[f726,f437])).
fof(f14126,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl220),
  inference(cnf_transformation,[],[f14126_D])).
fof(f14126_D,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f14127,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X15)) )),
  inference(forward_demodulation,[],[f14090,f487])).
fof(f14090,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X15)) )),
  inference(superposition,[],[f726,f480])).
fof(f14125,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) )),
  inference(forward_demodulation,[],[f14089,f487])).
fof(f14089,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) )),
  inference(superposition,[],[f726,f483])).
fof(f13787,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12))))) ) | $spl218),
  inference(cnf_transformation,[],[f13787_D])).
fof(f13787_D,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f13788,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f13763,f604])).
fof(f13763,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f705,f437])).
fof(f705,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7))))) )),
  inference(superposition,[],[f679,f462])).
fof(f13786,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f13762,f604])).
fof(f13762,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f705,f437])).
fof(f13780,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) | $spl216),
  inference(cnf_transformation,[],[f13780_D])).
fof(f13780_D,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f13781,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X15)) )),
  inference(forward_demodulation,[],[f13730,f487])).
fof(f13730,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X15)) )),
  inference(superposition,[],[f705,f480])).
fof(f13779,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) )),
  inference(forward_demodulation,[],[f13729,f487])).
fof(f13729,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) )),
  inference(superposition,[],[f705,f483])).
fof(f13278,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | $spl209),
  inference(cnf_transformation,[],[f13278_D])).
fof(f13278_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) <=> ~$spl209),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl209])])).
fof(f13673,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) | $spl185),
  inference(resolution,[],[f13093,f6146])).
fof(f13093,plain,(
  ( ! [X10,X9] : (sP2(X9,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')),X10)) ) | $spl185),
  inference(resolution,[],[f12956,f416])).
fof(f13279,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) ) | $spl210),
  inference(cnf_transformation,[],[f13279_D])).
fof(f13279_D,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f13672,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) ) | $spl185),
  inference(resolution,[],[f13093,f6024])).
fof(f13276,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | $spl206),
  inference(cnf_transformation,[],[f13276_D])).
fof(f13276_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f13671,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) ) | $spl185),
  inference(resolution,[],[f13093,f6100])).
fof(f13665,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl214),
  inference(cnf_transformation,[],[f13665_D])).
fof(f13665_D,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f13666,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f13627,f604])).
fof(f13627,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f704,f437])).
fof(f704,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f679,f462])).
fof(f13664,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f13626,f604])).
fof(f13626,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f704,f437])).
fof(f13650,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) ) | $spl212),
  inference(cnf_transformation,[],[f13650_D])).
fof(f13650_D,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f13653,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X15)) )),
  inference(subsumption_resolution,[],[f13652,f672])).
fof(f13652,plain,(
  ( ! [X14,X15] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X15)) )),
  inference(forward_demodulation,[],[f13651,f492])).
fof(f492,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_distrib|split|20')).
fof(f13651,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X15)) )),
  inference(forward_demodulation,[],[f13594,f487])).
fof(f13594,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X15)) )),
  inference(superposition,[],[f704,f480])).
fof(f13649,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) )),
  inference(subsumption_resolution,[],[f13648,f672])).
fof(f13648,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) )),
  inference(forward_demodulation,[],[f13647,f492])).
fof(f13647,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) )),
  inference(forward_demodulation,[],[f13593,f487])).
fof(f13593,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) )),
  inference(superposition,[],[f704,f483])).
fof(f13423,plain,(
  $false | ($spl3 | $spl184)),
  inference(subsumption_resolution,[],[f13422,f496])).
fof(f13422,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | ($spl3 | $spl184)),
  inference(subsumption_resolution,[],[f13291,f761])).
fof(f13291,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | $spl184),
  inference(resolution,[],[f12955,f642])).
fof(f12955,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | $spl184),
  inference(cnf_transformation,[],[f12955_D])).
fof(f12955_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f13421,plain,(
  $false | ($spl3 | $spl184)),
  inference(subsumption_resolution,[],[f13420,f496])).
fof(f13420,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | ($spl3 | $spl184)),
  inference(subsumption_resolution,[],[f13290,f761])).
fof(f13290,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK10) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | $spl184),
  inference(resolution,[],[f12955,f642])).
fof(f13235,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) | $spl185),
  inference(resolution,[],[f13093,f6146])).
fof(f13234,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) ) | $spl185),
  inference(resolution,[],[f13093,f6024])).
fof(f13233,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) ) | $spl185),
  inference(resolution,[],[f13093,f6100])).
fof(f13082,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',sK10)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num',X12))))) ) | $spl204),
  inference(cnf_transformation,[],[f13082_D])).
fof(f13082_D,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',sK10)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num',X12))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f13083,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14)))))) | ~p(s(bool,X15))) )),
  inference(subsumption_resolution,[],[f12922,f605])).
fof(f12922,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num',X14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14)))))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f720,f437])).
fof(f720,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num',X12))))) )),
  inference(superposition,[],[f679,f642])).
fof(f13081,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num',X12)))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(subsumption_resolution,[],[f12921,f605])).
fof(f12921,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num',X12)))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(superposition,[],[f720,f437])).
fof(f13079,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num',X4))))) ) | $spl202),
  inference(cnf_transformation,[],[f13079_D])).
fof(f13079_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num',X4))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f12910,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num',X6)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(superposition,[],[f720,f480])).
fof(f12909,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num',X4)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) )),
  inference(superposition,[],[f720,f483])).
fof(f13051,plain,(
  ( ! [X82] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X82) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X82))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f13050,f458])).
fof(f13050,plain,(
  ( ! [X82] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X82))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X82)))) )),
  inference(subsumption_resolution,[],[f13049,f677])).
fof(f677,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f590])).
fof(f590,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X11] : ! [X7] : (s(X11,X7) = s(X11,X7) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.REFL_CLAUSE')).
fof(f13049,plain,(
  ( ! [X82] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X82))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X82)))) )),
  inference(forward_demodulation,[],[f12898,f609])).
fof(f12898,plain,(
  ( ! [X82] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X82))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X82)))) )),
  inference(superposition,[],[f720,f445])).
fof(f13046,plain,(
  ( ! [X81] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X81) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X81))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f13045,f458])).
fof(f13045,plain,(
  ( ! [X81] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X81))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) )),
  inference(subsumption_resolution,[],[f13044,f677])).
fof(f13044,plain,(
  ( ! [X81] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X81))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) )),
  inference(forward_demodulation,[],[f12897,f609])).
fof(f12897,plain,(
  ( ! [X81] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X81))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) )),
  inference(superposition,[],[f720,f447])).
fof(f12987,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl199),
  inference(cnf_transformation,[],[f12987_D])).
fof(f12987_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f12989,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f12988,f458])).
fof(f12988,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f12872,f458])).
fof(f12872,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(superposition,[],[f720,f459])).
fof(f12981,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f12980,f458])).
fof(f12980,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',sK10)))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(forward_demodulation,[],[f12871,f458])).
fof(f12871,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(superposition,[],[f720,f460])).
fof(f12956,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | $spl185),
  inference(cnf_transformation,[],[f12956_D])).
fof(f12956_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f12964,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X41))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))))))) )),
  inference(subsumption_resolution,[],[f12865,f496])).
fof(f12865,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X41))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X40))))))) )),
  inference(superposition,[],[f720,f477])).
fof(f12962,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X39))))))) )),
  inference(subsumption_resolution,[],[f12864,f496])).
fof(f12864,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X38))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X39))))))) )),
  inference(superposition,[],[f720,f476])).
fof(f12961,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X36))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X36))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X36))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X37)))))) )),
  inference(subsumption_resolution,[],[f12863,f496])).
fof(f12863,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X36))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X36))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X36))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X37)))))) )),
  inference(superposition,[],[f720,f479])).
fof(f12960,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X34))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X34))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X35))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X34)))))) )),
  inference(subsumption_resolution,[],[f12862,f496])).
fof(f12862,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X34))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X34))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X35))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X34)))))) )),
  inference(superposition,[],[f720,f478])).
fof(f12958,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X33))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32)))))) )),
  inference(subsumption_resolution,[],[f12861,f496])).
fof(f12861,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X33))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X32)))))) )),
  inference(superposition,[],[f720,f482])).
fof(f12954,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X31)))))) )),
  inference(subsumption_resolution,[],[f12860,f496])).
fof(f12860,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK10),s('type.num.num','const.num.|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X31)))))) )),
  inference(superposition,[],[f720,f481])).
fof(f1147,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl25),
  inference(cnf_transformation,[],[f1147_D])).
fof(f1147_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f12630,plain,(
  ( ! [X9] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | ~p(s(bool,X9))) ) | $spl1),
  inference(subsumption_resolution,[],[f12593,f605])).
fof(f12593,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | ~p(s(bool,X9))) ) | $spl1),
  inference(superposition,[],[f759,f437])).
fof(f12629,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl1),
  inference(subsumption_resolution,[],[f12592,f605])).
fof(f12592,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl1),
  inference(superposition,[],[f759,f437])).
fof(f12027,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl181),
  inference(cnf_transformation,[],[f12027_D])).
fof(f12027_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f12130,plain,(
  ( ! [X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | ~p(s(bool,X11))) ) | $spl25),
  inference(resolution,[],[f7048,f6040])).
fof(f6040,plain,(
  ( ! [X0,X1] : (~sP2(X1,X0,'F_0') | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f6034,f661])).
fof(f661,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f404])).
fof(f404,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f6034,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~sP2(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f6031,f428])).
fof(f7048,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X9)) ) | $spl25),
  inference(resolution,[],[f7020,f416])).
fof(f7020,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl25),
  inference(forward_demodulation,[],[f6879,f487])).
fof(f6879,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl25),
  inference(resolution,[],[f1147,f463])).
fof(f7328,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl115),
  inference(cnf_transformation,[],[f7328_D])).
fof(f7328_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f12129,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,X10))) ) | $spl25),
  inference(resolution,[],[f7048,f6146])).
fof(f7329,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl116),
  inference(cnf_transformation,[],[f7329_D])).
fof(f7329_D,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f12128,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl25),
  inference(resolution,[],[f7048,f6024])).
fof(f7326,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl112),
  inference(cnf_transformation,[],[f7326_D])).
fof(f7326_D,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f12127,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl25),
  inference(resolution,[],[f7048,f6100])).
fof(f12171,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X2)) ) | $spl182),
  inference(cnf_transformation,[],[f12171_D])).
fof(f12171_D,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X2)) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f12170,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X2)) ) | $spl25),
  inference(subsumption_resolution,[],[f12124,f425])).
fof(f12124,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | ~sP1(X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X2) | ~sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X3,X2)) ) | $spl25),
  inference(resolution,[],[f7048,f411])).
fof(f12123,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | p(s(bool,X1))) ) | $spl25),
  inference(resolution,[],[f7048,f414])).
fof(f12026,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl178),
  inference(cnf_transformation,[],[f12026_D])).
fof(f12026_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f12028,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl180),
  inference(cnf_transformation,[],[f12028_D])).
fof(f12028_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f12024,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl0),
  inference(forward_demodulation,[],[f12023,f487])).
fof(f12023,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl0),
  inference(subsumption_resolution,[],[f11822,f677])).
fof(f11822,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl0),
  inference(superposition,[],[f462,f758])).
fof(f758,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl0),
  inference(cnf_transformation,[],[f758_D])).
fof(f758_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f11744,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) ) | $spl76),
  inference(superposition,[],[f3701,f480])).
fof(f11743,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl76),
  inference(superposition,[],[f3701,f483])).
fof(f11738,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl76),
  inference(superposition,[],[f3701,f480])).
fof(f11737,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl76),
  inference(superposition,[],[f3701,f483])).
fof(f11657,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) | $spl172),
  inference(cnf_transformation,[],[f11657_D])).
fof(f11657_D,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f11616,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl74),
  inference(superposition,[],[f3700,f480])).
fof(f11615,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl74),
  inference(superposition,[],[f3700,f483])).
fof(f11548,plain,(
  $false | ($spl6 | $spl166)),
  inference(trivial_inequality_removal,[],[f11526])).
fof(f11526,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl6 | $spl166)),
  inference(superposition,[],[f11310,f764])).
fof(f11547,plain,(
  $false | $spl166),
  inference(trivial_inequality_removal,[],[f11532])).
fof(f11532,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl166),
  inference(superposition,[],[f11310,f495])).
fof(f495,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f104])).
fof(f104,axiom,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_distrib|split|26')).
fof(f11546,plain,(
  $false | $spl166),
  inference(trivial_inequality_removal,[],[f11541])).
fof(f11541,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl166),
  inference(superposition,[],[f11310,f609])).
fof(f11545,plain,(
  $false | $spl166),
  inference(equality_resolution,[],[f11310])).
fof(f11485,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl170),
  inference(cnf_transformation,[],[f11485_D])).
fof(f11485_D,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f11486,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f11411,f604])).
fof(f11411,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X21))) )),
  inference(superposition,[],[f685,f437])).
fof(f685,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))))) )),
  inference(superposition,[],[f679,f477])).
fof(f11484,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,X18) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f11410,f604])).
fof(f11410,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,X18) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17)))))) | ~p(s(bool,X18)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f685,f437])).
fof(f11254,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) | $spl160),
  inference(cnf_transformation,[],[f11254_D])).
fof(f11254_D,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f11479,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X10)) )),
  inference(subsumption_resolution,[],[f11378,f477])).
fof(f11378,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X10)) )),
  inference(superposition,[],[f685,f480])).
fof(f11478,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(subsumption_resolution,[],[f11377,f477])).
fof(f11377,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6)) )),
  inference(superposition,[],[f685,f483])).
fof(f11440,plain,(
  ( ! [X77] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X77))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(forward_demodulation,[],[f11340,f458])).
fof(f11340,plain,(
  ( ! [X77] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X77))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f685,f459])).
fof(f11438,plain,(
  ( ! [X76] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(forward_demodulation,[],[f11339,f458])).
fof(f11339,plain,(
  ( ! [X76] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f685,f460])).
fof(f11432,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))))) )),
  inference(subsumption_resolution,[],[f11431,f477])).
fof(f11431,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X8) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))))) )),
  inference(forward_demodulation,[],[f11315,f458])).
fof(f11315,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) )),
  inference(superposition,[],[f685,f447])).
fof(f11430,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))))) )),
  inference(subsumption_resolution,[],[f11429,f477])).
fof(f11429,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))))) )),
  inference(forward_demodulation,[],[f11314,f458])).
fof(f11314,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))) )),
  inference(superposition,[],[f685,f445])).
fof(f11309,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl164),
  inference(cnf_transformation,[],[f11309_D])).
fof(f11309_D,plain,(
  ( ! [X19,X17] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f11311,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f11234,f604])).
fof(f11234,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X21))) )),
  inference(superposition,[],[f684,f437])).
fof(f684,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))))) )),
  inference(superposition,[],[f679,f476])).
fof(f11308,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,X18) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f11233,f604])).
fof(f11233,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,X18) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19)))))) | ~p(s(bool,X18)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f684,f437])).
fof(f11303,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X10)) )),
  inference(subsumption_resolution,[],[f11201,f476])).
fof(f11201,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X10)) )),
  inference(superposition,[],[f684,f480])).
fof(f11302,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(subsumption_resolution,[],[f11200,f476])).
fof(f11200,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6)) )),
  inference(superposition,[],[f684,f483])).
fof(f11264,plain,(
  ( ! [X77] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X77)))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(forward_demodulation,[],[f11163,f458])).
fof(f11163,plain,(
  ( ! [X77] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X77))))))) )),
  inference(superposition,[],[f684,f459])).
fof(f11262,plain,(
  ( ! [X76] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76)))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(forward_demodulation,[],[f11162,f458])).
fof(f11162,plain,(
  ( ! [X76] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))))))) )),
  inference(superposition,[],[f684,f460])).
fof(f3945,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))))) ) | $spl78),
  inference(cnf_transformation,[],[f3945_D])).
fof(f3945_D,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f10120,plain,(
  ( ! [X61,X63] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X61))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X63))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X61))))))) ) | $spl146),
  inference(cnf_transformation,[],[f10120_D])).
fof(f10120_D,plain,(
  ( ! [X61,X63] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X61))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X63))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X61))))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f11156,plain,(
  ( ! [X61,X62,X63] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X61))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X62)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X63))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X61))))))) )),
  inference(superposition,[],[f684,f477])).
fof(f11256,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))))) )),
  inference(subsumption_resolution,[],[f11255,f476])).
fof(f11255,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X8) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))))) )),
  inference(forward_demodulation,[],[f11138,f458])).
fof(f11138,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) )),
  inference(superposition,[],[f684,f447])).
fof(f11253,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) )),
  inference(subsumption_resolution,[],[f11252,f476])).
fof(f11252,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) )),
  inference(forward_demodulation,[],[f11137,f458])).
fof(f11137,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))) )),
  inference(superposition,[],[f684,f445])).
fof(f9180,plain,(
  ( ! [X17] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X17)) ) | $spl126),
  inference(cnf_transformation,[],[f9180_D])).
fof(f9180_D,plain,(
  ( ! [X17] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X17)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f11084,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl44),
  inference(superposition,[],[f1772,f480])).
fof(f11083,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl44),
  inference(superposition,[],[f1772,f483])).
fof(f10470,plain,(
  $false | $spl27),
  inference(trivial_inequality_removal,[],[f10354])).
fof(f10354,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl27),
  inference(resolution,[],[f1150,f464])).
fof(f464,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) != s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f321])).
fof(f10351,plain,(
  ( ! [X19,X17] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl158),
  inference(cnf_transformation,[],[f10351_D])).
fof(f10351_D,plain,(
  ( ! [X19,X17] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f10352,plain,(
  ( ! [X21,X22,X20] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X21))) )),
  inference(subsumption_resolution,[],[f10277,f605])).
fof(f10277,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X21))) )),
  inference(superposition,[],[f681,f437])).
fof(f681,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f679,f482])).
fof(f10350,plain,(
  ( ! [X19,X17,X18] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))) | ~p(s(bool,X18)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f10276,f605])).
fof(f10276,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,X18) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))) | ~p(s(bool,X18)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f681,f437])).
fof(f10169,plain,(
  ( ! [X6,X8] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) ) | $spl152),
  inference(cnf_transformation,[],[f10169_D])).
fof(f10169_D,plain,(
  ( ! [X6,X8] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f10345,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X10)) )),
  inference(subsumption_resolution,[],[f10244,f478])).
fof(f10244,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X10)) )),
  inference(superposition,[],[f681,f480])).
fof(f10344,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(subsumption_resolution,[],[f10243,f478])).
fof(f10243,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6)) )),
  inference(superposition,[],[f681,f483])).
fof(f10306,plain,(
  ( ! [X77] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X77))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(forward_demodulation,[],[f10206,f458])).
fof(f10206,plain,(
  ( ! [X77] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X77))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f681,f459])).
fof(f10304,plain,(
  ( ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(forward_demodulation,[],[f10205,f458])).
fof(f10205,plain,(
  ( ! [X76] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f681,f460])).
fof(f1772,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl44),
  inference(cnf_transformation,[],[f1772_D])).
fof(f1772_D,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f10199,plain,(
  ( ! [X61,X62,X63] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X61))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X62))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X63))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X61))))))) )),
  inference(superposition,[],[f681,f477])).
fof(f10119,plain,(
  ( ! [X60,X58] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X58))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X58))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X60))))))) ) | $spl144),
  inference(cnf_transformation,[],[f10119_D])).
fof(f10119_D,plain,(
  ( ! [X60,X58] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X58))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X58))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X60))))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f10198,plain,(
  ( ! [X59,X60,X58] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X58))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X59))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X58))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X60))))))) )),
  inference(superposition,[],[f681,f476])).
fof(f10298,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) )),
  inference(subsumption_resolution,[],[f10297,f478])).
fof(f10297,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X8) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) )),
  inference(forward_demodulation,[],[f10181,f458])).
fof(f10181,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) )),
  inference(superposition,[],[f681,f447])).
fof(f10296,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6)))))) )),
  inference(subsumption_resolution,[],[f10295,f478])).
fof(f10295,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6)))))) )),
  inference(forward_demodulation,[],[f10180,f458])).
fof(f10180,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))) )),
  inference(superposition,[],[f681,f445])).
fof(f10176,plain,(
  ( ! [X19,X17] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl154),
  inference(cnf_transformation,[],[f10176_D])).
fof(f10176_D,plain,(
  ( ! [X19,X17] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f10177,plain,(
  ( ! [X21,X22,X20] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X21))) )),
  inference(subsumption_resolution,[],[f10100,f605])).
fof(f10100,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X21))) )),
  inference(superposition,[],[f680,f437])).
fof(f680,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f679,f481])).
fof(f10175,plain,(
  ( ! [X19,X17,X18] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))) | ~p(s(bool,X18)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f10099,f605])).
fof(f10099,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,X18) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))) | ~p(s(bool,X18)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f680,f437])).
fof(f10170,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X10)) )),
  inference(subsumption_resolution,[],[f10067,f479])).
fof(f10067,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X10)) )),
  inference(superposition,[],[f680,f480])).
fof(f10168,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(subsumption_resolution,[],[f10066,f479])).
fof(f10066,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6)) )),
  inference(superposition,[],[f680,f483])).
fof(f10130,plain,(
  ( ! [X77] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X77))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(forward_demodulation,[],[f10029,f458])).
fof(f10029,plain,(
  ( ! [X77] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X77))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f680,f459])).
fof(f10126,plain,(
  ( ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(forward_demodulation,[],[f10028,f458])).
fof(f10028,plain,(
  ( ! [X76] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f680,f460])).
fof(f10117,plain,(
  ( ! [X50] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X50)))))) ) | $spl140),
  inference(cnf_transformation,[],[f10117_D])).
fof(f10117_D,plain,(
  ( ! [X50] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X50)))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f10022,plain,(
  ( ! [X61,X62,X63] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X61))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X62))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X63))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X61))))))) )),
  inference(superposition,[],[f680,f477])).
fof(f10021,plain,(
  ( ! [X59,X60,X58] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X58))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X59))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X58))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X60))))))) )),
  inference(superposition,[],[f680,f476])).
fof(f10118,plain,(
  ( ! [X51,X49] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X49))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X51))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X49)))))) ) | $spl142),
  inference(cnf_transformation,[],[f10118_D])).
fof(f10118_D,plain,(
  ( ! [X51,X49] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X49))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X51))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X49)))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f10019,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X52))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X53))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X54))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X52)))))) )),
  inference(superposition,[],[f680,f478])).
fof(f10018,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X49))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X50))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X51))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X49)))))) )),
  inference(superposition,[],[f680,f482])).
fof(f10116,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))) ) | $spl138),
  inference(cnf_transformation,[],[f10116_D])).
fof(f10116_D,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f10004,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) )),
  inference(superposition,[],[f680,f447])).
fof(f10003,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))) )),
  inference(superposition,[],[f680,f445])).
fof(f1767,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) | $spl42),
  inference(cnf_transformation,[],[f1767_D])).
fof(f1767_D,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f9953,plain,(
  ( ! [X19,X18] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X19) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X18)) )),
  inference(superposition,[],[f724,f480])).
fof(f724,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4)))) )),
  inference(superposition,[],[f679,f501])).
fof(f501,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X7,X15] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15))) <=> s('type.integer.int',X7) = s('type.integer.int',X15))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.integer.INT_EQ_NEG')).
fof(f9952,plain,(
  ( ! [X17,X16] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X16))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X16) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X17)) )),
  inference(superposition,[],[f724,f483])).
fof(f9995,plain,(
  ( ! [X13] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13)) ) | $spl134),
  inference(cnf_transformation,[],[f9995_D])).
fof(f9995_D,plain,(
  ( ! [X13] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f9951,plain,(
  ( ! [X14,X15] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X15)) )),
  inference(superposition,[],[f724,f480])).
fof(f9950,plain,(
  ( ! [X12,X13] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13))))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X12)) )),
  inference(superposition,[],[f724,f483])).
fof(f3700,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl74),
  inference(cnf_transformation,[],[f3700_D])).
fof(f3700_D,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f5092,plain,(
  ( ! [X14,X13] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,X13))) ) | $spl96),
  inference(cnf_transformation,[],[f5092_D])).
fof(f5092_D,plain,(
  ( ! [X14,X13] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,X13))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f9836,plain,(
  ( ! [X33,X31,X32] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X31,X32) | p(s(bool,X31)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X33))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl25),
  inference(superposition,[],[f7087,f482])).
fof(f7087,plain,(
  ( ! [X83,X82] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')),X82,X83) | p(s(bool,X82))) ) | $spl25),
  inference(resolution,[],[f7020,f532])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f343])).
fof(f343,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f207,f295,f294])).
fof(f295,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f145])).
fof(f145,axiom,(
  ! [X18,X19,X9] : ((p(s(bool,X9)) <=> (p(s(bool,X19)) & p(s(bool,X18)))) <=> ((p(s(bool,X9)) | ~p(s(bool,X19)) | ~p(s(bool,X18))) & (p(s(bool,X19)) | ~p(s(bool,X9))) & (p(s(bool,X18)) | ~p(s(bool,X9)))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.sat.sat_23_prove_ltb')).
fof(f9835,plain,(
  ( ! [X30,X28,X29] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X28,X29) | p(s(bool,X28)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X30)))))) ) | $spl25),
  inference(superposition,[],[f7087,f481])).
fof(f9834,plain,(
  ( ! [X26,X27,X25] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | p(s(bool,X25)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27)))))) ) | $spl25),
  inference(superposition,[],[f7087,f479])).
fof(f9833,plain,(
  ( ! [X24,X23,X22] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | p(s(bool,X22)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl25),
  inference(superposition,[],[f7087,f478])).
fof(f9832,plain,(
  ( ! [X21,X19,X20] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | p(s(bool,X19)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl25),
  inference(superposition,[],[f7087,f477])).
fof(f9831,plain,(
  ( ! [X17,X18,X16] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | p(s(bool,X16)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))))))) ) | $spl25),
  inference(superposition,[],[f7087,f476])).
fof(f9359,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl133),
  inference(cnf_transformation,[],[f9359_D])).
fof(f9359_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f9360,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f9301,f458])).
fof(f9301,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(superposition,[],[f716,f459])).
fof(f716,plain,(
  ( ! [X8] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) )),
  inference(superposition,[],[f679,f484])).
fof(f484,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f328])).
fof(f9357,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f9300,f458])).
fof(f9300,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(superposition,[],[f716,f460])).
fof(f9349,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) | $spl129),
  inference(cnf_transformation,[],[f9349_D])).
fof(f9349_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f9351,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5)))))) )),
  inference(forward_demodulation,[],[f9282,f458])).
fof(f9282,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5)))) )),
  inference(superposition,[],[f716,f447])).
fof(f9347,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4)))))) )),
  inference(forward_demodulation,[],[f9281,f458])).
fof(f9281,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4)))) )),
  inference(superposition,[],[f716,f445])).
fof(f9132,plain,(
  ( ! [X19,X18] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X19))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X19) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X18)) )),
  inference(superposition,[],[f702,f480])).
fof(f702,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f679,f501])).
fof(f9131,plain,(
  ( ! [X17,X16] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X17))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X16))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X16) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X17)) )),
  inference(superposition,[],[f702,f483])).
fof(f8577,plain,(
  ( ! [X81] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X81))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) ) | $spl124),
  inference(cnf_transformation,[],[f8577_D])).
fof(f8577_D,plain,(
  ( ! [X81] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X81))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f8495,plain,(
  ( ! [X82] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X82))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X82)))) )),
  inference(superposition,[],[f694,f445])).
fof(f694,plain,(
  ( ! [X20] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f679,f484])).
fof(f8494,plain,(
  ( ! [X81] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X81))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) )),
  inference(superposition,[],[f694,f447])).
fof(f8546,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl123),
  inference(cnf_transformation,[],[f8546_D])).
fof(f8546_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f8547,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f8469,f458])).
fof(f8469,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(superposition,[],[f694,f459])).
fof(f8544,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f8468,f458])).
fof(f8468,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(superposition,[],[f694,f460])).
fof(f7920,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl120),
  inference(cnf_transformation,[],[f7920_D])).
fof(f7920_D,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f7921,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X15))) )),
  inference(subsumption_resolution,[],[f7833,f605])).
fof(f7833,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f751,f437])).
fof(f751,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num','const.num.|48|_0'))))) )),
  inference(subsumption_resolution,[],[f698,f496])).
fof(f698,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) )),
  inference(superposition,[],[f679,f642])).
fof(f7919,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f7832,f605])).
fof(f7832,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X13)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f751,f437])).
fof(f7914,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(subsumption_resolution,[],[f7800,f641])).
fof(f641,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f393])).
fof(f7800,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(superposition,[],[f751,f480])).
fof(f7912,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) )),
  inference(subsumption_resolution,[],[f7799,f641])).
fof(f7799,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) )),
  inference(superposition,[],[f751,f483])).
fof(f7854,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(subsumption_resolution,[],[f7853,f641])).
fof(f7853,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f7737,f458])).
fof(f7737,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5)))) )),
  inference(superposition,[],[f751,f447])).
fof(f7852,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(subsumption_resolution,[],[f7851,f641])).
fof(f7851,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f7736,f458])).
fof(f7736,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4)))) )),
  inference(superposition,[],[f751,f445])).
fof(f7540,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f7444,f772])).
fof(f7444,plain,(
  ( ! [X204,X205] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X204),s('type.num.num',X205))))) ) | $spl16),
  inference(resolution,[],[f772,f639])).
fof(f639,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f391])).
fof(f391,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f137])).
fof(f137,axiom,(
  ! [X6,X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X6)))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.numeral.numeral_lte|split|5')).
fof(f7539,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f7443,f772])).
fof(f7443,plain,(
  ( ! [X202,X203] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X202),s('type.num.num',X203))))) ) | $spl16),
  inference(resolution,[],[f772,f618])).
fof(f7536,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f7438,f772])).
fof(f7438,plain,(
  ( ! [X194,X193] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X193))),s('type.num.num',X194))))) ) | $spl16),
  inference(resolution,[],[f772,f618])).
fof(f7535,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f7437,f772])).
fof(f7437,plain,(
  ( ! [X191,X192] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X191))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X192))))))) ) | $spl16),
  inference(resolution,[],[f772,f639])).
fof(f7530,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f7423,f772])).
fof(f7423,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f772,f612])).
fof(f612,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f381])).
fof(f381,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f237])).
fof(f237,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f51])).
fof(f51,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.NOT_CLAUSES|split|1')).
fof(f7529,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f7422,f772])).
fof(f7422,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f772,f610])).
fof(f610,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f380])).
fof(f380,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f236])).
fof(f236,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f52])).
fof(f52,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f7424,plain,(
  $false | $spl16),
  inference(resolution,[],[f772,f614])).
fof(f614,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f36])).
fof(f36,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.TRUTH')).
fof(f7425,plain,(
  $false | $spl16),
  inference(resolution,[],[f772,f615])).
fof(f7426,plain,(
  $false | $spl16),
  inference(resolution,[],[f772,f677])).
fof(f7440,plain,(
  $false | $spl16),
  inference(resolution,[],[f772,f496])).
fof(f7448,plain,(
  $false | ($spl8 | $spl16)),
  inference(resolution,[],[f772,f766])).
fof(f7449,plain,(
  $false | ($spl4 | $spl16)),
  inference(resolution,[],[f772,f763])).
fof(f7283,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,X10))) ) | $spl25),
  inference(resolution,[],[f7048,f6146])).
fof(f7282,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl25),
  inference(resolution,[],[f7048,f6024])).
fof(f7281,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl25),
  inference(resolution,[],[f7048,f6100])).
fof(f1150,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl27),
  inference(cnf_transformation,[],[f1150_D])).
fof(f1150_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f6988,plain,(
  ( ! [X9] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl25),
  inference(superposition,[],[f1147,f482])).
fof(f6987,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) ) | $spl25),
  inference(superposition,[],[f1147,f481])).
fof(f6986,plain,(
  ( ! [X7] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))) ) | $spl25),
  inference(superposition,[],[f1147,f479])).
fof(f6985,plain,(
  ( ! [X6] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl25),
  inference(superposition,[],[f1147,f478])).
fof(f6984,plain,(
  ( ! [X5] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl25),
  inference(superposition,[],[f1147,f477])).
fof(f6983,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl25),
  inference(superposition,[],[f1147,f476])).
fof(f5643,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl109),
  inference(cnf_transformation,[],[f5643_D])).
fof(f5643_D,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f5646,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num',sK10) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f5600,f458])).
fof(f5600,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(superposition,[],[f722,f459])).
fof(f722,plain,(
  ( ! [X0] : (s('type.num.num',sK10) != s('type.num.num',X0) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f679,f485])).
fof(f5641,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num',sK10) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f5599,f458])).
fof(f5599,plain,(
  s('type.num.num',sK10) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(superposition,[],[f722,f460])).
fof(f5503,plain,(
  ( ! [X81] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X81) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) ) | $spl106),
  inference(cnf_transformation,[],[f5503_D])).
fof(f5503_D,plain,(
  ( ! [X81] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X81) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f5427,plain,(
  ( ! [X82] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X82) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X82)))) )),
  inference(superposition,[],[f700,f445])).
fof(f700,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f679,f485])).
fof(f5426,plain,(
  ( ! [X81] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X81) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))) )),
  inference(superposition,[],[f700,f447])).
fof(f5468,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl101),
  inference(cnf_transformation,[],[f5468_D])).
fof(f5468_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f5471,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f5401,f458])).
fof(f5401,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(superposition,[],[f700,f459])).
fof(f5466,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f5400,f458])).
fof(f5400,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(superposition,[],[f700,f460])).
fof(f3947,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK10))))) ) | $spl80),
  inference(cnf_transformation,[],[f3947_D])).
fof(f3947_D,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK10))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f5344,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) ) | $spl8),
  inference(forward_demodulation,[],[f5294,f487])).
fof(f5294,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) ) | $spl8),
  inference(superposition,[],[f766,f480])).
fof(f5343,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl8),
  inference(forward_demodulation,[],[f5293,f487])).
fof(f5293,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl8),
  inference(superposition,[],[f766,f483])).
fof(f5020,plain,(
  ( ! [X5] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) | $spl94),
  inference(cnf_transformation,[],[f5020_D])).
fof(f5020_D,plain,(
  ( ! [X5] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f5288,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl8),
  inference(superposition,[],[f766,f480])).
fof(f5287,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl8),
  inference(superposition,[],[f766,f483])).
fof(f766,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl8),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f5149,plain,(
  ( ! [X14,X13] : (sP7(X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | p(s(bool,X13))) ) | $spl98),
  inference(cnf_transformation,[],[f5149_D])).
fof(f5149_D,plain,(
  ( ! [X14,X13] : (sP7(X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | p(s(bool,X13))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f5126,plain,(
  ( ! [X14,X15,X13] : (sP7(X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | p(s(bool,X13)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1224,f477])).
fof(f1224,plain,(
  ( ! [X85,X84] : (sP7(X84,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X85) | p(s(bool,X84))) )),
  inference(resolution,[],[f1130,f532])).
fof(f5064,plain,(
  ( ! [X14,X15,X13] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,X13)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1223,f477])).
fof(f1223,plain,(
  ( ! [X83,X82] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X82,X83) | p(s(bool,X82))) )),
  inference(resolution,[],[f1130,f532])).
fof(f5022,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) ) | $spl8),
  inference(forward_demodulation,[],[f4971,f487])).
fof(f4971,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) ) | $spl8),
  inference(superposition,[],[f766,f480])).
fof(f5021,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl8),
  inference(forward_demodulation,[],[f4970,f487])).
fof(f4970,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl8),
  inference(superposition,[],[f766,f483])).
fof(f4965,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl8),
  inference(superposition,[],[f766,f480])).
fof(f4964,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl8),
  inference(superposition,[],[f766,f483])).
fof(f4826,plain,(
  ( ! [X14,X13] : (sP7(X13,X14,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X13))) ) | $spl92),
  inference(cnf_transformation,[],[f4826_D])).
fof(f4826_D,plain,(
  ( ! [X14,X13] : (sP7(X13,X14,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X13))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4803,plain,(
  ( ! [X14,X15,X13] : (sP7(X13,X14,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X13)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1222,f477])).
fof(f1222,plain,(
  ( ! [X80,X81] : (sP7(X80,X81,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,X80))) )),
  inference(resolution,[],[f1130,f531])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f344])).
fof(f4769,plain,(
  ( ! [X14,X13] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,X14))) ) | $spl90),
  inference(cnf_transformation,[],[f4769_D])).
fof(f4769_D,plain,(
  ( ! [X14,X13] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,X14))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f4741,plain,(
  ( ! [X14,X15,X13] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,X14)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1221,f477])).
fof(f1221,plain,(
  ( ! [X78,X79] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X78,X79) | p(s(bool,X79))) )),
  inference(resolution,[],[f1130,f531])).
fof(f4707,plain,(
  ( ! [X14,X13] : (~sP7(X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | ~p(s(bool,X13))) ) | $spl88),
  inference(cnf_transformation,[],[f4707_D])).
fof(f4707_D,plain,(
  ( ! [X14,X13] : (~sP7(X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | ~p(s(bool,X13))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f4682,plain,(
  ( ! [X14,X15,X13] : (~sP7(X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | ~p(s(bool,X13)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1215,f477])).
fof(f1215,plain,(
  ( ! [X66,X67] : (~sP7(X66,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X67) | ~p(s(bool,X66))) )),
  inference(resolution,[],[f1130,f522])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f339])).
fof(f339,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f338])).
fof(f338,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f295])).
fof(f4649,plain,(
  ( ! [X14,X13] : (~sP7(X13,X14,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X13))) ) | $spl86),
  inference(cnf_transformation,[],[f4649_D])).
fof(f4649_D,plain,(
  ( ! [X14,X13] : (~sP7(X13,X14,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X13))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f4624,plain,(
  ( ! [X14,X15,X13] : (~sP7(X13,X14,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X13)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1214,f477])).
fof(f1214,plain,(
  ( ! [X64,X65] : (~sP7(X64,X65,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | ~p(s(bool,X64))) )),
  inference(resolution,[],[f1130,f521])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f340])).
fof(f4587,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl4),
  inference(forward_demodulation,[],[f4557,f486])).
fof(f486,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X6,X8] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X6))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.integer.INT_LT_CALCULATE|split|1')).
fof(f4557,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl4),
  inference(superposition,[],[f763,f480])).
fof(f4586,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl4),
  inference(forward_demodulation,[],[f4556,f486])).
fof(f4556,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl4),
  inference(superposition,[],[f763,f483])).
fof(f4466,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl4),
  inference(superposition,[],[f763,f477])).
fof(f4388,plain,(
  ( ! [X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) | $spl84),
  inference(cnf_transformation,[],[f4388_D])).
fof(f4388_D,plain,(
  ( ! [X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f4360,plain,(
  ( ! [X14,X15,X13] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | ~p(s(bool,X13)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1199,f477])).
fof(f1199,plain,(
  ( ! [X35,X34] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X34,X35) | ~p(s(bool,X34))) )),
  inference(resolution,[],[f1130,f504])).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP4(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP4(X0,X1,X2) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f331])).
fof(f331,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1)) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f763,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))))))) ) | $spl4),
  inference(cnf_transformation,[],[f763_D])).
fof(f763_D,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f4359,plain,(
  ( ! [X12,X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(superposition,[],[f1199,f476])).
fof(f4327,plain,(
  ( ! [X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11))) ) | $spl82),
  inference(cnf_transformation,[],[f4327_D])).
fof(f4327_D,plain,(
  ( ! [X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f4299,plain,(
  ( ! [X14,X15,X13] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | ~p(s(bool,X14)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1198,f477])).
fof(f1198,plain,(
  ( ! [X33,X32] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X33,X32) | ~p(s(bool,X32))) )),
  inference(resolution,[],[f1130,f503])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f332])).
fof(f4298,plain,(
  ( ! [X12,X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(superposition,[],[f1198,f476])).
fof(f4154,plain,(
  ( ! [X9] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X9))) ) | $spl7),
  inference(subsumption_resolution,[],[f4133,f605])).
fof(f4133,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X9))) ) | $spl7),
  inference(superposition,[],[f765,f437])).
fof(f4153,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f4132,f605])).
fof(f4132,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl7),
  inference(superposition,[],[f765,f437])).
fof(f3948,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl4),
  inference(forward_demodulation,[],[f3916,f486])).
fof(f3916,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl4),
  inference(superposition,[],[f763,f480])).
fof(f3946,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK10)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl4),
  inference(forward_demodulation,[],[f3915,f486])).
fof(f3915,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl4),
  inference(superposition,[],[f763,f483])).
fof(f3825,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl4),
  inference(superposition,[],[f763,f477])).
fof(f3709,plain,(
  ( ! [X9] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | ~p(s(bool,X9))) ) | $spl1),
  inference(subsumption_resolution,[],[f3672,f605])).
fof(f3672,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | ~p(s(bool,X9))) ) | $spl1),
  inference(superposition,[],[f759,f437])).
fof(f3708,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl1),
  inference(subsumption_resolution,[],[f3671,f605])).
fof(f3671,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl1),
  inference(superposition,[],[f759,f437])).
fof(f765,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl7),
  inference(cnf_transformation,[],[f765_D])).
fof(f765_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f3654,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl1),
  inference(superposition,[],[f759,f482])).
fof(f3653,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))) ) | $spl1),
  inference(superposition,[],[f759,f481])).
fof(f3652,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7)))))) ) | $spl1),
  inference(superposition,[],[f759,f479])).
fof(f3651,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl1),
  inference(superposition,[],[f759,f478])).
fof(f3650,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl1),
  inference(superposition,[],[f759,f477])).
fof(f3649,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))))) ) | $spl1),
  inference(superposition,[],[f759,f476])).
fof(f3616,plain,(
  ( ! [X6] : (~sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6) | p(s(bool,X6))) ) | $spl68),
  inference(cnf_transformation,[],[f3616_D])).
fof(f3616_D,plain,(
  ( ! [X6] : (~sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6) | p(s(bool,X6))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3593,plain,(
  ( ! [X8,X9] : (~sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1306,f477])).
fof(f1306,plain,(
  ( ! [X101] : (~sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X101) | p(s(bool,X101))) )),
  inference(subsumption_resolution,[],[f1239,f649])).
fof(f1239,plain,(
  ( ! [X101] : (p(s(bool,'F_0')) | p(s(bool,X101)) | ~sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X101)) )),
  inference(resolution,[],[f1130,f629])).
fof(f629,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9(X0,X1)) )),
  inference(cnf_transformation,[],[f388])).
fof(f3592,plain,(
  ( ! [X6,X7] : (~sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6) | p(s(bool,X6)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) )),
  inference(superposition,[],[f1306,f476])).
fof(f3561,plain,(
  ( ! [X6] : (~sP9(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6))) ) | $spl66),
  inference(cnf_transformation,[],[f3561_D])).
fof(f3561_D,plain,(
  ( ! [X6] : (~sP9(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3537,plain,(
  ( ! [X8,X9] : (~sP9(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1305,f477])).
fof(f1305,plain,(
  ( ! [X100] : (~sP9(X100,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,X100))) )),
  inference(subsumption_resolution,[],[f1238,f649])).
fof(f1238,plain,(
  ( ! [X100] : (p(s(bool,'F_0')) | p(s(bool,X100)) | ~sP9(X100,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)))) )),
  inference(resolution,[],[f1130,f629])).
fof(f3536,plain,(
  ( ! [X6,X7] : (~sP9(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) )),
  inference(superposition,[],[f1305,f476])).
fof(f3415,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl64),
  inference(cnf_transformation,[],[f3415_D])).
fof(f3415_D,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3201,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f744,f477])).
fof(f744,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(trivial_inequality_removal,[],[f738])).
fof(f738,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))),
  inference(superposition,[],[f679,f608])).
fof(f3200,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))))))) )),
  inference(superposition,[],[f744,f476])).
fof(f2787,plain,(
  ( ! [X6] : (sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6) | ~p(s(bool,X6))) ) | $spl62),
  inference(cnf_transformation,[],[f2787_D])).
fof(f2787_D,plain,(
  ( ! [X6] : (sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6) | ~p(s(bool,X6))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2764,plain,(
  ( ! [X8,X9] : (sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | ~p(s(bool,X8)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1231,f477])).
fof(f1231,plain,(
  ( ! [X95] : (sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X95) | ~p(s(bool,X95))) )),
  inference(resolution,[],[f1130,f568])).
fof(f568,plain,(
  ( ! [X0,X1] : (sP8(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0,X1] : ((~sP8(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1)))),
  inference(rectify,[],[f356])).
fof(f356,plain,(
  ! [X1,X0] : ((~sP8(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP8(X1,X0)))),
  inference(flattening,[],[f355])).
fof(f355,plain,(
  ! [X1,X0] : ((~sP8(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP8(X1,X0)))),
  inference(nnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X1,X0] : (sP8(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f2763,plain,(
  ( ! [X6,X7] : (sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6) | ~p(s(bool,X6)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) )),
  inference(superposition,[],[f1231,f476])).
fof(f2438,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl56),
  inference(cnf_transformation,[],[f2438_D])).
fof(f2438_D,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2726,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) ) | $spl12),
  inference(superposition,[],[f768,f480])).
fof(f2725,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl12),
  inference(superposition,[],[f768,f483])).
fof(f2437,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) | $spl54),
  inference(cnf_transformation,[],[f2437_D])).
fof(f2437_D,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2720,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl12),
  inference(superposition,[],[f768,f480])).
fof(f2719,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl12),
  inference(superposition,[],[f768,f483])).
fof(f768,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5)))))) ) | $spl12),
  inference(cnf_transformation,[],[f768_D])).
fof(f768_D,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5)))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2653,plain,(
  ( ! [X6] : (sP8(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X6))) ) | $spl60),
  inference(cnf_transformation,[],[f2653_D])).
fof(f2653_D,plain,(
  ( ! [X6] : (sP8(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X6))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2633,plain,(
  ( ! [X14,X15] : (sP8(X14,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X14)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15)))))) )),
  inference(superposition,[],[f1230,f481])).
fof(f1230,plain,(
  ( ! [X94] : (sP8(X94,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | ~p(s(bool,X94))) )),
  inference(resolution,[],[f1130,f565])).
fof(f565,plain,(
  ( ! [X0,X1] : (sP8(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f357])).
fof(f2632,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X12)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) )),
  inference(superposition,[],[f1230,f479])).
fof(f2630,plain,(
  ( ! [X8,X9] : (sP8(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1230,f477])).
fof(f2629,plain,(
  ( ! [X6,X7] : (sP8(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X6)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) )),
  inference(superposition,[],[f1230,f476])).
fof(f2598,plain,(
  ( ! [X6] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6) | p(s(bool,X6))) ) | $spl58),
  inference(cnf_transformation,[],[f2598_D])).
fof(f2598_D,plain,(
  ( ! [X6] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6) | p(s(bool,X6))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2578,plain,(
  ( ! [X14,X15] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | p(s(bool,X14)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15)))))) )),
  inference(superposition,[],[f1229,f481])).
fof(f1229,plain,(
  ( ! [X93] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X93) | p(s(bool,X93))) )),
  inference(resolution,[],[f1130,f563])).
fof(f563,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP8(X0,X1)) )),
  inference(cnf_transformation,[],[f357])).
fof(f2577,plain,(
  ( ! [X12,X13] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X12) | p(s(bool,X12)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) )),
  inference(superposition,[],[f1229,f479])).
fof(f2575,plain,(
  ( ! [X8,X9] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1229,f477])).
fof(f2574,plain,(
  ( ! [X6,X7] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6) | p(s(bool,X6)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) )),
  inference(superposition,[],[f1229,f476])).
fof(f2430,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) ) | $spl12),
  inference(superposition,[],[f768,f480])).
fof(f2429,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl12),
  inference(superposition,[],[f768,f483])).
fof(f2424,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl12),
  inference(superposition,[],[f768,f480])).
fof(f2423,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl12),
  inference(superposition,[],[f768,f483])).
fof(f2357,plain,(
  ( ! [X6] : (~sP8(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6))) ) | $spl52),
  inference(cnf_transformation,[],[f2357_D])).
fof(f2357_D,plain,(
  ( ! [X6] : (~sP8(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2337,plain,(
  ( ! [X14,X15] : (~sP8(X14,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X14)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15)))))) )),
  inference(superposition,[],[f1228,f481])).
fof(f1228,plain,(
  ( ! [X92] : (~sP8(X92,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10))) | p(s(bool,X92))) )),
  inference(resolution,[],[f1130,f563])).
fof(f2336,plain,(
  ( ! [X12,X13] : (~sP8(X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X12)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) )),
  inference(superposition,[],[f1228,f479])).
fof(f2334,plain,(
  ( ! [X8,X9] : (~sP8(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1228,f477])).
fof(f2333,plain,(
  ( ! [X6,X7] : (~sP8(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) )),
  inference(superposition,[],[f1228,f476])).
fof(f2143,plain,(
  ( ! [X14,X15] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15)))))) ) | $spl10),
  inference(superposition,[],[f767,f481])).
fof(f2142,plain,(
  ( ! [X12,X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) | $spl10),
  inference(superposition,[],[f767,f479])).
fof(f2140,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl10),
  inference(superposition,[],[f767,f477])).
fof(f2139,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) | $spl10),
  inference(superposition,[],[f767,f476])).
fof(f2130,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl10),
  inference(superposition,[],[f767,f480])).
fof(f2129,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl10),
  inference(superposition,[],[f767,f483])).
fof(f767,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) ) | $spl10),
  inference(cnf_transformation,[],[f767_D])).
fof(f767_D,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2038,plain,(
  ( ! [X10,X11] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl50),
  inference(cnf_transformation,[],[f2038_D])).
fof(f2038_D,plain,(
  ( ! [X10,X11] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2020,plain,(
  ( ! [X26,X27,X25] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1220,f482])).
fof(f1220,plain,(
  ( ! [X76,X77] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X76,X77)) )),
  inference(resolution,[],[f1130,f530])).
fof(f2019,plain,(
  ( ! [X24,X23,X22] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f1220,f481])).
fof(f2018,plain,(
  ( ! [X21,X19,X20] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) )),
  inference(superposition,[],[f1220,f479])).
fof(f2017,plain,(
  ( ! [X17,X18,X16] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1220,f478])).
fof(f2016,plain,(
  ( ! [X14,X15,X13] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1220,f477])).
fof(f2015,plain,(
  ( ! [X12,X10,X11] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(superposition,[],[f1220,f476])).
fof(f1985,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl48),
  inference(cnf_transformation,[],[f1985_D])).
fof(f1985_D,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1967,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1219,f482])).
fof(f1219,plain,(
  ( ! [X74,X75] : (sP6(X74,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X75)) )),
  inference(resolution,[],[f1130,f529])).
fof(f1966,plain,(
  ( ! [X24,X23,X22] : (sP6(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f1219,f481])).
fof(f1965,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) )),
  inference(superposition,[],[f1219,f479])).
fof(f1964,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1219,f478])).
fof(f1963,plain,(
  ( ! [X14,X15,X13] : (sP6(X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1219,f477])).
fof(f1962,plain,(
  ( ! [X12,X10,X11] : (sP6(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(superposition,[],[f1219,f476])).
fof(f1914,plain,(
  $false | $spl40),
  inference(trivial_inequality_removal,[],[f1906])).
fof(f1906,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | $spl40),
  inference(superposition,[],[f1766,f493])).
fof(f1913,plain,(
  $false | $spl40),
  inference(trivial_inequality_removal,[],[f1907])).
fof(f1907,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | $spl40),
  inference(superposition,[],[f1766,f490])).
fof(f1912,plain,(
  $false | $spl40),
  inference(trivial_inequality_removal,[],[f1908])).
fof(f1908,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | $spl40),
  inference(superposition,[],[f1766,f494])).
fof(f1911,plain,(
  $false | $spl40),
  inference(trivial_inequality_removal,[],[f1909])).
fof(f1909,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | $spl40),
  inference(superposition,[],[f1766,f491])).
fof(f1910,plain,(
  $false | $spl40),
  inference(equality_resolution,[],[f1766])).
fof(f1826,plain,(
  ( ! [X10,X11] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl46),
  inference(cnf_transformation,[],[f1826_D])).
fof(f1826_D,plain,(
  ( ! [X10,X11] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1807,plain,(
  ( ! [X26,X27,X25] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1206,f482])).
fof(f1206,plain,(
  ( ! [X48,X49] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X48,X49)) )),
  inference(resolution,[],[f1130,f513])).
fof(f1806,plain,(
  ( ! [X24,X23,X22] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f1206,f481])).
fof(f1805,plain,(
  ( ! [X21,X19,X20] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) )),
  inference(superposition,[],[f1206,f479])).
fof(f1804,plain,(
  ( ! [X17,X18,X16] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1206,f478])).
fof(f1803,plain,(
  ( ! [X14,X15,X13] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1206,f477])).
fof(f1802,plain,(
  ( ! [X12,X10,X11] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(superposition,[],[f1206,f476])).
fof(f1744,plain,(
  ( ! [X14,X15] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15)))))) ) | $spl10),
  inference(superposition,[],[f767,f481])).
fof(f1743,plain,(
  ( ! [X12,X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) ) | $spl10),
  inference(superposition,[],[f767,f479])).
fof(f1741,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl10),
  inference(superposition,[],[f767,f477])).
fof(f1740,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) ) | $spl10),
  inference(superposition,[],[f767,f476])).
fof(f1731,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) ) | $spl10),
  inference(superposition,[],[f767,f480])).
fof(f1730,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X4)) ) | $spl10),
  inference(superposition,[],[f767,f483])).
fof(f1673,plain,(
  ( ! [X10,X11] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl38),
  inference(cnf_transformation,[],[f1673_D])).
fof(f1673_D,plain,(
  ( ! [X10,X11] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1653,plain,(
  ( ! [X26,X27,X25] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1188,f482])).
fof(f1188,plain,(
  ( ! [X17,X16] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X16,X17)) )),
  inference(resolution,[],[f1130,f421])).
fof(f1652,plain,(
  ( ! [X24,X23,X22] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f1188,f481])).
fof(f1651,plain,(
  ( ! [X21,X19,X20] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) )),
  inference(superposition,[],[f1188,f479])).
fof(f1650,plain,(
  ( ! [X17,X18,X16] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1188,f478])).
fof(f1649,plain,(
  ( ! [X14,X15,X13] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1188,f477])).
fof(f1648,plain,(
  ( ! [X12,X10,X11] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(superposition,[],[f1188,f476])).
fof(f1619,plain,(
  ( ! [X10,X11] : (sP1(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl36),
  inference(cnf_transformation,[],[f1619_D])).
fof(f1619_D,plain,(
  ( ! [X10,X11] : (sP1(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1600,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1187,f482])).
fof(f1187,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X15)) )),
  inference(resolution,[],[f1130,f420])).
fof(f1599,plain,(
  ( ! [X24,X23,X22] : (sP1(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f1187,f481])).
fof(f1598,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) )),
  inference(superposition,[],[f1187,f479])).
fof(f1597,plain,(
  ( ! [X17,X18,X16] : (sP1(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1187,f478])).
fof(f1596,plain,(
  ( ! [X14,X15,X13] : (sP1(X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1187,f477])).
fof(f1595,plain,(
  ( ! [X12,X10,X11] : (sP1(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(superposition,[],[f1187,f476])).
fof(f1566,plain,(
  ( ! [X10,X11] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl34),
  inference(cnf_transformation,[],[f1566_D])).
fof(f1566_D,plain,(
  ( ! [X10,X11] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1546,plain,(
  ( ! [X26,X27,X25] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X25,X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1185,f482])).
fof(f1185,plain,(
  ( ! [X10,X11] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X10,X11)) )),
  inference(resolution,[],[f1130,f417])).
fof(f1545,plain,(
  ( ! [X24,X23,X22] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X22,X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f1185,f481])).
fof(f1544,plain,(
  ( ! [X21,X19,X20] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X19,X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) )),
  inference(superposition,[],[f1185,f479])).
fof(f1543,plain,(
  ( ! [X17,X18,X16] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X16,X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1185,f478])).
fof(f1542,plain,(
  ( ! [X14,X15,X13] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X13,X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1185,f477])).
fof(f1541,plain,(
  ( ! [X12,X10,X11] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(superposition,[],[f1185,f476])).
fof(f1512,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl32),
  inference(cnf_transformation,[],[f1512_D])).
fof(f1512_D,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1494,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X26) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X27))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1184,f482])).
fof(f1184,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK10)),X9)) )),
  inference(resolution,[],[f1130,f416])).
fof(f1493,plain,(
  ( ! [X24,X23,X22] : (sP2(X22,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f1184,f481])).
fof(f1492,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21)))))) )),
  inference(superposition,[],[f1184,f479])).
fof(f1491,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X17) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X18))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f1184,f478])).
fof(f1490,plain,(
  ( ! [X14,X15,X13] : (sP2(X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X15))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f1184,f477])).
fof(f1489,plain,(
  ( ! [X12,X10,X11] : (sP2(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X12))))))) )),
  inference(superposition,[],[f1184,f476])).
fof(f1179,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl31),
  inference(cnf_transformation,[],[f1179_D])).
fof(f1179_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f1418,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) ) | $spl29),
  inference(superposition,[],[f1371,f482])).
fof(f1417,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6)))))) ) | $spl29),
  inference(superposition,[],[f1371,f481])).
fof(f1416,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5)))))) ) | $spl29),
  inference(superposition,[],[f1371,f479])).
fof(f1415,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) ) | $spl29),
  inference(superposition,[],[f1371,f478])).
fof(f1414,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) ) | $spl29),
  inference(superposition,[],[f1371,f477])).
fof(f1413,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))))))) ) | $spl29),
  inference(superposition,[],[f1371,f476])).
fof(f1177,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK10))) | $spl29),
  inference(cnf_transformation,[],[f1177_D])).
fof(f1177_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK10))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f1155,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK10))) | $spl3),
  inference(superposition,[],[f761,f447])).
fof(f1154,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK10))) | $spl3),
  inference(superposition,[],[f761,f445])).
fof(f1098,plain,(
  ( ! [X7] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f745,f482])).
fof(f1097,plain,(
  ( ! [X6] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f745,f481])).
fof(f1096,plain,(
  ( ! [X5] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5)))))) )),
  inference(superposition,[],[f745,f479])).
fof(f1095,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f745,f478])).
fof(f1094,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f745,f477])).
fof(f1093,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))))))) )),
  inference(superposition,[],[f745,f476])).
fof(f761,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) | $spl3),
  inference(cnf_transformation,[],[f761_D])).
fof(f761_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1148,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(forward_demodulation,[],[f1092,f458])).
fof(f1092,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK10)))),
  inference(superposition,[],[f745,f447])).
fof(f1145,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(forward_demodulation,[],[f1091,f458])).
fof(f1091,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK10)))),
  inference(superposition,[],[f745,f445])).
fof(f979,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f906,f810])).
fof(f906,plain,(
  ( ! [X204,X203] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X203),s('type.num.num',X204))))) ) | $spl20),
  inference(resolution,[],[f810,f638])).
fof(f638,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f391])).
fof(f978,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f905,f810])).
fof(f905,plain,(
  ( ! [X202,X201] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X201),s('type.num.num',X202))))) ) | $spl20),
  inference(resolution,[],[f810,f619])).
fof(f977,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f900,f810])).
fof(f900,plain,(
  ( ! [X189,X190] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X189))),s('type.num.num',X190))))) ) | $spl20),
  inference(resolution,[],[f810,f619])).
fof(f976,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f899,f810])).
fof(f899,plain,(
  ( ! [X187,X188] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X187))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X188))))))) ) | $spl20),
  inference(resolution,[],[f810,f638])).
fof(f973,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f887,f810])).
fof(f887,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(resolution,[],[f810,f613])).
fof(f613,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f381])).
fof(f972,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f886,f810])).
fof(f886,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(resolution,[],[f810,f611])).
fof(f611,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f380])).
fof(f895,plain,(
  $false | $spl20),
  inference(resolution,[],[f810,f672])).
fof(f811,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) | $spl22),
  inference(cnf_transformation,[],[f811_D])).
fof(f811_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f809,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f804])).
fof(f804,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f783,f429])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f313])).
fof(f783,plain,(
  ( ! [X6,X5] : (sP3(X6,X5,'F_0') | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) )),
  inference(superposition,[],[f774,f426])).
fof(f770,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | $spl14),
  inference(cnf_transformation,[],[f770_D])).
fof(f770_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f776,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl19),
  inference(cnf_transformation,[],[f776_D])).
fof(f776_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f743,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))),
  inference(superposition,[],[f679,f664])).
fof(f664,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.integer.INT_LE_REDUCE|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f742,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))),
  inference(superposition,[],[f679,f662])).
fof(f662,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f404])).
fof(f740,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))),
  inference(superposition,[],[f679,f646])).
fof(f771,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | $spl15),
  inference(cnf_transformation,[],[f771_D])).
fof(f771_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f773,plain,(
  ( ! [X9] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f733,f605])).
fof(f733,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) | ~p(s(bool,X9))) )),
  inference(superposition,[],[f679,f437])).
fof(f769,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(subsumption_resolution,[],[f732,f605])).
fof(f732,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f679,f437])).
fof(f715,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f679,f482])).
fof(f714,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f679,f481])).
fof(f713,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5)))))) )),
  inference(superposition,[],[f679,f479])).
fof(f712,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10)))))) )),
  inference(superposition,[],[f679,f478])).
fof(f711,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))))) )),
  inference(superposition,[],[f679,f477])).
fof(f710,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))))))) )),
  inference(superposition,[],[f679,f476])).
fof(f759,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl1),
  inference(cnf_transformation,[],[f759_D])).
fof(f759_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f762,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f709,f458])).
fof(f709,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK10)))),
  inference(superposition,[],[f679,f447])).
fof(f757,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK10) | s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f708,f458])).
fof(f708,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK10)))),
  inference(superposition,[],[f679,f445])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 465
% Passive clauses: 8821
% Generated clauses: 23449
% Final active clauses: 403
% Final passive clauses: 6463
% Input formulas: 146
% Initial clauses: 267
% 
% Selected by SInE selection: 107
% SInE iterations: 1
% 
% Duplicate literals: 230
% Trivial inequalities: 274
% Fw subsumption resolutions: 1855
% Fw demodulations: 3927
% Bw demodulations: 7
% 
% Simple tautologies: 487
% Equational tautologies: 2
% Forward subsumptions: 7009
% 
% Binary resolution: 3424
% Forward superposition: 12029
% Backward superposition: 1358
% Equality resolution: 10
% 
% Split clauses: 531
% Split components: 1078
% SAT solver clauses: 574
% SAT solver unit clauses: 32
% SAT solver binary clauses: 212
% 
% Sat splits: 531
% Sat splitting refutations: 40
% 
% Memory used [KB]: 7931
% Time elapsed: 0.431 s
% ------------------------------
% ------------------------------
% Success in time 12.182 s
