Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.combin
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.combin
% SZS output start Proof for thm.combin
fof(f4643,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4638,f1649,f1696,f1392,f1201,f62,f82,f1146,f1247])).
fof(f1247,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)))) ) | $spl42),
  inference(cnf_transformation,[],[f1247_D])).
fof(f1247_D,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1146,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(duplicate_literal_removal,[],[f1143])).
fof(f1143,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(equality_factoring,[],[f162])).
fof(f162,plain,(
  ( ! [X28,X29,X27] : (s(bool,X27) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X27),s(bool,X28)))) )),
  inference(superposition,[],[f84,f69])).
fof(f69,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X19,X20] : s(X6,X19) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X19),s(X6,X20)))),
  file('thm.combin.SAME_KEY_UPDATE_DIFFER','thm.bool.COND_CLAUSES|split|0')).
fof(f84,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f63,f63])).
fof(f63,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.combin.SAME_KEY_UPDATE_DIFFER',aHL_BOOL_CASES)).
fof(f82,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f82_D])).
fof(f82_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f62,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f21])).
fof(f21,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.combin.SAME_KEY_UPDATE_DIFFER',aHL_FALSITY)).
fof(f1201,plain,(
  ( ! [X94,X92,X90,X93,X91,X89,X96] : (s(X93,X89) != s(X93,X92) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK8(X89,X90,X91,X92,X93,X94)),s(bool,X96)))) ) | $spl24),
  inference(cnf_transformation,[],[f1201_D])).
fof(f1201_D,plain,(
  ( ! [X94,X92,X90,X93,X91,X89,X96] : (s(X93,X89) != s(X93,X92) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK8(X89,X90,X91,X92,X93,X94)),s(bool,X96)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1392,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK8(X71,X72,X73,X74,X75,X76)),s(bool,X77))) | s(X75,X71) != s(X75,X74)) ) | ($spl3 | $spl42)),
  inference(trivial_inequality_removal,[],[f1386])).
fof(f1386,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK8(X71,X72,X73,X74,X75,X76)),s(bool,X77))) | s(X75,X71) != s(X75,X74)) ) | ($spl3 | $spl42)),
  inference(superposition,[],[f1247,f142])).
fof(f142,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,sK8(X0,X1,X2,X3,X4,X5)) | s(X4,X0) != s(X4,X3)) ) | $spl3),
  inference(subsumption_resolution,[],[f141,f82])).
fof(f141,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'F_0')) | s(X4,X0) != s(X4,X3) | s(bool,'T_0') = s(bool,sK8(X0,X1,X2,X3,X4,X5))) )),
  inference(superposition,[],[f76,f63])).
fof(f76,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,sK8(X5,X4,X3,X2,X1,X0))) | s(X1,X2) != s(X1,X5)) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,sK8(X5,X4,X3,X2,X1,X0))) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,sK8(X5,X4,X3,X2,X1,X0)))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,sK8(X5,X4,X3,X2,X1,X0)),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((~p(s(bool,X6)) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,X6))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : (((~p(s(bool,X6)) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,X6)))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(X1,X2) = s(X1,X5)) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(X1,X2) = s(X1,X5)) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X7,X6] : ! [X13,X14,X5,X21] : ? [X22] : ((p(s(bool,X22)) <=> s(X6,X13) = s(X6,X21)) & s(X7,'const.combin.UPDATE_4'(s(X6,X13),s(X7,X14),s(fun(X6,X7),X5),s(X6,X21))) = s(X7,'const.bool.COND_3'(s(bool,X22),s(X7,X14),s(X7,happ(s(fun(X6,X7),X5),s(X6,X21))))))),
  file('thm.combin.SAME_KEY_UPDATE_DIFFER','thm.combin.UPDATE_def')).
fof(f1696,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,sK8(X19,X20,X21,X19,X22,X23))) ) | $spl50),
  inference(cnf_transformation,[],[f1696_D])).
fof(f1696_D,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,sK8(X19,X20,X21,X19,X22,X23))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1649,plain,(
  ( ! [X88,X87,X85,X86,X84] : (s(bool,'T_0') = s(bool,sK8(X84,X85,X86,X84,X87,X88))) ) | $spl24),
  inference(superposition,[],[f1595,f69])).
fof(f1595,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK8(X0,X1,X2,X0,X3,X4)),s(bool,X5)))) ) | $spl24),
  inference(equality_resolution,[],[f1201])).
fof(f4638,plain,(
  $false | $spl50),
  inference(subsumption_resolution,[],[f4637,f55])).
fof(f55,plain,(
  s(sK1,sK4) != s(sK1,sK5)),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  s(sK1,sK4) != s(sK1,sK5) & s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f42])).
fof(f42,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s(X1,X4) != s(X1,X5) & s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))))),
  inference(rectify,[],[f39])).
fof(f39,plain,(
  ? [X2,X3,X4,X7,X8,X9] : (s(X3,X8) != s(X3,X9) & s(fun(X2,X3),'const.combin.UPDATE_3'(s(X2,X7),s(X3,X8),s(fun(X2,X3),X4))) = s(fun(X2,X3),'const.combin.UPDATE_3'(s(X2,X7),s(X3,X9),s(fun(X2,X3),X4))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X2,X3,X4,X7,X8,X9] : (s(X3,X8) != s(X3,X9) => s(fun(X2,X3),'const.combin.UPDATE_3'(s(X2,X7),s(X3,X8),s(fun(X2,X3),X4))) != s(fun(X2,X3),'const.combin.UPDATE_3'(s(X2,X7),s(X3,X9),s(fun(X2,X3),X4))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ~! [X2,X3] : ! [X4,X7,X8,X9] : (~s(X3,X8) = s(X3,X9) => ~s(fun(X2,X3),'const.combin.UPDATE_3'(s(X2,X7),s(X3,X8),s(fun(X2,X3),X4))) = s(fun(X2,X3),'const.combin.UPDATE_3'(s(X2,X7),s(X3,X9),s(fun(X2,X3),X4))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ! [X10,X11,X12,X13,X14,X15] : (~s(X9,X14) = s(X9,X15) => ~s(fun(X8,X9),'const.combin.UPDATE_3'(s(X8,X13),s(X9,X14),s(fun(X8,X9),X10))) = s(fun(X8,X9),'const.combin.UPDATE_3'(s(X8,X13),s(X9,X15),s(fun(X8,X9),X10))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12,X13,X14,X15] : (~s(X9,X14) = s(X9,X15) => ~s(fun(X8,X9),'const.combin.UPDATE_3'(s(X8,X13),s(X9,X14),s(fun(X8,X9),X10))) = s(fun(X8,X9),'const.combin.UPDATE_3'(s(X8,X13),s(X9,X15),s(fun(X8,X9),X10))))),
  file('thm.combin.SAME_KEY_UPDATE_DIFFER','thm.combin.SAME_KEY_UPDATE_DIFFER')).
fof(f4637,plain,(
  s(sK1,sK4) = s(sK1,sK5) | $spl50),
  inference(forward_demodulation,[],[f4613,f3877])).
fof(f3877,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(X23,X21) = s(X23,'const.combin.UPDATE_4'(s(X22,X19),s(X23,X21),s(fun(X22,X23),X20),s(X22,X19)))) ) | $spl50),
  inference(forward_demodulation,[],[f3687,f69])).
fof(f3687,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(X23,'const.combin.UPDATE_4'(s(X22,X19),s(X23,X21),s(fun(X22,X23),X20),s(X22,X19))) = s(X23,'const.bool.COND_3'(s(bool,'T_0'),s(X23,X21),s(X23,happ(s(fun(X22,X23),X20),s(X22,X19)))))) ) | $spl50),
  inference(superposition,[],[f77,f1696])).
fof(f77,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,sK8(X5,X4,X3,X2,X1,X0)),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f4613,plain,(
  s(sK1,sK5) = s(sK1,'const.combin.UPDATE_4'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2),s(sK0,sK3))) | $spl50),
  inference(superposition,[],[f1535,f3877])).
fof(f1535,plain,(
  ( ! [X0] : (s(sK1,'const.combin.UPDATE_4'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2),s(sK0,X0))) = s(sK1,'const.combin.UPDATE_4'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2),s(sK0,X0)))) )),
  inference(forward_demodulation,[],[f1468,f74])).
fof(f74,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.combin.UPDATE_3'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.combin.UPDATE_3'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4))),s(X1,X5)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.combin.UPDATE_3'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4))),s(X1,X5)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7,X6] : ! [X23,X24,X25,X26] : s(X7,'const.combin.UPDATE_4'(s(X6,X23),s(X7,X24),s(fun(X6,X7),X25),s(X6,X26))) = s(X7,happ(s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X23),s(X7,X24),s(fun(X6,X7),X25))),s(X6,X26)))),
  file('thm.combin.SAME_KEY_UPDATE_DIFFER','ARITY_3_4_const.combin.UPDATE')).
fof(f1468,plain,(
  ( ! [X0] : (s(sK1,'const.combin.UPDATE_4'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2),s(sK0,X0))) = s(sK1,happ(s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))),s(sK0,X0)))) )),
  inference(superposition,[],[f74,f56])).
fof(f56,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2)))),
  inference(cnf_transformation,[],[f43])).
% SZS output end Proof for thm.combin
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 49
% Passive clauses: 953
% Generated clauses: 4401
% Final active clauses: 46
% Final passive clauses: 884
% Input formulas: 15
% Initial clauses: 23
% 
% Duplicate literals: 341
% Trivial inequalities: 61
% Fw subsumption resolutions: 215
% Fw demodulations: 240
% Bw demodulations: 1
% 
% Simple tautologies: 93
% Equational tautologies: 24
% Forward subsumptions: 2297
% 
% Binary resolution: 10
% Forward superposition: 1990
% Backward superposition: 1269
% Self superposition: 128
% Equality factoring: 60
% Equality resolution: 9
% 
% Split clauses: 121
% Split components: 242
% SAT solver clauses: 195
% SAT solver unit clauses: 24
% SAT solver binary clauses: 84
% SAT solver learnt clauses: 1
% 
% Sat splits: 121
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 193
% TWLsolver calls for satisfiability: 21
% 
% Memory used [KB]: 4221
% Time elapsed: 0.156 s
% ------------------------------
% ------------------------------
% Success in time 0.158 s
