Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.probability
% (91314)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 923
% Passive clauses: 7007
% Generated clauses: 28715
% Final active clauses: 598
% Final passive clauses: 3458
% Input formulas: 107
% Initial clauses: 249
% 
% Duplicate literals: 205
% Trivial inequalities: 114
% Fw subsumption resolutions: 1665
% Fw demodulations: 1783
% Bw demodulations: 182
% 
% Simple tautologies: 182
% Equational tautologies: 27
% Forward subsumptions: 18023
% 
% Binary resolution: 5741
% Factoring: 958
% Forward superposition: 7439
% Backward superposition: 10033
% Self superposition: 84
% Equality factoring: 2
% Equality resolution: 16
% 
% Split clauses: 271
% Split components: 544
% SAT solver clauses: 509
% SAT solver unit clauses: 11
% SAT solver binary clauses: 58
% 
% Sat splits: 271
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 499
% TWLsolver calls for satisfiability: 406
% 
% Memory used [KB]: 13560
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.probability
% (91336)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 761
% Passive clauses: 2922
% Generated clauses: 22815
% Final active clauses: 559
% Final passive clauses: 1719
% Input formulas: 107
% Initial clauses: 249
% 
% Duplicate literals: 134
% Trivial inequalities: 139
% Fw subsumption resolutions: 1446
% Fw demodulations: 2533
% Bw demodulations: 474
% Condensations: 148
% Global subsumptions: 17
% 
% Simple tautologies: 221
% Equational tautologies: 9
% Forward subsumptions: 11723
% Fw demodulations to eq. taut.: 25
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 3492
% Factoring: 306
% Forward superposition: 5968
% Backward superposition: 4215
% Self superposition: 79
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 180
% Split components: 360
% SAT solver clauses: 3644
% SAT solver unit clauses: 312
% SAT solver binary clauses: 97
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 19
% 
% Sat splits: 180
% Sat splitting refutations: 52
% 
% TWLsolver clauses: 3574
% TWLsolver calls for satisfiability: 1922
% 
% Memory used [KB]: 16886
% Time elapsed: 0.728 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.probability
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.probability
% SZS output start Proof for thm.probability
fof(f109149,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2054,f2062,f2055,f2065,f2056,f2066,f2058,f2069,f2067,f2059,f2064,f2071,f2060,f2075,f2073,f2282,f9408,f9412,f9410,f10963,f10960,f10964,f10966,f11449,f11446,f11450,f11452,f15516,f15582,f15923,f15922,f15925,f15929,f15932,f15935,f15938,f15941,f15944,f15661,f15974,f15926,f15976,f15930,f15978,f15933,f15980,f15936,f15982,f15939,f15984,f15942,f15986,f15945,f15737,f15962,f15960,f16354,f16355,f16361,f16386,f16391,f16393,f16395,f16401,f17200,f17201,f17202,f17203,f17204,f17231,f17232,f17233,f17234,f17235,f17236,f17237,f17238,f17239,f17240,f17267,f17268,f17269,f17270,f17271,f17272,f17273,f17274,f17275,f17276,f17277,f17278,f17279,f17280,f17281,f17282,f17284,f17298,f17300,f17314,f17315,f17316,f17317,f17318,f17319,f17320,f17321,f17322,f17323,f17324,f17325,f17326,f17327,f17328,f17329,f17332,f17333,f17334,f17335,f17336,f17337,f17338,f17339,f17340,f17341,f17342,f17343,f17344,f17436,f17437,f17438,f17439,f17440,f17441,f17442,f17443,f17444,f17445,f17446,f17447,f17448,f17449,f17450,f17451,f17452,f17453,f17454,f17455,f17468,f17469,f17548,f24439,f24441,f24440,f24449,f27014,f27242,f27015,f27243,f27016,f27244,f27017,f27245,f27018,f27246,f27019,f27247,f27024,f27250,f27025,f27251,f27026,f27252,f27027,f27253,f27028,f27254,f27029,f27255,f27037,f27257,f27279,f27280,f27281,f27282,f27283,f27284,f27285,f27286,f27287,f27288,f27289,f27241,f27290,f27291,f27292,f27293,f27294,f27295,f27296,f27297,f27298,f27299,f27300,f27301,f27249,f27302,f27306,f27310,f27308,f27259,f27349,f27575,f27350,f27576,f27351,f27577,f27352,f27578,f27353,f27579,f27354,f27580,f27359,f27583,f27360,f27584,f27361,f27585,f27362,f27586,f27363,f27587,f27364,f27588,f27372,f27590,f27612,f27613,f27614,f27615,f27616,f27617,f27618,f27619,f27620,f27621,f27622,f27574,f27623,f27624,f27625,f27626,f27627,f27628,f27629,f27630,f27631,f27632,f27633,f27634,f27582,f27635,f27639,f27643,f27641,f27592,f29642,f29643,f29646,f29647,f29637,f29678,f29645,f29638,f29679,f29649,f29866,f29867,f29868,f29869,f29861,f29890,f27307,f29862,f29891,f27640,f36790,f36792,f36751,f36828,f36794,f41513,f41514,f41515,f41518,f41521,f41522,f41526,f54176,f54261,f54259,f54177,f54265,f54263,f81581,f81584,f81582,f88956,f88957,f88964,f88961,f88958,f88959,f88967,f88960,f88969,f88962,f88966,f91247,f91250,f91248,f91458,f91461,f91459,f91666,f91669,f91667,f91845,f91848,f91846,f92044,f92047,f92045,f92189,f92192,f92190,f104566,f104567,f104579,f104577,f105757,f105760,f105758,f106185,f106188,f106186,f106589,f106592,f15927,f106590,f109148])).
fof(f109148,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f109147,f94509])).
fof(f94509,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X24),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X23))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X24),s('type.realax.real',X23))),s('type.realax.real',X25)))) ) | $spl12),
  inference(backward_demodulation,[],[f94269,f69203])).
fof(f69203,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X24),s('type.realax.real',X23))),s('type.realax.real',X25))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X24),s('type.realax.real',X25))),s('type.realax.real',X23)))) ) | $spl12),
  inference(forward_demodulation,[],[f69202,f6402])).
fof(f6402,plain,(
  ( ! [X8,X7] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X7),s('type.realax.real',X8)))))) )),
  inference(forward_demodulation,[],[f6401,f5024])).
fof(f5024,plain,(
  ( ! [X4,X5] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X5))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))),s('type.realax.real',X4)))) )),
  inference(superposition,[],[f475,f474])).
fof(f474,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X13,X14] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X14))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X14),s('type.realax.real',X13)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.realax.REAL_ADD_SYM')).
fof(f475,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X13,X14] : s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X13),s('type.realax.real',X14))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.real_sub')).
fof(f6401,plain,(
  ( ! [X8,X7] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))),s('type.realax.real',X8))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X7),s('type.realax.real',X8)))))) )),
  inference(forward_demodulation,[],[f6331,f5021])).
fof(f5021,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f475,f392])).
fof(f392,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X13] : s('type.realax.real',X13) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_NEGNEG')).
fof(f6331,plain,(
  ( ! [X8,X7] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X7),s('type.realax.real',X8))))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))) )),
  inference(superposition,[],[f647,f475])).
fof(f647,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(forward_demodulation,[],[f486,f475])).
fof(f486,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X13,X14] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X14))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_NEG_ADD')).
fof(f69202,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X24),s('type.realax.real',X25))),s('type.realax.real',X23))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X23),s('type.realax.real',X24))))),s('type.realax.real',X25)))) ) | $spl12),
  inference(forward_demodulation,[],[f68761,f8847])).
fof(f8847,plain,(
  ( ! [X6,X7] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X6))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) ) | $spl12),
  inference(forward_demodulation,[],[f8823,f482])).
fof(f482,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X13,X14] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X13))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_LE_NEG')).
fof(f8823,plain,(
  ( ! [X6,X7] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) ) | $spl12),
  inference(superposition,[],[f2320,f475])).
fof(f2320,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f489])).
fof(f489,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X13,X14] : s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real',X14))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X14)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_LE_LNEG')).
fof(f2292,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0') | $spl12),
  inference(backward_demodulation,[],[f2071,f376])).
fof(f376,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f51])).
fof(f51,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.prim_rec.PRE|split|0')).
fof(f68761,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X23),s('type.realax.real',X24))))),s('type.realax.real',X25))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X23),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X24),s('type.realax.real',X25)))))))) ) | $spl12),
  inference(superposition,[],[f2320,f34699])).
fof(f34699,plain,(
  ( ! [X12,X10,X11] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X10),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X11),s('type.realax.real',X12))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X10),s('type.realax.real',X11))),s('type.realax.real',X12)))) )),
  inference(backward_demodulation,[],[f34692,f8619])).
fof(f8619,plain,(
  ( ! [X12,X10,X11] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X12))),s('type.realax.real',X11))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X10),s('type.realax.real',X11))),s('type.realax.real',X12)))) )),
  inference(backward_demodulation,[],[f8618,f8613])).
fof(f8613,plain,(
  ( ! [X12,X10,X11] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X12),s('type.realax.real',X11))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X10),s('type.realax.real',X11))),s('type.realax.real',X12)))) )),
  inference(forward_demodulation,[],[f8498,f5024])).
fof(f8498,plain,(
  ( ! [X12,X10,X11] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X10),s('type.realax.real',X11))),s('type.realax.real',X12))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11))),s('type.realax.real',X12)))))) )),
  inference(superposition,[],[f540,f475])).
fof(f540,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X13,X14,X21] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X14),s('type.realax.real',X21))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X14))),s('type.realax.real',X21)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.realax.REAL_ADD_ASSOC')).
fof(f8618,plain,(
  ( ! [X14,X15,X13] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X14))),s('type.realax.real',X15))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X14),s('type.realax.real',X15)))))) )),
  inference(forward_demodulation,[],[f8522,f475])).
fof(f8522,plain,(
  ( ! [X14,X15,X13] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X14))),s('type.realax.real',X15))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X14),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X15)))))))) )),
  inference(superposition,[],[f540,f475])).
fof(f34692,plain,(
  ( ! [X14,X15,X13] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X13),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X15),s('type.realax.real',X14))))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X14))),s('type.realax.real',X15)))) )),
  inference(backward_demodulation,[],[f34494,f8618])).
fof(f34494,plain,(
  ( ! [X4,X2,X3] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X3),s('type.realax.real',X2)))))) )),
  inference(superposition,[],[f475,f6402])).
fof(f94269,plain,(
  ( ! [X19,X20,X18] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X19),s('type.realax.real',X20))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X18),s('type.realax.real',X19))),s('type.realax.real',X20)))) ) | $spl12),
  inference(superposition,[],[f8851,f392])).
fof(f8851,plain,(
  ( ! [X12,X13,X11] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11))),s('type.realax.real',X12))),s('type.realax.real',X13))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X12),s('type.realax.real',X13)))))) ) | $spl12),
  inference(forward_demodulation,[],[f8850,f647])).
fof(f8850,plain,(
  ( ! [X12,X13,X11] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X12),s('type.realax.real',X13))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X11),s('type.realax.real',X12))))),s('type.realax.real',X13)))) ) | $spl12),
  inference(forward_demodulation,[],[f8826,f2320])).
fof(f8826,plain,(
  ( ! [X12,X13,X11] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X11),s('type.realax.real',X12))))),s('type.realax.real',X13))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X11),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X12),s('type.realax.real',X13)))))))) ) | $spl12),
  inference(superposition,[],[f2320,f540])).
fof(f109147,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK13),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) | $spl12),
  inference(forward_demodulation,[],[f109146,f94269])).
fof(f109146,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | $spl12),
  inference(unit_resulting_resolution,[],[f373,f37496,f574])).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f352])).
fof(f352,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f226,f272,f271,f270,f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f270,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f271,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f272,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X9,X22,X7] : ((p(s(bool,X7)) <=> s(bool,X9) = s(bool,X22)) <=> ((p(s(bool,X7)) | p(s(bool,X22)) | p(s(bool,X9))) & (p(s(bool,X7)) | ~p(s(bool,X9)) | ~p(s(bool,X22))) & (p(s(bool,X22)) | ~p(s(bool,X9)) | ~p(s(bool,X7))) & (p(s(bool,X9)) | ~p(s(bool,X22)) | ~p(s(bool,X7)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.sat.sat_2|49|_prove_ltb')).
fof(f37496,plain,(
  ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13)),'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | $spl12),
  inference(backward_demodulation,[],[f37300,f13151])).
fof(f13151,plain,(
  ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13)),'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | $spl12),
  inference(backward_demodulation,[],[f13149,f9404])).
fof(f9404,plain,(
  ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13)),'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | $spl12),
  inference(unit_resulting_resolution,[],[f373,f2325,f575])).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f352])).
fof(f2325,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) | $spl12),
  inference(backward_demodulation,[],[f2292,f640])).
fof(f640,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13)))),
  inference(backward_demodulation,[],[f377,f372])).
fof(f372,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13)))),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) & p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) & s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f233])).
fof(f233,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) & s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X3))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))),s(fun(X0,bool),X2))))),s('type.realax.real',X3))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) & s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X3))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))),s(fun(X0,bool),X2))))),s('type.realax.real',X3))))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X3))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))),s(fun(X0,bool),X2))))),s('type.realax.real',X3))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ~! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X3))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))),s(fun(X0,bool),X2))))),s('type.realax.real',X3))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))) => s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X9))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,bool),X8))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))),s(fun(X6,bool),X8))))),s('type.realax.real',X9))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))) => s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X9))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,bool),X8))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))),s(fun(X6,bool),X8))))),s('type.realax.real',X9))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.probability.PROB_COMPL_LE|49|')).
fof(f377,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f49])).
fof(f49,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|13')).
fof(f13149,plain,(
  s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))) | $spl12),
  inference(subsumption_resolution,[],[f13148,f373])).
fof(f13148,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))) | $spl12),
  inference(forward_demodulation,[],[f13147,f698])).
fof(f698,plain,(
  s(bool,'T_0') = s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),
  inference(unit_resulting_resolution,[],[f370,f420])).
fof(f420,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.bool.EQ_CLAUSES|split|1')).
fof(f370,plain,(
  p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))))),
  inference(cnf_transformation,[],[f280])).
fof(f13147,plain,(
  s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) | $spl12),
  inference(subsumption_resolution,[],[f13128,f373])).
fof(f13128,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) | $spl12),
  inference(superposition,[],[f2329,f1840])).
fof(f1840,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f371,f420])).
fof(f371,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))))))),
  inference(cnf_transformation,[],[f280])).
fof(f2329,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))),s(fun(X0,bool),X2))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f646])).
fof(f646,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))),s(fun(X0,bool),X2))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) )),
  inference(forward_demodulation,[],[f471,f377])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))),s(fun(X0,bool),X2))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | ~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) | s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.probability.prob_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))) => s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,bool),X8))))) = s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))),s(fun(X6,bool),X8))))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.probability.PROB_COMPL')).
fof(f37300,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl12),
  inference(backward_demodulation,[],[f37299,f2293])).
fof(f2293,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(backward_demodulation,[],[f2292,f377])).
fof(f37299,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(subsumption_resolution,[],[f37274,f3733])).
fof(f3733,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(superposition,[],[f400,f2448])).
fof(f2448,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl12),
  inference(backward_demodulation,[],[f2292,f2090])).
fof(f2090,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f2078,f395])).
fof(f395,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X10] : s('type.num.num',X10) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|7')).
fof(f2078,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f395,f428])).
fof(f428,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X10] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X10))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|17')).
fof(f400,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X10] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|14')).
fof(f37274,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(superposition,[],[f4262,f2293])).
fof(f4262,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(superposition,[],[f493,f2293])).
fof(f493,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X10,X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))) <=> s('type.num.num',X10) = s('type.num.num',X11))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|19')).
fof(f373,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.probability.PROB_COMPL_LE|49|',aHL_TRUTH)).
fof(f106590,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X3) | p(s(bool,X3))) ) | $spl242),
  inference(cnf_transformation,[],[f106590_D])).
fof(f106590_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X3) | p(s(bool,X3))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f15927,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(cnf_transformation,[],[f15927_D])).
fof(f15927_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f106592,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | p(s(bool,X3)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X3)) ) | $spl12),
  inference(condensation,[],[f106591])).
fof(f106591,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X31)))))) | p(s(bool,X32)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(subsumption_resolution,[],[f106504,f15620])).
fof(f15620,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f373,f1089])).
fof(f1089,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f413,f413])).
fof(f413,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.probability.PROB_COMPL_LE|49|',aHL_BOOL_CASES)).
fof(f106504,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X31)))))) | p(s(bool,X32)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(superposition,[],[f8923,f1089])).
fof(f8923,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X29)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X28)))))) | p(s(bool,X30)) | ~sP5(X29,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X27),s('type.realax.real',X28))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X30)) ) | $spl12),
  inference(superposition,[],[f557,f2321])).
fof(f2321,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f490])).
fof(f490,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X13,X14] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X14))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_LE_RNEG')).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f342])).
fof(f342,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f341])).
fof(f341,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f272])).
fof(f106589,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | p(s(bool,X3)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X3)) ) | $spl12),
  inference(condensation,[],[f106588])).
fof(f106588,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26)))))) | p(s(bool,X27)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(subsumption_resolution,[],[f106503,f15620])).
fof(f106503,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26)))))) | p(s(bool,X27)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(superposition,[],[f8923,f1089])).
fof(f106186,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),X3) | p(s(bool,X3))) ) | $spl240),
  inference(cnf_transformation,[],[f106186_D])).
fof(f106186_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),X3) | p(s(bool,X3))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f106188,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | p(s(bool,X3)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),X3)) ) | $spl12),
  inference(condensation,[],[f106187])).
fof(f106187,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30))),s('type.realax.real',X31)))) | p(s(bool,X32)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31)))),X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(subsumption_resolution,[],[f106100,f15620])).
fof(f106100,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30))),s('type.realax.real',X31)))) | p(s(bool,X32)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31)))),X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(superposition,[],[f8840,f1089])).
fof(f8840,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X29)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X27))),s('type.realax.real',X28)))) | p(s(bool,X30)) | ~sP5(X29,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X27),s('type.realax.real',X28)))),X30)) ) | $spl12),
  inference(superposition,[],[f557,f2320])).
fof(f106185,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | p(s(bool,X3)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),X3)) ) | $spl12),
  inference(condensation,[],[f106184])).
fof(f106184,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X25))),s('type.realax.real',X26)))) | p(s(bool,X27)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26)))),X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(subsumption_resolution,[],[f106099,f15620])).
fof(f106099,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X25))),s('type.realax.real',X26)))) | p(s(bool,X27)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26)))),X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(superposition,[],[f8840,f1089])).
fof(f105758,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~sP5('T_0','const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X2)))),X3) | p(s(bool,X3))) ) | $spl238),
  inference(cnf_transformation,[],[f105758_D])).
fof(f105758_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~sP5('T_0','const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X2)))),X3) | p(s(bool,X3))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f105760,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,X3)) | ~sP5('T_0','const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X2)))),X3)) )),
  inference(condensation,[],[f105759])).
fof(f105759,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real',X37)))) | p(s(bool,X38)) | ~sP5('T_0','const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X39),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X39),s('type.realax.real',X37)))),X38) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f105672,f15620])).
fof(f105672,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real',X37)))) | p(s(bool,X38)) | ~sP5('T_0','const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X39),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X39),s('type.realax.real',X37)))),X38) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f8764,f1089])).
fof(f8764,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,X52)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',X51)))) | p(s(bool,X53)) | ~sP5(X52,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X49),s('type.realax.real',X50))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X49),s('type.realax.real',X51)))),X53)) )),
  inference(superposition,[],[f557,f541])).
fof(f541,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X13,X14,X21] : s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',X21))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X14))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X21)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_LT_LADD')).
fof(f105757,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,X3)) | ~sP5('T_0','const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X2)))),X3)) )),
  inference(condensation,[],[f105756])).
fof(f105756,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) | p(s(bool,X32)) | ~sP5('T_0','const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X33),s('type.realax.real',X30))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X33),s('type.realax.real',X31)))),X32) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f105671,f15620])).
fof(f105671,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) | p(s(bool,X32)) | ~sP5('T_0','const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X33),s('type.realax.real',X30))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X33),s('type.realax.real',X31)))),X32) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f8764,f1089])).
fof(f104577,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK14(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK14(X0,X1,bool,bool))))) ) | $spl234),
  inference(cnf_transformation,[],[f104577_D])).
fof(f104577_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK14(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK14(X0,X1,bool,bool))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f104579,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl237),
  inference(cnf_transformation,[],[f104579_D])).
fof(f104579_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f104567,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK14(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK14(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f8951,f1089])).
fof(f8951,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK14(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK14(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f539,f413])).
fof(f539,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK14(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK14(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK14(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK14(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.probability.PROB_COMPL_LE|49|',aHL_EXT)).
fof(f104566,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK14(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK14(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f8951,f1089])).
fof(f92190,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0',X3)) ) | $spl232),
  inference(cnf_transformation,[],[f92190_D])).
fof(f92190_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0',X3)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f92192,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0',X3)) ) | $spl12),
  inference(condensation,[],[f92191])).
fof(f92191,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X31)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(subsumption_resolution,[],[f92114,f15620])).
fof(f92114,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X31)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(superposition,[],[f8925,f1089])).
fof(f8925,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X35),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X36)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X35),s('type.realax.real',X36))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X37,X38)) ) | $spl12),
  inference(superposition,[],[f607,f2321])).
fof(f607,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f366])).
fof(f366,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f230,f278,f277])).
fof(f277,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f278,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X9,X22,X7] : ((p(s(bool,X7)) <=> (p(s(bool,X22)) | p(s(bool,X9)))) <=> ((p(s(bool,X7)) | ~p(s(bool,X22))) & (p(s(bool,X7)) | ~p(s(bool,X9))) & (p(s(bool,X22)) | p(s(bool,X9)) | ~p(s(bool,X7)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.sat.sat_25_prove_ltb')).
fof(f92189,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0',X3)) ) | $spl12),
  inference(condensation,[],[f92188])).
fof(f92188,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(subsumption_resolution,[],[f92113,f15620])).
fof(f92113,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(superposition,[],[f8925,f1089])).
fof(f92045,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X3,'T_0')) ) | $spl230),
  inference(cnf_transformation,[],[f92045_D])).
fof(f92045_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X3,'T_0')) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f92047,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X3,'T_0')) ) | $spl12),
  inference(condensation,[],[f92046])).
fof(f92046,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X31)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(subsumption_resolution,[],[f91969,f15620])).
fof(f91969,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X31)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(superposition,[],[f8924,f1089])).
fof(f8924,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X31),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X32)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X31),s('type.realax.real',X32))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X34,X33)) ) | $spl12),
  inference(superposition,[],[f606,f2321])).
fof(f606,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f367])).
fof(f92044,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X3,'T_0')) ) | $spl12),
  inference(condensation,[],[f92043])).
fof(f92043,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(subsumption_resolution,[],[f91968,f15620])).
fof(f91968,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26)))))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(superposition,[],[f8924,f1089])).
fof(f91846,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),'T_0',X3)) ) | $spl228),
  inference(cnf_transformation,[],[f91846_D])).
fof(f91846_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),'T_0',X3)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f91848,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),'T_0',X3)) ) | $spl12),
  inference(condensation,[],[f91847])).
fof(f91847,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30))),s('type.realax.real',X31)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31)))),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(subsumption_resolution,[],[f91770,f15620])).
fof(f91770,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30))),s('type.realax.real',X31)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31)))),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(superposition,[],[f8842,f1089])).
fof(f8842,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X35))),s('type.realax.real',X36)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X35),s('type.realax.real',X36)))),X37,X38)) ) | $spl12),
  inference(superposition,[],[f607,f2320])).
fof(f91845,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),'T_0',X3)) ) | $spl12),
  inference(condensation,[],[f91844])).
fof(f91844,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X25))),s('type.realax.real',X26)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26)))),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(subsumption_resolution,[],[f91769,f15620])).
fof(f91769,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X25))),s('type.realax.real',X26)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26)))),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(superposition,[],[f8842,f1089])).
fof(f91667,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),X3,'T_0')) ) | $spl226),
  inference(cnf_transformation,[],[f91667_D])).
fof(f91667_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),X3,'T_0')) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f91669,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),X3,'T_0')) ) | $spl12),
  inference(condensation,[],[f91668])).
fof(f91668,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30))),s('type.realax.real',X31)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31)))),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(subsumption_resolution,[],[f91591,f15620])).
fof(f91591,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30))),s('type.realax.real',X31)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X30),s('type.realax.real',X31)))),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl12),
  inference(superposition,[],[f8841,f1089])).
fof(f8841,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X31))),s('type.realax.real',X32)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X31),s('type.realax.real',X32)))),X34,X33)) ) | $spl12),
  inference(superposition,[],[f606,f2320])).
fof(f91666,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2)))),X3,'T_0')) ) | $spl12),
  inference(condensation,[],[f91665])).
fof(f91665,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X25))),s('type.realax.real',X26)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26)))),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(subsumption_resolution,[],[f91590,f15620])).
fof(f91590,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X25))),s('type.realax.real',X26)))) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X25),s('type.realax.real',X26)))),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl12),
  inference(superposition,[],[f8841,f1089])).
fof(f91459,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2)))),'T_0',X4)) ) | $spl224),
  inference(cnf_transformation,[],[f91459_D])).
fof(f91459_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2)))),'T_0',X4)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f91461,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2)))),'T_0',X4)) )),
  inference(condensation,[],[f91460])).
fof(f91460,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real',X37)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X38),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X38),s('type.realax.real',X37)))),'T_0',X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f91381,f15620])).
fof(f91381,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real',X37)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X38),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X38),s('type.realax.real',X37)))),'T_0',X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f8766,f1089])).
fof(f8766,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,X62)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X60),s('type.realax.real',X61)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X59),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X59),s('type.realax.real',X61)))),X62,X63)) )),
  inference(superposition,[],[f607,f541])).
fof(f91458,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2)))),'T_0',X4)) )),
  inference(condensation,[],[f91457])).
fof(f91457,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X32),s('type.realax.real',X30))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X32),s('type.realax.real',X31)))),'T_0',X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f91380,f15620])).
fof(f91380,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X32),s('type.realax.real',X30))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X32),s('type.realax.real',X31)))),'T_0',X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f8766,f1089])).
fof(f91248,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2)))),X4,'T_0')) ) | $spl222),
  inference(cnf_transformation,[],[f91248_D])).
fof(f91248_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2)))),X4,'T_0')) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f91250,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2)))),X4,'T_0')) )),
  inference(condensation,[],[f91249])).
fof(f91249,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real',X37)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X38),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X38),s('type.realax.real',X37)))),X39,'T_0') | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f91170,f15620])).
fof(f91170,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real',X37)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X38),s('type.realax.real',X36))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X38),s('type.realax.real',X37)))),X39,'T_0') | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f8765,f1089])).
fof(f8765,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,X57)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X55),s('type.realax.real',X56)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X54),s('type.realax.real',X55))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X54),s('type.realax.real',X56)))),X58,X57)) )),
  inference(superposition,[],[f606,f541])).
fof(f91247,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2)))),X4,'T_0')) )),
  inference(condensation,[],[f91246])).
fof(f91246,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X32),s('type.realax.real',X30))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X32),s('type.realax.real',X31)))),X33,'T_0') | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f91169,f15620])).
fof(f91169,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X32),s('type.realax.real',X30))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X32),s('type.realax.real',X31)))),X33,'T_0') | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f8765,f1089])).
fof(f88966,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl216),
  inference(cnf_transformation,[],[f88966_D])).
fof(f88966_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f88962,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK14(X0,X1,bool,bool))) ) | $spl212),
  inference(cnf_transformation,[],[f88962_D])).
fof(f88962_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK14(X0,X1,bool,bool))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f88969,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl221),
  inference(cnf_transformation,[],[f88969_D])).
fof(f88969_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl221),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl221])])).
fof(f88960,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK14(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f8950,f413])).
fof(f8950,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK14(X0,X1,X2,bool))) )),
  inference(superposition,[],[f539,f413])).
fof(f88967,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl218),
  inference(cnf_transformation,[],[f88967_D])).
fof(f88967_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f88959,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK14(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f8950,f1089])).
fof(f88958,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK14(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f8950,f1089])).
fof(f88961,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl210),
  inference(cnf_transformation,[],[f88961_D])).
fof(f88961_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f88964,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl215),
  inference(cnf_transformation,[],[f88964_D])).
fof(f88964_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl215),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl215])])).
fof(f88957,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK14(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f8950,f1089])).
fof(f88956,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK14(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f8950,f1089])).
fof(f81582,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))),X3) | p(s(bool,X3))) ) | $spl208),
  inference(cnf_transformation,[],[f81582_D])).
fof(f81582_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))),X3) | p(s(bool,X3))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f81584,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,X3)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))),X3)) )),
  inference(condensation,[],[f81583])).
fof(f81583,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) | p(s(bool,X32)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X31))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30)))),X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f81509,f15620])).
fof(f81509,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) | p(s(bool,X32)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X31))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30)))),X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f6249,f1089])).
fof(f6249,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,X24)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real',X22)))) | p(s(bool,X25)) | ~sP5(X24,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X22))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X23)))),X25)) )),
  inference(superposition,[],[f557,f482])).
fof(f81581,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,X3)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))),X3)) )),
  inference(condensation,[],[f81580])).
fof(f81580,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real',X26)))) | p(s(bool,X27)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X25)))),X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f81508,f15620])).
fof(f81508,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real',X26)))) | p(s(bool,X27)) | ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X25)))),X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f6249,f1089])).
fof(f54263,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl204),
  inference(cnf_transformation,[],[f54263_D])).
fof(f54263_D,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f54265,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl206),
  inference(cnf_transformation,[],[f54265_D])).
fof(f54265_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f54177,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl12),
  inference(superposition,[],[f5270,f37307])).
fof(f37307,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl12),
  inference(backward_demodulation,[],[f37300,f3219])).
fof(f3219,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl12),
  inference(superposition,[],[f401,f2293])).
fof(f401,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X10] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X10))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|16')).
fof(f5270,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)))) ) | $spl12),
  inference(superposition,[],[f2313,f413])).
fof(f2313,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f412])).
fof(f412,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X10] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X10))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.num.|48|_0')))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|24')).
fof(f54259,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl200),
  inference(cnf_transformation,[],[f54259_D])).
fof(f54259_D,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f54261,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl202),
  inference(cnf_transformation,[],[f54261_D])).
fof(f54261_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f54176,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl12),
  inference(superposition,[],[f5270,f37306])).
fof(f37306,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl12),
  inference(backward_demodulation,[],[f37300,f3138])).
fof(f3138,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl12),
  inference(superposition,[],[f400,f2293])).
fof(f41526,plain,(
  $false | ($spl12 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f41525,f41512])).
fof(f41512,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl21)),
  inference(trivial_inequality_removal,[],[f41511])).
fof(f41511,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl21)),
  inference(superposition,[],[f37333,f413])).
fof(f37333,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl21)),
  inference(backward_demodulation,[],[f37300,f9412])).
fof(f41525,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f17581,f37300])).
fof(f17581,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl30)),
  inference(forward_demodulation,[],[f13150,f15922])).
fof(f13150,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | $spl12),
  inference(backward_demodulation,[],[f13149,f2325])).
fof(f41522,plain,(
  $false | ($spl12 | $spl21 | $spl53)),
  inference(subsumption_resolution,[],[f41512,f37547])).
fof(f37547,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl53)),
  inference(backward_demodulation,[],[f37300,f15962])).
fof(f41521,plain,(
  $false | ($spl12 | $spl21 | $spl53)),
  inference(subsumption_resolution,[],[f41520,f37547])).
fof(f41520,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl21)),
  inference(condensation,[],[f41519])).
fof(f41519,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12)))))) ) | ($spl12 | $spl21)),
  inference(subsumption_resolution,[],[f41510,f413])).
fof(f41510,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12)))))) ) | ($spl12 | $spl21)),
  inference(superposition,[],[f37333,f1089])).
fof(f41518,plain,(
  $false | ($spl12 | $spl21 | $spl53)),
  inference(subsumption_resolution,[],[f41517,f37547])).
fof(f41517,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl21)),
  inference(condensation,[],[f41516])).
fof(f41516,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl12 | $spl21)),
  inference(subsumption_resolution,[],[f41509,f413])).
fof(f41509,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl12 | $spl21)),
  inference(superposition,[],[f37333,f1089])).
fof(f41515,plain,(
  $false | ($spl12 | $spl21 | $spl53)),
  inference(subsumption_resolution,[],[f41500,f37547])).
fof(f41500,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl21)),
  inference(unit_resulting_resolution,[],[f37333,f413])).
fof(f41514,plain,(
  $false | ($spl12 | $spl21 | $spl53)),
  inference(subsumption_resolution,[],[f41501,f37547])).
fof(f41501,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl21)),
  inference(unit_resulting_resolution,[],[f693,f37333,f1089])).
fof(f693,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f641,f419])).
fof(f419,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f283])).
fof(f641,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f379,f373])).
fof(f379,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f24])).
fof(f24,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.bool.NOT_CLAUSES|split|1')).
fof(f41513,plain,(
  $false | ($spl12 | $spl21 | $spl53)),
  inference(subsumption_resolution,[],[f41504,f37547])).
fof(f41504,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl21)),
  inference(unit_resulting_resolution,[],[f693,f37333,f1089])).
fof(f36794,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl196),
  inference(cnf_transformation,[],[f36794_D])).
fof(f36794_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f36828,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl198),
  inference(cnf_transformation,[],[f36828_D])).
fof(f36828_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f36751,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl12),
  inference(superposition,[],[f1332,f3219])).
fof(f1332,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f397,f413])).
fof(f397,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X10] : s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X10))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|33')).
fof(f36792,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl194),
  inference(cnf_transformation,[],[f36792_D])).
fof(f36792_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f36790,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl12),
  inference(forward_demodulation,[],[f36729,f397])).
fof(f36729,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl12),
  inference(superposition,[],[f1332,f3219])).
fof(f27640,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl173),
  inference(cnf_transformation,[],[f27640_D])).
fof(f27640_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f29891,plain,(
  ( ! [X3] : (sP0(X3,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl192),
  inference(cnf_transformation,[],[f29891_D])).
fof(f29891_D,plain,(
  ( ! [X3] : (sP0(X3,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f29862,plain,(
  ( ! [X3] : (sP0(X3,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))))) ) | $spl12),
  inference(superposition,[],[f5189,f3219])).
fof(f5189,plain,(
  ( ! [X8,X7] : (sP0(X8,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.ZERO_0'))))) ) | $spl12),
  inference(superposition,[],[f505,f2312])).
fof(f2312,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f411])).
fof(f411,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X10] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.num.|48|_0')))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|27')).
fof(f505,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f323])).
fof(f323,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f27307,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl115),
  inference(cnf_transformation,[],[f27307_D])).
fof(f27307_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f29890,plain,(
  ( ! [X2] : (sP0(X2,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl190),
  inference(cnf_transformation,[],[f29890_D])).
fof(f29890_D,plain,(
  ( ! [X2] : (sP0(X2,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f29861,plain,(
  ( ! [X2] : (sP0(X2,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))))) ) | $spl12),
  inference(superposition,[],[f5189,f3138])).
fof(f29869,plain,(
  ( ! [X3] : (sP0(X3,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl188),
  inference(cnf_transformation,[],[f29869_D])).
fof(f29869_D,plain,(
  ( ! [X3] : (sP0(X3,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0')))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f29868,plain,(
  ( ! [X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | sP0(X3,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(forward_demodulation,[],[f29840,f2312])).
fof(f29840,plain,(
  ( ! [X3] : (sP0(X3,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0'))))) ) | $spl12),
  inference(superposition,[],[f5189,f3219])).
fof(f29867,plain,(
  ( ! [X2] : (sP0(X2,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl186),
  inference(cnf_transformation,[],[f29867_D])).
fof(f29867_D,plain,(
  ( ! [X2] : (sP0(X2,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0')))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f29866,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | sP0(X2,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(forward_demodulation,[],[f29839,f2312])).
fof(f29839,plain,(
  ( ! [X2] : (sP0(X2,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0'))))) ) | $spl12),
  inference(superposition,[],[f5189,f3138])).
fof(f29649,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl181),
  inference(cnf_transformation,[],[f29649_D])).
fof(f29649_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f29679,plain,(
  ( ! [X3] : (sP0(X3,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl184),
  inference(cnf_transformation,[],[f29679_D])).
fof(f29679_D,plain,(
  ( ! [X3] : (sP0(X3,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f29638,plain,(
  ( ! [X3] : (sP0(X3,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f5113,f3219])).
fof(f5113,plain,(
  ( ! [X6,X7] : (sP0(X7,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X6))))) ) | $spl12),
  inference(superposition,[],[f505,f2311])).
fof(f2311,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f410])).
fof(f410,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X10] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X10))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|21')).
fof(f29645,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl177),
  inference(cnf_transformation,[],[f29645_D])).
fof(f29645_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f29678,plain,(
  ( ! [X2] : (sP0(X2,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl182),
  inference(cnf_transformation,[],[f29678_D])).
fof(f29678_D,plain,(
  ( ! [X2] : (sP0(X2,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f29637,plain,(
  ( ! [X2] : (sP0(X2,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f5113,f3138])).
fof(f29647,plain,(
  ( ! [X3] : (sP0(X3,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl178),
  inference(cnf_transformation,[],[f29647_D])).
fof(f29647_D,plain,(
  ( ! [X3] : (sP0(X3,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f29646,plain,(
  ( ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP0(X3,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl12),
  inference(forward_demodulation,[],[f29616,f2311])).
fof(f29616,plain,(
  ( ! [X3] : (sP0(X3,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f5113,f3219])).
fof(f29643,plain,(
  ( ! [X2] : (sP0(X2,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl174),
  inference(cnf_transformation,[],[f29643_D])).
fof(f29643_D,plain,(
  ( ! [X2] : (sP0(X2,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f29642,plain,(
  ( ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP0(X2,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl12),
  inference(forward_demodulation,[],[f29615,f2311])).
fof(f29615,plain,(
  ( ! [X2] : (sP0(X2,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f5113,f3138])).
fof(f27592,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl147),
  inference(cnf_transformation,[],[f27592_D])).
fof(f27592_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f27641,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl172),
  inference(cnf_transformation,[],[f27641_D])).
fof(f27641_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f27643,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(forward_demodulation,[],[f27518,f2312])).
fof(f27518,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl12),
  inference(superposition,[],[f2454,f3219])).
fof(f2454,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))))) ) | $spl12),
  inference(forward_demodulation,[],[f2319,f2292])).
fof(f2319,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f434])).
fof(f434,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X10))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.arithmetic.LESS_EQ_|48|')).
fof(f27639,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(forward_demodulation,[],[f27516,f5828])).
fof(f5828,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) ) | $spl12),
  inference(superposition,[],[f478,f2448])).
fof(f478,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X10,X11] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X10))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|31')).
fof(f27516,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl12),
  inference(superposition,[],[f2452,f3219])).
fof(f2452,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))))) ) | $spl12),
  inference(forward_demodulation,[],[f2317,f2292])).
fof(f2317,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f432])).
fof(f432,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X10] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X10))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|30')).
fof(f27635,plain,(
  ( ! [X80,X81] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X80,X81) | ~p(s(bool,X80))) ) | $spl170),
  inference(cnf_transformation,[],[f27635_D])).
fof(f27635_D,plain,(
  ( ! [X80,X81] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X80,X81) | ~p(s(bool,X80))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f27582,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl131),
  inference(cnf_transformation,[],[f27582_D])).
fof(f27582_D,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f27634,plain,(
  ( ! [X80,X81] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X80,X81) | ~p(s(bool,X80))) ) | $spl12),
  inference(forward_demodulation,[],[f27446,f397])).
fof(f27446,plain,(
  ( ! [X80,X81] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X80,X81) | ~p(s(bool,X80)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1343,f3219])).
fof(f1343,plain,(
  ( ! [X17,X18,X16] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16)))),X17,X18) | ~p(s(bool,X17)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X16))))) )),
  inference(superposition,[],[f607,f397])).
fof(f27633,plain,(
  ( ! [X78,X79] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X78,X79) | ~p(s(bool,X79))) ) | $spl168),
  inference(cnf_transformation,[],[f27633_D])).
fof(f27633_D,plain,(
  ( ! [X78,X79] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X78,X79) | ~p(s(bool,X79))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f27632,plain,(
  ( ! [X78,X79] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X78,X79) | ~p(s(bool,X79))) ) | $spl12),
  inference(forward_demodulation,[],[f27445,f397])).
fof(f27445,plain,(
  ( ! [X78,X79] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X78,X79) | ~p(s(bool,X79)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1342,f3219])).
fof(f1342,plain,(
  ( ! [X14,X15,X13] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13)))),X15,X14) | ~p(s(bool,X14)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X13))))) )),
  inference(superposition,[],[f606,f397])).
fof(f27631,plain,(
  ( ! [X77] : (sP1(X77,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X77))) ) | $spl166),
  inference(cnf_transformation,[],[f27631_D])).
fof(f27631_D,plain,(
  ( ! [X77] : (sP1(X77,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X77))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f27630,plain,(
  ( ! [X77] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP1(X77,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X77))) ) | $spl12),
  inference(forward_demodulation,[],[f27444,f397])).
fof(f27444,plain,(
  ( ! [X77] : (sP1(X77,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X77)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1341,f3219])).
fof(f1341,plain,(
  ( ! [X12,X11] : (sP1(X12,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))))) | p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X11))))) )),
  inference(superposition,[],[f524,f397])).
fof(f524,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f330])).
fof(f330,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f329])).
fof(f329,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f27629,plain,(
  ( ! [X76] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X76) | p(s(bool,X76))) ) | $spl164),
  inference(cnf_transformation,[],[f27629_D])).
fof(f27629_D,plain,(
  ( ! [X76] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X76) | p(s(bool,X76))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f27628,plain,(
  ( ! [X76] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X76) | p(s(bool,X76))) ) | $spl12),
  inference(forward_demodulation,[],[f27443,f397])).
fof(f27443,plain,(
  ( ! [X76] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X76) | p(s(bool,X76)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1340,f3219])).
fof(f1340,plain,(
  ( ! [X10,X9] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))),X10) | p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X9))))) )),
  inference(superposition,[],[f521,f397])).
fof(f521,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f331])).
fof(f27627,plain,(
  ( ! [X75] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X75)) ) | $spl162),
  inference(cnf_transformation,[],[f27627_D])).
fof(f27627_D,plain,(
  ( ! [X75] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X75)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f27626,plain,(
  ( ! [X75] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X75)) ) | $spl12),
  inference(forward_demodulation,[],[f27442,f397])).
fof(f27442,plain,(
  ( ! [X75] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X75) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1339,f3219])).
fof(f1339,plain,(
  ( ! [X8,X7] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))),X8) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X7))))) )),
  inference(superposition,[],[f506,f397])).
fof(f506,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f324])).
fof(f27625,plain,(
  ( ! [X74] : (sP0(X74,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl160),
  inference(cnf_transformation,[],[f27625_D])).
fof(f27625_D,plain,(
  ( ! [X74] : (sP0(X74,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f27624,plain,(
  ( ! [X74] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP0(X74,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl12),
  inference(forward_demodulation,[],[f27441,f397])).
fof(f27441,plain,(
  ( ! [X74] : (sP0(X74,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1338,f3219])).
fof(f1338,plain,(
  ( ! [X6,X5] : (sP0(X6,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X5))))) )),
  inference(superposition,[],[f505,f397])).
fof(f27623,plain,(
  ( ! [X72,X73] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X72,X73) | ~p(s(bool,X72))) ) | $spl158),
  inference(cnf_transformation,[],[f27623_D])).
fof(f27623_D,plain,(
  ( ! [X72,X73] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X72,X73) | ~p(s(bool,X72))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f27574,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl117),
  inference(cnf_transformation,[],[f27574_D])).
fof(f27574_D,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f27622,plain,(
  ( ! [X72,X73] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X72,X73) | ~p(s(bool,X72))) ) | $spl12),
  inference(forward_demodulation,[],[f27436,f396])).
fof(f396,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X10] : s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X10))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|32')).
fof(f27436,plain,(
  ( ! [X72,X73] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X72,X73) | ~p(s(bool,X72)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1307,f3219])).
fof(f1307,plain,(
  ( ! [X17,X18,X16] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16)))),X17,X18) | ~p(s(bool,X17)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X16))))) )),
  inference(superposition,[],[f607,f396])).
fof(f27621,plain,(
  ( ! [X70,X71] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X70,X71) | ~p(s(bool,X71))) ) | $spl156),
  inference(cnf_transformation,[],[f27621_D])).
fof(f27621_D,plain,(
  ( ! [X70,X71] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X70,X71) | ~p(s(bool,X71))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f27620,plain,(
  ( ! [X70,X71] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X70,X71) | ~p(s(bool,X71))) ) | $spl12),
  inference(forward_demodulation,[],[f27435,f396])).
fof(f27435,plain,(
  ( ! [X70,X71] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X70,X71) | ~p(s(bool,X71)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1306,f3219])).
fof(f1306,plain,(
  ( ! [X14,X15,X13] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13)))),X15,X14) | ~p(s(bool,X14)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X13))))) )),
  inference(superposition,[],[f606,f396])).
fof(f27619,plain,(
  ( ! [X69] : (sP1(X69,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X69))) ) | $spl154),
  inference(cnf_transformation,[],[f27619_D])).
fof(f27619_D,plain,(
  ( ! [X69] : (sP1(X69,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X69))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f27618,plain,(
  ( ! [X69] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP1(X69,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X69))) ) | $spl12),
  inference(forward_demodulation,[],[f27434,f396])).
fof(f27434,plain,(
  ( ! [X69] : (sP1(X69,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X69)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1305,f3219])).
fof(f1305,plain,(
  ( ! [X12,X11] : (sP1(X12,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))))) | p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X11))))) )),
  inference(superposition,[],[f524,f396])).
fof(f27617,plain,(
  ( ! [X68] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X68) | p(s(bool,X68))) ) | $spl152),
  inference(cnf_transformation,[],[f27617_D])).
fof(f27617_D,plain,(
  ( ! [X68] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X68) | p(s(bool,X68))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f27616,plain,(
  ( ! [X68] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X68) | p(s(bool,X68))) ) | $spl12),
  inference(forward_demodulation,[],[f27433,f396])).
fof(f27433,plain,(
  ( ! [X68] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X68) | p(s(bool,X68)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1304,f3219])).
fof(f1304,plain,(
  ( ! [X10,X9] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))),X10) | p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X9))))) )),
  inference(superposition,[],[f521,f396])).
fof(f27615,plain,(
  ( ! [X67] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X67)) ) | $spl150),
  inference(cnf_transformation,[],[f27615_D])).
fof(f27615_D,plain,(
  ( ! [X67] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X67)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f27614,plain,(
  ( ! [X67] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X67)) ) | $spl12),
  inference(forward_demodulation,[],[f27432,f396])).
fof(f27432,plain,(
  ( ! [X67] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X67) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1303,f3219])).
fof(f1303,plain,(
  ( ! [X8,X7] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))),X8) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X7))))) )),
  inference(superposition,[],[f506,f396])).
fof(f27613,plain,(
  ( ! [X66] : (sP0(X66,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl148),
  inference(cnf_transformation,[],[f27613_D])).
fof(f27613_D,plain,(
  ( ! [X66] : (sP0(X66,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f27612,plain,(
  ( ! [X66] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP0(X66,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl12),
  inference(forward_demodulation,[],[f27431,f396])).
fof(f27431,plain,(
  ( ! [X66] : (sP0(X66,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1302,f3219])).
fof(f1302,plain,(
  ( ! [X6,X5] : (sP0(X6,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X5))))) )),
  inference(superposition,[],[f505,f396])).
fof(f27590,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl144),
  inference(cnf_transformation,[],[f27590_D])).
fof(f27590_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f27372,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(superposition,[],[f2314,f3219])).
fof(f2314,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f427])).
fof(f427,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f287])).
fof(f27588,plain,(
  ( ! [X30,X31] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X30,X31) | ~p(s(bool,X30))) ) | $spl142),
  inference(cnf_transformation,[],[f27588_D])).
fof(f27588_D,plain,(
  ( ! [X30,X31] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X30,X31) | ~p(s(bool,X30))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f27364,plain,(
  ( ! [X30,X31] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X30,X31) | ~p(s(bool,X30)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1343,f3219])).
fof(f27587,plain,(
  ( ! [X28,X29] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X28,X29) | ~p(s(bool,X29))) ) | $spl140),
  inference(cnf_transformation,[],[f27587_D])).
fof(f27587_D,plain,(
  ( ! [X28,X29] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X28,X29) | ~p(s(bool,X29))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f27363,plain,(
  ( ! [X28,X29] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X28,X29) | ~p(s(bool,X29)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1342,f3219])).
fof(f27586,plain,(
  ( ! [X27] : (sP1(X27,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X27))) ) | $spl138),
  inference(cnf_transformation,[],[f27586_D])).
fof(f27586_D,plain,(
  ( ! [X27] : (sP1(X27,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X27))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f27362,plain,(
  ( ! [X27] : (sP1(X27,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X27)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1341,f3219])).
fof(f27585,plain,(
  ( ! [X26] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X26) | p(s(bool,X26))) ) | $spl136),
  inference(cnf_transformation,[],[f27585_D])).
fof(f27585_D,plain,(
  ( ! [X26] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X26) | p(s(bool,X26))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f27361,plain,(
  ( ! [X26] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X26) | p(s(bool,X26)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1340,f3219])).
fof(f27584,plain,(
  ( ! [X25] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X25)) ) | $spl134),
  inference(cnf_transformation,[],[f27584_D])).
fof(f27584_D,plain,(
  ( ! [X25] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X25)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f27360,plain,(
  ( ! [X25] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X25) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1339,f3219])).
fof(f27583,plain,(
  ( ! [X24] : (sP0(X24,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl132),
  inference(cnf_transformation,[],[f27583_D])).
fof(f27583_D,plain,(
  ( ! [X24] : (sP0(X24,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f27359,plain,(
  ( ! [X24] : (sP0(X24,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1338,f3219])).
fof(f27580,plain,(
  ( ! [X23,X22] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X22,X23) | ~p(s(bool,X22))) ) | $spl128),
  inference(cnf_transformation,[],[f27580_D])).
fof(f27580_D,plain,(
  ( ! [X23,X22] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X22,X23) | ~p(s(bool,X22))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f27354,plain,(
  ( ! [X23,X22] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X22,X23) | ~p(s(bool,X22)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1307,f3219])).
fof(f27579,plain,(
  ( ! [X21,X20] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X20,X21) | ~p(s(bool,X21))) ) | $spl126),
  inference(cnf_transformation,[],[f27579_D])).
fof(f27579_D,plain,(
  ( ! [X21,X20] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X20,X21) | ~p(s(bool,X21))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f27353,plain,(
  ( ! [X21,X20] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X20,X21) | ~p(s(bool,X21)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1306,f3219])).
fof(f27578,plain,(
  ( ! [X19] : (sP1(X19,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X19))) ) | $spl124),
  inference(cnf_transformation,[],[f27578_D])).
fof(f27578_D,plain,(
  ( ! [X19] : (sP1(X19,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X19))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f27352,plain,(
  ( ! [X19] : (sP1(X19,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X19)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1305,f3219])).
fof(f27577,plain,(
  ( ! [X18] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X18) | p(s(bool,X18))) ) | $spl122),
  inference(cnf_transformation,[],[f27577_D])).
fof(f27577_D,plain,(
  ( ! [X18] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X18) | p(s(bool,X18))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f27351,plain,(
  ( ! [X18] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X18) | p(s(bool,X18)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1304,f3219])).
fof(f27576,plain,(
  ( ! [X17] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X17)) ) | $spl120),
  inference(cnf_transformation,[],[f27576_D])).
fof(f27576_D,plain,(
  ( ! [X17] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X17)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f27350,plain,(
  ( ! [X17] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X17) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1303,f3219])).
fof(f27575,plain,(
  ( ! [X16] : (sP0(X16,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl118),
  inference(cnf_transformation,[],[f27575_D])).
fof(f27575_D,plain,(
  ( ! [X16] : (sP0(X16,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f27349,plain,(
  ( ! [X16] : (sP0(X16,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1302,f3219])).
fof(f27259,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl89),
  inference(cnf_transformation,[],[f27259_D])).
fof(f27259_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f27308,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl114),
  inference(cnf_transformation,[],[f27308_D])).
fof(f27308_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f27310,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(forward_demodulation,[],[f27183,f2312])).
fof(f27183,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl12),
  inference(superposition,[],[f2454,f3138])).
fof(f27306,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(forward_demodulation,[],[f27181,f5828])).
fof(f27181,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl12),
  inference(superposition,[],[f2452,f3138])).
fof(f27302,plain,(
  ( ! [X80,X81] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X80,X81) | ~p(s(bool,X80))) ) | $spl112),
  inference(cnf_transformation,[],[f27302_D])).
fof(f27302_D,plain,(
  ( ! [X80,X81] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X80,X81) | ~p(s(bool,X80))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f27249,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl73),
  inference(cnf_transformation,[],[f27249_D])).
fof(f27249_D,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f27301,plain,(
  ( ! [X80,X81] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X80,X81) | ~p(s(bool,X80))) ) | $spl12),
  inference(forward_demodulation,[],[f27111,f397])).
fof(f27111,plain,(
  ( ! [X80,X81] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X80,X81) | ~p(s(bool,X80)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1343,f3138])).
fof(f27300,plain,(
  ( ! [X78,X79] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X78,X79) | ~p(s(bool,X79))) ) | $spl110),
  inference(cnf_transformation,[],[f27300_D])).
fof(f27300_D,plain,(
  ( ! [X78,X79] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X78,X79) | ~p(s(bool,X79))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f27299,plain,(
  ( ! [X78,X79] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X78,X79) | ~p(s(bool,X79))) ) | $spl12),
  inference(forward_demodulation,[],[f27110,f397])).
fof(f27110,plain,(
  ( ! [X78,X79] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X78,X79) | ~p(s(bool,X79)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1342,f3138])).
fof(f27298,plain,(
  ( ! [X77] : (sP1(X77,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X77))) ) | $spl108),
  inference(cnf_transformation,[],[f27298_D])).
fof(f27298_D,plain,(
  ( ! [X77] : (sP1(X77,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X77))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f27297,plain,(
  ( ! [X77] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP1(X77,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X77))) ) | $spl12),
  inference(forward_demodulation,[],[f27109,f397])).
fof(f27109,plain,(
  ( ! [X77] : (sP1(X77,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X77)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1341,f3138])).
fof(f27296,plain,(
  ( ! [X76] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X76) | p(s(bool,X76))) ) | $spl106),
  inference(cnf_transformation,[],[f27296_D])).
fof(f27296_D,plain,(
  ( ! [X76] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X76) | p(s(bool,X76))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f27295,plain,(
  ( ! [X76] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X76) | p(s(bool,X76))) ) | $spl12),
  inference(forward_demodulation,[],[f27108,f397])).
fof(f27108,plain,(
  ( ! [X76] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X76) | p(s(bool,X76)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1340,f3138])).
fof(f27294,plain,(
  ( ! [X75] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X75)) ) | $spl104),
  inference(cnf_transformation,[],[f27294_D])).
fof(f27294_D,plain,(
  ( ! [X75] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X75)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f27293,plain,(
  ( ! [X75] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X75)) ) | $spl12),
  inference(forward_demodulation,[],[f27107,f397])).
fof(f27107,plain,(
  ( ! [X75] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X75) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1339,f3138])).
fof(f27292,plain,(
  ( ! [X74] : (sP0(X74,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl102),
  inference(cnf_transformation,[],[f27292_D])).
fof(f27292_D,plain,(
  ( ! [X74] : (sP0(X74,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f27291,plain,(
  ( ! [X74] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP0(X74,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl12),
  inference(forward_demodulation,[],[f27106,f397])).
fof(f27106,plain,(
  ( ! [X74] : (sP0(X74,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1338,f3138])).
fof(f27290,plain,(
  ( ! [X72,X73] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X72,X73) | ~p(s(bool,X72))) ) | $spl100),
  inference(cnf_transformation,[],[f27290_D])).
fof(f27290_D,plain,(
  ( ! [X72,X73] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X72,X73) | ~p(s(bool,X72))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f27241,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl59),
  inference(cnf_transformation,[],[f27241_D])).
fof(f27241_D,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f27289,plain,(
  ( ! [X72,X73] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X72,X73) | ~p(s(bool,X72))) ) | $spl12),
  inference(forward_demodulation,[],[f27101,f396])).
fof(f27101,plain,(
  ( ! [X72,X73] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X72,X73) | ~p(s(bool,X72)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1307,f3138])).
fof(f27288,plain,(
  ( ! [X70,X71] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X70,X71) | ~p(s(bool,X71))) ) | $spl98),
  inference(cnf_transformation,[],[f27288_D])).
fof(f27288_D,plain,(
  ( ! [X70,X71] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X70,X71) | ~p(s(bool,X71))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f27287,plain,(
  ( ! [X70,X71] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X70,X71) | ~p(s(bool,X71))) ) | $spl12),
  inference(forward_demodulation,[],[f27100,f396])).
fof(f27100,plain,(
  ( ! [X70,X71] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X70,X71) | ~p(s(bool,X71)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1306,f3138])).
fof(f27286,plain,(
  ( ! [X69] : (sP1(X69,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X69))) ) | $spl96),
  inference(cnf_transformation,[],[f27286_D])).
fof(f27286_D,plain,(
  ( ! [X69] : (sP1(X69,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X69))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f27285,plain,(
  ( ! [X69] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP1(X69,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X69))) ) | $spl12),
  inference(forward_demodulation,[],[f27099,f396])).
fof(f27099,plain,(
  ( ! [X69] : (sP1(X69,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,X69)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1305,f3138])).
fof(f27284,plain,(
  ( ! [X68] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X68) | p(s(bool,X68))) ) | $spl94),
  inference(cnf_transformation,[],[f27284_D])).
fof(f27284_D,plain,(
  ( ! [X68] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X68) | p(s(bool,X68))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f27283,plain,(
  ( ! [X68] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X68) | p(s(bool,X68))) ) | $spl12),
  inference(forward_demodulation,[],[f27098,f396])).
fof(f27098,plain,(
  ( ! [X68] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X68) | p(s(bool,X68)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1304,f3138])).
fof(f27282,plain,(
  ( ! [X67] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X67)) ) | $spl92),
  inference(cnf_transformation,[],[f27282_D])).
fof(f27282_D,plain,(
  ( ! [X67] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X67)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f27281,plain,(
  ( ! [X67] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X67)) ) | $spl12),
  inference(forward_demodulation,[],[f27097,f396])).
fof(f27097,plain,(
  ( ! [X67] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X67) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1303,f3138])).
fof(f27280,plain,(
  ( ! [X66] : (sP0(X66,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl90),
  inference(cnf_transformation,[],[f27280_D])).
fof(f27280_D,plain,(
  ( ! [X66] : (sP0(X66,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f27279,plain,(
  ( ! [X66] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | sP0(X66,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl12),
  inference(forward_demodulation,[],[f27096,f396])).
fof(f27096,plain,(
  ( ! [X66] : (sP0(X66,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl12),
  inference(superposition,[],[f1302,f3138])).
fof(f27257,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl86),
  inference(cnf_transformation,[],[f27257_D])).
fof(f27257_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f27037,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(superposition,[],[f2314,f3138])).
fof(f27255,plain,(
  ( ! [X30,X31] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X30,X31) | ~p(s(bool,X30))) ) | $spl84),
  inference(cnf_transformation,[],[f27255_D])).
fof(f27255_D,plain,(
  ( ! [X30,X31] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X30,X31) | ~p(s(bool,X30))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f27029,plain,(
  ( ! [X30,X31] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X30,X31) | ~p(s(bool,X30)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1343,f3138])).
fof(f27254,plain,(
  ( ! [X28,X29] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X28,X29) | ~p(s(bool,X29))) ) | $spl82),
  inference(cnf_transformation,[],[f27254_D])).
fof(f27254_D,plain,(
  ( ! [X28,X29] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X28,X29) | ~p(s(bool,X29))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f27028,plain,(
  ( ! [X28,X29] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X28,X29) | ~p(s(bool,X29)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1342,f3138])).
fof(f27253,plain,(
  ( ! [X27] : (sP1(X27,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X27))) ) | $spl80),
  inference(cnf_transformation,[],[f27253_D])).
fof(f27253_D,plain,(
  ( ! [X27] : (sP1(X27,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X27))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f27027,plain,(
  ( ! [X27] : (sP1(X27,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X27)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1341,f3138])).
fof(f27252,plain,(
  ( ! [X26] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X26) | p(s(bool,X26))) ) | $spl78),
  inference(cnf_transformation,[],[f27252_D])).
fof(f27252_D,plain,(
  ( ! [X26] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X26) | p(s(bool,X26))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f27026,plain,(
  ( ! [X26] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X26) | p(s(bool,X26)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1340,f3138])).
fof(f27251,plain,(
  ( ! [X25] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X25)) ) | $spl76),
  inference(cnf_transformation,[],[f27251_D])).
fof(f27251_D,plain,(
  ( ! [X25] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X25)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f27025,plain,(
  ( ! [X25] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X25) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1339,f3138])).
fof(f27250,plain,(
  ( ! [X24] : (sP0(X24,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl74),
  inference(cnf_transformation,[],[f27250_D])).
fof(f27250_D,plain,(
  ( ! [X24] : (sP0(X24,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f27024,plain,(
  ( ! [X24] : (sP0(X24,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1338,f3138])).
fof(f27247,plain,(
  ( ! [X23,X22] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X22,X23) | ~p(s(bool,X22))) ) | $spl70),
  inference(cnf_transformation,[],[f27247_D])).
fof(f27247_D,plain,(
  ( ! [X23,X22] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X22,X23) | ~p(s(bool,X22))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f27019,plain,(
  ( ! [X23,X22] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X22,X23) | ~p(s(bool,X22)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1307,f3138])).
fof(f27246,plain,(
  ( ! [X21,X20] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X20,X21) | ~p(s(bool,X21))) ) | $spl68),
  inference(cnf_transformation,[],[f27246_D])).
fof(f27246_D,plain,(
  ( ! [X21,X20] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X20,X21) | ~p(s(bool,X21))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f27018,plain,(
  ( ! [X21,X20] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X20,X21) | ~p(s(bool,X21)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1306,f3138])).
fof(f27245,plain,(
  ( ! [X19] : (sP1(X19,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X19))) ) | $spl66),
  inference(cnf_transformation,[],[f27245_D])).
fof(f27245_D,plain,(
  ( ! [X19] : (sP1(X19,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X19))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f27017,plain,(
  ( ! [X19] : (sP1(X19,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,X19)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1305,f3138])).
fof(f27244,plain,(
  ( ! [X18] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X18) | p(s(bool,X18))) ) | $spl64),
  inference(cnf_transformation,[],[f27244_D])).
fof(f27244_D,plain,(
  ( ! [X18] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X18) | p(s(bool,X18))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f27016,plain,(
  ( ! [X18] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X18) | p(s(bool,X18)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1304,f3138])).
fof(f27243,plain,(
  ( ! [X17] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X17)) ) | $spl62),
  inference(cnf_transformation,[],[f27243_D])).
fof(f27243_D,plain,(
  ( ! [X17] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X17)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f27015,plain,(
  ( ! [X17] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),X17) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1303,f3138])).
fof(f27242,plain,(
  ( ! [X16] : (sP0(X16,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl60),
  inference(cnf_transformation,[],[f27242_D])).
fof(f27242_D,plain,(
  ( ! [X16] : (sP0(X16,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f27014,plain,(
  ( ! [X16] : (sP0(X16,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1302,f3138])).
fof(f24449,plain,(
  $false | $spl54),
  inference(subsumption_resolution,[],[f24448,f641])).
fof(f24448,plain,(
  p(s(bool,'F_0')) | $spl54),
  inference(forward_demodulation,[],[f24447,f2147])).
fof(f2147,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f2102,f716])).
fof(f716,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f383,f424])).
fof(f424,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.bool.EQ_CLAUSES|split|3')).
fof(f383,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f70])).
fof(f70,axiom,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|34')).
fof(f2102,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f396,f2090])).
fof(f24447,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))) | $spl54),
  inference(unit_resulting_resolution,[],[f641,f24440,f648])).
fof(f648,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f504,f641])).
fof(f504,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f324])).
fof(f24440,plain,(
  ( ! [X31] : (sP0(X31,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl54),
  inference(cnf_transformation,[],[f24440_D])).
fof(f24440_D,plain,(
  ( ! [X31] : (sP0(X31,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f24441,plain,(
  ( ! [X30] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X30))))))))) ) | $spl56),
  inference(cnf_transformation,[],[f24441_D])).
fof(f24441_D,plain,(
  ( ! [X30] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X30))))))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f24439,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X30)))))))) | sP0(X31,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(forward_demodulation,[],[f24438,f396])).
fof(f24438,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X30)))))))))) | sP0(X31,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(forward_demodulation,[],[f24437,f7338])).
fof(f7338,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) ) | $spl12),
  inference(superposition,[],[f485,f2448])).
fof(f485,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X10,X11] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X10),s('type.num.num',X11))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|12')).
fof(f24437,plain,(
  ( ! [X30,X31] : (sP0(X31,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X30))))))))))) ) | $spl12),
  inference(forward_demodulation,[],[f24409,f2448])).
fof(f24409,plain,(
  ( ! [X30,X31] : (sP0(X31,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X30))))))))))) ) | $spl12),
  inference(superposition,[],[f1302,f2310])).
fof(f2310,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f407])).
fof(f407,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X10] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10)))))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|9')).
fof(f17548,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16263,f373])).
fof(f16263,plain,(
  ~p(s(bool,'T_0')) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f15927,f10932])).
fof(f10932,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f10569,f2451])).
fof(f2451,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | $spl12),
  inference(forward_demodulation,[],[f2316,f2292])).
fof(f2316,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f431])).
fof(f431,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f289])).
fof(f10569,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f5199,f2454])).
fof(f5199,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))) | ($spl12 | $spl17)),
  inference(backward_demodulation,[],[f5180,f2450])).
fof(f2450,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))) | ($spl12 | $spl17)),
  inference(backward_demodulation,[],[f2292,f2266])).
fof(f2266,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl17),
  inference(unit_resulting_resolution,[],[f2075,f433])).
fof(f433,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f290])).
fof(f5180,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl12),
  inference(superposition,[],[f2312,f2293])).
fof(f17469,plain,(
  $false | ($spl12 | $spl36 | $spl53)),
  inference(subsumption_resolution,[],[f17176,f15927])).
fof(f17176,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | ($spl12 | $spl36 | $spl53)),
  inference(backward_demodulation,[],[f16414,f15962])).
fof(f16414,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | ($spl12 | $spl36)),
  inference(subsumption_resolution,[],[f16169,f373])).
fof(f16169,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | ($spl12 | $spl36)),
  inference(backward_demodulation,[],[f15927,f2453])).
fof(f2453,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | $spl12),
  inference(forward_demodulation,[],[f2318,f2292])).
fof(f2318,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f433])).
fof(f17468,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17162,f2382])).
fof(f2382,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'T_0')) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f1160])).
fof(f1160,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f388,f574])).
fof(f388,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X10] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|29')).
fof(f17162,plain,(
  ~sP5('T_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f15290])).
fof(f15290,plain,(
  ~sP5('T_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0') | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f9839,f556])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f343])).
fof(f9839,plain,(
  ~sP4('T_0','T_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f373,f373,f2449,f561])).
fof(f561,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f345])).
fof(f345,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f344])).
fof(f344,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f271])).
fof(f2449,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | ($spl12 | $spl17)),
  inference(backward_demodulation,[],[f2292,f2149])).
fof(f2149,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl17),
  inference(unit_resulting_resolution,[],[f2075,f431])).
fof(f17455,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f17150,f2340])).
fof(f2340,plain,(
  ( ! [X0,X1] : (sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X0,X1)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f670])).
fof(f670,plain,(
  ( ! [X0,X1] : (sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f375,f572])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f350])).
fof(f350,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f269])).
fof(f375,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f71])).
fof(f71,axiom,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|35')).
fof(f17150,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0','F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f12504])).
fof(f12504,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0','F_0') | ($spl12 | $spl29)),
  inference(unit_resulting_resolution,[],[f641,f663,f657,f12034,f559])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f343])).
fof(f12034,plain,(
  ~sP5('F_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl29)),
  inference(unit_resulting_resolution,[],[f641,f12029,f576])).
fof(f576,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f352])).
fof(f12029,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl29)),
  inference(backward_demodulation,[],[f11677,f2408])).
fof(f2408,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(backward_demodulation,[],[f2292,f1331])).
fof(f1331,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f397,f377])).
fof(f11677,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f11452,f424])).
fof(f657,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f641,f564])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f346])).
fof(f663,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f641,f568])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f348])).
fof(f348,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f347])).
fof(f347,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f270])).
fof(f17454,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f17149,f2339])).
fof(f2339,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X1)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f668])).
fof(f668,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f375,f571])).
fof(f571,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f351])).
fof(f17149,plain,(
  ~sP2('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f12487])).
fof(f12487,plain,(
  ~sP2('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | ($spl12 | $spl29)),
  inference(unit_resulting_resolution,[],[f641,f661,f657,f12031,f560])).
fof(f560,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f343])).
fof(f12031,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | ($spl12 | $spl29)),
  inference(unit_resulting_resolution,[],[f641,f12029,f576])).
fof(f661,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f641,f567])).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f349])).
fof(f17453,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17146,f13200])).
fof(f13200,plain,(
  ( ! [X0] : (~sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'F_0')) ) | $spl12),
  inference(unit_resulting_resolution,[],[f641,f373,f2411])).
fof(f2411,plain,(
  ( ! [X10,X8,X9] : (~sP5(X9,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X8)),X10) | ~p(s(bool,X9)) | p(s(bool,X10))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f1437])).
fof(f1437,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | ~sP5(X9,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)),X10)) )),
  inference(subsumption_resolution,[],[f1424,f373])).
fof(f1424,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | p(s(bool,X10)) | ~sP5(X9,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)),X10)) )),
  inference(superposition,[],[f557,f387])).
fof(f387,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X10] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|26')).
fof(f17146,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f12464])).
fof(f12464,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f373,f10926,f574])).
fof(f10926,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17)),
  inference(backward_demodulation,[],[f10572,f5180])).
fof(f10572,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f5199,f424])).
fof(f17452,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f17142,f4786])).
fof(f4786,plain,(
  sP5('F_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl12),
  inference(unit_resulting_resolution,[],[f2340,f657,f663,f641,f559])).
fof(f17142,plain,(
  ~sP5('F_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f12034])).
fof(f17451,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f17140,f4869])).
fof(f4869,plain,(
  sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | $spl12),
  inference(unit_resulting_resolution,[],[f2339,f2332,f661,f641,f560])).
fof(f2332,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f654])).
fof(f654,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f375,f562])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f346])).
fof(f17140,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f12031])).
fof(f17450,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17124,f2370])).
fof(f2370,plain,(
  ( ! [X4,X5] : (sP0(X5,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f925])).
fof(f925,plain,(
  ( ! [X4,X5] : (sP0(X5,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f913,f373])).
fof(f913,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | sP0(X5,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f505,f388])).
fof(f17124,plain,(
  ~sP0('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10060])).
fof(f10060,plain,(
  ~sP0('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f2449,f2449,f648])).
fof(f17449,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17123,f2370])).
fof(f17123,plain,(
  ~sP0('F_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10059])).
fof(f10059,plain,(
  ~sP0('F_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f641,f2449,f648])).
fof(f17448,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17122,f2371])).
fof(f2371,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.ZERO_0')),X7)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f926])).
fof(f926,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')),X7)) )),
  inference(subsumption_resolution,[],[f914,f373])).
fof(f914,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')),X7)) )),
  inference(superposition,[],[f506,f388])).
fof(f17122,plain,(
  ~sP0('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10051])).
fof(f10051,plain,(
  ~sP0('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f641,f2449,f648])).
fof(f17447,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17117,f2297])).
fof(f2297,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f388])).
fof(f17117,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10046])).
fof(f10046,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f373,f2449,f630])).
fof(f630,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f629])).
fof(f629,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f618])).
fof(f618,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f368])).
fof(f368,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X13,X19,X14,X20] : ((s(bool,X13) = s(bool,X19) & (p(s(bool,X19)) => s(bool,X14) = s(bool,X20))) => ((p(s(bool,X13)) => p(s(bool,X14))) <=> (p(s(bool,X19)) => p(s(bool,X20)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.bool.IMP_CONG')).
fof(f17446,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17099,f10279])).
fof(f10279,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'T_0',X1)) ) | $spl12),
  inference(unit_resulting_resolution,[],[f373,f2375])).
fof(f2375,plain,(
  ( ! [X17,X15,X16] : (sP9('const.arithmetic.>=_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.ZERO_0')),X16,X17) | ~p(s(bool,X16))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f930])).
fof(f930,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X16)) | sP9('const.arithmetic.>=_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')),X16,X17)) )),
  inference(subsumption_resolution,[],[f918,f373])).
fof(f918,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | sP9('const.arithmetic.>=_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')),X16,X17)) )),
  inference(superposition,[],[f607,f388])).
fof(f17099,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'T_0',X0)) ) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f9970])).
fof(f9970,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0',X0)) ) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f373,f2449,f595])).
fof(f595,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f362])).
fof(f362,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f361])).
fof(f361,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f360])).
fof(f360,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f278])).
fof(f17445,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17098,f10233])).
fof(f10233,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),X1,'T_0')) ) | $spl12),
  inference(unit_resulting_resolution,[],[f373,f2374])).
fof(f2374,plain,(
  ( ! [X14,X12,X13] : (sP9('const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.ZERO_0')),X14,X13) | ~p(s(bool,X13))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f929])).
fof(f929,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X13)) | sP9('const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')),X14,X13)) )),
  inference(subsumption_resolution,[],[f917,f373])).
fof(f917,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | sP9('const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')),X14,X13)) )),
  inference(superposition,[],[f606,f388])).
fof(f17098,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X0,'T_0')) ) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f9969])).
fof(f9969,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X0,'T_0')) ) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f373,f2449,f594])).
fof(f594,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f362])).
fof(f17444,plain,(
  $false | ($spl12 | $spl17 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f17067,f17381])).
fof(f17381,plain,(
  sP5('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | ($spl12 | $spl17 | $spl29 | $spl36)),
  inference(forward_demodulation,[],[f16901,f16414])).
fof(f16901,plain,(
  sP5('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | ($spl12 | $spl17 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f15343])).
fof(f15343,plain,(
  sP5('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0') | ($spl12 | $spl17 | $spl29)),
  inference(unit_resulting_resolution,[],[f373,f669,f11731,f9845,f558])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f343])).
fof(f9845,plain,(
  ( ! [X0,X1] : (sP4('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X0,X1)) ) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f2449,f564])).
fof(f11731,plain,(
  ( ! [X0,X1] : (sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X0,X1)) ) | $spl29),
  inference(unit_resulting_resolution,[],[f11452,f568])).
fof(f669,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f373,f572])).
fof(f17067,plain,(
  ~sP5('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f9904])).
fof(f9904,plain,(
  ~sP5('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f2146,f2449,f576])).
fof(f2146,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f2101,f697])).
fof(f697,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f375,f420])).
fof(f2101,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f397,f2090])).
fof(f17443,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17047,f2410])).
fof(f2410,plain,(
  ( ! [X0] : (sP4('T_0','T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f1373])).
fof(f1373,plain,(
  ( ! [X0] : (sP4('T_0','T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f1160,f556])).
fof(f17047,plain,(
  ~sP4('T_0','T_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f9839])).
fof(f17442,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17033,f13261])).
fof(f13261,plain,(
  ( ! [X0] : (~sP5('T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) ) | $spl12),
  inference(unit_resulting_resolution,[],[f641,f373,f2412])).
fof(f2412,plain,(
  ( ! [X19,X17,X18] : (~sP5(X18,'const.arithmetic.>=_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.ZERO_0')),X19) | ~p(s(bool,X18)) | p(s(bool,X19))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f1438])).
fof(f1438,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | ~sP5(X18,'const.arithmetic.>=_2'(s('type.num.num',X17),s('type.num.num','const.num.|48|_0')),X19)) )),
  inference(subsumption_resolution,[],[f1427,f373])).
fof(f1427,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | p(s(bool,X19)) | ~sP5(X18,'const.arithmetic.>=_2'(s('type.num.num',X17),s('type.num.num','const.num.|48|_0')),X19)) )),
  inference(superposition,[],[f557,f388])).
fof(f17033,plain,(
  sP5('T_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f9815])).
fof(f9815,plain,(
  sP5('T_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f665,f655,f661,f2449,f559])).
fof(f655,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f641,f563])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f346])).
fof(f665,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f570])).
fof(f570,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f351])).
fof(f17441,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17032,f2388])).
fof(f2388,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f1167])).
fof(f1167,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f373,f388,f574])).
fof(f17032,plain,(
  ~sP5('T_0','T_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f9806])).
fof(f9806,plain,(
  ~sP5('T_0','T_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f373,f373,f2449,f557])).
fof(f17440,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17028,f9716])).
fof(f9716,plain,(
  ( ! [X0] : (sP1('F_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f641,f2373])).
fof(f2373,plain,(
  ( ! [X10,X11] : (sP1(X11,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.ZERO_0'))) | p(s(bool,X11))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f928])).
fof(f928,plain,(
  ( ! [X10,X11] : (p(s(bool,X11)) | sP1(X11,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f916,f373])).
fof(f916,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | sP1(X11,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f524,f388])).
fof(f17028,plain,(
  ~sP1('F_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f9802])).
fof(f9802,plain,(
  ~sP1('F_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f641,f2449,f519])).
fof(f519,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f331])).
fof(f17439,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17027,f9696])).
fof(f9696,plain,(
  ( ! [X0] : (sP1('const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) ) | $spl12),
  inference(unit_resulting_resolution,[],[f641,f2372])).
fof(f2372,plain,(
  ( ! [X8,X9] : (sP1('const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.ZERO_0')),X9) | p(s(bool,X9))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f927])).
fof(f927,plain,(
  ( ! [X8,X9] : (p(s(bool,X9)) | sP1('const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')),X9)) )),
  inference(subsumption_resolution,[],[f915,f373])).
fof(f915,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | sP1('const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')),X9)) )),
  inference(superposition,[],[f521,f388])).
fof(f17027,plain,(
  ~sP1('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f9800])).
fof(f9800,plain,(
  ~sP1('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f641,f2449,f519])).
fof(f17438,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17024,f17389])).
fof(f17389,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl12 | $spl36)),
  inference(subsumption_resolution,[],[f16909,f16414])).
fof(f16909,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl12 | $spl36)),
  inference(backward_demodulation,[],[f16414,f493])).
fof(f17024,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) ) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f4950])).
fof(f4950,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f2447,f2456])).
fof(f2456,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X1)) ) | $spl12),
  inference(forward_demodulation,[],[f2323,f2292])).
fof(f2323,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f496])).
fof(f496,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(flattening,[],[f317])).
fof(f317,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11,X10] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X10))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X11) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.arithmetic.ADD_EQ_|48|')).
fof(f2447,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17)),
  inference(backward_demodulation,[],[f2292,f2075])).
fof(f17437,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f17023,f17389])).
fof(f17023,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) ) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f4903])).
fof(f4903,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0)))) ) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f2447,f2455])).
fof(f2455,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | $spl12),
  inference(forward_demodulation,[],[f2322,f2292])).
fof(f2322,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f495])).
fof(f495,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f318])).
fof(f17436,plain,(
  $false | ($spl11 | $spl12 | $spl36)),
  inference(subsumption_resolution,[],[f17021,f17389])).
fof(f17021,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl11 | $spl12 | $spl36)),
  inference(backward_demodulation,[],[f16414,f2445])).
fof(f2445,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl11 | $spl12)),
  inference(backward_demodulation,[],[f2292,f2069])).
fof(f17344,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16875,f4805])).
fof(f4805,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0','F_0') | $spl12),
  inference(unit_resulting_resolution,[],[f2338,f655,f661,f641,f559])).
fof(f2338,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f666])).
fof(f666,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f375,f570])).
fof(f16875,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0','F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f13053])).
fof(f13053,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0','F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f11760,f554])).
fof(f554,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f343])).
fof(f11760,plain,(
  ~sP2('F_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f641,f641,f11452,f569])).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f351])).
fof(f17343,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16874,f4869])).
fof(f16874,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f13035])).
fof(f13035,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f11746,f554])).
fof(f11746,plain,(
  ~sP2('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f641,f641,f11452,f569])).
fof(f17342,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16873,f4786])).
fof(f16873,plain,(
  ~sP5('F_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f13030])).
fof(f13030,plain,(
  ~sP5('F_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f11732,f554])).
fof(f11732,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0','F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f641,f641,f11452,f569])).
fof(f17341,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16872,f2339])).
fof(f16872,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f13013])).
fof(f13013,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | ($spl12 | $spl29)),
  inference(forward_demodulation,[],[f13006,f2293])).
fof(f13006,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f641,f1321,f657,f11731,f560])).
fof(f1321,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num',X0)),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f641,f397,f576])).
fof(f17340,plain,(
  $false | ($spl12 | $spl25 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16870,f16605])).
fof(f16605,plain,(
  ~sP5('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | ($spl12 | $spl25 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11100])).
fof(f11100,plain,(
  ~sP5('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | $spl25),
  inference(unit_resulting_resolution,[],[f2146,f10966,f576])).
fof(f16870,plain,(
  sP5('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | ($spl12 | $spl25 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f13009])).
fof(f13009,plain,(
  sP5('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0') | ($spl25 | $spl29)),
  inference(unit_resulting_resolution,[],[f373,f669,f11040,f11731,f558])).
fof(f11040,plain,(
  ( ! [X0,X1] : (sP4('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X0,X1)) ) | $spl25),
  inference(unit_resulting_resolution,[],[f10966,f564])).
fof(f17339,plain,(
  $false | ($spl12 | $spl25 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16869,f16605])).
fof(f16869,plain,(
  sP5('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f13008])).
fof(f13008,plain,(
  sP5('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0') | ($spl12 | $spl29)),
  inference(unit_resulting_resolution,[],[f373,f669,f2334,f11731,f558])).
fof(f2334,plain,(
  ( ! [X0,X1] : (sP4('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),X0,X1)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f658])).
fof(f658,plain,(
  ( ! [X0,X1] : (sP4('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f383,f564])).
fof(f17338,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16868,f2168])).
fof(f2168,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f641,f2146,f576])).
fof(f16868,plain,(
  sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f13007])).
fof(f13007,plain,(
  sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f373,f669,f657,f11731,f558])).
fof(f17337,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16867,f9544])).
fof(f9544,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl12),
  inference(superposition,[],[f1319,f2448])).
fof(f1319,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f641,f397,f576])).
fof(f16867,plain,(
  sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f13005])).
fof(f13005,plain,(
  sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29)),
  inference(unit_resulting_resolution,[],[f641,f2340,f657,f11731,f560])).
fof(f17336,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16863,f2340])).
fof(f16863,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f12999])).
fof(f12999,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | ($spl12 | $spl29)),
  inference(forward_demodulation,[],[f12985,f2293])).
fof(f12985,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),'F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f641,f1319,f657,f11730,f560])).
fof(f11730,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X1)) ) | $spl29),
  inference(unit_resulting_resolution,[],[f11452,f567])).
fof(f17335,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16862,f9544])).
fof(f16862,plain,(
  sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f12984])).
fof(f12984,plain,(
  sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl29)),
  inference(unit_resulting_resolution,[],[f641,f2339,f657,f11730,f560])).
fof(f17334,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16858,f2170])).
fof(f2170,plain,(
  ~sP5('F_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f641,f2146,f576])).
fof(f16858,plain,(
  sP5('F_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f12980])).
fof(f12980,plain,(
  sP5('F_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f641,f667,f657,f11730,f560])).
fof(f667,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f373,f571])).
fof(f17333,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16857,f4738])).
fof(f4738,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0','T_0') | $spl12),
  inference(unit_resulting_resolution,[],[f669,f653,f2335,f373,f558])).
fof(f2335,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f660])).
fof(f660,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f375,f566])).
fof(f566,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f349])).
fof(f653,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f562])).
fof(f16857,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0','T_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f12960])).
fof(f12960,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0','T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f11725,f555])).
fof(f555,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f343])).
fof(f11725,plain,(
  ~sP3('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f373,f373,f11452,f565])).
fof(f565,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f349])).
fof(f17332,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16846,f2169])).
fof(f2169,plain,(
  sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f373,f2146,f574])).
fof(f16846,plain,(
  ~sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f12922])).
fof(f12922,plain,(
  ~sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f11718,f556])).
fof(f11718,plain,(
  ~sP4('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f373,f373,f11452,f561])).
fof(f17329,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16813,f16692])).
fof(f16692,plain,(
  ~sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11681])).
fof(f11681,plain,(
  ~sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f11452,f11452,f519])).
fof(f16813,plain,(
  sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11956])).
fof(f11956,plain,(
  sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl29),
  inference(unit_resulting_resolution,[],[f11452,f2197])).
fof(f2197,plain,(
  ( ! [X3] : (sP1(X3,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f2181,f373])).
fof(f2181,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | sP1(X3,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f524,f2146])).
fof(f17328,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16812,f16692])).
fof(f16812,plain,(
  sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11955])).
fof(f11955,plain,(
  sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f11452,f2196])).
fof(f2196,plain,(
  ( ! [X2] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f2180,f373])).
fof(f2180,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X2)) )),
  inference(superposition,[],[f521,f2146])).
fof(f17327,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16809,f2194])).
fof(f2194,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(subsumption_resolution,[],[f2178,f373])).
fof(f2178,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f505,f2146])).
fof(f16809,plain,(
  ~sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11945])).
fof(f11945,plain,(
  ~sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f11452,f11452,f648])).
fof(f17326,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16808,f2194])).
fof(f16808,plain,(
  ~sP0('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11944])).
fof(f11944,plain,(
  ~sP0('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f641,f11452,f648])).
fof(f17325,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16807,f2195])).
fof(f2195,plain,(
  ( ! [X1] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X1)) )),
  inference(subsumption_resolution,[],[f2179,f373])).
fof(f2179,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X1)) )),
  inference(superposition,[],[f506,f2146])).
fof(f16807,plain,(
  ~sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11935])).
fof(f11935,plain,(
  ~sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f641,f11452,f648])).
fof(f17324,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16802,f2146])).
fof(f16802,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11929])).
fof(f11929,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f373,f11452,f630])).
fof(f17323,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16801,f2357])).
fof(f2357,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f816])).
fof(f816,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f678,f373,f631])).
fof(f631,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f597])).
fof(f597,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f362])).
fof(f678,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f375,f602])).
fof(f602,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f364])).
fof(f364,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f363])).
fof(f363,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f277])).
fof(f16801,plain,(
  ~sP9('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11924])).
fof(f11924,plain,(
  ~sP9('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f373,f11452,f11452,f608])).
fof(f608,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f367])).
fof(f17322,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16800,f2357])).
fof(f16800,plain,(
  ~sP9('T_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11923])).
fof(f11923,plain,(
  ~sP9('T_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f373,f641,f11452,f608])).
fof(f17321,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16799,f2358])).
fof(f2358,plain,(
  ( ! [X0] : (sP9('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X0)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f820])).
fof(f820,plain,(
  ( ! [X0] : (sP9('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f680,f373,f631])).
fof(f680,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f375,f603])).
fof(f603,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f365])).
fof(f16799,plain,(
  ~sP9('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11917])).
fof(f11917,plain,(
  ~sP9('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f641,f373,f11452,f608])).
fof(f17320,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16790,f2344])).
fof(f2344,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f678])).
fof(f16790,plain,(
  ~sP8('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11893])).
fof(f11893,plain,(
  ~sP8('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f373,f11452,f11452,f601])).
fof(f601,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f365])).
fof(f17319,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16789,f2344])).
fof(f16789,plain,(
  ~sP8('T_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11892])).
fof(f11892,plain,(
  ~sP8('T_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f373,f641,f11452,f601])).
fof(f17318,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16788,f2345])).
fof(f2345,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X1)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f680])).
fof(f16788,plain,(
  ~sP8('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11884])).
fof(f11884,plain,(
  ~sP8('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f373,f641,f11452,f601])).
fof(f17317,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16784,f2362])).
fof(f2362,plain,(
  ( ! [X0] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0',X0)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f870])).
fof(f870,plain,(
  ( ! [X0] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f373,f797])).
fof(f797,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f795,f373])).
fof(f795,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(superposition,[],[f607,f697])).
fof(f16784,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0',X0)) ) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11853])).
fof(f11853,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0',X0)) ) | $spl29),
  inference(unit_resulting_resolution,[],[f373,f11452,f595])).
fof(f17316,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16783,f2355])).
fof(f2355,plain,(
  ( ! [X0] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X0,'T_0')) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f791])).
fof(f791,plain,(
  ( ! [X0] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f787])).
fof(f787,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1,X0) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f785,f373])).
fof(f785,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1,X0)) )),
  inference(superposition,[],[f606,f697])).
fof(f16783,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X0,'T_0')) ) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11852])).
fof(f11852,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X0,'T_0')) ) | $spl29),
  inference(unit_resulting_resolution,[],[f373,f11452,f594])).
fof(f17315,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16782,f2415])).
fof(f2415,plain,(
  sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0','T_0') | $spl12),
  inference(backward_demodulation,[],[f2292,f1540])).
fof(f1540,plain,(
  sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f672,f373,f373,f580])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f355])).
fof(f355,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f354])).
fof(f354,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f353])).
fof(f353,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f672,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f375,f585])).
fof(f585,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f356])).
fof(f356,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f16782,plain,(
  ~sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0','T_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11851])).
fof(f11851,plain,(
  ~sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'T_0','T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f373,f373,f11452,f593])).
fof(f593,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f358])).
fof(f358,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f228,f275,f274])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X9,X22,X7] : ((p(s(bool,X7)) <=> (p(s(bool,X22)) & p(s(bool,X9)))) <=> ((p(s(bool,X7)) | ~p(s(bool,X22)) | ~p(s(bool,X9))) & (p(s(bool,X22)) | ~p(s(bool,X7))) & (p(s(bool,X9)) | ~p(s(bool,X7)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.sat.sat_23_prove_ltb')).
fof(f17314,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16775,f2341])).
fof(f2341,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f672])).
fof(f16775,plain,(
  ~sP6('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11832])).
fof(f11832,plain,(
  ~sP6('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f373,f373,f11452,f584])).
fof(f584,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f357])).
fof(f17300,plain,(
  $false | ($spl12 | $spl25 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f17299,f4810])).
fof(f4810,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl12),
  inference(unit_resulting_resolution,[],[f2338,f2333,f2336,f641,f559])).
fof(f2336,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),X1)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f662])).
fof(f662,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f383,f567])).
fof(f2333,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),X1)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f656])).
fof(f656,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f383,f563])).
fof(f17299,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl25 | $spl29 | $spl36)),
  inference(forward_demodulation,[],[f16758,f16414])).
fof(f16758,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl25 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11811])).
fof(f11811,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl25 | $spl29)),
  inference(unit_resulting_resolution,[],[f11343,f11452,f576])).
fof(f11343,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl25)),
  inference(backward_demodulation,[],[f10993,f2407])).
fof(f2407,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(backward_demodulation,[],[f2292,f1295])).
fof(f1295,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f396,f377])).
fof(f10993,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl25),
  inference(unit_resulting_resolution,[],[f10966,f424])).
fof(f17298,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16757,f4810])).
fof(f16757,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11810])).
fof(f11810,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl29),
  inference(unit_resulting_resolution,[],[f2147,f11452,f576])).
fof(f17284,plain,(
  $false | ($spl12 | $spl25 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f17283,f16573])).
fof(f16573,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl25 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11013])).
fof(f11013,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | ($spl12 | $spl25)),
  inference(unit_resulting_resolution,[],[f2338,f655,f661,f10966,f559])).
fof(f17283,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl25 | $spl29 | $spl36)),
  inference(forward_demodulation,[],[f16735,f16414])).
fof(f16735,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | ($spl12 | $spl25 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11788])).
fof(f11788,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | ($spl12 | $spl25 | $spl29)),
  inference(unit_resulting_resolution,[],[f11343,f11452,f576])).
fof(f17282,plain,(
  $false | ($spl12 | $spl25 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16734,f16573])).
fof(f16734,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11787])).
fof(f11787,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f2147,f11452,f576])).
fof(f17281,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16731,f4742])).
fof(f4742,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | $spl12),
  inference(unit_resulting_resolution,[],[f669,f2332,f2335,f373,f558])).
fof(f16731,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11784])).
fof(f11784,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f2146,f11452,f576])).
fof(f17280,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16728,f4741])).
fof(f4741,plain,(
  ( ! [X0] : (sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0')) ) | $spl12),
  inference(unit_resulting_resolution,[],[f2338,f3847,f2335,f373,f558])).
fof(f3847,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f2808,f562])).
fof(f2808,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl12),
  inference(forward_demodulation,[],[f2800,f2305])).
fof(f2305,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f399])).
fof(f399,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X13] : s('type.realax.real',X13) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_ADD_RID')).
fof(f2800,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f2305,f503])).
fof(f503,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))))),
  inference(flattening,[],[f321])).
fof(f321,plain,(
  ! [X0,X1] : (((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) <=> s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X13,X14] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X14)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X13))))) <=> s('type.realax.real',X13) = s('type.realax.real',X14))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_LE_ANTISYM')).
fof(f16728,plain,(
  ( ! [X0] : (~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0')) ) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11781])).
fof(f11781,plain,(
  ( ! [X0] : (~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0')) ) | ($spl12 | $spl29)),
  inference(unit_resulting_resolution,[],[f3819,f11452,f576])).
fof(f3819,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f2808,f420])).
fof(f17279,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16727,f4739])).
fof(f4739,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0') | $spl12),
  inference(unit_resulting_resolution,[],[f669,f685,f2335,f373,f558])).
fof(f685,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) )),
  inference(unit_resulting_resolution,[],[f370,f562])).
fof(f16727,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11780])).
fof(f11780,plain,(
  ~sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f698,f11452,f576])).
fof(f17278,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16722,f2338])).
fof(f16722,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11763])).
fof(f11763,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f11452,f11452,f11452,f569])).
fof(f17277,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16721,f2338])).
fof(f16721,plain,(
  ~sP2('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11762])).
fof(f11762,plain,(
  ~sP2('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f641,f11452,f11452,f569])).
fof(f17276,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16720,f2338])).
fof(f16720,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11761])).
fof(f11761,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f11452,f641,f11452,f569])).
fof(f17275,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16719,f2338])).
fof(f16719,plain,(
  ~sP2('F_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11760])).
fof(f17274,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16718,f2339])).
fof(f16718,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11747])).
fof(f11747,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f11452,f641,f11452,f569])).
fof(f17273,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16717,f2339])).
fof(f16717,plain,(
  ~sP2('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11746])).
fof(f17272,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16716,f2340])).
fof(f16716,plain,(
  ~sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0','F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11732])).
fof(f17271,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16713,f2335])).
fof(f16713,plain,(
  ~sP3('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11725])).
fof(f17270,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16710,f2332])).
fof(f16710,plain,(
  ~sP4('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11718])).
fof(f17269,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16695,f2171])).
fof(f2171,plain,(
  sP5('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f373,f2146,f574])).
fof(f16695,plain,(
  ~sP5('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11684])).
fof(f11684,plain,(
  ~sP5('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f373,f373,f11452,f557])).
fof(f17268,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16691,f2353])).
fof(f2353,plain,(
  sP1('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl12),
  inference(backward_demodulation,[],[f2292,f763])).
fof(f763,plain,(
  sP1('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f641,f745])).
fof(f745,plain,(
  ( ! [X0] : (sP1(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f743,f373])).
fof(f743,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f524,f697])).
fof(f16691,plain,(
  ~sP1('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11680])).
fof(f11680,plain,(
  ~sP1('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(unit_resulting_resolution,[],[f641,f11452,f519])).
fof(f17267,plain,(
  $false | ($spl12 | $spl29 | $spl36)),
  inference(subsumption_resolution,[],[f16690,f2351])).
fof(f2351,plain,(
  sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | $spl12),
  inference(backward_demodulation,[],[f2292,f739])).
fof(f739,plain,(
  sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f641,f736])).
fof(f736,plain,(
  ( ! [X0] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f734,f373])).
fof(f734,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(superposition,[],[f521,f697])).
fof(f16690,plain,(
  ~sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl29 | $spl36)),
  inference(backward_demodulation,[],[f16414,f11678])).
fof(f11678,plain,(
  ~sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f641,f11452,f519])).
fof(f17240,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16563,f16414])).
fof(f16563,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10936])).
fof(f10936,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl17)),
  inference(forward_demodulation,[],[f10935,f2448])).
fof(f10935,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f10569,f493])).
fof(f17239,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16562,f16414])).
fof(f16562,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) ) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10930])).
fof(f10930,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0)))) ) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f10569,f2455])).
fof(f17238,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16561,f16414])).
fof(f16561,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) ) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10929])).
fof(f10929,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f10569,f2456])).
fof(f17237,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16552,f2363])).
fof(f2363,plain,(
  ( ! [X4,X5] : (sP0(X5,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X4)))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f894])).
fof(f894,plain,(
  ( ! [X4,X5] : (sP0(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) )),
  inference(subsumption_resolution,[],[f882,f373])).
fof(f882,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | sP0(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f505,f387])).
fof(f16552,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10840])).
fof(f10840,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f5199,f5199,f648])).
fof(f17236,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16551,f2363])).
fof(f16551,plain,(
  ~sP0('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10839])).
fof(f10839,plain,(
  ~sP0('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f641,f5199,f648])).
fof(f17235,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16550,f2364])).
fof(f2364,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X6)),X7)) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f895])).
fof(f895,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)),X7)) )),
  inference(subsumption_resolution,[],[f883,f373])).
fof(f883,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)),X7)) )),
  inference(superposition,[],[f506,f387])).
fof(f16550,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10830])).
fof(f10830,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f641,f5199,f648])).
fof(f17234,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16545,f2296])).
fof(f2296,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f387])).
fof(f16545,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10821])).
fof(f10821,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f373,f5199,f630])).
fof(f17233,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16538,f13318])).
fof(f13318,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'F_0','F_0')) ) | $spl12),
  inference(unit_resulting_resolution,[],[f641,f641,f2418])).
fof(f2418,plain,(
  ( ! [X10,X8,X9] : (~sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X8)),X10,X9) | p(s(bool,X9)) | p(s(bool,X10))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f1753])).
fof(f1753,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | ~sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)),X10,X9)) )),
  inference(subsumption_resolution,[],[f1739,f373])).
fof(f1739,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | p(s(bool,X10)) | ~sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)),X10,X9)) )),
  inference(superposition,[],[f608,f387])).
fof(f16538,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'F_0','F_0') | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10805])).
fof(f10805,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0','F_0') | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f641,f641,f5199,f605])).
fof(f605,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f367])).
fof(f17232,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16527,f10187])).
fof(f10187,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'T_0',X1)) ) | $spl12),
  inference(unit_resulting_resolution,[],[f373,f2368])).
fof(f2368,plain,(
  ( ! [X17,X15,X16] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X15)),X16,X17) | ~p(s(bool,X16))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f899])).
fof(f899,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X16)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X15)),X16,X17)) )),
  inference(subsumption_resolution,[],[f887,f373])).
fof(f887,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X15)),X16,X17)) )),
  inference(superposition,[],[f607,f387])).
fof(f16527,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'T_0',X0)) ) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10745])).
fof(f10745,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')),'T_0',X0)) ) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f373,f5199,f595])).
fof(f17231,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16526,f10141])).
fof(f10141,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),X1,'T_0')) ) | $spl12),
  inference(unit_resulting_resolution,[],[f373,f2367])).
fof(f2367,plain,(
  ( ! [X14,X12,X13] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X12)),X14,X13) | ~p(s(bool,X13))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f898])).
fof(f898,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X13)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)),X14,X13)) )),
  inference(subsumption_resolution,[],[f886,f373])).
fof(f886,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)),X14,X13)) )),
  inference(superposition,[],[f606,f387])).
fof(f16526,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),X0,'T_0')) ) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10744])).
fof(f10744,plain,(
  ( ! [X0] : (~sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')),X0,'T_0')) ) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f373,f5199,f594])).
fof(f17204,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16455,f2409])).
fof(f2409,plain,(
  ( ! [X0] : (sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f1360])).
fof(f1360,plain,(
  ( ! [X0] : (sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f1159,f556])).
fof(f1159,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f387,f574])).
fof(f16455,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10612])).
fof(f10612,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f373,f373,f5199,f561])).
fof(f17203,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16441,f13200])).
fof(f16441,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10588])).
fof(f10588,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f665,f655,f661,f5199,f559])).
fof(f17202,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16440,f2387])).
fof(f2387,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f1166])).
fof(f1166,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f373,f387,f574])).
fof(f16440,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10579])).
fof(f10579,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f373,f373,f5199,f557])).
fof(f17201,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16436,f2378])).
fof(f2378,plain,(
  ( ! [X0] : (sP1('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f1131])).
fof(f1131,plain,(
  ( ! [X0] : (sP1('F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f641,f897])).
fof(f897,plain,(
  ( ! [X10,X11] : (sP1(X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10))) | p(s(bool,X11))) )),
  inference(subsumption_resolution,[],[f885,f373])).
fof(f885,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | sP1(X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))) )),
  inference(superposition,[],[f524,f387])).
fof(f16436,plain,(
  ~sP1('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10575])).
fof(f10575,plain,(
  ~sP1('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f641,f5199,f519])).
fof(f17200,plain,(
  $false | ($spl12 | $spl17 | $spl36)),
  inference(subsumption_resolution,[],[f16435,f2377])).
fof(f2377,plain,(
  ( ! [X0] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'F_0')) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f1069])).
fof(f1069,plain,(
  ( ! [X0] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f641,f896])).
fof(f896,plain,(
  ( ! [X8,X9] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)),X9) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f884,f373])).
fof(f884,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)),X9)) )),
  inference(superposition,[],[f521,f387])).
fof(f16435,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl17 | $spl36)),
  inference(backward_demodulation,[],[f16414,f10573])).
fof(f10573,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0') | ($spl12 | $spl17)),
  inference(unit_resulting_resolution,[],[f641,f5199,f519])).
fof(f16401,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f16400,f16382])).
fof(f16382,plain,(
  ( ! [X18] : (p(s(bool,X18))) ) | $spl36),
  inference(global_subsumption,[],[f372,f371,f370,f373,f375,f376,f377,f640,f641,f383,f384,f385,f386,f387,f388,f389,f390,f391,f392,f393,f394,f395,f396,f397,f398,f399,f400,f401,f402,f403,f643,f405,f406,f407,f644,f645,f410,f411,f412,f413,f420,f419,f424,f423,f428,f427,f432,f431,f434,f433,f646,f474,f475,f476,f477,f478,f479,f480,f481,f482,f483,f484,f485,f647,f487,f488,f489,f490,f492,f491,f494,f493,f497,f496,f495,f500,f499,f498,f503,f501,f506,f505,f648,f524,f521,f520,f519,f530,f529,f539,f540,f541,f542,f543,f544,f545,f546,f547,f560,f559,f558,f557,f556,f555,f554,f564,f563,f562,f561,f568,f567,f566,f565,f572,f571,f570,f569,f576,f575,f574,f573,f632,f580,f579,f578,f577,f587,f586,f585,f584,f593,f590,f589,f588,f600,f599,f598,f631,f596,f595,f594,f604,f603,f602,f601,f608,f607,f606,f605,f630,f650,f649,f652,f651,f654,f653,f656,f655,f658,f657,f660,f659,f662,f661,f664,f663,f666,f665,f668,f667,f670,f669,f672,f671,f674,f673,f676,f675,f678,f677,f680,f679,f682,f681,f692,f691,f690,f689,f688,f687,f686,f685,f684,f683,f694,f693,f698,f697,f716,f729,f730,f733,f736,f739,f742,f745,f756,f763,f764,f767,f770,f776,f780,f781,f784,f787,f791,f794,f797,f820,f819,f817,f816,f823,f833,f831,f870,f894,f895,f896,f897,f898,f899,f900,f925,f926,f927,f928,f929,f930,f931,f965,f966,f967,f968,f969,f970,f985,f1002,f1009,f1069,f1089,f1131,f1151,f1150,f1149,f1148,f1171,f1170,f1169,f1168,f1167,f1166,f1165,f1164,f1163,f1162,f1161,f1160,f1159,f1158,f1184,f1183,f1180,f1179,f1207,f1206,f1205,f1204,f1203,f1202,f1201,f1200,f1199,f1198,f1197,f1196,f1195,f1194,f1286,f1285,f1284,f1283,f1295,f1296,f1302,f1303,f1304,f1305,f1306,f1307,f1309,f1322,f1321,f1320,f1319,f1331,f1332,f1338,f1339,f1340,f1341,f1342,f1343,f1345,f1360,f1373,f1436,f1437,f1438,f1429,f1439,f1431,f1475,f1488,f1504,f1514,f1522,f1535,f1541,f1540,f1539,f1551,f1561,f1572,f1674,f1673,f1685,f1696,f1752,f1753,f1754,f1744,f1755,f1746,f2010,f2009,f2008,f2007,f2006,f2005,f2004,f2003,f2002,f2001,f2000,f1999,f1998,f1994,f1993,f1989,f1988,f1987,f1980,f1968,f1962,f1959,f1958,f1954,f1953,f1952,f1951,f1950,f1947,f1945,f1940,f1939,f1933,f1932,f1925,f1923,f1922,f1921,f1920,f1915,f1913,f1910,f1909,f1908,f1907,f1906,f1905,f1904,f1903,f1902,f1901,f1900,f1899,f1898,f1897,f1896,f1895,f1894,f1893,f1892,f1891,f1890,f1889,f1888,f1887,f1886,f1885,f1884,f1883,f1877,f1876,f1869,f1868,f1862,f1861,f1854,f1853,f1852,f1849,f1848,f1847,f1846,f1844,f1843,f1842,f1840,f2052,f2053,f2054,f2055,f2056,f2057,f2058,f2059,f2060,f2090,f2146,f2147,f2171,f2170,f2169,f2168,f2194,f2195,f2196,f2197,f2198,f2199,f2200,f2201,f2204,f3293,f3311,f3312,f4106,f4116,f4169,f4179,f4205,f4559,f4734,f4733,f4728,f4721,f4718,f4804,f4803,f4785,f4868,f5021,f5024,f5409,f5419,f5421,f5618,f5617,f5616,f5615,f5653,f5661,f5662,f5663,f5664,f5665,f5666,f5667,f5668,f5670,f5716,f5715,f5714,f5713,f5751,f5757,f5758,f5759,f5760,f5761,f5762,f5763,f5764,f5766,f5811,f5810,f5809,f5808,f5846,f5852,f5853,f5854,f5855,f5856,f5857,f5858,f5859,f5861,f5900,f5899,f5898,f5897,f5931,f5937,f5938,f5939,f5940,f5941,f5942,f5943,f5944,f5946,f5985,f5984,f5983,f5982,f6020,f6027,f6028,f6029,f6030,f6031,f6032,f6033,f6034,f6036,f6212,f6211,f6210,f6209,f6225,f6238,f6239,f6245,f6246,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6396,f6328,f6402,f6333,f6334,f6335,f6336,f6337,f6450,f6463,f6531,f6532,f7053,f7263,f7191,f7195,f7376,f7539,f7685,f7699,f7790,f7804,f7859,f7873,f7920,f7934,f7943,f7944,f8495,f8613,f8500,f8519,f8618,f8619,f8621,f8685,f8684,f8683,f8682,f8702,f8705,f8707,f8770,f8741,f8744,f8746,f8749,f8750,f8751,f8752,f8753,f8754,f8760,f8761,f8762,f8763,f8764,f8765,f8766,f8767,f8769,f8779,f8950,f8951,f8965,f9017,f8980,f8992,f9034,f9160,f9168,f9176,f9186,f9190,f9193,f9372,f9375,f9264,f9268,f9269,f9270,f9271,f9455,f9516,f9567,f9632,f12531,f12560,f12557,f12586,f12632,f12648,f12645,f12667,f12702,f12732,f13126,f13166,f13167,f13168,f13169,f13170,f13171,f13174,f15148,f15223,f15261,f15309,f15370,f15373,f15376,f15380,f15383,f15385,f15386,f15390,f15392,f15427,f15437,f15440,f15442,f15446,f15449,f15450,f15451,f15918,f15454,f15456,f15551,f15552,f15619,f15620,f15925,f15929,f15932,f15935,f15938,f15941,f15944,f15946,f15947,f15948,f15949,f15950,f15951,f15952,f15953,f15954,f15955,f15957,f15958,f15959,f15965,f15966,f15967,f15969,f15970,f15971,f15972,f15693,f15694,f15974,f15976,f15978,f15980,f15982,f15984,f15986,f15927,f16364,f16379,f16380,f16381])).
fof(f16381,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X18)) | sP1('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X16))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X17)))),X18)) ) | $spl36),
  inference(subsumption_resolution,[],[f16033,f373])).
fof(f16033,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP1('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X16))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X17)))),X18)) ) | $spl36),
  inference(backward_demodulation,[],[f15927,f6247])).
fof(f16380,plain,(
  ( ! [X14,X15,X13] : (sP0('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14)))),X15)) ) | $spl36),
  inference(subsumption_resolution,[],[f16032,f373])).
fof(f16032,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | sP0('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14)))),X15)) ) | $spl36),
  inference(backward_demodulation,[],[f15927,f6246])).
fof(f16379,plain,(
  ( ! [X12,X10,X11] : (sP0(X12,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11)))))) ) | $spl36),
  inference(subsumption_resolution,[],[f16031,f373])).
fof(f16031,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | sP0(X12,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11)))))) ) | $spl36),
  inference(backward_demodulation,[],[f15927,f6245])).
fof(f16364,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl36),
  inference(subsumption_resolution,[],[f16363,f373])).
fof(f16363,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl36),
  inference(forward_demodulation,[],[f16362,f15927])).
fof(f16362,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl36),
  inference(subsumption_resolution,[],[f16024,f373])).
fof(f16024,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl36),
  inference(backward_demodulation,[],[f15927,f501])).
fof(f15694,plain,(
  ( ! [X163,X161,X164,X162] : (s(bool,'T_0') != s(bool,f(s(X161,sK14(X162,X163,bool,X161)))) | s(fun(X161,bool),X162) = s(fun(X161,bool),X163) | s(bool,X164) = s(bool,g(s(X161,sK14(X162,X163,bool,X161)))) | s(bool,'T_0') = s(bool,X164)) )),
  inference(superposition,[],[f539,f1089])).
fof(f15693,plain,(
  ( ! [X158,X159,X157,X160] : (s(X159,f(s(bool,'T_0'))) != s(X159,g(s(bool,'T_0'))) | s(fun(bool,X159),X157) = s(fun(bool,X159),X158) | s(bool,X160) = s(bool,sK14(X157,X158,X159,bool)) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f539,f1089])).
fof(f15972,plain,(
  ( ! [X149,X147,X148] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X148),s('type.realax.real',X147)))) | s(bool,X149) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X147),s('type.realax.real',X148))) | s(bool,'T_0') = s(bool,X149)) )),
  inference(subsumption_resolution,[],[f15691,f373])).
fof(f15691,plain,(
  ( ! [X149,X147,X148] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X148),s('type.realax.real',X147)))) | s(bool,X149) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X147),s('type.realax.real',X148))) | s(bool,'T_0') = s(bool,X149)) )),
  inference(superposition,[],[f491,f1089])).
fof(f15971,plain,(
  ( ! [X142,X140,X143,X141] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X141),s('type.realax.real',X142)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X140),s('type.realax.real',X142)))) | s(bool,X143) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X140),s('type.realax.real',X141))) | s(bool,'T_0') = s(bool,X143)) )),
  inference(subsumption_resolution,[],[f15689,f373])).
fof(f15689,plain,(
  ( ! [X142,X140,X143,X141] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X141),s('type.realax.real',X142)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X140),s('type.realax.real',X142)))) | s(bool,X143) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X140),s('type.realax.real',X141))) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f544,f1089])).
fof(f15970,plain,(
  ( ! [X138,X136,X139,X137] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X137),s('type.realax.real',X138)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X136),s('type.realax.real',X138)))) | s(bool,X139) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X136),s('type.realax.real',X137))) | s(bool,'T_0') = s(bool,X139)) )),
  inference(subsumption_resolution,[],[f15688,f373])).
fof(f15688,plain,(
  ( ! [X138,X136,X139,X137] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X137),s('type.realax.real',X138)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X136),s('type.realax.real',X138)))) | s(bool,X139) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X136),s('type.realax.real',X137))) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f545,f1089])).
fof(f15969,plain,(
  ( ! [X134,X132,X135,X133] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X134),s('type.realax.real',X133)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X134),s('type.realax.real',X132)))) | s(bool,X135) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X132),s('type.realax.real',X133))) | s(bool,'T_0') = s(bool,X135)) )),
  inference(subsumption_resolution,[],[f15687,f373])).
fof(f15687,plain,(
  ( ! [X134,X132,X135,X133] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X134),s('type.realax.real',X133)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X134),s('type.realax.real',X132)))) | s(bool,X135) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X132),s('type.realax.real',X133))) | s(bool,'T_0') = s(bool,X135)) )),
  inference(superposition,[],[f546,f1089])).
fof(f15967,plain,(
  ( ! [X92,X93,X91] : (s('type.num.num',X91) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X92))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X91),s('type.num.num',X92)))) | s(bool,X93) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X91),s('type.num.num','const.num.SUC_1'(s('type.num.num',X92))))) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f15668,f373])).
fof(f15668,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,'T_0')) | s('type.num.num',X91) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X92))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X91),s('type.num.num',X92)))) | s(bool,X93) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X91),s('type.num.num','const.num.SUC_1'(s('type.num.num',X92))))) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f498,f1089])).
fof(f15966,plain,(
  ( ! [X90,X88,X89] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X89),s('type.num.num',X88)))) | s(bool,X90) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X88),s('type.num.num',X89))) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f15667,f373])).
fof(f15667,plain,(
  ( ! [X90,X88,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X89),s('type.num.num',X88)))) | s(bool,X90) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X88),s('type.num.num',X89))) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f530,f1089])).
fof(f15965,plain,(
  ( ! [X87,X85,X86] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X85),s('type.num.num','const.num.SUC_1'(s('type.num.num',X86)))))) | s(bool,X87) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X85),s('type.num.num',X86))) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f15666,f373])).
fof(f15666,plain,(
  ( ! [X87,X85,X86] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X85),s('type.num.num','const.num.SUC_1'(s('type.num.num',X86)))))) | s(bool,X87) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X85),s('type.num.num',X86))) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f500,f1089])).
fof(f15959,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X65),s('type.realax.real',X64)))) | s('type.realax.real',X64) = s('type.realax.real',X65) | s(bool,X66) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X64),s('type.realax.real',X65))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f15658,f373])).
fof(f15658,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X65),s('type.realax.real',X64)))) | s('type.realax.real',X64) = s('type.realax.real',X65) | s(bool,X66) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X64),s('type.realax.real',X65))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f501,f1089])).
fof(f15958,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X60),s('type.realax.real',X62)))) | s(bool,X63) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X60),s('type.realax.real',X61))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f15657,f373])).
fof(f15657,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X60),s('type.realax.real',X62)))) | s(bool,X63) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X60),s('type.realax.real',X61))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f547,f1089])).
fof(f15957,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X58),s('type.realax.real',X56))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X58),s('type.realax.real',X57)))))) | s(bool,X59) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X56),s('type.realax.real',X57))) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f15656,f373])).
fof(f15656,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X58),s('type.realax.real',X56))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X58),s('type.realax.real',X57)))))) | s(bool,X59) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X56),s('type.realax.real',X57))) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f543,f1089])).
fof(f15955,plain,(
  ( ! [X43,X41,X44,X42] : (~sP8(X41,X42,X43) | sP9(X41,X42,X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f15649,f373])).
fof(f15649,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | ~sP8(X41,X42,X43) | sP9(X41,X42,X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f631,f1089])).
fof(f15954,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | p(s(bool,X36)) | ~sP9(X34,X36,X35) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f15647,f373])).
fof(f15647,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'T_0')) | p(s(bool,X35)) | p(s(bool,X36)) | ~sP9(X34,X36,X35) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f608,f1089])).
fof(f15953,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X31)) | sP9(X30,X31,X32) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f15646,f373])).
fof(f15646,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | sP9(X30,X31,X32) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f607,f1089])).
fof(f15952,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,X27)) | sP9(X26,X28,X27) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f15645,f373])).
fof(f15645,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | sP9(X26,X28,X27) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f606,f1089])).
fof(f15951,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,X23)) | p(s(bool,X24)) | ~sP5(X23,X22,X24) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f15644,f373])).
fof(f15644,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP5(X23,X22,X24) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f557,f1089])).
fof(f15950,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X20)) | sP1(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f15643,f373])).
fof(f15643,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | p(s(bool,X20)) | sP1(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f524,f1089])).
fof(f15949,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X17)) | sP1(X16,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f15642,f373])).
fof(f15642,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | p(s(bool,X17)) | sP1(X16,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f521,f1089])).
fof(f15948,plain,(
  ( ! [X14,X15,X13] : (sP0(X13,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f15641,f373])).
fof(f15641,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | sP0(X13,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f506,f1089])).
fof(f15947,plain,(
  ( ! [X12,X10,X11] : (sP0(X11,X10) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f15640,f373])).
fof(f15640,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | sP0(X11,X10) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f505,f1089])).
fof(f15946,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f15638,f373])).
fof(f15638,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f423,f1089])).
fof(f15619,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f419,f1089])).
fof(f15552,plain,(
  ( ! [X173,X171,X174,X172] : (s(bool,X174) != s(bool,f(s(X171,sK14(X172,X173,bool,X171)))) | s(fun(X171,bool),X172) = s(fun(X171,bool),X173) | s(bool,'T_0') = s(bool,X174) | s(bool,'T_0') = s(bool,g(s(X171,sK14(X172,X173,bool,X171))))) )),
  inference(superposition,[],[f539,f1089])).
fof(f15551,plain,(
  ( ! [X167,X169,X170,X168] : (s(X169,f(s(bool,X170))) != s(X169,g(s(bool,X170))) | s(fun(bool,X169),X167) = s(fun(bool,X169),X168) | s(bool,'T_0') = s(bool,X170) | s(bool,'T_0') = s(bool,sK14(X167,X168,X169,bool))) )),
  inference(superposition,[],[f539,f1089])).
fof(f15456,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X60),s('type.realax.real',X61))) | s(bool,X62) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X59),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X59),s('type.realax.real',X61))))) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f1089,f541])).
fof(f15454,plain,(
  ( ! [X57,X56] : (s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X56))) | s(bool,X57) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56))))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f1089,f397])).
fof(f15918,plain,(
  ( ! [X54] : (s(bool,X54) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f15452,f693])).
fof(f15452,plain,(
  ( ! [X54] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X54) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f1089,f2147])).
fof(f15451,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X52))) | s(bool,X53) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52))))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f1089,f396])).
fof(f15450,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X50),s('type.num.num',X49))) | s(bool,X51) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f1089,f478])).
fof(f15449,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X46),s('type.num.num',X45))) | s(bool,X47) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46))))) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1089,f477])).
fof(f15446,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X38),s('type.num.num',X39))) | s(bool,X40) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X38))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1089,f480])).
fof(f15442,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X30),s('type.num.num',X31))) | s(bool,X32) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f1089,f476])).
fof(f15440,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real',X24))) | s(bool,X26) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X24))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X25))))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f1089,f482])).
fof(f15437,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X15),s('type.num.num',X16))) | s(bool,X17) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X15))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X16))))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f1089,f479])).
fof(f15427,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1089,f1089])).
fof(f15392,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,X74) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X72),s('type.realax.real',X73))) | s(bool,'T_0') = s(bool,X74) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X71),s('type.realax.real',X72))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X71),s('type.realax.real',X73)))))) )),
  inference(superposition,[],[f1089,f541])).
fof(f15390,plain,(
  ( ! [X68,X69] : (s(bool,X69) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X68))) | s(bool,'T_0') = s(bool,X69) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68)))))) )),
  inference(superposition,[],[f1089,f397])).
fof(f15386,plain,(
  ( ! [X64,X63] : (s(bool,X64) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X63))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X63)))))) )),
  inference(superposition,[],[f1089,f396])).
fof(f15385,plain,(
  ( ! [X61,X62,X60] : (s(bool,X62) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X61),s('type.num.num',X60))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X61)))))) )),
  inference(superposition,[],[f1089,f478])).
fof(f15383,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X56),s('type.num.num',X55))) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))) )),
  inference(superposition,[],[f1089,f477])).
fof(f15380,plain,(
  ( ! [X50,X48,X49] : (s(bool,X50) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X48),s('type.num.num',X49))) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(superposition,[],[f1089,f480])).
fof(f15376,plain,(
  ( ! [X41,X42,X40] : (s(bool,X42) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X40),s('type.num.num',X41))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41)))))) )),
  inference(superposition,[],[f1089,f476])).
fof(f15373,plain,(
  ( ! [X35,X33,X34] : (s(bool,X35) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X34),s('type.realax.real',X33))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X33))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X34)))))) )),
  inference(superposition,[],[f1089,f482])).
fof(f15370,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num',X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X24))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f1089,f479])).
fof(f15309,plain,(
  ( ! [X0] : (sP4('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f1904,f556])).
fof(f15261,plain,(
  ( ! [X0] : (sP4('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f1903,f556])).
fof(f15223,plain,(
  ( ! [X0] : (sP4('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.EVEN_1'(s('type.num.num',X0)),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1895,f556])).
fof(f15148,plain,(
  ( ! [X0] : (sP4('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.ODD_1'(s('type.num.num',X0)),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1894,f556])).
fof(f13174,plain,(
  ( ! [X14,X13] : (~sP8('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X13,X14) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X13,X14)) )),
  inference(subsumption_resolution,[],[f13143,f373])).
fof(f13143,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~sP8('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X13,X14) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X13,X14)) )),
  inference(superposition,[],[f631,f1840])).
fof(f13171,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | ~sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X11,X10)) )),
  inference(subsumption_resolution,[],[f13141,f373])).
fof(f13141,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X10)) | p(s(bool,X11)) | ~sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X11,X10)) )),
  inference(superposition,[],[f608,f1840])).
fof(f13170,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X8,X9)) )),
  inference(subsumption_resolution,[],[f13140,f373])).
fof(f13140,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X8,X9)) )),
  inference(superposition,[],[f607,f1840])).
fof(f13169,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X7,X6)) )),
  inference(subsumption_resolution,[],[f13139,f373])).
fof(f13139,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X7,X6)) )),
  inference(superposition,[],[f606,f1840])).
fof(f13168,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | ~sP5(X4,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X5)) )),
  inference(subsumption_resolution,[],[f13138,f373])).
fof(f13138,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~sP5(X4,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X5)) )),
  inference(superposition,[],[f557,f1840])).
fof(f13167,plain,(
  ( ! [X3] : (p(s(bool,X3)) | sP1(X3,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(subsumption_resolution,[],[f13137,f373])).
fof(f13137,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | sP1(X3,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(superposition,[],[f524,f1840])).
fof(f13166,plain,(
  ( ! [X2] : (p(s(bool,X2)) | sP1('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X2)) )),
  inference(subsumption_resolution,[],[f13136,f373])).
fof(f13136,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X2)) )),
  inference(superposition,[],[f521,f1840])).
fof(f13126,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f373,f1840,f574])).
fof(f12732,plain,(
  sP5('F_0','F_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f641,f663,f657,f1886,f559])).
fof(f12702,plain,(
  sP5('F_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0')),
  inference(unit_resulting_resolution,[],[f641,f661,f657,f1885,f560])).
fof(f12667,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f641,f663,f655,f1884,f559])).
fof(f12645,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f373,f669,f653,f1883,f558])).
fof(f12648,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0')),
  inference(unit_resulting_resolution,[],[f373,f669,f1868,f1883,f558])).
fof(f12632,plain,(
  ~sP5('F_0','T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f1876,f555])).
fof(f12586,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1869,f555])).
fof(f12557,plain,(
  sP5('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0')),
  inference(unit_resulting_resolution,[],[f373,f669,f659,f1868,f558])).
fof(f12560,plain,(
  sP5('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0')),
  inference(unit_resulting_resolution,[],[f373,f669,f686,f1868,f558])).
fof(f12531,plain,(
  ~sP5('T_0','F_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f1861,f556])).
fof(f9632,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f1322,f556])).
fof(f9567,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.EVEN_1'(s('type.num.num',X0)),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1320,f556])).
fof(f9516,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f1286,f556])).
fof(f9455,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num',X0)),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1284,f556])).
fof(f9271,plain,(
  ( ! [X19,X20,X18] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X20))),s('type.realax.real',X19))))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X18),s('type.realax.real',X19))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X20),s('type.realax.real',X19)))))) )),
  inference(superposition,[],[f647,f542])).
fof(f9270,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real',X15))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X15))),s('type.realax.real',X17))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X14),s('type.realax.real',X16))),s('type.realax.real',X15)))))) )),
  inference(superposition,[],[f541,f542])).
fof(f9269,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real',X11))),s('type.realax.real',X13))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X12))),s('type.realax.real',X11))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real',X11))),s('type.realax.real',X13)))))) )),
  inference(superposition,[],[f541,f542])).
fof(f9268,plain,(
  ( ! [X6,X8,X7,X9] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X8))),s('type.realax.real',X7))),s('type.realax.real',X9))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X7))),s('type.realax.real',X9)))))) )),
  inference(superposition,[],[f540,f542])).
fof(f9264,plain,(
  ( ! [X4,X5,X3] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X5))),s('type.realax.real',X4))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4)))))) )),
  inference(superposition,[],[f542,f474])).
fof(f9375,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X6))),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X7))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))))) )),
  inference(forward_demodulation,[],[f9263,f475])).
fof(f9263,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X7))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))))) )),
  inference(superposition,[],[f542,f481])).
fof(f9372,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X6))),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X7)))))) )),
  inference(forward_demodulation,[],[f9227,f5024])).
fof(f9227,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))),s('type.realax.real',X8))),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X7)))))) )),
  inference(superposition,[],[f542,f481])).
fof(f9193,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))))) )),
  inference(forward_demodulation,[],[f9076,f400])).
fof(f9076,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) | s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))))) )),
  inference(superposition,[],[f498,f400])).
fof(f9190,plain,(
  ( ! [X35,X36] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X35)))))))) | s('type.num.num',X36) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X35))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X35))))))))))) )),
  inference(forward_demodulation,[],[f9189,f401])).
fof(f9189,plain,(
  ( ! [X35,X36] : (s('type.num.num',X36) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X35))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X35)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))))))) )),
  inference(forward_demodulation,[],[f9188,f400])).
fof(f9188,plain,(
  ( ! [X35,X36] : (s('type.num.num',X36) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X35))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X35)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))))))) )),
  inference(forward_demodulation,[],[f9187,f401])).
fof(f9187,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X35)))))))))) | s('type.num.num',X36) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))))))) )),
  inference(forward_demodulation,[],[f9074,f400])).
fof(f9074,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X35)))))))))) | s('type.num.num',X36) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))))))) )),
  inference(superposition,[],[f498,f401])).
fof(f9186,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X32),s('type.num.num',X33)))))))) | s('type.num.num',X34) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X32),s('type.num.num',X33))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X32),s('type.num.num',X33))))))))))) )),
  inference(forward_demodulation,[],[f9185,f485])).
fof(f9185,plain,(
  ( ! [X33,X34,X32] : (s('type.num.num',X34) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X32),s('type.num.num',X33))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X32),s('type.num.num',X33)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))))))))) )),
  inference(forward_demodulation,[],[f9184,f400])).
fof(f9184,plain,(
  ( ! [X33,X34,X32] : (s('type.num.num',X34) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X32),s('type.num.num',X33))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X32),s('type.num.num',X33)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))))))))) )),
  inference(forward_demodulation,[],[f9183,f485])).
fof(f9183,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X32),s('type.num.num',X33)))))))))) | s('type.num.num',X34) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))))))))) )),
  inference(forward_demodulation,[],[f9073,f400])).
fof(f9073,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X32),s('type.num.num',X33)))))))))) | s('type.num.num',X34) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))))))))) )),
  inference(superposition,[],[f498,f485])).
fof(f9176,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X23),s('type.num.num',X24)))))))) | s('type.num.num',X25) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X23),s('type.num.num',X24))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X23),s('type.num.num',X24))))))))))) )),
  inference(forward_demodulation,[],[f9175,f483])).
fof(f9175,plain,(
  ( ! [X24,X23,X25] : (s('type.num.num',X25) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X23),s('type.num.num',X24))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X23),s('type.num.num',X24)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))))))))) )),
  inference(forward_demodulation,[],[f9174,f400])).
fof(f9174,plain,(
  ( ! [X24,X23,X25] : (s('type.num.num',X25) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X23),s('type.num.num',X24))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X23),s('type.num.num',X24)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))))))))) )),
  inference(forward_demodulation,[],[f9173,f483])).
fof(f9173,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X23),s('type.num.num',X24)))))))))) | s('type.num.num',X25) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))))))))) )),
  inference(forward_demodulation,[],[f9069,f400])).
fof(f9069,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X23),s('type.num.num',X24)))))))))) | s('type.num.num',X25) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))))))))) )),
  inference(superposition,[],[f498,f483])).
fof(f9168,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num',X17)))))))) | s('type.num.num',X18) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num',X17))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num',X17))))))))))) )),
  inference(forward_demodulation,[],[f9167,f487])).
fof(f9167,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num',X18) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num',X17))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num',X17)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))))))))) )),
  inference(forward_demodulation,[],[f9166,f400])).
fof(f9166,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num',X18) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num',X17))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num',X17)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))))))))) )),
  inference(forward_demodulation,[],[f9165,f487])).
fof(f9165,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num',X17)))))))))) | s('type.num.num',X18) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))))))))) )),
  inference(forward_demodulation,[],[f9066,f400])).
fof(f9066,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num',X17)))))))))) | s('type.num.num',X18) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))))))))) )),
  inference(superposition,[],[f498,f487])).
fof(f9160,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))))) | s('type.num.num',X11) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))))))))) )),
  inference(forward_demodulation,[],[f9159,f400])).
fof(f9159,plain,(
  ( ! [X10,X11] : (s('type.num.num',X11) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))))))))) )),
  inference(forward_demodulation,[],[f9158,f400])).
fof(f9158,plain,(
  ( ! [X10,X11] : (s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))))))))) )),
  inference(forward_demodulation,[],[f9157,f400])).
fof(f9157,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))))))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))))))))) )),
  inference(forward_demodulation,[],[f9063,f400])).
fof(f9063,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))))))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))))))))) )),
  inference(superposition,[],[f498,f400])).
fof(f9034,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X6)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X6))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X7))))))) )),
  inference(forward_demodulation,[],[f9033,f475])).
fof(f9033,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X7)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X6))))) )),
  inference(forward_demodulation,[],[f8994,f475])).
fof(f8994,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X6)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))))) )),
  inference(superposition,[],[f543,f482])).
fof(f8992,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))))))) )),
  inference(superposition,[],[f543,f479])).
fof(f8980,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X11),s('type.realax.real',X10)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X12))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X11))))))))) )),
  inference(superposition,[],[f543,f474])).
fof(f9017,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X5))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X4)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))),s('type.realax.real',X4))))))) )),
  inference(forward_demodulation,[],[f8978,f475])).
fof(f8978,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))),s('type.realax.real',X4)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X5))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X4))))))))))) )),
  inference(superposition,[],[f543,f647])).
fof(f8965,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X11),s('type.realax.real',X10))),s('type.realax.real',X12)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X11))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X12))))))) )),
  inference(superposition,[],[f543,f474])).
fof(f8779,plain,(
  ~sP3('T_0','const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'F_0')),
  inference(unit_resulting_resolution,[],[f641,f689,f657,f1201,f560])).
fof(f8769,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',X75)))) | ~sP8('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X73),s('type.realax.real',X74))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X73),s('type.realax.real',X75)))),X76,X77) | sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X73),s('type.realax.real',X74))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X73),s('type.realax.real',X75)))),X76,X77)) )),
  inference(superposition,[],[f631,f541])).
fof(f8767,plain,(
  ( ! [X68,X66,X64,X67,X65] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X65),s('type.realax.real',X66)))) | p(s(bool,X67)) | p(s(bool,X68)) | ~sP9('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real',X65))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X64),s('type.realax.real',X66)))),X68,X67)) )),
  inference(superposition,[],[f608,f541])).
fof(f8763,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real',X47)))) | p(s(bool,X48)) | sP1(X48,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real',X46))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real',X47)))))) )),
  inference(superposition,[],[f524,f541])).
fof(f8762,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X42),s('type.realax.real',X43)))) | p(s(bool,X44)) | sP1('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X41),s('type.realax.real',X42))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X41),s('type.realax.real',X43)))),X44)) )),
  inference(superposition,[],[f521,f541])).
fof(f8761,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X38),s('type.realax.real',X39)))) | sP0('const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X37),s('type.realax.real',X38))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X37),s('type.realax.real',X39)))),X40)) )),
  inference(superposition,[],[f506,f541])).
fof(f8760,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',X35)))) | sP0(X36,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X33),s('type.realax.real',X34))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X33),s('type.realax.real',X35)))))) )),
  inference(superposition,[],[f505,f541])).
fof(f8754,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real',X17)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X15),s('type.realax.real',X17))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X15),s('type.realax.real',X16))))))) )),
  inference(superposition,[],[f491,f541])).
fof(f8753,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X14)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X12),s('type.realax.real',X14))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X12),s('type.realax.real',X13))))))) )),
  inference(superposition,[],[f492,f541])).
fof(f8752,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X9),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real',X10))),s('type.realax.real',X11)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real',X11))))) )),
  inference(superposition,[],[f544,f541])).
fof(f8751,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X6))),s('type.realax.real',X7)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s('type.realax.real',X7))))) )),
  inference(superposition,[],[f545,f541])).
fof(f8750,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))) )),
  inference(superposition,[],[f546,f541])).
fof(f8749,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X5))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5)))) )),
  inference(superposition,[],[f541,f413])).
fof(f8746,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X16))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X15))),s('type.realax.real',X17))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X16)))))))) )),
  inference(superposition,[],[f541,f484])).
fof(f8744,plain,(
  ( ! [X12,X10,X11] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X12),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X12))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X10),s('type.realax.real',X11)))))) )),
  inference(superposition,[],[f541,f475])).
fof(f8741,plain,(
  ( ! [X4,X2,X3] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X3))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2)))))) )),
  inference(superposition,[],[f541,f474])).
fof(f8770,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',X21))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X19),s('type.realax.real',X20))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X19),s('type.realax.real',X21)))))))) )),
  inference(forward_demodulation,[],[f8708,f540])).
fof(f8708,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',X21))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X19),s('type.realax.real',X20))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X19))),s('type.realax.real',X21)))))) )),
  inference(superposition,[],[f541,f540])).
fof(f8707,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X16))),s('type.realax.real',X17))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X16))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X15))),s('type.realax.real',X17)))))) )),
  inference(superposition,[],[f541,f484])).
fof(f8705,plain,(
  ( ! [X12,X10,X11] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11))),s('type.realax.real',X12))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X10),s('type.realax.real',X11))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X12)))))) )),
  inference(superposition,[],[f541,f475])).
fof(f8702,plain,(
  ( ! [X4,X2,X3] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X4))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X4)))))) )),
  inference(superposition,[],[f541,f474])).
fof(f8682,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))),'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f641,f541,f576])).
fof(f8683,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))),'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f373,f541,f574])).
fof(f8684,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)),'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) )),
  inference(unit_resulting_resolution,[],[f641,f541,f576])).
fof(f8685,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)),'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) )),
  inference(unit_resulting_resolution,[],[f373,f541,f574])).
fof(f8621,plain,(
  ( ! [X4,X2,X3] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X4)))))) )),
  inference(forward_demodulation,[],[f8620,f647])).
fof(f8620,plain,(
  ( ! [X4,X2,X3] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X4))))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real',X4)))) )),
  inference(forward_demodulation,[],[f8524,f647])).
fof(f8524,plain,(
  ( ! [X4,X2,X3] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X3))))),s('type.realax.real',X4))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X4)))))))) )),
  inference(superposition,[],[f647,f540])).
fof(f8519,plain,(
  ( ! [X6,X7,X5] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X7))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) )),
  inference(superposition,[],[f540,f474])).
fof(f8500,plain,(
  ( ! [X17,X15,X16] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X16))))),s('type.realax.real',X17))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X15))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X16))),s('type.realax.real',X17)))))) )),
  inference(superposition,[],[f540,f484])).
fof(f8495,plain,(
  ( ! [X4,X2,X3] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X4))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2))),s('type.realax.real',X4)))) )),
  inference(superposition,[],[f540,f474])).
fof(f7944,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))),s('type.realax.real',X5))))) )),
  inference(superposition,[],[f547,f482])).
fof(f7943,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real',X2))))) )),
  inference(superposition,[],[f547,f479])).
fof(f7934,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X19),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X17)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X17),s('type.realax.real',X18))),s('type.realax.real',X20)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X19),s('type.realax.real',X20))))) )),
  inference(superposition,[],[f547,f474])).
fof(f7920,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X17))),s('type.realax.real',X19)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X19),s('type.realax.real',X20)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X17),s('type.realax.real',X18))),s('type.realax.real',X20))))) )),
  inference(superposition,[],[f547,f474])).
fof(f7873,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X17)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X17),s('type.realax.real',X18)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X20),s('type.realax.real',X19))))) )),
  inference(superposition,[],[f546,f474])).
fof(f7859,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X17))),s('type.realax.real',X19)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',X19)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X20),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X17),s('type.realax.real',X18))))))) )),
  inference(superposition,[],[f546,f474])).
fof(f7804,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X17)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X17),s('type.realax.real',X18))),s('type.realax.real',X20)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real',X20))))) )),
  inference(superposition,[],[f545,f474])).
fof(f7790,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X17))),s('type.realax.real',X19)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real',X20)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X17),s('type.realax.real',X18))),s('type.realax.real',X20))))) )),
  inference(superposition,[],[f545,f474])).
fof(f7699,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X17)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X17),s('type.realax.real',X18))),s('type.realax.real',X20)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real',X20))))) )),
  inference(superposition,[],[f544,f474])).
fof(f7685,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X18),s('type.realax.real',X17))),s('type.realax.real',X19)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X19),s('type.realax.real',X20)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X17),s('type.realax.real',X18))),s('type.realax.real',X20))))) )),
  inference(superposition,[],[f544,f474])).
fof(f7539,plain,(
  ( ! [X57,X58,X56] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X58))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X56),s('type.num.num',X57))))))) | s('type.num.num',X58) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X57)))))) )),
  inference(superposition,[],[f493,f487])).
fof(f7376,plain,(
  ( ! [X52,X50,X51] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X50),s('type.num.num',X51))))))) | s('type.num.num',X52) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))))) )),
  inference(superposition,[],[f493,f485])).
fof(f7195,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X2))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f647,f484])).
fof(f7191,plain,(
  ( ! [X4,X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X5))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X4)))))) )),
  inference(superposition,[],[f484,f474])).
fof(f7263,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) )),
  inference(forward_demodulation,[],[f7262,f479])).
fof(f7262,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))))) )),
  inference(forward_demodulation,[],[f7156,f484])).
fof(f7156,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))))))) )),
  inference(unit_resulting_resolution,[],[f484,f503])).
fof(f7053,plain,(
  ( ! [X39,X38,X40] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X38),s('type.num.num',X39))))))) | s('type.num.num',X40) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X38))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))))) )),
  inference(superposition,[],[f493,f483])).
fof(f6532,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))) )),
  inference(forward_demodulation,[],[f6473,f482])).
fof(f6473,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))) )),
  inference(superposition,[],[f501,f482])).
fof(f6531,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f6472,f479])).
fof(f6472,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f501,f479])).
fof(f6463,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X9)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X10))),s('type.realax.real',X11)))) | s('type.realax.real',X11) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X10)))) )),
  inference(superposition,[],[f501,f474])).
fof(f6450,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X9))),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X10)))))) | s('type.realax.real',X11) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X10)))) )),
  inference(superposition,[],[f501,f474])).
fof(f6337,plain,(
  ( ! [X12,X13,X11] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X11),s('type.realax.real',X12))),s('type.realax.real',X13))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11))),s('type.realax.real',X12)))))) )),
  inference(superposition,[],[f482,f647])).
fof(f6336,plain,(
  ( ! [X10,X8,X9] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real',X9))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real',X9))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10)))))) )),
  inference(superposition,[],[f482,f647])).
fof(f6335,plain,(
  ( ! [X6,X7,X5] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X5),s('type.realax.real',X6))),s('type.realax.real',X7))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))),s('type.realax.real',X6))),s('type.realax.real',X7)))) )),
  inference(superposition,[],[f481,f647])).
fof(f6334,plain,(
  ( ! [X4,X2,X3] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real',X3)))))) )),
  inference(superposition,[],[f475,f647])).
fof(f6333,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))))) )),
  inference(superposition,[],[f392,f647])).
fof(f6328,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X3),s('type.realax.real',X2)))))) )),
  inference(superposition,[],[f647,f474])).
fof(f6396,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))))))))) )),
  inference(forward_demodulation,[],[f6299,f6225])).
fof(f6299,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))) )),
  inference(unit_resulting_resolution,[],[f647,f503])).
fof(f6254,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X42),s('type.realax.real',X41)))) | ~sP8('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X41))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X42)))),X43,X44) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X41))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X42)))),X43,X44)) )),
  inference(superposition,[],[f631,f482])).
fof(f6252,plain,(
  ( ! [X37,X35,X36,X34] : (~sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X34))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X35)))),X37,X36) | p(s(bool,X36)) | p(s(bool,X37)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X35),s('type.realax.real',X34))))) )),
  inference(superposition,[],[f608,f482])).
fof(f6251,plain,(
  ( ! [X30,X33,X31,X32] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X31)))),X32,X33) | ~p(s(bool,X32)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X31),s('type.realax.real',X30))))) )),
  inference(superposition,[],[f607,f482])).
fof(f6250,plain,(
  ( ! [X28,X26,X29,X27] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X27)))),X29,X28) | ~p(s(bool,X28)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real',X26))))) )),
  inference(superposition,[],[f606,f482])).
fof(f6248,plain,(
  ( ! [X21,X19,X20] : (sP1(X21,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X19))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X20))))) | p(s(bool,X21)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X20),s('type.realax.real',X19))))) )),
  inference(superposition,[],[f524,f482])).
fof(f6247,plain,(
  ( ! [X17,X18,X16] : (sP1('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X16))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X17)))),X18) | p(s(bool,X18)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X16))))) )),
  inference(superposition,[],[f521,f482])).
fof(f6246,plain,(
  ( ! [X14,X15,X13] : (sP0('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14)))),X15) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X13))))) )),
  inference(superposition,[],[f506,f482])).
fof(f6245,plain,(
  ( ! [X12,X10,X11] : (sP0(X12,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',X10))))) )),
  inference(superposition,[],[f505,f482])).
fof(f6239,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f482,f413])).
fof(f6238,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X0)))) )),
  inference(superposition,[],[f482,f392])).
fof(f6225,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f482,f392])).
fof(f6209,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f641,f482,f576])).
fof(f6210,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f373,f482,f574])).
fof(f6211,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(unit_resulting_resolution,[],[f641,f482,f576])).
fof(f6212,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(unit_resulting_resolution,[],[f373,f482,f574])).
fof(f6036,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X43),s('type.num.num',X44)))) | ~sP8('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))),X45,X46) | sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))),X45,X46)) )),
  inference(superposition,[],[f631,f480])).
fof(f6034,plain,(
  ( ! [X39,X37,X38,X36] : (~sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X37)))),X39,X38) | p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X37))))) )),
  inference(superposition,[],[f608,f480])).
fof(f6033,plain,(
  ( ! [X35,X33,X34,X32] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33)))),X34,X35) | ~p(s(bool,X34)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X33))))) )),
  inference(superposition,[],[f607,f480])).
fof(f6032,plain,(
  ( ! [X30,X28,X31,X29] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29)))),X31,X30) | ~p(s(bool,X30)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))))) )),
  inference(superposition,[],[f606,f480])).
fof(f6031,plain,(
  ( ! [X26,X24,X27,X25] : (~sP5(X26,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25)))),X27) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num',X25))))) )),
  inference(superposition,[],[f557,f480])).
fof(f6030,plain,(
  ( ! [X23,X21,X22] : (sP1(X23,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))))) | p(s(bool,X23)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num',X22))))) )),
  inference(superposition,[],[f524,f480])).
fof(f6029,plain,(
  ( ! [X19,X20,X18] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19)))),X20) | p(s(bool,X20)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',X19))))) )),
  inference(superposition,[],[f521,f480])).
fof(f6028,plain,(
  ( ! [X17,X15,X16] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16)))),X17) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X15),s('type.num.num',X16))))) )),
  inference(superposition,[],[f506,f480])).
fof(f6027,plain,(
  ( ! [X14,X12,X13] : (sP0(X14,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num',X13))))) )),
  inference(superposition,[],[f505,f480])).
fof(f6020,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) )),
  inference(superposition,[],[f480,f413])).
fof(f5982,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f641,f480,f576])).
fof(f5983,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f373,f480,f574])).
fof(f5984,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f641,f480,f576])).
fof(f5985,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f373,f480,f574])).
fof(f5946,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X41),s('type.num.num',X42)))) | ~sP8('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X41))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X42)))),X43,X44) | sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X41))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X42)))),X43,X44)) )),
  inference(superposition,[],[f631,f479])).
fof(f5944,plain,(
  ( ! [X37,X35,X36,X34] : (~sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X34))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X35)))),X37,X36) | p(s(bool,X36)) | p(s(bool,X37)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num',X35))))) )),
  inference(superposition,[],[f608,f479])).
fof(f5943,plain,(
  ( ! [X30,X33,X31,X32] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X30))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X31)))),X32,X33) | ~p(s(bool,X32)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X30),s('type.num.num',X31))))) )),
  inference(superposition,[],[f607,f479])).
fof(f5942,plain,(
  ( ! [X28,X26,X29,X27] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X26))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X27)))),X29,X28) | ~p(s(bool,X28)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f606,f479])).
fof(f5941,plain,(
  ( ! [X24,X23,X25,X22] : (~sP5(X24,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X22))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X23)))),X25) | ~p(s(bool,X24)) | p(s(bool,X25)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X22),s('type.num.num',X23))))) )),
  inference(superposition,[],[f557,f479])).
fof(f5940,plain,(
  ( ! [X21,X19,X20] : (sP1(X21,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X19))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X20))))) | p(s(bool,X21)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X20))))) )),
  inference(superposition,[],[f524,f479])).
fof(f5939,plain,(
  ( ! [X17,X18,X16] : (sP1('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X16))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X17)))),X18) | p(s(bool,X18)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X16),s('type.num.num',X17))))) )),
  inference(superposition,[],[f521,f479])).
fof(f5938,plain,(
  ( ! [X14,X15,X13] : (sP0('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X13))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X14)))),X15) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num',X14))))) )),
  inference(superposition,[],[f506,f479])).
fof(f5937,plain,(
  ( ! [X12,X10,X11] : (sP0(X12,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X10))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X11))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X11))))) )),
  inference(superposition,[],[f505,f479])).
fof(f5931,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X3))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3)))) )),
  inference(superposition,[],[f479,f413])).
fof(f5897,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f641,f479,f576])).
fof(f5898,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f373,f479,f574])).
fof(f5899,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f641,f479,f576])).
fof(f5900,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f373,f479,f574])).
fof(f5861,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X42),s('type.num.num',X41)))) | ~sP8('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42)))),X43,X44) | sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42)))),X43,X44)) )),
  inference(superposition,[],[f631,f478])).
fof(f5859,plain,(
  ( ! [X37,X35,X36,X34] : (~sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35)))),X37,X36) | p(s(bool,X36)) | p(s(bool,X37)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X35),s('type.num.num',X34))))) )),
  inference(superposition,[],[f608,f478])).
fof(f5858,plain,(
  ( ! [X30,X33,X31,X32] : (sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31)))),X32,X33) | ~p(s(bool,X32)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f607,f478])).
fof(f5857,plain,(
  ( ! [X28,X26,X29,X27] : (sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27)))),X29,X28) | ~p(s(bool,X28)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X27),s('type.num.num',X26))))) )),
  inference(superposition,[],[f606,f478])).
fof(f5856,plain,(
  ( ! [X24,X23,X25,X22] : (~sP5(X24,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))),X25) | ~p(s(bool,X24)) | p(s(bool,X25)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X23),s('type.num.num',X22))))) )),
  inference(superposition,[],[f557,f478])).
fof(f5855,plain,(
  ( ! [X21,X19,X20] : (sP1(X21,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))))) | p(s(bool,X21)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X20),s('type.num.num',X19))))) )),
  inference(superposition,[],[f524,f478])).
fof(f5854,plain,(
  ( ! [X17,X18,X16] : (sP1('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))),X18) | p(s(bool,X18)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X17),s('type.num.num',X16))))) )),
  inference(superposition,[],[f521,f478])).
fof(f5853,plain,(
  ( ! [X14,X15,X13] : (sP0('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14)))),X15) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X13))))) )),
  inference(superposition,[],[f506,f478])).
fof(f5852,plain,(
  ( ! [X12,X10,X11] : (sP0(X12,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X10))))) )),
  inference(superposition,[],[f505,f478])).
fof(f5846,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X2)))) )),
  inference(superposition,[],[f478,f413])).
fof(f5808,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f641,f478,f576])).
fof(f5809,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f373,f478,f574])).
fof(f5810,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f641,f478,f576])).
fof(f5811,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f373,f478,f574])).
fof(f5766,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X42),s('type.num.num',X41)))) | ~sP8('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42)))),X43,X44) | sP9('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42)))),X43,X44)) )),
  inference(superposition,[],[f631,f477])).
fof(f5764,plain,(
  ( ! [X37,X35,X36,X34] : (~sP9('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35)))),X37,X36) | p(s(bool,X36)) | p(s(bool,X37)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num',X34))))) )),
  inference(superposition,[],[f608,f477])).
fof(f5763,plain,(
  ( ! [X30,X33,X31,X32] : (sP9('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31)))),X32,X33) | ~p(s(bool,X32)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X31),s('type.num.num',X30))))) )),
  inference(superposition,[],[f607,f477])).
fof(f5762,plain,(
  ( ! [X28,X26,X29,X27] : (sP9('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27)))),X29,X28) | ~p(s(bool,X28)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num',X26))))) )),
  inference(superposition,[],[f606,f477])).
fof(f5761,plain,(
  ( ! [X24,X23,X25,X22] : (~sP5(X24,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))),X25) | ~p(s(bool,X24)) | p(s(bool,X25)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num',X22))))) )),
  inference(superposition,[],[f557,f477])).
fof(f5760,plain,(
  ( ! [X21,X19,X20] : (sP1(X21,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))))) | p(s(bool,X21)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num',X19))))) )),
  inference(superposition,[],[f524,f477])).
fof(f5759,plain,(
  ( ! [X17,X18,X16] : (sP1('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))),X18) | p(s(bool,X18)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X17),s('type.num.num',X16))))) )),
  inference(superposition,[],[f521,f477])).
fof(f5758,plain,(
  ( ! [X14,X15,X13] : (sP0('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14)))),X15) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num',X13))))) )),
  inference(superposition,[],[f506,f477])).
fof(f5757,plain,(
  ( ! [X12,X10,X11] : (sP0(X12,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X10))))) )),
  inference(superposition,[],[f505,f477])).
fof(f5751,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) )),
  inference(superposition,[],[f477,f413])).
fof(f5713,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f641,f477,f576])).
fof(f5714,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f373,f477,f574])).
fof(f5715,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f641,f477,f576])).
fof(f5716,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f373,f477,f574])).
fof(f5670,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X45),s('type.num.num',X46)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46)))),X47,X48) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46)))),X47,X48)) )),
  inference(superposition,[],[f631,f476])).
fof(f5668,plain,(
  ( ! [X39,X41,X38,X40] : (~sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X38))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))),X41,X40) | p(s(bool,X40)) | p(s(bool,X41)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X38),s('type.num.num',X39))))) )),
  inference(superposition,[],[f608,f476])).
fof(f5667,plain,(
  ( ! [X37,X35,X36,X34] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35)))),X36,X37) | ~p(s(bool,X36)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num',X35))))) )),
  inference(superposition,[],[f607,f476])).
fof(f5666,plain,(
  ( ! [X30,X33,X31,X32] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31)))),X33,X32) | ~p(s(bool,X32)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X30),s('type.num.num',X31))))) )),
  inference(superposition,[],[f606,f476])).
fof(f5665,plain,(
  ( ! [X28,X26,X29,X27] : (~sP5(X28,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27)))),X29) | ~p(s(bool,X28)) | p(s(bool,X29)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X26),s('type.num.num',X27))))) )),
  inference(superposition,[],[f557,f476])).
fof(f5664,plain,(
  ( ! [X24,X23,X25] : (sP1(X25,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))))) | p(s(bool,X25)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X23),s('type.num.num',X24))))) )),
  inference(superposition,[],[f524,f476])).
fof(f5663,plain,(
  ( ! [X21,X22,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))),X22) | p(s(bool,X22)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X20),s('type.num.num',X21))))) )),
  inference(superposition,[],[f521,f476])).
fof(f5662,plain,(
  ( ! [X19,X17,X18] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18)))),X19) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X17),s('type.num.num',X18))))) )),
  inference(superposition,[],[f506,f476])).
fof(f5661,plain,(
  ( ! [X14,X15,X16] : (sP0(X16,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X15))))) )),
  inference(superposition,[],[f505,f476])).
fof(f5653,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3)))) )),
  inference(superposition,[],[f476,f413])).
fof(f5615,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f641,f476,f576])).
fof(f5616,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f373,f476,f574])).
fof(f5617,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f641,f476,f576])).
fof(f5618,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f373,f476,f574])).
fof(f5421,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))) )),
  inference(forward_demodulation,[],[f5378,f400])).
fof(f5378,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))))))) )),
  inference(superposition,[],[f499,f400])).
fof(f5419,plain,(
  ( ! [X24,X23] : (s('type.num.num',X24) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X23))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X23))))))))))) )),
  inference(forward_demodulation,[],[f5418,f400])).
fof(f5418,plain,(
  ( ! [X24,X23] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X23)))))))))) | s('type.num.num',X24) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X23)))))))) )),
  inference(forward_demodulation,[],[f5417,f401])).
fof(f5417,plain,(
  ( ! [X24,X23] : (s('type.num.num',X24) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X23))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))))))))))) )),
  inference(forward_demodulation,[],[f5376,f400])).
fof(f5376,plain,(
  ( ! [X24,X23] : (s('type.num.num',X24) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X23))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))))))))))) )),
  inference(superposition,[],[f499,f401])).
fof(f5409,plain,(
  ( ! [X8,X7] : (s('type.num.num',X8) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))))))) )),
  inference(forward_demodulation,[],[f5408,f400])).
fof(f5408,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))))) | s('type.num.num',X8) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))) )),
  inference(forward_demodulation,[],[f5407,f400])).
fof(f5407,plain,(
  ( ! [X8,X7] : (s('type.num.num',X8) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))))))))))) )),
  inference(forward_demodulation,[],[f5368,f400])).
fof(f5368,plain,(
  ( ! [X8,X7] : (s('type.num.num',X8) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))))))))))) )),
  inference(superposition,[],[f499,f400])).
fof(f4868,plain,(
  sP5('F_0','const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'F_0')),
  inference(unit_resulting_resolution,[],[f688,f685,f661,f641,f560])).
fof(f4785,plain,(
  sP5('F_0','F_0','const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),
  inference(unit_resulting_resolution,[],[f689,f657,f663,f641,f559])).
fof(f4803,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f665,f655,f661,f641,f559])).
fof(f4804,plain,(
  sP5('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f687,f655,f661,f641,f559])).
fof(f4718,plain,(
  ~sP3('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f669,f1194,f657,f373,f558])).
fof(f4721,plain,(
  ~sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f669,f2168,f657,f373,f558])).
fof(f4728,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f669,f653,f659,f373,f558])).
fof(f4733,plain,(
  sP5('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f669,f653,f686,f373,f558])).
fof(f4734,plain,(
  sP5('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0')),
  inference(unit_resulting_resolution,[],[f669,f685,f686,f373,f558])).
fof(f4559,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(subsumption_resolution,[],[f4502,f641])).
fof(f4502,plain,(
  ( ! [X2,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(superposition,[],[f529,f413])).
fof(f4205,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(subsumption_resolution,[],[f4186,f641])).
fof(f4186,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f492,f413])).
fof(f4179,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X6)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real',X8))))) )),
  inference(superposition,[],[f492,f474])).
fof(f4169,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s('type.realax.real',X8)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X7))))))) )),
  inference(superposition,[],[f492,f474])).
fof(f4116,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X6)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real',X8))))) )),
  inference(superposition,[],[f491,f474])).
fof(f4106,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s('type.realax.real',X8)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X7))))))) )),
  inference(superposition,[],[f491,f474])).
fof(f3312,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s('type.realax.real',X9)))) | s('type.realax.real',X9) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real',X7)))) )),
  inference(superposition,[],[f503,f474])).
fof(f3311,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X5)))))) | s('type.realax.real',X6) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X5),s('type.realax.real',X4)))) )),
  inference(superposition,[],[f503,f474])).
fof(f3293,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0))))))) )),
  inference(unit_resulting_resolution,[],[f474,f503])).
fof(f2204,plain,(
  ( ! [X14,X13] : (~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X13,X14) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X13,X14)) )),
  inference(subsumption_resolution,[],[f2187,f373])).
fof(f2187,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X13,X14) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X13,X14)) )),
  inference(superposition,[],[f631,f2146])).
fof(f2201,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | ~sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X11,X10)) )),
  inference(subsumption_resolution,[],[f2185,f373])).
fof(f2185,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X10)) | p(s(bool,X11)) | ~sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X11,X10)) )),
  inference(superposition,[],[f608,f2146])).
fof(f2200,plain,(
  ( ! [X8,X9] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X8,X9) | ~p(s(bool,X8))) )),
  inference(subsumption_resolution,[],[f2184,f373])).
fof(f2184,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X8,X9)) )),
  inference(superposition,[],[f607,f2146])).
fof(f2199,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X7,X6) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f2183,f373])).
fof(f2183,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X7,X6)) )),
  inference(superposition,[],[f606,f2146])).
fof(f2198,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | ~sP5(X4,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X5)) )),
  inference(subsumption_resolution,[],[f2182,f373])).
fof(f2182,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~sP5(X4,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X5)) )),
  inference(superposition,[],[f557,f2146])).
fof(f2057,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f427,f395])).
fof(f2053,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f427,f394])).
fof(f2052,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f427,f393])).
fof(f1842,plain,(
  ( ! [X0] : (sP0(X0,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(unit_resulting_resolution,[],[f371,f505])).
fof(f1843,plain,(
  ( ! [X0] : (sP0('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0)) )),
  inference(unit_resulting_resolution,[],[f371,f506])).
fof(f1844,plain,(
  ~sP1('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0')),
  inference(unit_resulting_resolution,[],[f373,f371,f520])).
fof(f1846,plain,(
  ~sP1('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f373,f371,f520])).
fof(f1847,plain,(
  ~sP1('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f371,f371,f520])).
fof(f1848,plain,(
  sP1('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0')),
  inference(unit_resulting_resolution,[],[f641,f371,f521])).
fof(f1849,plain,(
  sP1('F_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f641,f371,f524])).
fof(f1852,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0')),
  inference(unit_resulting_resolution,[],[f373,f641,f371,f557])).
fof(f1853,plain,(
  ~sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0')),
  inference(unit_resulting_resolution,[],[f371,f641,f371,f557])).
fof(f1854,plain,(
  ~sP4('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f373,f641,f371,f561])).
fof(f1861,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0')),
  inference(unit_resulting_resolution,[],[f373,f641,f371,f561])).
fof(f1862,plain,(
  ~sP4('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0')),
  inference(unit_resulting_resolution,[],[f371,f641,f371,f561])).
fof(f1868,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(unit_resulting_resolution,[],[f371,f562])).
fof(f1869,plain,(
  ~sP3('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f373,f641,f371,f565])).
fof(f1876,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0')),
  inference(unit_resulting_resolution,[],[f373,f641,f371,f565])).
fof(f1877,plain,(
  ~sP3('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0')),
  inference(unit_resulting_resolution,[],[f371,f641,f371,f565])).
fof(f1883,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(unit_resulting_resolution,[],[f371,f566])).
fof(f1884,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(unit_resulting_resolution,[],[f371,f570])).
fof(f1885,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X1)) )),
  inference(unit_resulting_resolution,[],[f371,f571])).
fof(f1886,plain,(
  ( ! [X0,X1] : (sP2('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f371,f572])).
fof(f1887,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0')),
  inference(unit_resulting_resolution,[],[f698,f371,f574])).
fof(f1888,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f387,f371,f574])).
fof(f1889,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f388,f371,f574])).
fof(f1890,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f697,f371,f574])).
fof(f1891,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f716,f371,f574])).
fof(f1892,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f385,f371,f574])).
fof(f1893,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f386,f371,f574])).
fof(f1894,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f396,f371,f574])).
fof(f1895,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f397,f371,f574])).
fof(f1896,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),
  inference(unit_resulting_resolution,[],[f698,f371,f574])).
fof(f1897,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f387,f371,f574])).
fof(f1898,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f388,f371,f574])).
fof(f1899,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f697,f371,f574])).
fof(f1900,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f716,f371,f574])).
fof(f1901,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f385,f371,f574])).
fof(f1902,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f386,f371,f574])).
fof(f1903,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.ODD_1'(s('type.num.num',X0)),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f396,f371,f574])).
fof(f1904,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.EVEN_1'(s('type.num.num',X0)),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f397,f371,f574])).
fof(f1905,plain,(
  ~sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f693,f371,f575])).
fof(f1906,plain,(
  ~sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f694,f371,f575])).
fof(f1907,plain,(
  ~sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f693,f371,f575])).
fof(f1908,plain,(
  ~sP5('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f694,f371,f575])).
fof(f1909,plain,(
  ( ! [X0] : (~sP7('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f641,f371,f578])).
fof(f1910,plain,(
  ( ! [X0] : (~sP7('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f641,f371,f579])).
fof(f1913,plain,(
  sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0')),
  inference(unit_resulting_resolution,[],[f373,f672,f371,f580])).
fof(f1915,plain,(
  sP7('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0')),
  inference(unit_resulting_resolution,[],[f373,f690,f371,f580])).
fof(f1920,plain,(
  sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f373,f672,f371,f580])).
fof(f1921,plain,(
  sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f371,f672,f371,f580])).
fof(f1922,plain,(
  sP7('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f373,f690,f371,f580])).
fof(f1923,plain,(
  sP7('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f371,f690,f371,f580])).
fof(f1925,plain,(
  ~sP6('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f373,f641,f371,f584])).
fof(f1932,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0')),
  inference(unit_resulting_resolution,[],[f373,f641,f371,f584])).
fof(f1933,plain,(
  ~sP6('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0')),
  inference(unit_resulting_resolution,[],[f371,f641,f371,f584])).
fof(f1939,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(unit_resulting_resolution,[],[f371,f585])).
fof(f1940,plain,(
  sP7('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0')),
  inference(unit_resulting_resolution,[],[f373,f373,f371,f590])).
fof(f1945,plain,(
  sP7('T_0','T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f373,f373,f371,f590])).
fof(f1947,plain,(
  sP7('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f373,f371,f371,f590])).
fof(f1950,plain,(
  sP7('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f373,f373,f371,f590])).
fof(f1951,plain,(
  sP7('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0')),
  inference(unit_resulting_resolution,[],[f371,f373,f371,f590])).
fof(f1952,plain,(
  sP7('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f373,f371,f371,f590])).
fof(f1953,plain,(
  sP7('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f371,f371,f371,f590])).
fof(f1954,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0')),
  inference(unit_resulting_resolution,[],[f641,f373,f371,f593])).
fof(f1958,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f641,f373,f371,f593])).
fof(f1959,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f641,f371,f371,f593])).
fof(f1962,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(unit_resulting_resolution,[],[f641,f371,f594])).
fof(f1968,plain,(
  ( ! [X0] : (~sP9('F_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0)) )),
  inference(unit_resulting_resolution,[],[f641,f371,f595])).
fof(f1980,plain,(
  ~sP8('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f641,f641,f371,f601])).
fof(f1987,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(unit_resulting_resolution,[],[f371,f602])).
fof(f1988,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X1)) )),
  inference(unit_resulting_resolution,[],[f371,f603])).
fof(f1989,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(unit_resulting_resolution,[],[f373,f371,f606])).
fof(f1993,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(unit_resulting_resolution,[],[f371,f371,f606])).
fof(f1994,plain,(
  ( ! [X0] : (sP9('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0)) )),
  inference(unit_resulting_resolution,[],[f373,f371,f607])).
fof(f1998,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0)) )),
  inference(unit_resulting_resolution,[],[f371,f371,f607])).
fof(f1999,plain,(
  ~sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f641,f641,f371,f608])).
fof(f2000,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))),
  inference(unit_resulting_resolution,[],[f641,f371,f630])).
fof(f2001,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f677,f371,f631])).
fof(f2002,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f678,f371,f631])).
fof(f2003,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) )),
  inference(unit_resulting_resolution,[],[f691,f371,f631])).
fof(f2004,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f679,f371,f631])).
fof(f2005,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X0)) )),
  inference(unit_resulting_resolution,[],[f692,f371,f631])).
fof(f2006,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f680,f371,f631])).
fof(f2007,plain,(
  ( ! [X0] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(unit_resulting_resolution,[],[f371,f787])).
fof(f2008,plain,(
  ( ! [X0] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0)) )),
  inference(unit_resulting_resolution,[],[f371,f797])).
fof(f2009,plain,(
  ( ! [X0] : (sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X0,'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))))) )),
  inference(unit_resulting_resolution,[],[f371,f967])).
fof(f2010,plain,(
  ( ! [X0] : (sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'const.bool.IN_2'(s(fun(sK10,bool),sK12),s(fun(fun(sK10,bool),bool),'const.probability.events_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),X0)) )),
  inference(unit_resulting_resolution,[],[f371,f968])).
fof(f1746,plain,(
  ( ! [X28,X29,X27] : (~sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27)))),X29,X28) | p(s(bool,X28)) | p(s(bool,X29)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X27))))) )),
  inference(superposition,[],[f608,f397])).
fof(f1755,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | ~sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X26,X25)) )),
  inference(subsumption_resolution,[],[f1745,f373])).
fof(f1745,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | p(s(bool,X25)) | p(s(bool,X26)) | ~sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X26,X25)) )),
  inference(superposition,[],[f608,f697])).
fof(f1744,plain,(
  ( ! [X24,X23,X22] : (~sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22)))),X24,X23) | p(s(bool,X23)) | p(s(bool,X24)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X22))))) )),
  inference(superposition,[],[f608,f396])).
fof(f1754,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | ~sP9('const.arithmetic.>=_2'(s('type.num.num',X17),s('type.num.num','const.num.|48|_0')),X19,X18)) )),
  inference(subsumption_resolution,[],[f1742,f373])).
fof(f1742,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | p(s(bool,X19)) | ~sP9('const.arithmetic.>=_2'(s('type.num.num',X17),s('type.num.num','const.num.|48|_0')),X19,X18)) )),
  inference(superposition,[],[f608,f388])).
fof(f1752,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | ~sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X7,X6)) )),
  inference(subsumption_resolution,[],[f1738,f373])).
fof(f1738,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X7,X6)) )),
  inference(superposition,[],[f608,f698])).
fof(f1696,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1685,f596])).
fof(f1685,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f373,f641,f641,f601])).
fof(f1673,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f681,f641,f641,f600])).
fof(f1674,plain,(
  sP9('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f682,f641,f641,f600])).
fof(f1572,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1171,f556])).
fof(f1561,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1551,f577])).
fof(f1551,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f373,f641,f373,f584])).
fof(f1539,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f671,f373,f373,f580])).
fof(f1541,plain,(
  sP7('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f690,f373,f373,f580])).
fof(f1535,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1522,f554])).
fof(f1522,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f641,f641,f641,f569])).
fof(f1514,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1504,f555])).
fof(f1504,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f373,f641,f373,f565])).
fof(f1488,plain,(
  ( ! [X0] : (sP4('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1170,f556])).
fof(f1475,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f373,f641,f373,f561])).
fof(f1431,plain,(
  ( ! [X28,X29,X27] : (~sP5(X28,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27)))),X29) | ~p(s(bool,X28)) | p(s(bool,X29)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X27))))) )),
  inference(superposition,[],[f557,f397])).
fof(f1439,plain,(
  ( ! [X26,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | ~sP5(X25,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X26)) )),
  inference(subsumption_resolution,[],[f1430,f373])).
fof(f1430,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP5(X25,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X26)) )),
  inference(superposition,[],[f557,f697])).
fof(f1429,plain,(
  ( ! [X24,X23,X22] : (~sP5(X23,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22)))),X24) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X22))))) )),
  inference(superposition,[],[f557,f396])).
fof(f1436,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | ~sP5(X6,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X7)) )),
  inference(subsumption_resolution,[],[f1423,f373])).
fof(f1423,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP5(X6,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X7)) )),
  inference(superposition,[],[f557,f698])).
fof(f1345,plain,(
  ( ! [X23,X21,X22] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))),X22,X23) | ~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))),X22,X23) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X21))))) )),
  inference(superposition,[],[f631,f397])).
fof(f1320,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f373,f397,f574])).
fof(f1322,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num',X0)),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f373,f397,f574])).
fof(f1309,plain,(
  ( ! [X23,X21,X22] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))),X22,X23) | ~sP8('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))),X22,X23) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X21))))) )),
  inference(superposition,[],[f631,f396])).
fof(f1296,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f396,f413])).
fof(f1283,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f641,f396,f576])).
fof(f1284,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f373,f396,f574])).
fof(f1285,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num',X0)),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f641,f396,f576])).
fof(f1286,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num',X0)),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f373,f396,f574])).
fof(f1194,plain,(
  ~sP5('F_0','const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0')),
  inference(unit_resulting_resolution,[],[f641,f698,f576])).
fof(f1195,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f641,f387,f576])).
fof(f1196,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f641,f388,f576])).
fof(f1197,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f641,f697,f576])).
fof(f1198,plain,(
  ~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f641,f716,f576])).
fof(f1199,plain,(
  ( ! [X0] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f641,f385,f576])).
fof(f1200,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f641,f386,f576])).
fof(f1201,plain,(
  ~sP5('F_0','T_0','const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),
  inference(unit_resulting_resolution,[],[f641,f698,f576])).
fof(f1202,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f641,f387,f576])).
fof(f1203,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f641,f388,f576])).
fof(f1204,plain,(
  ~sP5('F_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f641,f697,f576])).
fof(f1205,plain,(
  ~sP5('F_0','F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f641,f716,f576])).
fof(f1206,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f641,f385,f576])).
fof(f1207,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f641,f386,f576])).
fof(f1179,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f373,f693,f575])).
fof(f1180,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f373,f694,f575])).
fof(f1183,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f373,f693,f575])).
fof(f1184,plain,(
  ~sP5('T_0','T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f373,f694,f575])).
fof(f1158,plain,(
  sP5('T_0','const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0')),
  inference(unit_resulting_resolution,[],[f373,f698,f574])).
fof(f1161,plain,(
  sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f373,f697,f574])).
fof(f1162,plain,(
  sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f373,f716,f574])).
fof(f1163,plain,(
  ( ! [X0] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f373,f385,f574])).
fof(f1164,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f373,f386,f574])).
fof(f1165,plain,(
  sP5('T_0','T_0','const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),
  inference(unit_resulting_resolution,[],[f373,f698,f574])).
fof(f1168,plain,(
  sP5('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f373,f697,f574])).
fof(f1169,plain,(
  sP5('T_0','F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f373,f716,f574])).
fof(f1170,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f373,f385,f574])).
fof(f1171,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f373,f386,f574])).
fof(f1148,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f641,f693,f573])).
fof(f1149,plain,(
  sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f641,f694,f573])).
fof(f1150,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f641,f693,f573])).
fof(f1151,plain,(
  sP5('F_0','T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f641,f694,f573])).
fof(f1009,plain,(
  ( ! [X0] : (sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f373,f968])).
fof(f1002,plain,(
  ( ! [X0] : (sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f967])).
fof(f985,plain,(
  sP1('F_0','const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))),
  inference(unit_resulting_resolution,[],[f641,f966])).
fof(f970,plain,(
  sP1('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),'F_0')),
  inference(unit_resulting_resolution,[],[f641,f965])).
fof(f969,plain,(
  ( ! [X8,X9] : (~sP8('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X8,X9) | sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X8,X9)) )),
  inference(subsumption_resolution,[],[f959,f373])).
fof(f959,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~sP8('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X8,X9) | sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X8,X9)) )),
  inference(superposition,[],[f631,f698])).
fof(f968,plain,(
  ( ! [X6,X7] : (sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X6,X7) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f958,f373])).
fof(f958,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X6,X7)) )),
  inference(superposition,[],[f607,f698])).
fof(f967,plain,(
  ( ! [X4,X5] : (sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X5,X4) | ~p(s(bool,X4))) )),
  inference(subsumption_resolution,[],[f957,f373])).
fof(f957,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X4)) | sP9('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X5,X4)) )),
  inference(superposition,[],[f606,f698])).
fof(f966,plain,(
  ( ! [X3] : (sP1(X3,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f956,f373])).
fof(f956,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | sP1(X3,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) )),
  inference(superposition,[],[f524,f698])).
fof(f965,plain,(
  ( ! [X2] : (sP1('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f955,f373])).
fof(f955,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X2)) )),
  inference(superposition,[],[f521,f698])).
fof(f931,plain,(
  ( ! [X19,X20,X18] : (~sP8('const.arithmetic.>=_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0')),X19,X20) | sP9('const.arithmetic.>=_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0')),X19,X20)) )),
  inference(subsumption_resolution,[],[f919,f373])).
fof(f919,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'T_0')) | ~sP8('const.arithmetic.>=_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0')),X19,X20) | sP9('const.arithmetic.>=_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0')),X19,X20)) )),
  inference(superposition,[],[f631,f388])).
fof(f900,plain,(
  ( ! [X19,X20,X18] : (~sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)),X19,X20) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)),X19,X20)) )),
  inference(subsumption_resolution,[],[f888,f373])).
fof(f888,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'T_0')) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)),X19,X20) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)),X19,X20)) )),
  inference(superposition,[],[f631,f387])).
fof(f831,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f674,f641,f632])).
fof(f833,plain,(
  ( ! [X0] : (sP7('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f676,f641,f632])).
fof(f823,plain,(
  ( ! [X0,X1] : (~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(subsumption_resolution,[],[f821,f373])).
fof(f821,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(superposition,[],[f631,f697])).
fof(f817,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) )),
  inference(unit_resulting_resolution,[],[f691,f373,f631])).
fof(f819,plain,(
  ( ! [X0] : (sP9('T_0','const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X0)) )),
  inference(unit_resulting_resolution,[],[f692,f373,f631])).
fof(f794,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f373,f373,f607])).
fof(f784,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f373,f606])).
fof(f781,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f373,f641,f595])).
fof(f780,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f641,f594])).
fof(f776,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f641,f641,f589])).
fof(f770,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f641,f641,f588])).
fof(f767,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f373,f641,f579])).
fof(f764,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f373,f641,f578])).
fof(f756,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f641,f641,f648])).
fof(f742,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f641,f373,f524])).
fof(f733,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f641,f373,f521])).
fof(f730,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f373,f373,f520])).
fof(f729,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f641,f641,f519])).
fof(f694,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f383,f419])).
fof(f683,plain,(
  ( ! [X0] : (sP0(X0,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) )),
  inference(unit_resulting_resolution,[],[f370,f505])).
fof(f684,plain,(
  ( ! [X0] : (sP0('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X0)) )),
  inference(unit_resulting_resolution,[],[f370,f506])).
fof(f686,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) )),
  inference(unit_resulting_resolution,[],[f370,f566])).
fof(f687,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) )),
  inference(unit_resulting_resolution,[],[f370,f570])).
fof(f688,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X1)) )),
  inference(unit_resulting_resolution,[],[f370,f571])).
fof(f689,plain,(
  ( ! [X0,X1] : (sP2('const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f370,f572])).
fof(f690,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) )),
  inference(unit_resulting_resolution,[],[f370,f585])).
fof(f691,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)))) )),
  inference(unit_resulting_resolution,[],[f370,f602])).
fof(f692,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.probability.prob_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11)),X1)) )),
  inference(unit_resulting_resolution,[],[f370,f603])).
fof(f681,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f641,f604])).
fof(f682,plain,(
  ( ! [X0,X1] : (sP8('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f383,f604])).
fof(f679,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f373,f603])).
fof(f677,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f602])).
fof(f675,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f641,f587])).
fof(f676,plain,(
  ( ! [X0,X1] : (sP6('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f383,f587])).
fof(f673,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f641,f586])).
fof(f674,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f383,f586])).
fof(f671,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f585])).
fof(f664,plain,(
  ( ! [X0,X1] : (sP3('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f383,f568])).
fof(f659,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f566])).
fof(f651,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f373,f506])).
fof(f652,plain,(
  ( ! [X0] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f375,f506])).
fof(f649,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f373,f505])).
fof(f650,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f375,f505])).
fof(f604,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f365])).
fof(f596,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f362])).
fof(f598,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f362])).
fof(f599,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f362])).
fof(f600,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f362])).
fof(f588,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f359])).
fof(f589,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f359])).
fof(f590,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f359])).
fof(f586,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f357])).
fof(f587,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f357])).
fof(f577,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f355])).
fof(f578,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f355])).
fof(f579,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f355])).
fof(f632,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f583])).
fof(f583,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f355])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f352])).
fof(f547,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(ennf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X13,X14,X21] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X14)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X21))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X21)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_LE_TRANS')).
fof(f546,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(ennf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X13,X14,X21] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X14)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',X21))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X21)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_LET_TRANS')).
fof(f545,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2))))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(flattening,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(ennf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X13,X14,X21] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X14)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',X21))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X21)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.realax.REAL_LT_TRANS')).
fof(f544,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(ennf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X13,X14,X21] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X14)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X21))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X21)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_LTE_TRANS')).
fof(f543,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))))))) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))))),
  inference(ennf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X13,X14,X21] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X21)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X14))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X21)))))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_LE_LADD_IMP')).
fof(f542,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X13,X14,X21] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X14))),s('type.realax.real',X21))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X21))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X21)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_RDISTRIB')).
fof(f529,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11,X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X10)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X11)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.arithmetic.NOT_LESS_EQUAL')).
fof(f530,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f333])).
fof(f520,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f331])).
fof(f501,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) )),
  inference(cnf_transformation,[],[f322])).
fof(f498,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) & ((s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))) & ((s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) <=> (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11,X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))) <=> (s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X10))))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.arithmetic.arithmetic_2|49|7_prove_ltb')).
fof(f499,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f320])).
fof(f500,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f320])).
fof(f497,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f318])).
fof(f494,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f316])).
fof(f491,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) & (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))))),
  inference(nnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X14,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X14)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X13)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.real_lt')).
fof(f492,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f315])).
fof(f488,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X10,X11] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X11))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|2')).
fof(f487,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X10,X11] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num',X11))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|5')).
fof(f484,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X11,X10] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X10))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X11))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X10)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_ADD')).
fof(f483,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X10,X11] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X10),s('type.num.num',X11))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|8')).
fof(f481,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X13,X14] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real',X14)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_MUL_LNEG')).
fof(f480,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X10,X11] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X11))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|22')).
fof(f479,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X11,X10] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X10))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X11))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X10)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_LE')).
fof(f477,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X10,X11] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X10))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|25')).
fof(f476,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X10,X11] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X11))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|28')).
fof(f423,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f285])).
fof(f645,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X0)))) )),
  inference(forward_demodulation,[],[f409,f377])).
fof(f409,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X13] : s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X13)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_MUL_LID')).
fof(f644,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f408,f377])).
fof(f408,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X13] : s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_MUL_RID')).
fof(f406,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X10] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|10')).
fof(f405,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X13] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))),s('type.realax.real',X13)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_ADD_LINV')).
fof(f643,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f404,f377])).
fof(f404,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X10] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|11')).
fof(f403,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X13] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_MUL_LZERO')).
fof(f402,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X13] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13)))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_ADD_RINV')).
fof(f398,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X13] : s('type.realax.real',X13) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.real.REAL_ADD_LID')).
fof(f394,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X10] : s('type.num.num',X10) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|1')).
fof(f393,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X10] : s('type.num.num',X10) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.arithmetic.ADD|split|0')).
fof(f391,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X10] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.arithmetic.MULT|split|0')).
fof(f390,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X10] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|4')).
fof(f389,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X10] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|6')).
fof(f386,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X10] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|23')).
fof(f385,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X10] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.numeral.numeral_distrib|split|20')).
fof(f384,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10))))),
  file('thm.probability.PROB_COMPL_LE|49|','thm.arithmetic.ZERO_LESS_EQ')).
fof(f16400,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl36),
  inference(subsumption_resolution,[],[f16132,f373])).
fof(f16132,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl36),
  inference(backward_demodulation,[],[f15927,f530])).
fof(f16395,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f16394,f16382])).
fof(f16394,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real',X17))))) ) | $spl36),
  inference(subsumption_resolution,[],[f16111,f373])).
fof(f16111,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real',X17))))) ) | $spl36),
  inference(backward_demodulation,[],[f15927,f8754])).
fof(f16393,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f16392,f16382])).
fof(f16392,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X6))))))) ) | $spl36),
  inference(subsumption_resolution,[],[f16105,f373])).
fof(f16105,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X7),s('type.realax.real',X6))))))) ) | $spl36),
  inference(backward_demodulation,[],[f15927,f4116])).
fof(f16391,plain,(
  $false | ($spl12 | $spl36)),
  inference(subsumption_resolution,[],[f16101,f15927])).
fof(f16101,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) | ($spl12 | $spl36)),
  inference(backward_demodulation,[],[f15927,f13150])).
fof(f16386,plain,(
  $false | ($spl12 | $spl36)),
  inference(subsumption_resolution,[],[f16385,f16382])).
fof(f16385,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X7))))))) ) | ($spl12 | $spl36)),
  inference(subsumption_resolution,[],[f16068,f373])).
fof(f16068,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X7))))))) ) | ($spl12 | $spl36)),
  inference(backward_demodulation,[],[f15927,f14193])).
fof(f14193,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X7)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X7))))))) ) | $spl12),
  inference(backward_demodulation,[],[f14076,f4106])).
fof(f14076,plain,(
  ( ! [X12,X10,X11] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X12),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X11),s('type.realax.real',X10))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X12))),s('type.realax.real',X11)))) ) | $spl12),
  inference(superposition,[],[f541,f8612])).
fof(f8612,plain,(
  ( ! [X8,X9] : (s('type.realax.real',X9) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X9),s('type.realax.real',X8)))))) ) | $spl12),
  inference(forward_demodulation,[],[f8611,f2304])).
fof(f2304,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real',X0)))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f398])).
fof(f8611,plain,(
  ( ! [X8,X9] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real',X9))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X9),s('type.realax.real',X8)))))) ) | $spl12),
  inference(forward_demodulation,[],[f8497,f5024])).
fof(f8497,plain,(
  ( ! [X8,X9] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.realax.real',X9))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real',X9)))))) ) | $spl12),
  inference(superposition,[],[f540,f2306])).
fof(f2306,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) ) | $spl12),
  inference(backward_demodulation,[],[f2292,f402])).
fof(f16361,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f16360,f373])).
fof(f16360,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(forward_demodulation,[],[f16359,f15927])).
fof(f16359,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) ) | $spl36),
  inference(subsumption_resolution,[],[f16022,f373])).
fof(f16022,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) ) | $spl36),
  inference(backward_demodulation,[],[f15927,f491])).
fof(f16355,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f15989,f373])).
fof(f15989,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(backward_demodulation,[],[f15927,f641])).
fof(f16354,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f16353,f15927])).
fof(f16353,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl36),
  inference(subsumption_resolution,[],[f15987,f419])).
fof(f15987,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl36),
  inference(backward_demodulation,[],[f15927,f423])).
fof(f15960,plain,(
  ( ! [X74] : (s(bool,X74) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) | s(bool,'T_0') = s(bool,X74)) ) | $spl50),
  inference(cnf_transformation,[],[f15960_D])).
fof(f15960_D,plain,(
  ( ! [X74] : (s(bool,X74) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) | s(bool,'T_0') = s(bool,X74)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f15962,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | $spl53),
  inference(cnf_transformation,[],[f15962_D])).
fof(f15962_D,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f15737,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | s(bool,X74) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) | s(bool,'T_0') = s(bool,X74)) ) | $spl12),
  inference(superposition,[],[f13150,f1089])).
fof(f15945,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl48),
  inference(cnf_transformation,[],[f15945_D])).
fof(f15945_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f15986,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f15985])).
fof(f15985,plain,(
  ( ! [X47,X48,X46,X49] : (~sP8('T_0',X47,X48) | sP9('T_0',X47,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f15710,f15620])).
fof(f15710,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~sP8('T_0',X47,X48) | sP9('T_0',X47,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f631,f1089])).
fof(f15942,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl46),
  inference(cnf_transformation,[],[f15942_D])).
fof(f15942_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f15984,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f15983])).
fof(f15983,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X40)) | p(s(bool,X41)) | ~sP9('T_0',X41,X40) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f15708,f15620])).
fof(f15708,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | p(s(bool,X41)) | ~sP9('T_0',X41,X40) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f608,f1089])).
fof(f15939,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl44),
  inference(cnf_transformation,[],[f15939_D])).
fof(f15939_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f15982,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f15981])).
fof(f15981,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | sP9('T_0',X36,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f15707,f15620])).
fof(f15707,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | sP9('T_0',X36,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f607,f1089])).
fof(f15936,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f15936_D])).
fof(f15936_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f15980,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f15979])).
fof(f15979,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | sP9('T_0',X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f15706,f15620])).
fof(f15706,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X31)) | ~p(s(bool,X32)) | sP9('T_0',X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f606,f1089])).
fof(f15933,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl40),
  inference(cnf_transformation,[],[f15933_D])).
fof(f15933_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f15978,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f15977])).
fof(f15977,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP5(X28,'T_0',X29) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f15705,f15620])).
fof(f15705,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | ~sP5(X28,'T_0',X29) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f557,f1089])).
fof(f15930,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl38),
  inference(cnf_transformation,[],[f15930_D])).
fof(f15930_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f15976,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f15975])).
fof(f15975,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X25)) | sP1(X25,'T_0') | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f15704,f15620])).
fof(f15704,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP1(X25,'T_0') | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f524,f1089])).
fof(f15926,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl34),
  inference(cnf_transformation,[],[f15926_D])).
fof(f15926_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f15974,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f15973])).
fof(f15973,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | sP1('T_0',X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f15703,f15620])).
fof(f15703,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | p(s(bool,X22)) | sP1('T_0',X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f521,f1089])).
fof(f15661,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | s(bool,X74) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) | s(bool,'T_0') = s(bool,X74)) ) | $spl12),
  inference(superposition,[],[f13150,f1089])).
fof(f15944,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f15943])).
fof(f15943,plain,(
  ( ! [X47,X48,X46,X49] : (~sP8('T_0',X47,X48) | sP9('T_0',X47,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f15634,f15620])).
fof(f15634,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~sP8('T_0',X47,X48) | sP9('T_0',X47,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f631,f1089])).
fof(f15941,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f15940])).
fof(f15940,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X40)) | p(s(bool,X41)) | ~sP9('T_0',X41,X40) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f15632,f15620])).
fof(f15632,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | p(s(bool,X41)) | ~sP9('T_0',X41,X40) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f608,f1089])).
fof(f15938,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f15937])).
fof(f15937,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | sP9('T_0',X36,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f15631,f15620])).
fof(f15631,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | sP9('T_0',X36,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f607,f1089])).
fof(f15935,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f15934])).
fof(f15934,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | sP9('T_0',X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f15630,f15620])).
fof(f15630,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X31)) | ~p(s(bool,X32)) | sP9('T_0',X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f606,f1089])).
fof(f15932,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f15931])).
fof(f15931,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP5(X28,'T_0',X29) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f15629,f15620])).
fof(f15629,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | ~sP5(X28,'T_0',X29) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f557,f1089])).
fof(f15929,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f15928])).
fof(f15928,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X25)) | sP1(X25,'T_0') | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f15628,f15620])).
fof(f15628,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP1(X25,'T_0') | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f524,f1089])).
fof(f15925,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f15924])).
fof(f15924,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X22)) | sP1('T_0',X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f15627,f15620])).
fof(f15627,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | p(s(bool,X22)) | sP1('T_0',X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f521,f1089])).
fof(f15922,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) | $spl30),
  inference(cnf_transformation,[],[f15922_D])).
fof(f15922_D,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f15923,plain,(
  ( ! [X79] : (s(bool,X79) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | s(bool,'T_0') = s(bool,X79)) ) | $spl32),
  inference(cnf_transformation,[],[f15923_D])).
fof(f15923_D,plain,(
  ( ! [X79] : (s(bool,X79) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | s(bool,'T_0') = s(bool,X79)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f15582,plain,(
  ( ! [X79] : (s(bool,X79) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) | s(bool,'T_0') = s(bool,X79)) ) | $spl12),
  inference(superposition,[],[f13150,f1089])).
fof(f15516,plain,(
  ( ! [X79] : (s(bool,X79) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13)))) ) | $spl12),
  inference(superposition,[],[f13150,f1089])).
fof(f11452,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl29),
  inference(cnf_transformation,[],[f11452_D])).
fof(f11452_D,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f11450,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl26),
  inference(cnf_transformation,[],[f11450_D])).
fof(f11450_D,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f11446,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl12),
  inference(superposition,[],[f1338,f2293])).
fof(f11449,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl12),
  inference(forward_demodulation,[],[f11448,f397])).
fof(f11448,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(forward_demodulation,[],[f11427,f3733])).
fof(f11427,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1338,f2293])).
fof(f10966,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl25),
  inference(cnf_transformation,[],[f10966_D])).
fof(f10966_D,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f10964,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl22),
  inference(cnf_transformation,[],[f10964_D])).
fof(f10964_D,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f10960,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl12),
  inference(superposition,[],[f1302,f2293])).
fof(f10963,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | sP0(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl12),
  inference(forward_demodulation,[],[f10962,f396])).
fof(f10962,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(forward_demodulation,[],[f10941,f3733])).
fof(f10941,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) ) | $spl12),
  inference(superposition,[],[f1302,f2293])).
fof(f9410,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) | $spl18),
  inference(cnf_transformation,[],[f9410_D])).
fof(f9410_D,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f9412,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | $spl21),
  inference(cnf_transformation,[],[f9412_D])).
fof(f9412_D,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f9408,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.realax.real',sK13))),s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),sK12))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.probability.prob_2'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11),s(fun(sK10,bool),'const.pred_set.DIFF_2'(s(fun(sK10,bool),'const.probability.p_space_1'(s('type.pair.prod'(fun(sK10,bool),'type.pair.prod'(fun(fun(sK10,bool),bool),fun(fun(sK10,bool),'type.realax.real'))),sK11))),s(fun(sK10,bool),sK12))))),s('type.realax.real',sK13))) | $spl12),
  inference(superposition,[],[f2325,f413])).
fof(f2282,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f2281,f373])).
fof(f2281,plain,(
  ~p(s(bool,'T_0')) | $spl3),
  inference(forward_demodulation,[],[f2280,f387])).
fof(f2280,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl3),
  inference(forward_demodulation,[],[f2265,f411])).
fof(f2265,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl3),
  inference(unit_resulting_resolution,[],[f2064,f433])).
fof(f2073,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl14),
  inference(cnf_transformation,[],[f2073_D])).
fof(f2073_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2075,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl17),
  inference(cnf_transformation,[],[f2075_D])).
fof(f2075_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f2060,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f427,f377])).
fof(f2071,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) | $spl12),
  inference(cnf_transformation,[],[f2071_D])).
fof(f2071_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2064,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(cnf_transformation,[],[f2064_D])).
fof(f2064_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f2059,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f427,f376])).
fof(f2067,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(cnf_transformation,[],[f2067_D])).
fof(f2067_D,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2069,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl11),
  inference(cnf_transformation,[],[f2069_D])).
fof(f2069_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f2058,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f427,f643])).
fof(f2066,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) ) | $spl6),
  inference(cnf_transformation,[],[f2066_D])).
fof(f2066_D,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2056,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f427,f389])).
fof(f2065,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) ) | $spl4),
  inference(cnf_transformation,[],[f2065_D])).
fof(f2065_D,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f2055,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) )),
  inference(superposition,[],[f427,f391])).
fof(f2062,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(cnf_transformation,[],[f2062_D])).
fof(f2062_D,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2054,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f427,f390])).
% SZS output end Proof for thm.probability
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 2027
% Passive clauses: 12740
% Generated clauses: 108929
% Final active clauses: 1608
% Final passive clauses: 6840
% Input formulas: 107
% Initial clauses: 249
% Discarded non-redundant clauses: 2269
% 
% Duplicate literals: 210
% Trivial inequalities: 136
% Fw subsumption resolutions: 1299
% Fw demodulations: 33569
% Bw demodulations: 2655
% Condensations: 211
% Global subsumptions: 1
% 
% Simple tautologies: 270
% Equational tautologies: 78
% Forward subsumptions: 59908
% Fw demodulations to eq. taut.: 247
% Bw demodulations to eq. taut.: 26
% 
% Unit resulting resolution: 16248
% Forward superposition: 32395
% Backward superposition: 20552
% Self superposition: 176
% Equality factoring: 26
% Equality resolution: 30
% 
% Split clauses: 123
% Split components: 251
% SAT solver clauses: 13093
% SAT solver unit clauses: 1037
% SAT solver binary clauses: 4498
% 
% Sat splits: 123
% Sat splitting refutations: 104
% 
% Memory used [KB]: 97482
% Time elapsed: 7.269 s
% ------------------------------
% ------------------------------
% Success in time 8.607 s
