Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (113493)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 121
% Passive clauses: 3948
% Generated clauses: 14403
% Final active clauses: 113
% Final passive clauses: 3545
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 521
% Trivial inequalities: 211
% Fw subsumption resolutions: 1469
% Fw demodulations: 474
% Bw demodulations: 6
% 
% Simple tautologies: 297
% Equational tautologies: 90
% Forward subsumptions: 6351
% 
% Binary resolution: 282
% Factoring: 2
% Forward superposition: 4327
% Backward superposition: 6378
% Self superposition: 290
% Equality factoring: 197
% Equality resolution: 34
% 
% Split clauses: 491
% Split components: 1026
% SAT solver clauses: 1087
% SAT solver unit clauses: 42
% SAT solver binary clauses: 379
% 
% Sat splits: 491
% 
% TWLsolver clauses: 1057
% TWLsolver calls for satisfiability: 139
% 
% Memory used [KB]: 11769
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (113513)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 173
% Passive clauses: 1472
% Generated clauses: 17254
% Final active clauses: 159
% Final passive clauses: 1235
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 306
% Trivial inequalities: 196
% Fw subsumption resolutions: 657
% Fw demodulations: 2927
% Bw demodulations: 72
% Condensations: 695
% Global subsumptions: 12
% 
% Simple tautologies: 145
% Equational tautologies: 46
% Forward subsumptions: 8462
% 
% Binary resolution: 68
% Factoring: 2
% Forward superposition: 5211
% Backward superposition: 5423
% Self superposition: 139
% Equality factoring: 68
% Equality resolution: 38
% 
% Split clauses: 41
% Split components: 88
% SAT solver clauses: 2476
% SAT solver unit clauses: 158
% SAT solver binary clauses: 67
% SAT solver learnt clauses: 24
% 
% Sat splits: 41
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2428
% TWLsolver calls for satisfiability: 435
% 
% Memory used [KB]: 26481
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f105053,plain,(
  $false),
  inference(subsumption_resolution,[],[f105052,f330])).
fof(f330,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : ! [X1,X2,X3] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  file('thm.ringNorm.polynom_distinct|split|9','thm.ringNorm.polynom_distinct|split|8')).
fof(f105052,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(forward_demodulation,[],[f103639,f105051])).
fof(f105051,plain,(
  ( ! [X62] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.ringNorm.polynom'(sK0),'type.ringNorm.polynom'(sK0)),X62),s('type.ringNorm.polynom'(sK0),sK3)))) )),
  inference(forward_demodulation,[],[f103638,f245])).
fof(f245,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f185])).
fof(f185,plain,(
  ? [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X0),X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  file('thm.ringNorm.polynom_distinct|split|9','thm.ringNorm.polynom_distinct|split|9')).
fof(f103638,plain,(
  ( ! [X62] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.ringNorm.polynom'(sK0),'type.ringNorm.polynom'(sK0)),X62),s('type.ringNorm.polynom'(sK0),sK3)))) )),
  inference(superposition,[],[f49887,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X31,X32] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X31),s('type.ringNorm.polynom'(X6),X32))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X6),X31))),s('type.ringNorm.polynom'(X6),X32)))),
  file('thm.ringNorm.polynom_distinct|split|9','ARITY_0_2_const.ringNorm.Pmult')).
fof(f49887,plain,(
  ( ! [X6,X10,X11] : (s(X6,happ(s(fun('type.ringNorm.polynom'(sK0),X6),X11),s('type.ringNorm.polynom'(sK0),sK3))) = s(X6,happ(s(fun('type.ringNorm.polynom'(sK0),X6),happ(s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X6)),X10),s('type.ringNorm.polynom'(sK0),sK2))),s('type.ringNorm.polynom'(sK0),sK1)))) )),
  inference(superposition,[],[f12032,f350])).
fof(f350,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),X7),s('type.ringNorm.polynom'(X1),X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X1),X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X15,X6] : ! [X9,X13,X38,X39,X40,X41] : s(X15,happ(s(fun('type.ringNorm.polynom'(X6),X15),X41),s('type.ringNorm.polynom'(X6),X9))) = s(X15,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(X6),X9))),s(fun('type.quote.index',X15),X13),s(fun(X6,X15),X38),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X15)),X39),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X15)),X40),s(fun('type.ringNorm.polynom'(X6),X15),X41)))),
  file('thm.ringNorm.polynom_distinct|split|9','thm.ringNorm.polynom_case_def|split|4')).
fof(f12032,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),happ(s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X1),s('type.ringNorm.polynom'(sK0),sK2))),s('type.ringNorm.polynom'(sK0),sK1))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(sK0),'const.ringNorm.Popp_1'(s('type.ringNorm.polynom'(sK0),sK3))),s(fun('type.quote.index',X0),X2),s(fun(sK0,X0),X3),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X4),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X1),s(fun('type.ringNorm.polynom'(sK0),X0),X5)))) )),
  inference(superposition,[],[f352,f245])).
fof(f352,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X15,X6] : ! [X8,X7,X13,X38,X39,X40,X41] : s(X15,happ(s(fun('type.ringNorm.polynom'(X6),X15),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X15)),X40),s('type.ringNorm.polynom'(X6),X8))),s('type.ringNorm.polynom'(X6),X7))) = s(X15,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7))),s(fun('type.quote.index',X15),X13),s(fun(X6,X15),X38),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X15)),X39),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X15)),X40),s(fun('type.ringNorm.polynom'(X6),X15),X41)))),
  file('thm.ringNorm.polynom_distinct|split|9','thm.ringNorm.polynom_case_def|split|3')).
fof(f103639,plain,(
  ( ! [X63] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun('type.ringNorm.polynom'(sK0),'type.ringNorm.polynom'(sK0)),X63),s('type.ringNorm.polynom'(sK0),sK3)))) )),
  inference(superposition,[],[f49887,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6] : ! [X31,X32] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X31),s('type.ringNorm.polynom'(X6),X32))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X6),X31))),s('type.ringNorm.polynom'(X6),X32)))),
  file('thm.ringNorm.polynom_distinct|split|9','ARITY_0_2_const.ringNorm.Pplus')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 536
% Passive clauses: 14620
% Generated clauses: 117391
% Final active clauses: 311
% Final passive clauses: 1424
% Input formulas: 69
% Initial clauses: 125
% Discarded non-redundant clauses: 4506
% 
% Duplicate literals: 2520
% Trivial inequalities: 721
% Fw subsumption resolutions: 4374
% Fw demodulations: 13554
% Bw demodulations: 9097
% Condensations: 5552
% Global subsumptions: 249
% 
% Simple tautologies: 637
% Equational tautologies: 604
% Forward subsumptions: 69338
% Fw demodulations to eq. taut.: 2801
% Bw demodulations to eq. taut.: 1297
% 
% Unit resulting resolution: 19487
% Factoring: 8
% Forward superposition: 29102
% Backward superposition: 18316
% Self superposition: 307
% Equality factoring: 138
% Equality resolution: 60
% 
% Split clauses: 248
% Split components: 518
% SAT solver clauses: 15848
% SAT solver unit clauses: 628
% SAT solver binary clauses: 2756
% 
% Sat splits: 248
% Sat splitting refutations: 481
% 
% Memory used [KB]: 68826
% Time elapsed: 8.705 s
% ------------------------------
% ------------------------------
% Success in time 10.05 s
