Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real
% (58518)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 450
% Passive clauses: 3260
% Generated clauses: 20362
% Final active clauses: 390
% Final passive clauses: 2371
% Input formulas: 42
% Initial clauses: 132
% 
% Duplicate literals: 769
% Trivial inequalities: 163
% Fw subsumption resolutions: 2265
% Fw demodulations: 1941
% Bw demodulations: 53
% 
% Simple tautologies: 514
% Equational tautologies: 56
% Forward subsumptions: 10625
% Fw demodulations to eq. taut.: 16
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 7167
% Factoring: 25
% Forward superposition: 3559
% Backward superposition: 3838
% Self superposition: 140
% Equality factoring: 74
% Equality resolution: 21
% 
% Split clauses: 671
% Split components: 1366
% SAT solver clauses: 1072
% SAT solver unit clauses: 30
% SAT solver binary clauses: 145
% 
% Sat splits: 671
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1055
% TWLsolver calls for satisfiability: 264
% 
% Memory used [KB]: 8827
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% (58552)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 196
% Passive clauses: 731
% Generated clauses: 10723
% Final active clauses: 183
% Final passive clauses: 485
% Input formulas: 42
% Initial clauses: 132
% 
% Duplicate literals: 648
% Trivial inequalities: 91
% Fw subsumption resolutions: 1002
% Fw demodulations: 954
% Bw demodulations: 52
% Condensations: 995
% Global subsumptions: 14
% 
% Simple tautologies: 272
% Equational tautologies: 49
% Forward subsumptions: 4845
% Fw demodulations to eq. taut.: 20
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 626
% Factoring: 10
% Forward superposition: 2489
% Backward superposition: 2674
% Self superposition: 165
% Equality factoring: 100
% Equality resolution: 14
% 
% Split clauses: 60
% Split components: 129
% SAT solver clauses: 1270
% SAT solver unit clauses: 145
% SAT solver binary clauses: 62
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 10
% 
% Sat splits: 60
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1214
% TWLsolver calls for satisfiability: 1046
% 
% Memory used [KB]: 8571
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f10743,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f10742])).
fof(f10742,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK9),s('type.realax.real',sK8))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK9),s('type.realax.real',sK8)))),
  inference(forward_demodulation,[],[f10741,f6287])).
fof(f6287,plain,(
  ( ! [X6] : (s('type.realax.real',X6) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',sK7))),s('type.realax.real',sK7)))) )),
  inference(superposition,[],[f2066,f1128])).
fof(f1128,plain,(
  ( ! [X10,X11,X9] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X10))),s('type.realax.real',X11))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X10),s('type.realax.real',X11)))))) )),
  inference(forward_demodulation,[],[f1116,f251])).
fof(f251,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X7] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7)))))),
  file('thm.real.REAL_DIV_OUTER_CANCEL','thm.real.real_div')).
fof(f1116,plain,(
  ( ! [X10,X11,X9] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X10))),s('type.realax.real',X11))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X10),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11)))))))) )),
  inference(superposition,[],[f276,f251])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7,X8] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real',X8))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real',X8)))),
  file('thm.real.REAL_DIV_OUTER_CANCEL','thm.realax.REAL_MUL_ASSOC')).
fof(f2066,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK7)))))) )),
  inference(unit_resulting_resolution,[],[f199,f1488])).
fof(f1488,plain,(
  ( ! [X6,X5] : (s('type.realax.real',X6) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X5),s('type.realax.real',X5))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1258,f1357])).
fof(f1357,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK7))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1259,f251])).
fof(f1259,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f1256,f216])).
fof(f216,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : (~s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))))),
  file('thm.real.REAL_DIV_OUTER_CANCEL','thm.real.REAL_MUL_RINV')).
fof(f1256,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK7))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f1227,f251])).
fof(f1227,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK7),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK7)))))),
  inference(unit_resulting_resolution,[],[f199,f216])).
fof(f1258,plain,(
  ( ! [X1] : (s('type.realax.real',X1) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK7)))))) )),
  inference(backward_demodulation,[],[f1256,f1059])).
fof(f1059,plain,(
  ( ! [X1] : (s('type.realax.real',X1) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f208,f250])).
fof(f250,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X7] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real',X6)))),
  file('thm.real.REAL_DIV_OUTER_CANCEL','thm.realax.REAL_MUL_SYM')).
fof(f208,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : s('type.realax.real',X6) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X6)))),
  file('thm.real.REAL_DIV_OUTER_CANCEL','thm.real.REAL_MUL_LID')).
fof(f199,plain,(
  s('type.realax.real',sK7) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  s('type.realax.real',sK7) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK9),s('type.realax.real',sK8))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK8))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK9),s('type.realax.real',sK7)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9])],[f104])).
fof(f104,plain,(
  ? [X0,X1,X2] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ~! [X0,X1,X2] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ~! [X0,X1,X2] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X0))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (~s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X6))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (~s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X6))))))),
  file('thm.real.REAL_DIV_OUTER_CANCEL','thm.real.REAL_DIV_OUTER_CANCEL')).
fof(f10741,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK9),s('type.realax.real',sK8))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK9),s('type.realax.real',sK7))),s('type.realax.real',sK7))),s('type.realax.real',sK8)))),
  inference(forward_demodulation,[],[f10660,f4872])).
fof(f4872,plain,(
  ( ! [X26,X27,X25] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X25),s('type.realax.real',X26))),s('type.realax.real',X27))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X26),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X25),s('type.realax.real',X27)))))) )),
  inference(forward_demodulation,[],[f4782,f251])).
fof(f4782,plain,(
  ( ! [X26,X27,X25] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X25),s('type.realax.real',X26))),s('type.realax.real',X27))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X26),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X27)))))))) )),
  inference(superposition,[],[f1104,f251])).
fof(f1104,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f276,f250])).
fof(f10660,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK9),s('type.realax.real',sK8))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK7),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK9),s('type.realax.real',sK7))))),s('type.realax.real',sK8)))),
  inference(superposition,[],[f200,f1129])).
fof(f1129,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X8))),s('type.realax.real',X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real',X8)))) )),
  inference(backward_demodulation,[],[f1128,f1125])).
fof(f1125,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',X7))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real',X8)))) )),
  inference(forward_demodulation,[],[f1106,f1079])).
fof(f1079,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f251,f250])).
fof(f1106,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real',X8))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X8)))))) )),
  inference(superposition,[],[f276,f251])).
fof(f200,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK9),s('type.realax.real',sK8))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK7),s('type.realax.real',sK8))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK9),s('type.realax.real',sK7)))))),
  inference(cnf_transformation,[],[f147])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 209
% Passive clauses: 1181
% Generated clauses: 11064
% Final active clauses: 133
% Final passive clauses: 320
% Input formulas: 42
% Initial clauses: 132
% 
% Duplicate literals: 76
% Trivial inequalities: 12
% Fw subsumption resolutions: 105
% Fw demodulations: 2592
% Bw demodulations: 424
% Condensations: 241
% Global subsumptions: 7
% 
% Simple tautologies: 62
% Equational tautologies: 44
% Forward subsumptions: 6613
% Fw demodulations to eq. taut.: 36
% Bw demodulations to eq. taut.: 14
% 
% Unit resulting resolution: 345
% Factoring: 1
% Forward superposition: 3408
% Backward superposition: 2906
% Self superposition: 195
% Equality factoring: 26
% Equality resolution: 2
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 1251
% SAT solver unit clauses: 139
% SAT solver binary clauses: 329
% 
% Sat splits: 27
% Sat splitting refutations: 7
% 
% Memory used [KB]: 11257
% Time elapsed: 0.509 s
% ------------------------------
% ------------------------------
% Success in time 1.822 s
