Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ind_type
% (56153)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 333
% Passive clauses: 4442
% Generated clauses: 32631
% Final active clauses: 190
% Final passive clauses: 3725
% Input formulas: 84
% Initial clauses: 115
% 
% Duplicate literals: 931
% Trivial inequalities: 226
% Fw subsumption resolutions: 1622
% Fw demodulations: 6050
% Bw demodulations: 106
% 
% Simple tautologies: 311
% Equational tautologies: 56
% Forward subsumptions: 18749
% Fw demodulations to eq. taut.: 20
% Bw demodulations to eq. taut.: 8
% 
% Binary resolution: 225
% Factoring: 1
% Forward superposition: 11020
% Backward superposition: 12015
% Self superposition: 167
% Equality factoring: 90
% Equality resolution: 31
% 
% Split clauses: 183
% Split components: 372
% SAT solver clauses: 485
% SAT solver unit clauses: 6
% SAT solver binary clauses: 88
% 
% Sat splits: 183
% Sat splitting refutations: 15
% 
% TWLsolver clauses: 483
% TWLsolver calls for satisfiability: 114
% 
% Memory used [KB]: 17270
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ind_type
% (56162)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 232
% Passive clauses: 1311
% Generated clauses: 24344
% Final active clauses: 176
% Final passive clauses: 826
% Input formulas: 84
% Initial clauses: 115
% 
% Duplicate literals: 532
% Trivial inequalities: 174
% Fw subsumption resolutions: 1059
% Fw demodulations: 4228
% Bw demodulations: 261
% Condensations: 685
% Global subsumptions: 7
% 
% Simple tautologies: 190
% Equational tautologies: 43
% Forward subsumptions: 13353
% Fw demodulations to eq. taut.: 25
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 147
% Factoring: 2
% Forward superposition: 7219
% Backward superposition: 8132
% Self superposition: 143
% Equality factoring: 70
% Equality resolution: 21
% 
% Split clauses: 80
% Split components: 160
% SAT solver clauses: 2677
% SAT solver unit clauses: 178
% SAT solver binary clauses: 64
% SAT solver learnt clauses: 16
% SAT solver learnt literals: 6
% 
% Sat splits: 80
% Sat splitting refutations: 53
% 
% TWLsolver clauses: 2647
% TWLsolver calls for satisfiability: 553
% 
% Memory used [KB]: 18421
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ind_type
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ind_type
% SZS output start Proof for thm.ind_type
fof(f7739,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f217,f334,f333,f216,f335,f332,f956,f967,f965,f960,f972,f970,f1443,f1464,f1462,f1760,f1784,f1782,f2069,f2092,f2090,f2593,f2623,f2621,f3094,f3129,f3127,f3463,f3508,f3506,f4326,f4378,f4376,f5496,f5497,f5498,f5499,f5500,f5504,f5507,f5511,f5515,f5519,f5523,f5527,f5531,f5533,f5535,f5537,f5541,f5542,f5544,f5545,f5546,f5547,f5548,f5549,f5550,f5552,f5554,f5579,f5583,f5605,f5610,f5611,f5612,f5618,f5622,f5626,f5890,f5877,f5964,f6099,f6100,f6101,f6102,f6103,f6104,f6105,f6108,f6112,f6115,f6119,f6123,f6127,f6131,f6135,f6139,f6142,f6146,f6148,f6150,f6152,f6156,f6159,f6160,f6162,f6164,f6165,f6212,f6216,f6218,f6220,f6223,f6235,f6238,f6240,f6241,f6243,f6244,f6247,f6249,f6254,f6257,f6260,f6262,f6266,f7532,f7533,f7548,f7546,f7651,f7652,f7667,f7665,f7678,f7701,f7702,f7705,f7723,f7728,f7731,f7734,f7738])).
fof(f7738,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f7700,f218])).
fof(f218,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ind_type.FCONS_DEST',aHL_TRUTH)).
fof(f7700,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(trivial_inequality_removal,[],[f7699])).
fof(f7699,plain,(
  s(sK0,sK3) != s(sK0,sK3) | ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(superposition,[],[f5893,f287])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X13,X14] : s(X6,X13) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X13),s(X6,X14)))),
  file('thm.ind_type.FCONS_DEST','thm.bool.COND_CLAUSES|split|0')).
fof(f5893,plain,(
  ( ! [X4] : (s(sK0,sK3) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X4))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f5892,f289])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ! [X9,X8] : s(X6,X9) = s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X8),s('type.num.num','const.num.|48|_0')))),
  file('thm.ind_type.FCONS_DEST','thm.ind_type.FCONS|split|0')).
fof(f5892,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X4))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f5891,f230])).
fof(f230,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.SUB|split|0')).
fof(f5891,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | ~p(s(bool,X4))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f5646,f332])).
fof(f5646,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | ~p(s(bool,X4))) ) | $spl2),
  inference(forward_demodulation,[],[f334,f222])).
fof(f222,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f55])).
fof(f55,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|13')).
fof(f7734,plain,(
  $false | ($spl0 | $spl2 | $spl42)),
  inference(subsumption_resolution,[],[f7733,f218])).
fof(f7733,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl42)),
  inference(forward_demodulation,[],[f7732,f7546])).
fof(f7732,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl2 | $spl42)),
  inference(subsumption_resolution,[],[f7695,f287])).
fof(f7695,plain,(
  s(sK0,sK3) != s(sK0,'const.bool.COND_3'(s(bool,'T_0'),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl2 | $spl42)),
  inference(superposition,[],[f5893,f7546])).
fof(f7731,plain,(
  $false | ($spl0 | $spl2 | $spl16 | $spl42)),
  inference(subsumption_resolution,[],[f7730,f218])).
fof(f7730,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl16 | $spl42)),
  inference(forward_demodulation,[],[f7729,f7556])).
fof(f7556,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl42)),
  inference(backward_demodulation,[],[f7546,f6384])).
fof(f6384,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16)),
  inference(backward_demodulation,[],[f6317,f645])).
fof(f645,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f247,f222])).
fof(f247,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X7] : s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X7))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|33')).
fof(f6317,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | ($spl0 | $spl16)),
  inference(unit_resulting_resolution,[],[f6172,f272])).
fof(f272,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X7))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|18')).
fof(f6172,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16)),
  inference(forward_demodulation,[],[f1463,f332])).
fof(f1463,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) | $spl16),
  inference(cnf_transformation,[],[f1463_D])).
fof(f1463_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f7729,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl2 | $spl16 | $spl42)),
  inference(subsumption_resolution,[],[f7694,f287])).
fof(f7694,plain,(
  s(sK0,sK3) != s(sK0,'const.bool.COND_3'(s(bool,'T_0'),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl2 | $spl16 | $spl42)),
  inference(superposition,[],[f5893,f7556])).
fof(f7728,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f7727,f218])).
fof(f7727,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f7726,f348])).
fof(f348,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f220,f265])).
fof(f265,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.ind_type.FCONS_DEST','thm.bool.EQ_CLAUSES|split|0')).
fof(f220,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f77])).
fof(f77,axiom,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|35')).
fof(f7726,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f7693,f287])).
fof(f7693,plain,(
  s(sK0,sK3) != s(sK0,'const.bool.COND_3'(s(bool,'T_0'),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | ($spl0 | $spl2)),
  inference(superposition,[],[f5893,f348])).
fof(f7723,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f7722,f218])).
fof(f7722,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f7721,f235])).
fof(f235,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X7] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|29')).
fof(f7721,plain,(
  ( ! [X10] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0'))))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f7691,f287])).
fof(f7691,plain,(
  ( ! [X10] : (s(sK0,sK3) != s(sK0,'const.bool.COND_3'(s(bool,'T_0'),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0'))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f5893,f235])).
fof(f7705,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f7704,f218])).
fof(f7704,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f7703,f236])).
fof(f236,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X7] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|26')).
fof(f7703,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f7681,f287])).
fof(f7681,plain,(
  ( ! [X2] : (s(sK0,sK3) != s(sK0,'const.bool.COND_3'(s(bool,'T_0'),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f5893,f236])).
fof(f7702,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f7676,f287])).
fof(f7676,plain,(
  s(sK0,sK3) != s(sK0,'const.bool.COND_3'(s(bool,'T_0'),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f218,f5893])).
fof(f7701,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f7677,f218])).
fof(f7677,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f287,f5893])).
fof(f7678,plain,(
  $false | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f218,f287,f5893])).
fof(f7665,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl46),
  inference(cnf_transformation,[],[f7665_D])).
fof(f7665_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f7667,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl48),
  inference(cnf_transformation,[],[f7667_D])).
fof(f7667_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f7652,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16)),
  inference(superposition,[],[f261,f6385])).
fof(f6385,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16)),
  inference(backward_demodulation,[],[f6317,f670])).
fof(f670,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f248,f222])).
fof(f248,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X7] : s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X7))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|32')).
fof(f261,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.ind_type.FCONS_DEST',aHL_BOOL_CASES)).
fof(f7651,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16)),
  inference(superposition,[],[f6385,f261])).
fof(f7546,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl42),
  inference(cnf_transformation,[],[f7546_D])).
fof(f7546_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f7548,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl44),
  inference(cnf_transformation,[],[f7548_D])).
fof(f7548_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f7533,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16)),
  inference(superposition,[],[f261,f6384])).
fof(f7532,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16)),
  inference(superposition,[],[f6384,f261])).
fof(f6266,plain,(
  $false | ($spl0 | $spl16 | $spl37)),
  inference(subsumption_resolution,[],[f6265,f6172])).
fof(f6265,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl37)),
  inference(forward_demodulation,[],[f6264,f6172])).
fof(f6264,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl37)),
  inference(forward_demodulation,[],[f6263,f6172])).
fof(f6263,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl16 | $spl37)),
  inference(forward_demodulation,[],[f6204,f6172])).
fof(f6204,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl0 | $spl16 | $spl37)),
  inference(backward_demodulation,[],[f6172,f5682])).
fof(f5682,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))))))) | ($spl0 | $spl37)),
  inference(backward_demodulation,[],[f332,f3874])).
fof(f3874,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))))))) | $spl37),
  inference(unit_resulting_resolution,[],[f3599,f318])).
fof(f318,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X7,X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))) <=> s('type.num.num',X7) = s('type.num.num',X11))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|19')).
fof(f3599,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))))) | $spl37),
  inference(unit_resulting_resolution,[],[f3508,f318])).
fof(f6262,plain,(
  $false | ($spl0 | $spl16 | $spl37)),
  inference(subsumption_resolution,[],[f6261,f236])).
fof(f6261,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl0 | $spl16 | $spl37)),
  inference(forward_demodulation,[],[f6203,f6172])).
fof(f6203,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl0 | $spl16 | $spl37)),
  inference(backward_demodulation,[],[f6172,f5681])).
fof(f5681,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl0 | $spl37)),
  inference(backward_demodulation,[],[f332,f3724])).
fof(f3724,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl37),
  inference(unit_resulting_resolution,[],[f3600,f264])).
fof(f264,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f195])).
fof(f3600,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl37),
  inference(forward_demodulation,[],[f3597,f3170])).
fof(f3170,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f301,f994])).
fof(f994,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f979,f288])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X13,X14] : s(X6,X14) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X13),s(X6,X14)))),
  file('thm.ind_type.FCONS_DEST','thm.bool.COND_CLAUSES|split|1')).
fof(f979,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(unit_resulting_resolution,[],[f288,f273])).
fof(f273,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f199])).
fof(f301,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X7,X11] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X7))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|31')).
fof(f3597,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))))))) | $spl37),
  inference(unit_resulting_resolution,[],[f3508,f276])).
fof(f276,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X7] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|30')).
fof(f6260,plain,(
  $false | ($spl0 | $spl16 | $spl37)),
  inference(subsumption_resolution,[],[f6259,f344])).
fof(f344,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f336,f264])).
fof(f336,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f224,f218])).
fof(f224,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f27])).
fof(f27,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.ind_type.FCONS_DEST','thm.bool.NOT_CLAUSES|split|1')).
fof(f6259,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl16 | $spl37)),
  inference(forward_demodulation,[],[f6258,f236])).
fof(f6258,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl0 | $spl16 | $spl37)),
  inference(forward_demodulation,[],[f6202,f6172])).
fof(f6202,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl0 | $spl16 | $spl37)),
  inference(backward_demodulation,[],[f6172,f5680])).
fof(f5680,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl0 | $spl37)),
  inference(backward_demodulation,[],[f332,f3723])).
fof(f3723,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl37),
  inference(unit_resulting_resolution,[],[f3600,f269])).
fof(f269,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.ind_type.FCONS_DEST','thm.bool.EQ_CLAUSES|split|3')).
fof(f6257,plain,(
  $false | ($spl0 | $spl16 | $spl37)),
  inference(subsumption_resolution,[],[f6256,f6172])).
fof(f6256,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl37)),
  inference(forward_demodulation,[],[f6255,f6172])).
fof(f6255,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl37)),
  inference(forward_demodulation,[],[f6201,f6172])).
fof(f6201,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl16 | $spl37)),
  inference(backward_demodulation,[],[f6172,f5679])).
fof(f5679,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl0 | $spl37)),
  inference(backward_demodulation,[],[f332,f3599])).
fof(f6254,plain,(
  $false | ($spl0 | $spl13 | $spl16 | $spl21 | $spl37)),
  inference(subsumption_resolution,[],[f6253,f972])).
fof(f6253,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21 | $spl37)),
  inference(forward_demodulation,[],[f6252,f6172])).
fof(f6252,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21 | $spl37)),
  inference(forward_demodulation,[],[f6251,f6226])).
fof(f6226,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))) | ($spl0 | $spl16 | $spl21)),
  inference(forward_demodulation,[],[f6225,f230])).
fof(f6225,plain,(
  s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl21)),
  inference(forward_demodulation,[],[f6197,f6172])).
fof(f6197,plain,(
  s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl21)),
  inference(backward_demodulation,[],[f6172,f5674])).
fof(f5674,plain,(
  s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl21)),
  inference(backward_demodulation,[],[f332,f3440])).
fof(f3440,plain,(
  s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl21),
  inference(superposition,[],[f341,f1790])).
fof(f1790,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f1784,f262])).
fof(f262,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f124])).
fof(f124,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X11] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X11) | ? [X7] : s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.num_CASES')).
fof(f341,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f256,f222])).
fof(f256,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11] : s('type.num.num',X11) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.SUC_SUB|49|')).
fof(f6251,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl16 | $spl37)),
  inference(forward_demodulation,[],[f6200,f6172])).
fof(f6200,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl0 | $spl16 | $spl37)),
  inference(backward_demodulation,[],[f6172,f5678])).
fof(f5678,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))))))) | ($spl0 | $spl37)),
  inference(backward_demodulation,[],[f332,f3596])).
fof(f3596,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))))))) | $spl37),
  inference(unit_resulting_resolution,[],[f3508,f262])).
fof(f6249,plain,(
  $false | ($spl0 | $spl16 | $spl37)),
  inference(subsumption_resolution,[],[f6248,f6172])).
fof(f6248,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl37)),
  inference(forward_demodulation,[],[f6198,f6172])).
fof(f6198,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl37)),
  inference(backward_demodulation,[],[f6172,f5676])).
fof(f5676,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl37)),
  inference(backward_demodulation,[],[f332,f3508])).
fof(f6247,plain,(
  $false | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(subsumption_resolution,[],[f6246,f6172])).
fof(f6246,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(forward_demodulation,[],[f6245,f6172])).
fof(f6245,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(forward_demodulation,[],[f6234,f6172])).
fof(f6234,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(backward_demodulation,[],[f6226,f6224])).
fof(f6224,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))))))) | ($spl0 | $spl16 | $spl33)),
  inference(forward_demodulation,[],[f6196,f6172])).
fof(f6196,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))))))) | ($spl0 | $spl16 | $spl33)),
  inference(backward_demodulation,[],[f6172,f5673])).
fof(f5673,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))))))) | ($spl0 | $spl33)),
  inference(backward_demodulation,[],[f332,f3435])).
fof(f3435,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))))))) | $spl33),
  inference(unit_resulting_resolution,[],[f3141,f318])).
fof(f3141,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))))) | $spl33),
  inference(unit_resulting_resolution,[],[f3129,f318])).
fof(f6244,plain,(
  $false | ($spl0 | $spl13 | $spl16 | $spl21)),
  inference(subsumption_resolution,[],[f6233,f972])).
fof(f6233,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21)),
  inference(backward_demodulation,[],[f6226,f6213])).
fof(f6213,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl16 | $spl21)),
  inference(forward_demodulation,[],[f6190,f6172])).
fof(f6190,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl16 | $spl21)),
  inference(backward_demodulation,[],[f6172,f5656])).
fof(f5656,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl0 | $spl21)),
  inference(backward_demodulation,[],[f332,f1790])).
fof(f6243,plain,(
  $false | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(subsumption_resolution,[],[f6242,f6172])).
fof(f6242,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(forward_demodulation,[],[f6232,f6172])).
fof(f6232,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(backward_demodulation,[],[f6226,f6195])).
fof(f6195,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl0 | $spl16 | $spl33)),
  inference(backward_demodulation,[],[f6172,f5670])).
fof(f5670,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl0 | $spl33)),
  inference(backward_demodulation,[],[f332,f3141])).
fof(f6241,plain,(
  $false | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(subsumption_resolution,[],[f6231,f236])).
fof(f6231,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(backward_demodulation,[],[f6226,f5672])).
fof(f5672,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl0 | $spl33)),
  inference(backward_demodulation,[],[f332,f3237])).
fof(f3237,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl33),
  inference(unit_resulting_resolution,[],[f3215,f264])).
fof(f3215,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl33),
  inference(backward_demodulation,[],[f3170,f3139])).
fof(f3139,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))))) | $spl33),
  inference(unit_resulting_resolution,[],[f3129,f276])).
fof(f6240,plain,(
  $false | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(subsumption_resolution,[],[f6239,f344])).
fof(f6239,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(forward_demodulation,[],[f6230,f236])).
fof(f6230,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(backward_demodulation,[],[f6226,f5671])).
fof(f5671,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl0 | $spl33)),
  inference(backward_demodulation,[],[f332,f3236])).
fof(f3236,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))),s('type.num.num','const.arithmetic.ZERO_0'))) | $spl33),
  inference(unit_resulting_resolution,[],[f3215,f269])).
fof(f6238,plain,(
  $false | ($spl0 | $spl13 | $spl16 | $spl21 | $spl33)),
  inference(subsumption_resolution,[],[f6237,f972])).
fof(f6237,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(forward_demodulation,[],[f6236,f6172])).
fof(f6236,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(forward_demodulation,[],[f6229,f6226])).
fof(f6229,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(backward_demodulation,[],[f6226,f5669])).
fof(f5669,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl0 | $spl33)),
  inference(backward_demodulation,[],[f332,f3138])).
fof(f3138,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))))) | $spl33),
  inference(unit_resulting_resolution,[],[f3129,f262])).
fof(f6235,plain,(
  $false | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(subsumption_resolution,[],[f6227,f6172])).
fof(f6227,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21 | $spl33)),
  inference(backward_demodulation,[],[f6226,f5667])).
fof(f5667,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl33)),
  inference(backward_demodulation,[],[f332,f3129])).
fof(f6223,plain,(
  $false | ($spl0 | $spl16 | $spl21)),
  inference(subsumption_resolution,[],[f6222,f6172])).
fof(f6222,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21)),
  inference(forward_demodulation,[],[f6221,f6172])).
fof(f6221,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl21)),
  inference(forward_demodulation,[],[f6194,f6172])).
fof(f6194,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl16 | $spl21)),
  inference(backward_demodulation,[],[f6172,f5664])).
fof(f5664,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl0 | $spl21)),
  inference(backward_demodulation,[],[f332,f2760])).
fof(f2760,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f2277,f318])).
fof(f2277,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f1784,f318])).
fof(f6220,plain,(
  $false | ($spl0 | $spl16 | $spl21)),
  inference(subsumption_resolution,[],[f6219,f6172])).
fof(f6219,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21)),
  inference(forward_demodulation,[],[f6193,f6172])).
fof(f6193,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl21)),
  inference(backward_demodulation,[],[f6172,f5661])).
fof(f5661,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl21)),
  inference(backward_demodulation,[],[f332,f2277])).
fof(f6218,plain,(
  $false | ($spl0 | $spl16 | $spl21)),
  inference(subsumption_resolution,[],[f6217,f235])).
fof(f6217,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21)),
  inference(forward_demodulation,[],[f6192,f6172])).
fof(f6192,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl21)),
  inference(backward_demodulation,[],[f6172,f5658])).
fof(f5658,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl21)),
  inference(backward_demodulation,[],[f332,f1867])).
fof(f1867,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f1791,f264])).
fof(f1791,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f1784,f276])).
fof(f6216,plain,(
  $false | ($spl0 | $spl16 | $spl21)),
  inference(subsumption_resolution,[],[f6215,f344])).
fof(f6215,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl16 | $spl21)),
  inference(forward_demodulation,[],[f6214,f235])).
fof(f6214,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21)),
  inference(forward_demodulation,[],[f6191,f6172])).
fof(f6191,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl21)),
  inference(backward_demodulation,[],[f6172,f5657])).
fof(f5657,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl21)),
  inference(backward_demodulation,[],[f332,f1866])).
fof(f1866,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f1791,f269])).
fof(f6212,plain,(
  $false | ($spl0 | $spl16 | $spl21)),
  inference(subsumption_resolution,[],[f6188,f6172])).
fof(f6188,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl21)),
  inference(backward_demodulation,[],[f6172,f5654])).
fof(f5654,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl21)),
  inference(backward_demodulation,[],[f332,f1784])).
fof(f6165,plain,(
  $false | ($spl0 | $spl13 | $spl17)),
  inference(subsumption_resolution,[],[f6091,f218])).
fof(f6091,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl13 | $spl17)),
  inference(backward_demodulation,[],[f5968,f1064])).
fof(f1064,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl13),
  inference(unit_resulting_resolution,[],[f972,f276])).
fof(f5968,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f5967,f265])).
fof(f5967,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X0))) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f5966,f236])).
fof(f5966,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0'))) | p(s(bool,X0))) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f5965,f3170])).
fof(f5965,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f5955,f336])).
fof(f5955,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | s(bool,X0) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl17)),
  inference(superposition,[],[f311,f5652])).
fof(f5652,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f1579])).
fof(f1579,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))) | $spl17),
  inference(unit_resulting_resolution,[],[f1471,f269])).
fof(f1471,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) | $spl17),
  inference(unit_resulting_resolution,[],[f1464,f276])).
fof(f311,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.ind_type.FCONS_DEST','thm.bool.IMP_ANTISYM_AX')).
fof(f6164,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122,f6126,f6130,f6134,f6138,f6145,f6147,f6149,f6151,f6155,f6161,f6163])).
fof(f6163,plain,(
  ( ! [X134,X133] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X133),s('type.num.num',X134)))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6089,f218])).
fof(f6089,plain,(
  ( ! [X134,X133] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X133),s('type.num.num',X134)))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f5638])).
fof(f6161,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6087,f218])).
fof(f6087,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f276])).
fof(f6155,plain,(
  ( ! [X10] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6154,f218])).
fof(f6154,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6153,f5968])).
fof(f6153,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6057,f218])).
fof(f6057,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f3613])).
fof(f6151,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6054,f218])).
fof(f6054,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f4079])).
fof(f6149,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6053,f218])).
fof(f6053,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f4081])).
fof(f6147,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6052,f218])).
fof(f6052,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f3948])).
fof(f6145,plain,(
  ( ! [X249,X248,X247] : (s('type.num.num',X249) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248)))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6144,f218])).
fof(f6144,plain,(
  ( ! [X249,X248,X247] : (~p(s(bool,'T_0')) | s('type.num.num',X249) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248)))) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6143,f5968])).
fof(f6143,plain,(
  ( ! [X249,X248,X247] : (s('type.num.num',X249) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X249),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248))))))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6048,f218])).
fof(f6048,plain,(
  ( ! [X249,X248,X247] : (~p(s(bool,'T_0')) | s('type.num.num',X249) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X249),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248))))))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f5632])).
fof(f6138,plain,(
  ( ! [X10] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6137,f218])).
fof(f6137,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6136,f5968])).
fof(f6136,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6040,f218])).
fof(f6040,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f3641])).
fof(f6134,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6133,f218])).
fof(f6133,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6132,f5968])).
fof(f6132,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9)))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6038,f218])).
fof(f6038,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9)))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f4084])).
fof(f6130,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6129,f218])).
fof(f6129,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6128,f5968])).
fof(f6128,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6037,f218])).
fof(f6037,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f4083])).
fof(f6126,plain,(
  ( ! [X6,X5] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6125,f218])).
fof(f6125,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6124,f5968])).
fof(f6124,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6)))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6036,f218])).
fof(f6036,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6)))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f4082])).
fof(f6122,plain,(
  ( ! [X10,X11] : (s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6121,f218])).
fof(f6121,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6120,f5968])).
fof(f6120,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))),s('type.num.num',X11)))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6035,f218])).
fof(f6035,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))),s('type.num.num',X11)))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f3910])).
fof(f6118,plain,(
  ( ! [X10,X11] : (s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6117,f218])).
fof(f6117,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6116,f5968])).
fof(f6116,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10))))))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6034,f218])).
fof(f6034,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10))))))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f3882])).
fof(f6111,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6110,f218])).
fof(f6110,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6109,f5968])).
fof(f6109,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6022,f218])).
fof(f6022,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f317])).
fof(f6002,plain,(
  ( ! [X14,X12,X10,X13] : (s(X10,X14) = s(X10,'const.bool.COND_3'(s(bool,'T_0'),s(X10,happ(s(fun('type.num.num',X10),X12),s('type.num.num',X13))),s(X10,X14)))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f5311])).
fof(f6001,plain,(
  ( ! [X8,X7,X5,X9] : (s(X5,happ(s(fun('type.num.num',X5),X7),s('type.num.num',X8))) = s(X5,'const.bool.COND_3'(s(bool,'T_0'),s(X5,X9),s(X5,happ(s(fun('type.num.num',X5),X7),s('type.num.num',X8)))))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f5310])).
fof(f5986,plain,(
  ( ! [X43,X41,X42] : (s('type.num.num',X43) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X41),s('type.num.num',X42))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43)))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f2306])).
fof(f5982,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X11),s('type.num.num',X12))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f959])).
fof(f5975,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f288])).
fof(f6098,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f5973,f5968])).
fof(f5973,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f268])).
fof(f5843,plain,(
  ( ! [X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))) | s('type.num.num',X16) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.ZERO_0'))))) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f5825,f994])).
fof(f5825,plain,(
  ( ! [X16] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))) | s('type.num.num',X16) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.ZERO_0'))))) ) | ($spl0 | $spl17)),
  inference(superposition,[],[f318,f5740])).
fof(f5798,plain,(
  s('type.num.num',sK4('const.arithmetic.ZERO_0')) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl17)),
  inference(superposition,[],[f341,f5740])).
fof(f5797,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK4('const.arithmetic.ZERO_0')))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X2),s(fun('type.num.num',X0),X1),s('type.num.num','const.arithmetic.ZERO_0')))) ) | ($spl0 | $spl17)),
  inference(superposition,[],[f292,f5740])).
fof(f5740,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.ZERO_0')))) | ($spl0 | $spl17)),
  inference(unit_resulting_resolution,[],[f5650,f262])).
fof(f5683,plain,(
  ( ! [X8,X7,X9] : (s(X7,happ(s(fun('type.num.num',X7),X8),s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) = s(X7,'const.ind_type.FCONS_3'(s(X7,X9),s(fun('type.num.num',X7),X8),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f5304])).
fof(f5689,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num',X8) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f5666,f332])).
fof(f5666,plain,(
  ( ! [X8] : (s('type.num.num',X8) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f3106])).
fof(f5665,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f2991])).
fof(f5663,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f2355])).
fof(f5688,plain,(
  ( ! [X12] : (s('type.num.num',X12) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f5662,f332])).
fof(f5662,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X12) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f2291])).
fof(f5660,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f2276])).
fof(f5659,plain,(
  s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f1999])).
fof(f5653,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f1580])).
fof(f5651,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f1470])).
fof(f5650,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f1469])).
fof(f5649,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f332,f1464])).
fof(f242,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X7] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|20')).
fof(f243,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X7] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|23')).
fof(f245,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X7] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_lte|split|2')).
fof(f246,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X7] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_lte|split|1')).
fof(f268,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f363,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f228,f269])).
fof(f646,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f247,f261])).
fof(f671,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f248,f261])).
fof(f733,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(superposition,[],[f287,f261])).
fof(f816,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)))) )),
  inference(superposition,[],[f288,f261])).
fof(f959,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X11),s('type.num.num',X12)))) )),
  inference(superposition,[],[f272,f288])).
fof(f1042,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f1016,f363])).
fof(f1016,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f248,f994])).
fof(f1514,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f289,f261])).
fof(f2306,plain,(
  ( ! [X43,X41,X42] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))) | s('type.num.num',X43) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X41),s('type.num.num',X42)))) )),
  inference(superposition,[],[f318,f288])).
fof(f2542,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f257,f261])).
fof(f2558,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f2540,f246])).
fof(f2540,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f257,f222])).
fof(f2663,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f258,f261])).
fof(f2732,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f260,f261])).
fof(f2848,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f298,f261])).
fof(f2863,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(forward_demodulation,[],[f2819,f243])).
fof(f2819,plain,(
  ( ! [X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f298,f994])).
fof(f2954,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f299,f261])).
fof(f3044,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f300,f261])).
fof(f3200,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f301,f261])).
fof(f3300,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f302,f261])).
fof(f3408,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f303,f261])).
fof(f3571,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f304,f261])).
fof(f3198,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f301,f222])).
fof(f2924,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f299,f222])).
fof(f299,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X7,X11] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X11))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|28')).
fof(f300,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X7,X11] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X11))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X11)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_lte|split|3')).
fof(f303,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X7,X11] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X11))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_lte|split|6')).
fof(f304,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X7,X11] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X11))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_lte|split|4')).
fof(f316,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(flattening,[],[f209])).
fof(f209,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X11,X7] : (s('type.num.num',X7) = s('type.num.num',X11) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X7)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X11))))))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.EQ_LESS_EQ')).
fof(f317,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f320,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X7,X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X11)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X7)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_lte|split|5')).
fof(f321,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f212])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2)))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X11,X7,X12] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X12))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X7))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X12)))))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.ADD_MONO_LESS_EQ')).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11,X7,X12] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X7)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X12))))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X12)))))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.LESS_EQ_TRANS')).
fof(f1415,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X20)))) | s('type.num.num',X20) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X19)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19)) )),
  inference(superposition,[],[f316,f262])).
fof(f1416,plain,(
  ( ! [X21,X22] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X22),s('type.num.num',X21)))) | s('type.num.num',X22) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X21)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X21)) )),
  inference(superposition,[],[f316,f262])).
fof(f2955,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f316,f299])).
fof(f3045,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f316,f300])).
fof(f3409,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f316,f303])).
fof(f3572,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f316,f304])).
fof(f3861,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(subsumption_resolution,[],[f3806,f336])).
fof(f3806,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f321,f261])).
fof(f3882,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X11)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10))))))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) )),
  inference(superposition,[],[f317,f262])).
fof(f3910,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))),s('type.num.num',X11)))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) )),
  inference(superposition,[],[f317,f262])).
fof(f4082,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6)))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) )),
  inference(forward_demodulation,[],[f3938,f300])).
fof(f3938,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5)))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) )),
  inference(superposition,[],[f317,f300])).
fof(f4083,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) )),
  inference(forward_demodulation,[],[f3939,f299])).
fof(f3939,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) )),
  inference(superposition,[],[f317,f299])).
fof(f4084,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X10)))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))) )),
  inference(forward_demodulation,[],[f3940,f303])).
fof(f3940,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9)))))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))) )),
  inference(superposition,[],[f317,f303])).
fof(f3785,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))),s('type.num.num',X11)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) )),
  inference(superposition,[],[f321,f262])).
fof(f3641,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))),s('type.num.num',X11)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) )),
  inference(superposition,[],[f320,f262])).
fof(f257,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X7] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.num.|48|_0')))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|27')).
fof(f3948,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f3947,f994])).
fof(f3947,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(forward_demodulation,[],[f3946,f994])).
fof(f3946,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(subsumption_resolution,[],[f3945,f218])).
fof(f3945,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(forward_demodulation,[],[f3876,f236])).
fof(f3876,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f317,f994])).
fof(f4081,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))) )),
  inference(subsumption_resolution,[],[f4080,f218])).
fof(f4080,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3933,f236])).
fof(f3933,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))) )),
  inference(superposition,[],[f317,f257])).
fof(f4079,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(subsumption_resolution,[],[f3931,f218])).
fof(f3931,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f317,f231])).
fof(f2543,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f316,f257])).
fof(f3613,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X11)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) )),
  inference(superposition,[],[f320,f262])).
fof(f3757,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X11)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) )),
  inference(superposition,[],[f321,f262])).
fof(f3169,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f301,f222])).
fof(f2952,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f299,f222])).
fof(f302,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X7,X11] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X11))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|22')).
fof(f298,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X7,X11] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X7))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|25')).
fof(f3298,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f302,f222])).
fof(f2818,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f298,f222])).
fof(f260,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X7] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X7))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|21')).
fof(f2747,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(backward_demodulation,[],[f2730,f2661])).
fof(f2730,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f260,f222])).
fof(f258,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X7] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X7))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.num.|48|_0')))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|24')).
fof(f3269,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f302,f222])).
fof(f2846,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f298,f222])).
fof(f1425,plain,(
  ( ! [X18] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X18))))) )),
  inference(subsumption_resolution,[],[f1414,f277])).
fof(f1414,plain,(
  ( ! [X18] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X18)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18)) )),
  inference(superposition,[],[f277,f262])).
fof(f277,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f5304,plain,(
  ( ! [X8,X7,X9] : (s(X7,happ(s(fun('type.num.num',X7),X8),s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) = s(X7,'const.ind_type.FCONS_3'(s(X7,X9),s(fun('type.num.num',X7),X8),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) ) | $spl17),
  inference(superposition,[],[f292,f1470])).
fof(f5307,plain,(
  ( ! [X19,X17,X18,X16] : (s(X17,happ(s(fun('type.num.num',X17),X18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))))) = s(X17,'const.ind_type.FCONS_3'(s(X17,X19),s(fun('type.num.num',X17),X18),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16)))))))) )),
  inference(superposition,[],[f292,f255])).
fof(f5310,plain,(
  ( ! [X8,X7,X5,X9] : (s(X5,happ(s(fun('type.num.num',X5),X7),s('type.num.num',X8))) = s(X5,'const.bool.COND_3'(s(bool,'F_0'),s(X5,X9),s(X5,happ(s(fun('type.num.num',X5),X7),s('type.num.num',X8)))))) )),
  inference(superposition,[],[f288,f292])).
fof(f5311,plain,(
  ( ! [X14,X12,X10,X13] : (s(X10,X14) = s(X10,'const.bool.COND_3'(s(bool,'F_0'),s(X10,happ(s(fun('type.num.num',X10),X12),s('type.num.num',X13))),s(X10,X14)))) )),
  inference(superposition,[],[f288,f292])).
fof(f5314,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X26),s('type.num.num',X27))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X25),s(fun('type.num.num',bool),X26),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))) )),
  inference(superposition,[],[f261,f292])).
fof(f5643,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',X36))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',X36))))) )),
  inference(forward_demodulation,[],[f5317,f292])).
fof(f5317,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',X36)))) | s(bool,'F_0') != s(bool,'const.ind_type.FCONS_3'(s(bool,X34),s(fun('type.num.num',bool),X35),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36)))))) )),
  inference(superposition,[],[f268,f292])).
fof(f5642,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',X39))) | p(s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',X39))))) )),
  inference(forward_demodulation,[],[f5318,f292])).
fof(f5318,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',X39)))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X37),s(fun('type.num.num',bool),X38),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39)))))) )),
  inference(superposition,[],[f269,f292])).
fof(f5332,plain,(
  ( ! [X85,X86] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X85),s('type.num.num',X86))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f242,f292])).
fof(f5333,plain,(
  ( ! [X88,X89] : (s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X88),s('type.num.num',X89)))))) )),
  inference(superposition,[],[f243,f292])).
fof(f5335,plain,(
  ( ! [X94,X95] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X94),s('type.num.num',X95))))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f245,f292])).
fof(f5336,plain,(
  ( ! [X97,X98] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X97),s('type.num.num',X98))))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f246,f292])).
fof(f5337,plain,(
  ( ! [X101,X100] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X100),s('type.num.num',X101))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X100),s('type.num.num',X101)))))))) )),
  inference(superposition,[],[f247,f292])).
fof(f5338,plain,(
  ( ! [X103,X104] : (s(bool,'const.arithmetic.ODD_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X103),s('type.num.num',X104))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X103),s('type.num.num',X104)))))))) )),
  inference(superposition,[],[f248,f292])).
fof(f5343,plain,(
  ( ! [X118,X119] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X118),s('type.num.num',X119))),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X118),s('type.num.num',X119))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f257,f292])).
fof(f5344,plain,(
  ( ! [X121,X122] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X121),s('type.num.num',X122))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X121),s('type.num.num',X122))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f258,f292])).
fof(f5345,plain,(
  ( ! [X125,X124] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X124),s('type.num.num',X125))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X124),s('type.num.num',X125)))))))) )),
  inference(superposition,[],[f260,f292])).
fof(f5346,plain,(
  ( ! [X127,X128,X126] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X127),s('type.num.num',X128))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X126),s(fun('type.num.num','type.num.num'),X127),s('type.num.num','const.num.SUC_1'(s('type.num.num',X128)))))) )),
  inference(superposition,[],[f272,f292])).
fof(f5638,plain,(
  ( ! [X134,X133] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X133),s('type.num.num',X134))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X133),s('type.num.num',X134))))))) )),
  inference(forward_demodulation,[],[f5348,f292])).
fof(f5348,plain,(
  ( ! [X134,X132,X133] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X133),s('type.num.num',X134)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X132),s(fun('type.num.num','type.num.num'),X133),s('type.num.num','const.num.SUC_1'(s('type.num.num',X134)))))) )),
  inference(superposition,[],[f276,f292])).
fof(f5637,plain,(
  ( ! [X136,X137] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X136),s('type.num.num',X137))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X136),s('type.num.num',X137))))))) )),
  inference(forward_demodulation,[],[f5349,f292])).
fof(f5349,plain,(
  ( ! [X136,X135,X137] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X136),s('type.num.num',X137)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X135),s(fun('type.num.num','type.num.num'),X136),s('type.num.num','const.num.SUC_1'(s('type.num.num',X137)))))) )),
  inference(superposition,[],[f277,f292])).
fof(f5350,plain,(
  ( ! [X140,X141,X139] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X139),s('type.num.num',X140))),s('type.num.num',X141))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X141))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X139),s('type.num.num',X140)))))))) )),
  inference(superposition,[],[f298,f292])).
fof(f5351,plain,(
  ( ! [X144,X145,X143] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X145),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X143),s('type.num.num',X144))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X143),s('type.num.num',X144))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X145)))))) )),
  inference(superposition,[],[f298,f292])).
fof(f5352,plain,(
  ( ! [X149,X147,X148] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X149),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X147),s('type.num.num',X148))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X149))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X147),s('type.num.num',X148)))))))) )),
  inference(superposition,[],[f299,f292])).
fof(f5353,plain,(
  ( ! [X152,X151,X153] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X151),s('type.num.num',X152))),s('type.num.num',X153))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X151),s('type.num.num',X152))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X153)))))) )),
  inference(superposition,[],[f299,f292])).
fof(f5354,plain,(
  ( ! [X156,X157,X155] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X157),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X155),s('type.num.num',X156))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X157))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X155),s('type.num.num',X156)))))))) )),
  inference(superposition,[],[f300,f292])).
fof(f5355,plain,(
  ( ! [X161,X159,X160] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X159),s('type.num.num',X160))),s('type.num.num',X161))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X159),s('type.num.num',X160))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X161)))))) )),
  inference(superposition,[],[f300,f292])).
fof(f5356,plain,(
  ( ! [X165,X163,X164] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X163),s('type.num.num',X164))),s('type.num.num',X165))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X165))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X163),s('type.num.num',X164)))))))) )),
  inference(superposition,[],[f301,f292])).
fof(f5357,plain,(
  ( ! [X167,X169,X168] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X169),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X167),s('type.num.num',X168))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X167),s('type.num.num',X168))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X169)))))) )),
  inference(superposition,[],[f301,f292])).
fof(f5358,plain,(
  ( ! [X173,X171,X172] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X173),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X171),s('type.num.num',X172))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X173))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X171),s('type.num.num',X172)))))))) )),
  inference(superposition,[],[f302,f292])).
fof(f5359,plain,(
  ( ! [X177,X175,X176] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X175),s('type.num.num',X176))),s('type.num.num',X177))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X175),s('type.num.num',X176))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X177)))))) )),
  inference(superposition,[],[f302,f292])).
fof(f5360,plain,(
  ( ! [X180,X181,X179] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X181),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X179),s('type.num.num',X180))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X181))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X179),s('type.num.num',X180)))))))) )),
  inference(superposition,[],[f303,f292])).
fof(f5361,plain,(
  ( ! [X185,X184,X183] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X183),s('type.num.num',X184))),s('type.num.num',X185))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X183),s('type.num.num',X184))))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X185)))))) )),
  inference(superposition,[],[f303,f292])).
fof(f5362,plain,(
  ( ! [X189,X187,X188] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X189),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X187),s('type.num.num',X188))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X189))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X187),s('type.num.num',X188)))))))) )),
  inference(superposition,[],[f304,f292])).
fof(f5363,plain,(
  ( ! [X191,X192,X193] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X191),s('type.num.num',X192))),s('type.num.num',X193))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X191),s('type.num.num',X192))))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X193)))))) )),
  inference(superposition,[],[f304,f292])).
fof(f5636,plain,(
  ( ! [X237,X235,X236] : (s('type.num.num',X237) != s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X235),s('type.num.num',X236))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X235),s('type.num.num',X236))),s('type.num.num',X237))))) )),
  inference(forward_demodulation,[],[f5374,f292])).
fof(f5374,plain,(
  ( ! [X237,X235,X236,X234] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X235),s('type.num.num',X236))),s('type.num.num',X237)))) | s('type.num.num',X237) != s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X234),s(fun('type.num.num','type.num.num'),X235),s('type.num.num','const.num.SUC_1'(s('type.num.num',X236)))))) )),
  inference(superposition,[],[f316,f292])).
fof(f5635,plain,(
  ( ! [X241,X239,X240] : (s('type.num.num',X241) != s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X239),s('type.num.num',X240))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X241),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X239),s('type.num.num',X240))))))) )),
  inference(forward_demodulation,[],[f5375,f292])).
fof(f5375,plain,(
  ( ! [X241,X239,X240,X238] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X241),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X239),s('type.num.num',X240)))))) | s('type.num.num',X241) != s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X238),s(fun('type.num.num','type.num.num'),X239),s('type.num.num','const.num.SUC_1'(s('type.num.num',X240)))))) )),
  inference(superposition,[],[f316,f292])).
fof(f5632,plain,(
  ( ! [X249,X248,X247] : (s('type.num.num',X249) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X249),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248))),s('type.num.num',X249))))) )),
  inference(forward_demodulation,[],[f5631,f292])).
fof(f5631,plain,(
  ( ! [X249,X246,X248,X247] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X249),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248))),s('type.num.num',X249)))) | s('type.num.num',X249) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X246),s(fun('type.num.num','type.num.num'),X247),s('type.num.num','const.num.SUC_1'(s('type.num.num',X248)))))) )),
  inference(forward_demodulation,[],[f5377,f292])).
fof(f5377,plain,(
  ( ! [X249,X246,X248,X247] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X247),s('type.num.num',X248))),s('type.num.num',X249)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X249),s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X246),s(fun('type.num.num','type.num.num'),X247),s('type.num.num','const.num.SUC_1'(s('type.num.num',X248)))))))) | s('type.num.num',X249) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X246),s(fun('type.num.num','type.num.num'),X247),s('type.num.num','const.num.SUC_1'(s('type.num.num',X248)))))) )),
  inference(superposition,[],[f317,f292])).
fof(f5379,plain,(
  ( ! [X255,X257,X254,X256] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X257))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X255),s('type.num.num',X256))))) | s('type.num.num',X257) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X254),s(fun('type.num.num','type.num.num'),X255),s('type.num.num','const.num.SUC_1'(s('type.num.num',X256)))))) )),
  inference(superposition,[],[f318,f292])).
fof(f5630,plain,(
  ( ! [X269,X267,X268] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X267),s('type.num.num',X268))),s('type.num.num',X269)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X269))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X267),s('type.num.num',X268))))))))) )),
  inference(forward_demodulation,[],[f5382,f292])).
fof(f5382,plain,(
  ( ! [X269,X267,X268,X266] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X269))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X267),s('type.num.num',X268)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X266),s(fun('type.num.num','type.num.num'),X267),s('type.num.num','const.num.SUC_1'(s('type.num.num',X268))))),s('type.num.num',X269))))) )),
  inference(superposition,[],[f320,f292])).
fof(f5629,plain,(
  ( ! [X271,X273,X272] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X273),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X271),s('type.num.num',X272)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X271),s('type.num.num',X272))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X273))))))) )),
  inference(forward_demodulation,[],[f5383,f292])).
fof(f5383,plain,(
  ( ! [X271,X273,X270,X272] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X271),s('type.num.num',X272))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X273)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X273),s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X270),s(fun('type.num.num','type.num.num'),X271),s('type.num.num','const.num.SUC_1'(s('type.num.num',X272))))))))) )),
  inference(superposition,[],[f320,f292])).
fof(f5628,plain,(
  ( ! [X277,X275,X276] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X275),s('type.num.num',X276))),s('type.num.num',X277)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X277))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X275),s('type.num.num',X276))))))))) )),
  inference(forward_demodulation,[],[f5384,f292])).
fof(f5384,plain,(
  ( ! [X277,X275,X276,X274] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X277))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X275),s('type.num.num',X276)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X274),s(fun('type.num.num','type.num.num'),X275),s('type.num.num','const.num.SUC_1'(s('type.num.num',X276))))),s('type.num.num',X277))))) )),
  inference(superposition,[],[f321,f292])).
fof(f5627,plain,(
  ( ! [X280,X279,X281] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X281),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X279),s('type.num.num',X280)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X279),s('type.num.num',X280))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X281))))))) )),
  inference(forward_demodulation,[],[f5385,f292])).
fof(f5385,plain,(
  ( ! [X280,X279,X281,X278] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X279),s('type.num.num',X280))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X281)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X281),s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X278),s(fun('type.num.num','type.num.num'),X279),s('type.num.num','const.num.SUC_1'(s('type.num.num',X280))))))))) )),
  inference(superposition,[],[f321,f292])).
fof(f5378,plain,(
  ( ! [X253,X251,X252,X250] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X253))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X251),s('type.num.num',X252))))) | s('type.num.num',X253) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X250),s(fun('type.num.num','type.num.num'),X251),s('type.num.num','const.num.SUC_1'(s('type.num.num',X252)))))) )),
  inference(superposition,[],[f318,f292])).
fof(f5302,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,happ(s(fun('type.num.num',X1),X2),s('type.num.num',sK4(X0)))) = s(X1,'const.ind_type.FCONS_3'(s(X1,X3),s(fun('type.num.num',X1),X2),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f292,f262])).
fof(f5091,plain,(
  ( ! [X134,X136,X135] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X136))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X134),s('type.num.num',X135))))))) | s('type.num.num',X136) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X134))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X135)))))) )),
  inference(superposition,[],[f318,f309])).
fof(f5061,plain,(
  ( ! [X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num',X49))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(superposition,[],[f272,f309])).
fof(f5125,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f5124,f309])).
fof(f5124,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f5038,f255])).
fof(f5038,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f307,f309])).
fof(f5123,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f5122,f309])).
fof(f5122,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f5037,f255])).
fof(f5037,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f308,f309])).
fof(f5121,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f5036,f233])).
fof(f5036,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f309,f994])).
fof(f5120,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(forward_demodulation,[],[f5035,f339])).
fof(f5035,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f309,f222])).
fof(f5119,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f5000,f234])).
fof(f5000,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f309,f994])).
fof(f5118,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f4999,f338])).
fof(f4999,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f309,f222])).
fof(f4900,plain,(
  ( ! [X127,X128,X126] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X128))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X126),s('type.num.num','const.num.SUC_1'(s('type.num.num',X127))))))) | s('type.num.num',X128) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X126),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X126),s('type.num.num',X127)))))) )),
  inference(superposition,[],[f318,f308])).
fof(f4872,plain,(
  ( ! [X47,X46] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X46),s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X46),s('type.num.num',X47)))))) )),
  inference(superposition,[],[f272,f308])).
fof(f4843,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f308,f339])).
fof(f4691,plain,(
  ( ! [X121,X120,X119] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X121))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X119))),s('type.num.num',X120))))) | s('type.num.num',X121) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X119),s('type.num.num',X120))),s('type.num.num',X120)))) )),
  inference(superposition,[],[f318,f307])).
fof(f4665,plain,(
  ( ! [X45,X46] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))),s('type.num.num',X46))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X45),s('type.num.num',X46))),s('type.num.num',X46)))) )),
  inference(superposition,[],[f272,f307])).
fof(f4608,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X3)))) )),
  inference(superposition,[],[f307,f338])).
fof(f4717,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f4606,f339])).
fof(f4606,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f307,f339])).
fof(f4511,plain,(
  ( ! [X114,X112,X113] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X114))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X112),s('type.num.num',X113))))))) | s('type.num.num',X114) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X112))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X113)))))) )),
  inference(superposition,[],[f318,f306])).
fof(f4487,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X44),s('type.num.num',X45))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))))) )),
  inference(superposition,[],[f272,f306])).
fof(f4539,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f4464,f232])).
fof(f4464,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f306,f994])).
fof(f4463,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f306,f222])).
fof(f4538,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f4431,f230])).
fof(f4431,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f306,f994])).
fof(f4430,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f306,f222])).
fof(f4250,plain,(
  ( ! [X107,X108,X106] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X108))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X106),s('type.num.num',X107))))))) | s('type.num.num',X108) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X106))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X107)))))) )),
  inference(superposition,[],[f318,f305])).
fof(f4228,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X44),s('type.num.num',X45))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))))) )),
  inference(superposition,[],[f272,f305])).
fof(f4277,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f4205,f340])).
fof(f4205,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f305,f994])).
fof(f4204,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f305,f222])).
fof(f4174,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f305,f994])).
fof(f4173,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f305,f222])).
fof(f3106,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) | s('type.num.num',X8) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl17),
  inference(superposition,[],[f318,f1999])).
fof(f2991,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f2355,f318])).
fof(f2790,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f2763,f994])).
fof(f2763,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1336,f318])).
fof(f2678,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f2662,f243])).
fof(f2662,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f258,f994])).
fof(f2661,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f258,f222])).
fof(f2414,plain,(
  ( ! [X10,X9] : (s('type.num.num',X9) != s('type.num.num',X10) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X9)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) )),
  inference(superposition,[],[f319,f262])).
fof(f2355,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f2276,f318])).
fof(f2308,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f318,f222])).
fof(f2307,plain,(
  ( ! [X45,X46,X44] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46))) | s('type.num.num',X46) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X44),s(fun('type.num.num','type.num.num'),X45),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f318,f289])).
fof(f2305,plain,(
  ( ! [X39,X38,X40] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X38))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40))) | s('type.num.num',X40) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X38),s('type.num.num',X39)))) )),
  inference(superposition,[],[f318,f287])).
fof(f2303,plain,(
  ( ! [X35] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))) | s('type.num.num',X35) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f318,f221])).
fof(f2302,plain,(
  ( ! [X33,X34] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34))) | s('type.num.num',X34) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X33)))))))) )),
  inference(superposition,[],[f318,f254])).
fof(f2301,plain,(
  ( ! [X31,X32] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))) | s('type.num.num',X32) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X31)))))))) )),
  inference(superposition,[],[f318,f253])).
fof(f2300,plain,(
  ( ! [X30,X29] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X30) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f318,f340])).
fof(f2299,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))) | s('type.num.num',X28) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X27)))) )),
  inference(superposition,[],[f318,f342])).
fof(f2298,plain,(
  ( ! [X26,X25] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26))) | s('type.num.num',X26) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f318,f244])).
fof(f2297,plain,(
  ( ! [X24,X23] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))) | s('type.num.num',X24) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X23)))) )),
  inference(superposition,[],[f318,f241])).
fof(f2296,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))) | s('type.num.num',X22) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X21)))) )),
  inference(superposition,[],[f318,f338])).
fof(f2295,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))) | s('type.num.num',X20) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)))) )),
  inference(superposition,[],[f318,f234])).
fof(f2294,plain,(
  ( ! [X17,X18] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18))) | s('type.num.num',X18) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f318,f339])).
fof(f2293,plain,(
  ( ! [X15,X16] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))) | s('type.num.num',X16) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f318,f233])).
fof(f2291,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))) | s('type.num.num',X12) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) ) | $spl17),
  inference(superposition,[],[f318,f1470])).
fof(f2289,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) | s('type.num.num',X10) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X9)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) )),
  inference(superposition,[],[f318,f262])).
fof(f2287,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) | s('type.num.num',X7) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f318,f341])).
fof(f2286,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f318,f230])).
fof(f2285,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f318,f232])).
fof(f2284,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f318,f994])).
fof(f2276,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))) | $spl17),
  inference(unit_resulting_resolution,[],[f1464,f318])).
fof(f2147,plain,(
  ( ! [X19] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X19)))) )),
  inference(superposition,[],[f272,f342])).
fof(f2020,plain,(
  ( ! [X19] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f272,f341])).
fof(f2000,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f341,f255])).
fof(f1999,plain,(
  s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl17),
  inference(superposition,[],[f341,f1470])).
fof(f1997,plain,(
  ( ! [X0] : (s('type.num.num',sK4(X0)) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f341,f262])).
fof(f1895,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f255,f994])).
fof(f1894,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f255,f222])).
fof(f1816,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f254,f222])).
fof(f1641,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f1604,f221])).
fof(f1604,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f251,f994])).
fof(f1603,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f251,f222])).
fof(f1580,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))) | $spl17),
  inference(unit_resulting_resolution,[],[f1471,f264])).
fof(f1542,plain,(
  ( ! [X66,X65] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X65))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X65),s(fun('type.num.num','type.num.num'),X66),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f272,f289])).
fof(f1469,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',sK1) | $spl17),
  inference(unit_resulting_resolution,[],[f1464,f273])).
fof(f1470,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) | $spl17),
  inference(unit_resulting_resolution,[],[f1464,f262])).
fof(f1412,plain,(
  ( ! [X16] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X16) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X16)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f273,f262])).
fof(f1411,plain,(
  ( ! [X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X15)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X15)) )),
  inference(superposition,[],[f272,f262])).
fof(f1379,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1335,f272])).
fof(f1343,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f1279,f265])).
fof(f1337,plain,(
  ( ! [X4,X5] : (s('type.num.num',X4) != s('type.num.num',X5) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X5)))) )),
  inference(subsumption_resolution,[],[f1278,f336])).
fof(f1278,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s('type.num.num',X4) != s('type.num.num',X5) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X5)))) )),
  inference(superposition,[],[f316,f261])).
fof(f1336,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3)))) )),
  inference(subsumption_resolution,[],[f1277,f336])).
fof(f1277,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3)))) )),
  inference(superposition,[],[f316,f246])).
fof(f1335,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))) )),
  inference(subsumption_resolution,[],[f1276,f336])).
fof(f1276,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))) )),
  inference(superposition,[],[f316,f245])).
fof(f1279,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0))))) )),
  inference(forward_demodulation,[],[f1245,f288])).
fof(f1245,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X1),s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f288,f316])).
fof(f1187,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(subsumption_resolution,[],[f1142,f336])).
fof(f1142,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f277,f261])).
fof(f1043,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f1017,f348])).
fof(f1017,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f247,f994])).
fof(f958,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X9),s('type.num.num',X10)))) )),
  inference(superposition,[],[f272,f287])).
fof(f955,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f272,f244])).
fof(f954,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) )),
  inference(superposition,[],[f272,f241])).
fof(f953,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X5)))) )),
  inference(superposition,[],[f272,f338])).
fof(f951,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f272,f339])).
fof(f949,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f272,f232])).
fof(f768,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f288,f261])).
fof(f545,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f527,f336])).
fof(f527,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f311,f261])).
fof(f496,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f261,f261])).
fof(f345,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f228,f264])).
fof(f322,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.ind_type.FCONS_DEST',aHL_EXT)).
fof(f319,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f211])).
fof(f314,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f310,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X7,X11] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X11))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|2')).
fof(f309,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X7,X11] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X11))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|5')).
fof(f308,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11,X7] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num',X7)))))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.MULT_SUC')).
fof(f307,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11,X7] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X7))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num',X7))),s('type.num.num',X7)))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.MULT|split|1')).
fof(f306,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X7,X11] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X11))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|8')).
fof(f305,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X7,X11] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num',X11))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|12')).
fof(f295,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | ~p(s(bool,X2)) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X19,X20,X16,X21,X17,X22] : ((s(bool,X19) = s(bool,X20) & (p(s(bool,X20)) => s(X6,X16) = s(X6,X21)) & (~p(s(bool,X20)) => s(X6,X17) = s(X6,X22))) => s(X6,'const.bool.COND_3'(s(bool,X19),s(X6,X16),s(X6,X17))) = s(X6,'const.bool.COND_3'(s(bool,X20),s(X6,X21),s(X6,X22))))),
  file('thm.ind_type.FCONS_DEST','thm.bool.COND_CONG')).
fof(f296,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X5) != s(X0,X6) | p(s(bool,X2)) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f182])).
fof(f297,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X5) != s(X0,X6) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f182])).
fof(f293,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X6] : ! [X23,X24,X25] : s(X6,'const.ind_type.FCONS_3'(s(X6,X23),s(fun('type.num.num',X6),X24),s('type.num.num',X25))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(fun('type.num.num',X6),fun('type.num.num',X6)),happ(s(fun(X6,fun(fun('type.num.num',X6),fun('type.num.num',X6))),'const.ind_type.FCONS_0'),s(X6,X23))),s(fun('type.num.num',X6),X24))),s('type.num.num',X25)))),
  file('thm.ind_type.FCONS_DEST','ARITY_0_3_const.ind_type.FCONS')).
fof(f292,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6] : ! [X9,X8,X7] : s(X6,happ(s(fun('type.num.num',X6),X8),s('type.num.num',X7))) = s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))),
  file('thm.ind_type.FCONS_DEST','thm.ind_type.FCONS|split|1')).
fof(f342,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f259,f222])).
fof(f259,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X7] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X7)))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.SUC_ONE_ADD')).
fof(f255,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X7] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|14')).
fof(f254,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7)))))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|9')).
fof(f253,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7)))))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|10')).
fof(f340,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f252,f222])).
fof(f252,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X7] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|11')).
fof(f251,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X7] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X7))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|16')).
fof(f339,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f250,f222])).
fof(f250,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11] : s('type.num.num',X11) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.MULT_RIGHT_|49|')).
fof(f338,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f249,f222])).
fof(f249,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X11] : s('type.num.num',X11) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X11)))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.MULT_LEFT_|49|')).
fof(f244,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|1')).
fof(f241,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.ADD|split|0')).
fof(f234,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X11] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.MULT_CLAUSES|split|0')).
fof(f233,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0')))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.MULT_|48|')).
fof(f232,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11] : s('type.num.num',X11) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0')))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.SUB_|48||split|1')).
fof(f231,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X7] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X7)))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_lte|split|0')).
fof(f229,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))),
  file('thm.ind_type.FCONS_DEST','thm.arithmetic.ZERO_LESS_EQ')).
fof(f228,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f76])).
fof(f76,axiom,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.ind_type.FCONS_DEST','thm.numeral.numeral_distrib|split|34')).
fof(f221,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f57])).
fof(f57,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.ind_type.FCONS_DEST','thm.prim_rec.PRE|split|0')).
fof(f6162,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122,f6126,f6130,f6134,f6138,f6145,f6147,f6149,f6151,f6155,f6161])).
fof(f6160,plain,(
  $false | ($spl0 | $spl17 | $spl41)),
  inference(subsumption_resolution,[],[f6064,f218])).
fof(f6064,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl41)),
  inference(backward_demodulation,[],[f5968,f5145])).
fof(f5145,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl41),
  inference(unit_resulting_resolution,[],[f4737,f320])).
fof(f4737,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))))))))) | $spl41),
  inference(unit_resulting_resolution,[],[f4563,f321])).
fof(f4563,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl41),
  inference(forward_demodulation,[],[f4560,f3170])).
fof(f4560,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))))))))) | $spl41),
  inference(unit_resulting_resolution,[],[f4378,f276])).
fof(f6159,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6158,f218])).
fof(f6158,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6157,f5968])).
fof(f6157,plain,(
  ( ! [X271,X273,X272] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X273),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X271),s('type.num.num',X272))))))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6059,f218])).
fof(f6059,plain,(
  ( ! [X271,X273,X272] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X273),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X271),s('type.num.num',X272))))))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f5629])).
fof(f6156,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122,f6126,f6130,f6134,f6138,f6145,f6147,f6149,f6151,f6155])).
fof(f6152,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122,f6126,f6130,f6134,f6138,f6145,f6147,f6149,f6151])).
fof(f6150,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122,f6126,f6130,f6134,f6138,f6145,f6147,f6149])).
fof(f6148,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122,f6126,f6130,f6134,f6138,f6145,f6147])).
fof(f6146,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122,f6126,f6130,f6134,f6138,f6145])).
fof(f6142,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6141,f218])).
fof(f6141,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6140,f5968])).
fof(f6140,plain,(
  ( ! [X269,X267,X268] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X269))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X267),s('type.num.num',X268))))))))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6047,f218])).
fof(f6047,plain,(
  ( ! [X269,X267,X268] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X269))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X267),s('type.num.num',X268))))))))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f5630])).
fof(f6139,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122,f6126,f6130,f6134,f6138])).
fof(f6135,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122,f6126,f6130,f6134])).
fof(f6131,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122,f6126,f6130])).
fof(f6127,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122,f6126])).
fof(f6123,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118,f6122])).
fof(f6119,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111,f6118])).
fof(f6115,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6114,f218])).
fof(f6114,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6113,f5968])).
fof(f6113,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6023,f218])).
fof(f6023,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f320])).
fof(f6112,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002,f6111])).
fof(f6108,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6107,f218])).
fof(f6107,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f6106,f5968])).
fof(f6106,plain,(
  ( ! [X35,X36] : (~p(s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',X36))))) ) | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6009,f5968])).
fof(f6009,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',X36))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',X36))))) ) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f5643])).
fof(f6105,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001,f6002])).
fof(f6104,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986,f6001])).
fof(f6103,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982,f5986])).
fof(f6102,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975,f5982])).
fof(f6101,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f5976,f218])).
fof(f5976,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17)),
  inference(backward_demodulation,[],[f5968,f336])).
fof(f6100,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098,f5975])).
fof(f6099,plain,(
  $false | ($spl0 | $spl17)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f244,f338,f339,f251,f340,f253,f254,f255,f341,f342,f261,f262,f265,f264,f273,f272,f287,f289,f292,f293,f297,f296,f295,f305,f306,f307,f308,f309,f310,f314,f311,f319,f318,f322,f345,f344,f348,f496,f545,f768,f949,f951,f953,f954,f955,f956,f958,f960,f994,f1043,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1411,f1412,f1464,f1471,f1470,f1469,f1542,f1580,f1603,f1641,f1760,f1816,f1894,f1895,f1997,f1999,f2000,f2020,f2147,f2276,f2284,f2285,f2286,f2287,f2289,f2291,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2307,f2308,f2355,f2414,f2661,f2678,f2790,f2991,f3094,f3106,f4173,f4174,f4204,f4277,f4228,f4250,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5378,f5627,f5628,f5629,f5630,f5379,f5632,f5635,f5636,f5363,f5362,f5361,f5360,f5359,f5358,f5357,f5356,f5355,f5354,f5353,f5352,f5351,f5350,f5637,f5638,f5346,f5345,f5344,f5343,f5338,f5337,f5336,f5335,f5333,f5332,f5642,f5643,f5314,f5311,f5310,f5307,f5304,f645,f247,f670,f248,f277,f276,f1425,f2846,f3269,f258,f2730,f2747,f260,f2818,f3298,f298,f302,f2952,f3169,f3757,f3613,f2543,f3170,f4079,f4081,f3948,f257,f3641,f3785,f4084,f4083,f4082,f3910,f3882,f3861,f3572,f3409,f3045,f2955,f1416,f1415,f324,f323,f321,f320,f317,f316,f304,f303,f301,f300,f299,f2924,f3198,f3571,f3408,f3300,f3200,f3044,f2954,f2863,f2848,f2732,f2663,f2558,f2542,f2306,f1579,f1514,f1042,f959,f816,f733,f671,f646,f363,f336,f288,f269,f268,f246,f245,f243,f242,f332,f5649,f5650,f5651,f5652,f5653,f5659,f5660,f5688,f5663,f5665,f5689,f5683,f5740,f5797,f5798,f5843,f5968,f6098])).
fof(f5964,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f5963,f218])).
fof(f5963,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f5962,f236])).
fof(f5962,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | ($spl0 | $spl17)),
  inference(forward_demodulation,[],[f5950,f3170])).
fof(f5950,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17)),
  inference(unit_resulting_resolution,[],[f5652,f268])).
fof(f5877,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f287,f5692])).
fof(f5692,plain,(
  ( ! [X4] : (s(sK0,sK3) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f5691,f289])).
fof(f5691,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f5685,f230])).
fof(f5685,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num','const.num.|48|_0'))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | ($spl0 | $spl2 | $spl4)),
  inference(backward_demodulation,[],[f332,f5647])).
fof(f5647,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f5646,f343])).
fof(f343,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | p(s(bool,X4))) ) | $spl4),
  inference(forward_demodulation,[],[f335,f222])).
fof(f5890,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f5889])).
fof(f5889,plain,(
  s(sK0,sK3) != s(sK0,sK3) | ($spl0 | $spl2 | $spl4)),
  inference(superposition,[],[f5692,f287])).
fof(f5626,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399,f5307,f5551,f5553,f5346,f5582,f5604,f5609,f5378,f5379,f5625])).
fof(f5625,plain,(
  ( ! [X304,X303,X305] : (s('type.num.num',sK1) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X303),s(fun('type.num.num','type.num.num'),X304),s('type.num.num','const.num.SUC_1'(s('type.num.num',X305))))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X304),s('type.num.num',X305)))))) ) | $spl1),
  inference(forward_demodulation,[],[f5393,f1394])).
fof(f5393,plain,(
  ( ! [X304,X303,X305] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X304),s('type.num.num',X305))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(sK1)))) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X303),s(fun('type.num.num','type.num.num'),X304),s('type.num.num','const.num.SUC_1'(s('type.num.num',X305)))))) ) | $spl1),
  inference(superposition,[],[f2290,f292])).
fof(f5609,plain,(
  ( ! [X249,X246,X248,X247] : (s('type.num.num',X249) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X246),s(fun('type.num.num','type.num.num'),X247),s('type.num.num','const.num.SUC_1'(s('type.num.num',X248)))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5608,f218])).
fof(f5608,plain,(
  ( ! [X249,X246,X248,X247] : (~p(s(bool,'T_0')) | s('type.num.num',X249) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X246),s(fun('type.num.num','type.num.num'),X247),s('type.num.num','const.num.SUC_1'(s('type.num.num',X248)))))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5607,f5400])).
fof(f5607,plain,(
  ( ! [X249,X246,X248,X247] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X249),s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X246),s(fun('type.num.num','type.num.num'),X247),s('type.num.num','const.num.SUC_1'(s('type.num.num',X248)))))))) | s('type.num.num',X249) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X246),s(fun('type.num.num','type.num.num'),X247),s('type.num.num','const.num.SUC_1'(s('type.num.num',X248)))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5606,f218])).
fof(f5606,plain,(
  ( ! [X249,X246,X248,X247] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X249),s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X246),s(fun('type.num.num','type.num.num'),X247),s('type.num.num','const.num.SUC_1'(s('type.num.num',X248)))))))) | s('type.num.num',X249) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X246),s(fun('type.num.num','type.num.num'),X247),s('type.num.num','const.num.SUC_1'(s('type.num.num',X248)))))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5377,f5400])).
fof(f5604,plain,(
  ( ! [X243,X244,X242,X245] : (s('type.num.num',X245) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X242),s(fun('type.num.num','type.num.num'),X243),s('type.num.num','const.num.SUC_1'(s('type.num.num',X244)))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5603,f218])).
fof(f5603,plain,(
  ( ! [X243,X244,X242,X245] : (~p(s(bool,'T_0')) | s('type.num.num',X245) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X242),s(fun('type.num.num','type.num.num'),X243),s('type.num.num','const.num.SUC_1'(s('type.num.num',X244)))))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5602,f5400])).
fof(f5602,plain,(
  ( ! [X243,X244,X242,X245] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X242),s(fun('type.num.num','type.num.num'),X243),s('type.num.num','const.num.SUC_1'(s('type.num.num',X244))))),s('type.num.num',X245)))) | s('type.num.num',X245) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X242),s(fun('type.num.num','type.num.num'),X243),s('type.num.num','const.num.SUC_1'(s('type.num.num',X244)))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5601,f218])).
fof(f5601,plain,(
  ( ! [X243,X244,X242,X245] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X242),s(fun('type.num.num','type.num.num'),X243),s('type.num.num','const.num.SUC_1'(s('type.num.num',X244))))),s('type.num.num',X245)))) | s('type.num.num',X245) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X242),s(fun('type.num.num','type.num.num'),X243),s('type.num.num','const.num.SUC_1'(s('type.num.num',X244)))))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5376,f5400])).
fof(f5376,plain,(
  ( ! [X243,X244,X242,X245] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X245),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X243),s('type.num.num',X244)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X242),s(fun('type.num.num','type.num.num'),X243),s('type.num.num','const.num.SUC_1'(s('type.num.num',X244))))),s('type.num.num',X245)))) | s('type.num.num',X245) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X242),s(fun('type.num.num','type.num.num'),X243),s('type.num.num','const.num.SUC_1'(s('type.num.num',X244)))))) )),
  inference(superposition,[],[f317,f292])).
fof(f5582,plain,(
  ( ! [X134,X132,X133] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X132),s(fun('type.num.num','type.num.num'),X133),s('type.num.num','const.num.SUC_1'(s('type.num.num',X134)))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5581,f218])).
fof(f5581,plain,(
  ( ! [X134,X132,X133] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X132),s(fun('type.num.num','type.num.num'),X133),s('type.num.num','const.num.SUC_1'(s('type.num.num',X134)))))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5348,f5400])).
fof(f5553,plain,(
  ( ! [X14,X12,X10,X13] : (s(X10,X14) = s(X10,'const.bool.COND_3'(s(bool,'T_0'),s(X10,happ(s(fun('type.num.num',X10),X12),s('type.num.num',X13))),s(X10,X14)))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5311,f5400])).
fof(f5551,plain,(
  ( ! [X8,X7,X5,X9] : (s(X5,happ(s(fun('type.num.num',X5),X7),s('type.num.num',X8))) = s(X5,'const.bool.COND_3'(s(bool,'T_0'),s(X5,X9),s(X5,happ(s(fun('type.num.num',X5),X7),s('type.num.num',X8)))))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5310,f5400])).
fof(f5399,plain,(
  ( ! [X4] : (s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num',sK4(sK1)))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num',sK4(sK1)))))) | p(s(bool,X4))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5303,f2042])).
fof(f5543,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5488,f218])).
fof(f5488,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f276])).
fof(f5540,plain,(
  ( ! [X10] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5539,f218])).
fof(f5539,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5538,f5400])).
fof(f5538,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5472,f218])).
fof(f5472,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f3613])).
fof(f5536,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5469,f218])).
fof(f5469,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f4079])).
fof(f5534,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5468,f218])).
fof(f5468,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f4081])).
fof(f5532,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5467,f218])).
fof(f5467,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f3948])).
fof(f5530,plain,(
  ( ! [X10] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5529,f218])).
fof(f5529,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5528,f5400])).
fof(f5528,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5465,f218])).
fof(f5465,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f3641])).
fof(f5526,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5525,f218])).
fof(f5525,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5524,f5400])).
fof(f5524,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9)))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5463,f218])).
fof(f5463,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9)))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f4084])).
fof(f5522,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5521,f218])).
fof(f5521,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5520,f5400])).
fof(f5520,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5462,f218])).
fof(f5462,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f4083])).
fof(f5518,plain,(
  ( ! [X6,X5] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5517,f218])).
fof(f5517,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5516,f5400])).
fof(f5516,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6)))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5461,f218])).
fof(f5461,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6)))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f4082])).
fof(f5514,plain,(
  ( ! [X10,X11] : (s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5513,f218])).
fof(f5513,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5512,f5400])).
fof(f5512,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))),s('type.num.num',X11)))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5460,f218])).
fof(f5460,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))),s('type.num.num',X11)))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f3910])).
fof(f5510,plain,(
  ( ! [X10,X11] : (s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5509,f218])).
fof(f5509,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5508,f5400])).
fof(f5508,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10))))))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5459,f218])).
fof(f5459,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10))))))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f3882])).
fof(f5503,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5502,f218])).
fof(f5502,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5501,f5400])).
fof(f5501,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5447,f218])).
fof(f5447,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f317])).
fof(f5421,plain,(
  ( ! [X43,X41,X42] : (s('type.num.num',X43) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X41),s('type.num.num',X42))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43)))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f2306])).
fof(f5414,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X11),s('type.num.num',X12))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f959])).
fof(f5407,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f288])).
fof(f5495,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5405,f5400])).
fof(f5405,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f268])).
fof(f5400,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5398,f768])).
fof(f5398,plain,(
  ( ! [X1] : (s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num',sK4(sK1)))) != s(sK0,'const.bool.COND_3'(s(bool,X1),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num',sK4(sK1)))))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5303,f2209])).
fof(f5303,plain,(
  ( ! [X6,X4,X5] : (s(X4,happ(s(fun('type.num.num',X4),X5),s('type.num.num',sK4(sK1)))) = s(X4,'const.ind_type.FCONS_3'(s(X4,X6),s(fun('type.num.num',X4),X5),s('type.num.num',sK1)))) ) | $spl1),
  inference(superposition,[],[f292,f1394])).
fof(f4299,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2514,f318])).
fof(f2514,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2353,f318])).
fof(f2353,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2274,f318])).
fof(f2290,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(sK1))))) ) | $spl1),
  inference(superposition,[],[f318,f1394])).
fof(f2288,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))) | s('type.num.num',X8) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl1),
  inference(superposition,[],[f318,f1998])).
fof(f2274,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) | $spl1),
  inference(unit_resulting_resolution,[],[f333,f318])).
fof(f2209,plain,(
  ( ! [X1] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X1),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num',sK4(sK1)))))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f2184,f336])).
fof(f2184,plain,(
  ( ! [X1] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X1),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num',sK4(sK1)))))) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X1)) ) | ($spl1 | $spl4)),
  inference(superposition,[],[f2042,f261])).
fof(f2206,plain,(
  s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num',sK4(sK1)))) != s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f2182,f288])).
fof(f2182,plain,(
  s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,'F_0'),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num',sK4(sK1)))))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f336,f2042])).
fof(f2042,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num',sK4(sK1)))))) | p(s(bool,X4))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1998,f343])).
fof(f1998,plain,(
  s('type.num.num',sK4(sK1)) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl1),
  inference(superposition,[],[f341,f1394])).
fof(f1394,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(sK1)))) | $spl1),
  inference(unit_resulting_resolution,[],[f333,f262])).
fof(f1092,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1))) | $spl1),
  inference(unit_resulting_resolution,[],[f1063,f269])).
fof(f1093,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1))) | $spl1),
  inference(unit_resulting_resolution,[],[f1063,f264])).
fof(f1063,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1)))) | $spl1),
  inference(unit_resulting_resolution,[],[f333,f276])).
fof(f5622,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5621,f218])).
fof(f5621,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5620,f5400])).
fof(f5620,plain,(
  ( ! [X271,X273,X270,X272] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X273),s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X270),s(fun('type.num.num','type.num.num'),X271),s('type.num.num','const.num.SUC_1'(s('type.num.num',X272))))))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5619,f218])).
fof(f5619,plain,(
  ( ! [X271,X273,X270,X272] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X273),s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X270),s(fun('type.num.num','type.num.num'),X271),s('type.num.num','const.num.SUC_1'(s('type.num.num',X272))))))))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5383,f5400])).
fof(f5618,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5617,f218])).
fof(f5617,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5616,f5400])).
fof(f5616,plain,(
  ( ! [X269,X267,X268,X266] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X266),s(fun('type.num.num','type.num.num'),X267),s('type.num.num','const.num.SUC_1'(s('type.num.num',X268))))),s('type.num.num',X269))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5615,f218])).
fof(f5615,plain,(
  ( ! [X269,X267,X268,X266] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X266),s(fun('type.num.num','type.num.num'),X267),s('type.num.num','const.num.SUC_1'(s('type.num.num',X268))))),s('type.num.num',X269))))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5382,f5400])).
fof(f5612,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399,f5307,f5551,f5553,f5346,f5582,f5604,f5609,f5378,f5379])).
fof(f5611,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399,f5307,f5551,f5553,f5346,f5582,f5604,f5609,f5378])).
fof(f5610,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399,f5307,f5551,f5553,f5346,f5582,f5604,f5609])).
fof(f5605,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399,f5307,f5551,f5553,f5346,f5582,f5604])).
fof(f5583,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399,f5307,f5551,f5553,f5346,f5582])).
fof(f5579,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399,f5307,f5551,f5553,f5346])).
fof(f5554,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399,f5307,f5551,f5553])).
fof(f5552,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399,f5307,f5551])).
fof(f5550,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399,f5307])).
fof(f5549,plain,(
  $false | ($spl1 | $spl4 | $spl21)),
  inference(global_subsumption,[],[f5306,f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399])).
fof(f5306,plain,(
  ( ! [X14,X15,X13] : (s(X13,happ(s(fun('type.num.num',X13),X14),s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))))) = s(X13,'const.ind_type.FCONS_3'(s(X13,X15),s(fun('type.num.num',X13),X14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))))) ) | $spl21),
  inference(superposition,[],[f292,f1790])).
fof(f5548,plain,(
  $false | ($spl1 | $spl4 | $spl25)),
  inference(global_subsumption,[],[f5305,f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399])).
fof(f5305,plain,(
  ( ! [X12,X10,X11] : (s(X10,happ(s(fun('type.num.num',X10),X11),s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))))) = s(X10,'const.ind_type.FCONS_3'(s(X10,X12),s(fun('type.num.num',X10),X11),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))))) ) | $spl25),
  inference(superposition,[],[f292,f2099])).
fof(f2099,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))))) | $spl25),
  inference(unit_resulting_resolution,[],[f2092,f262])).
fof(f5547,plain,(
  $false | ($spl1 | $spl4 | $spl17)),
  inference(global_subsumption,[],[f5304,f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399])).
fof(f5546,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543,f5399])).
fof(f5545,plain,(
  $false | ($spl1 | $spl4 | $spl13)),
  inference(subsumption_resolution,[],[f5490,f218])).
fof(f5490,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4 | $spl13)),
  inference(backward_demodulation,[],[f5400,f1064])).
fof(f5544,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540,f5543])).
fof(f5542,plain,(
  $false | ($spl1 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f5476,f218])).
fof(f5476,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4 | $spl41)),
  inference(backward_demodulation,[],[f5400,f5145])).
fof(f5541,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536,f5540])).
fof(f5537,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534,f5536])).
fof(f5535,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532,f5534])).
fof(f5533,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530,f5532])).
fof(f5531,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526,f5530])).
fof(f5527,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522,f5526])).
fof(f5523,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518,f5522])).
fof(f5519,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514,f5518])).
fof(f5515,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510,f5514])).
fof(f5511,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503,f5510])).
fof(f5507,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5506,f218])).
fof(f5506,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f5505,f5400])).
fof(f5505,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5448,f218])).
fof(f5448,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f320])).
fof(f5504,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421,f5503])).
fof(f5500,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414,f5421])).
fof(f5499,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407,f5414])).
fof(f5498,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5408,f218])).
fof(f5408,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f5400,f336])).
fof(f5497,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495,f5407])).
fof(f5496,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f217,f216,f218,f220,f221,f222,f336,f228,f229,f230,f231,f232,f233,f234,f235,f236,f241,f242,f243,f244,f245,f246,f247,f248,f338,f339,f251,f340,f253,f254,f255,f341,f257,f258,f342,f260,f261,f262,f265,f264,f269,f268,f273,f272,f277,f276,f287,f288,f289,f292,f293,f297,f296,f295,f298,f299,f300,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f314,f311,f317,f316,f319,f318,f321,f320,f322,f323,f324,f343,f333,f345,f344,f348,f363,f496,f545,f645,f646,f670,f671,f733,f768,f816,f949,f951,f953,f954,f955,f956,f958,f959,f960,f994,f1042,f1043,f1063,f1093,f1092,f1187,f1279,f1335,f1336,f1337,f1343,f1379,f1394,f1411,f1412,f1425,f1415,f1416,f1443,f1514,f1542,f1603,f1641,f1816,f1894,f1895,f1997,f1998,f2042,f2000,f2020,f2069,f2147,f2206,f2209,f2274,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2293,f2294,f2295,f2296,f2297,f2298,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2307,f2308,f2353,f2414,f2514,f2558,f2542,f2543,f2661,f2678,f2663,f2730,f2747,f2732,f2790,f2818,f2863,f2846,f2848,f2924,f2952,f2954,f2955,f3044,f3045,f3169,f3170,f3198,f3200,f3269,f3298,f3300,f3408,f3409,f3571,f3572,f3613,f3641,f3757,f3785,f3861,f3948,f3882,f3910,f4079,f4081,f4082,f4083,f4084,f4173,f4174,f4204,f4277,f4228,f4250,f4299,f4430,f4538,f4463,f4539,f4487,f4511,f4717,f4608,f4665,f4691,f4843,f4872,f4900,f5118,f5119,f5120,f5121,f5123,f5125,f5061,f5091,f5302,f5303,f5400,f5495])).
fof(f4376,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl38),
  inference(cnf_transformation,[],[f4376_D])).
fof(f4376_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4378,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))))) | $spl41),
  inference(cnf_transformation,[],[f4378_D])).
fof(f4378_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f4326,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl25),
  inference(superposition,[],[f272,f2572])).
fof(f2572,plain,(
  s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl25),
  inference(superposition,[],[f341,f2099])).
fof(f3506,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))))) | $spl34),
  inference(cnf_transformation,[],[f3506_D])).
fof(f3506_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f3508,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) | $spl37),
  inference(cnf_transformation,[],[f3508_D])).
fof(f3508_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f3463,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))))) | $spl21),
  inference(superposition,[],[f272,f1790])).
fof(f3127,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl30),
  inference(cnf_transformation,[],[f3127_D])).
fof(f3127_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3129,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) | $spl33),
  inference(cnf_transformation,[],[f3129_D])).
fof(f3129_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f3094,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl17),
  inference(superposition,[],[f272,f1999])).
fof(f2621,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))))) | $spl26),
  inference(cnf_transformation,[],[f2621_D])).
fof(f2621_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2623,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))))) | $spl29),
  inference(cnf_transformation,[],[f2623_D])).
fof(f2623_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f2593,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1))))))) | $spl25),
  inference(superposition,[],[f272,f2099])).
fof(f2090,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl22),
  inference(cnf_transformation,[],[f2090_D])).
fof(f2090_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2092,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))) | $spl25),
  inference(cnf_transformation,[],[f2092_D])).
fof(f2092_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f2069,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4(sK1)))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl1),
  inference(superposition,[],[f272,f1998])).
fof(f1782,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) | $spl18),
  inference(cnf_transformation,[],[f1782_D])).
fof(f1782_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1784,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))) | $spl21),
  inference(cnf_transformation,[],[f1784_D])).
fof(f1784_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1760,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4('const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))))) | $spl17),
  inference(superposition,[],[f272,f1470])).
fof(f1462,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(sK1)))) | $spl14),
  inference(cnf_transformation,[],[f1462_D])).
fof(f1462_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(sK1)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1464,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) | $spl17),
  inference(cnf_transformation,[],[f1464_D])).
fof(f1464_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f1443,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4(sK1)))) | $spl1),
  inference(superposition,[],[f272,f1394])).
fof(f970,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl10),
  inference(cnf_transformation,[],[f970_D])).
fof(f970_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f972,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl13),
  inference(cnf_transformation,[],[f972_D])).
fof(f972_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f960,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f272,f222])).
fof(f965,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(cnf_transformation,[],[f965_D])).
fof(f965_D,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f967,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl9),
  inference(cnf_transformation,[],[f967_D])).
fof(f967_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f956,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f272,f340])).
fof(f332,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | $spl0),
  inference(cnf_transformation,[],[f332_D])).
fof(f332_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f335,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,X4))) ) | $spl4),
  inference(cnf_transformation,[],[f335_D])).
fof(f335_D,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,X4))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f216,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | p(s(bool,X4))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X4] : (((p(s(bool,X4)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1)) & (~p(s(bool,X4)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1))) | s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f190])).
fof(f190,plain,(
  ? [X0,X1,X2,X3] : ! [X4] : (((p(s(bool,X4)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & (~p(s(bool,X4)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) | s(X0,'const.ind_type.FCONS_3'(s(X0,X3),s(fun('type.num.num',X0),X2),s('type.num.num',X1))) != s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X3),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(nnf_transformation,[],[f178])).
fof(f178,plain,(
  ? [X0,X1,X2,X3] : ! [X4] : ((p(s(bool,X4)) <~> s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | s(X0,'const.ind_type.FCONS_3'(s(X0,X3),s(fun('type.num.num',X0),X2),s('type.num.num',X1))) != s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X3),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ~! [X0,X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & s(X0,'const.ind_type.FCONS_3'(s(X0,X3),s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X3),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ~! [X0] : ! [X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & s(X0,'const.ind_type.FCONS_3'(s(X0,X3),s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X3),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ? [X10] : ((p(s(bool,X10)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) & s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X8),s('type.num.num',X7))) = s(X6,'const.bool.COND_3'(s(bool,X10),s(X6,X9),s(X6,happ(s(fun('type.num.num',X6),X8),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ? [X10] : ((p(s(bool,X10)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) & s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X8),s('type.num.num',X7))) = s(X6,'const.bool.COND_3'(s(bool,X10),s(X6,X9),s(X6,happ(s(fun('type.num.num',X6),X8),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  file('thm.ind_type.FCONS_DEST','thm.ind_type.FCONS_DEST')).
fof(f333,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) | $spl1),
  inference(cnf_transformation,[],[f333_D])).
fof(f333_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f334,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,X4))) ) | $spl2),
  inference(cnf_transformation,[],[f334_D])).
fof(f334_D,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,X4))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f217,plain,(
  ( ! [X4] : (s(sK0,'const.ind_type.FCONS_3'(s(sK0,sK3),s(fun('type.num.num',sK0),sK2),s('type.num.num',sK1))) != s(sK0,'const.bool.COND_3'(s(bool,X4),s(sK0,sK3),s(sK0,happ(s(fun('type.num.num',sK0),sK2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) | ~p(s(bool,X4))) )),
  inference(cnf_transformation,[],[f191])).
% SZS output end Proof for thm.ind_type
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 194
% Passive clauses: 739
% Generated clauses: 7770
% Final active clauses: 83
% Final passive clauses: 214
% Input formulas: 84
% Initial clauses: 115
% 
% Duplicate literals: 1
% Trivial inequalities: 26
% Fw subsumption resolutions: 222
% Fw demodulations: 1475
% Bw demodulations: 407
% Condensations: 1
% Global subsumptions: 50
% 
% Simple tautologies: 23
% Equational tautologies: 37
% Forward subsumptions: 5000
% Fw demodulations to eq. taut.: 88
% 
% Unit resulting resolution: 591
% Factoring: 1
% Forward superposition: 2403
% Backward superposition: 2131
% Self superposition: 38
% Equality factoring: 2
% Equality resolution: 6
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 914
% SAT solver unit clauses: 227
% SAT solver binary clauses: 356
% 
% Sat splits: 15
% Sat splitting refutations: 93
% 
% Memory used [KB]: 6268
% Time elapsed: 0.194 s
% ------------------------------
% ------------------------------
% Success in time 1.514 s
