Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (24977)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1387
% Passive clauses: 8153
% Generated clauses: 39575
% Final active clauses: 1319
% Final passive clauses: 6461
% Input formulas: 261
% Initial clauses: 518
% 
% Duplicate literals: 186
% Trivial inequalities: 121
% Fw subsumption resolutions: 2417
% Fw demodulations: 3321
% Bw demodulations: 54
% 
% Simple tautologies: 264
% Equational tautologies: 32
% Forward subsumptions: 24811
% 
% Binary resolution: 8328
% Factoring: 633
% Forward superposition: 11633
% Backward superposition: 12160
% Self superposition: 46
% Equality factoring: 2
% Equality resolution: 26
% 
% Split clauses: 326
% Split components: 654
% SAT solver clauses: 547
% SAT solver unit clauses: 46
% SAT solver binary clauses: 98
% 
% Sat splits: 326
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 547
% TWLsolver calls for satisfiability: 449
% 
% Memory used [KB]: 18933
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (24991)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 708
% Passive clauses: 4048
% Generated clauses: 28134
% Final active clauses: 563
% Final passive clauses: 2282
% Input formulas: 261
% Initial clauses: 518
% 
% Duplicate literals: 101
% Trivial inequalities: 87
% Fw subsumption resolutions: 1212
% Fw demodulations: 3761
% Bw demodulations: 1165
% Condensations: 62
% Global subsumptions: 34
% 
% Simple tautologies: 285
% Equational tautologies: 16
% Forward subsumptions: 13253
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 10
% 
% Binary resolution: 2777
% Factoring: 112
% Forward superposition: 6684
% Backward superposition: 6240
% Self superposition: 61
% Equality factoring: 2
% Equality resolution: 24
% 
% Split clauses: 83
% Split components: 166
% SAT solver clauses: 4634
% SAT solver unit clauses: 676
% SAT solver binary clauses: 188
% SAT solver learnt clauses: 69
% SAT solver learnt literals: 12
% 
% Sat splits: 83
% Sat splitting refutations: 34
% 
% TWLsolver clauses: 4497
% TWLsolver calls for satisfiability: 5080
% 
% Memory used [KB]: 22387
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% (25008)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 909
% Passive clauses: 9846
% Generated clauses: 106005
% Final active clauses: 863
% Final passive clauses: 2014
% Input formulas: 261
% Initial clauses: 518
% Discarded non-redundant clauses: 2655
% 
% Duplicate literals: 10
% Trivial inequalities: 63
% Fw subsumption resolutions: 2594
% Fw demodulations: 27756
% Bw demodulations: 97
% Condensations: 19
% Global subsumptions: 2224
% 
% Simple tautologies: 112
% Equational tautologies: 59
% Forward subsumptions: 35804
% Fw demodulations to eq. taut.: 402
% Bw demodulations to eq. taut.: 10
% 
% Unit resulting resolution: 58344
% Factoring: 5
% Forward superposition: 8588
% Backward superposition: 5648
% Self superposition: 75
% Equality factoring: 2
% Equality resolution: 25
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 17093
% SAT solver unit clauses: 7444
% SAT solver binary clauses: 8590
% 
% Sat splits: 20
% Sat splitting refutations: 4980
% 
% Memory used [KB]: 241019
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f13422,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f862,f863,f1036,f1034,f868,f869,f870,f871,f872,f1039,f873,f875,f960,f1037,f974,f1053,f1052,f995,f996,f1060,f1007,f1008,f1062,f1015,f1066,f1069,f1068,f1211,f1213,f1215,f1238,f1239,f1240,f1241,f1242,f1244,f1363,f1245,f1325,f1362,f1366,f1490,f1489,f1488,f1477,f1470,f1576,f1577,f1578,f1579,f1580,f1581,f1614,f1613,f1617,f1800,f1799,f1777,f1774,f1773,f1771,f1770,f1860,f1861,f1899,f1915,f1929,f1931,f1928,f1947,f1050,f1961,f1963,f1960,f1976,f1977,f1978,f1979,f1980,f1982,f1983,f2063,f2108,f2293,f2295,f2297,f2340,f2372,f2383,f2384,f2385,f2386,f2387,f2389,f2103,f2390,f2470,f2102,f2512,f2517,f2518,f2519,f2520,f2521,f2647,f2522,f2524,f2609,f2646,f2655,f2658,f2666,f2667,f2668,f2669,f2670,f2672,f1040,f2791,f2673,f2753,f2790,f2793,f2794,f2827,f2829,f2908,f2909,f2910,f2942,f2943,f2944,f2976,f2977,f2978,f3010,f3011,f3012,f3045,f3046,f3047,f3079,f3080,f3081,f3113,f3114,f3115,f3148,f3149,f3150,f3183,f3184,f3185,f3218,f3219,f3220,f3226,f3228,f3229,f3231,f3233,f3318,f3357,f3362,f3363,f3373,f3375,f3376,f3377,f3380,f3460,f3499,f3503,f3504,f3537,f3538,f3540,f3541,f3546,f3548,f3549,f3551,f3553,f3638,f3696,f3701,f3737,f3738,f3772,f3773,f3807,f3808,f3826,f3827,f1058,f3832,f3833,f3863,f3861,f3850,f3851,f3867,f3865,f3969,f3971,f3972,f3973,f3976,f4056,f4108,f4109,f4114,f4149,f4150,f4185,f4186,f4220,f4221,f4255,f4256,f4291,f4292,f4294,f4295,f3690,f3688,f4300,f4302,f4303,f4305,f4307,f4392,f4445,f3698,f4448,f4484,f4485,f4520,f4521,f1056,f4644,f4651,f4708,f4712,f4717,f4646,f4882,f4884,f4886,f5021,f5027,f5087,f5091,f5023,f5118,f5120,f5121,f5123,f5125,f5210,f5267,f5280,f5282,f5283,f5284,f5287,f5367,f5409,f5478,f5480,f5481,f5482,f5485,f5565,f5618,f5620,f5621,f5623,f5625,f5710,f5794,f5797,f5782,f5828,f5829,f5823,f5860,f5861,f5855,f3697,f5873,f5875,f5876,f5877,f5880,f5960,f6025,f6027,f6028,f6029,f6032,f6112,f5606,f6227,f6229,f6230,f6231,f6234,f6314,f6379,f6384,f6539,f6540,f6545,f6547,f6548,f6550,f6552,f6637,f6713,f6880,f6882,f6883,f6884,f6887,f6967,f7030,f7032,f7033,f7136,f7137,f6703,f6701,f7142,f7144,f7145,f7147,f7149,f7234,f6704,f7303,f6716,f6714,f7469,f7471,f7472,f7473,f7476,f7556,f7619,f7622,f7793,f7846,f7800,f7847,f7801,f1033,f7848,f7842,f7849,f7843,f7850,f8017,f8019,f8020,f8021,f8024,f8104,f8177,f8182,f8412,f8413,f8414,f8415,f8567,f8577,f8579,f8580,f8581,f8584,f8586,f8587,f8589,f8591,f8676,f8691,f8722,f8723,f7270,f7269,f1064,f3859,f8735,f8737,f8738,f8740,f8742,f8827,f8884,f8913,f8914,f8901,f8915,f8944,f8945,f8932,f8946,f8981,f2339,f9048,f9049,f9050,f9051,f9084,f2371,f9095,f9097,f9098,f9099,f9102,f9182,f9222,f9294,f9295,f9296,f9297,f9300,f9302,f9303,f9305,f9307,f9392,f9438,f9439,f9483,f9485,f9486,f9487,f9490,f9570,f9623,f9666,f9668,f9669,f9671,f9673,f9758,f9773,f9803,f9804,f9875,f9876,f9877,f9878,f9882,f9884,f9885,f9886,f9889,f9969,f9992,f10014,f10015,f10070,f10071,f10058,f10072,f10101,f10102,f10089,f10103,f10112,f10114,f10115,f10117,f10119,f10204,f10261,f10274,f10276,f10277,f10278,f10281,f10361,f10401,f10408,f10410,f10411,f10413,f10415,f10500,f9431,f10545,f2657,f10546,f10552,f10554,f10555,f10556,f10559,f10639,f10678,f10680,f2656,f10681,f10722,f10725,f10716,f10728,f10729,f10736,f10738,f10739,f10741,f10743,f10828,f3691,f10885,f10914,f10915,f10902,f10916,f10945,f10946,f10933,f10947,f10976,f10979,f10978,f10964,f10980,f2828,f10981,f11020,f11021,f11014,f11024,f2830,f11025,f3539,f11054,f11055,f10724,f11048,f11058,f2911,f11059,f3739,f11069,f11071,f11072,f11073,f11076,f11156,f4105,f11209,f11214,f11216,f11217,f11219,f11221,f11306,f8717,f11321,f11351,f2660,f11352,f3365,f11360,f11362,f11363,f11364,f11367,f11447,f10010,f11470,f11490,f2796,f11491,f3364,f11519,f11528,f2945,f11529,f3774,f11596,f11597,f11598,f11599,f11625,f11634,f2979,f11635,f3809,f11661,f11670,f3013,f11671,f4151,f11697,f11707,f3048,f11708,f4187,f11733,f11743,f3082,f11744,f4222,f11769,f11779,f3116,f11780,f4257,f11847,f11848,f11849,f11850,f11875,f11886,f3151,f11887,f4293,f11912,f11923,f3186,f11924,f4486,f11949,f5796,f11960,f3221,f11961,f4522,f12029,f12030,f12031,f12032,f12100,f12101,f12102,f12103,f12114,f12116,f12117,f1038,f12118,f1042,f12121,f1041,f12201,f1043,f5406,f12242,f12314,f12315,f12316,f12317,f1072,f1070,f12346,f1930,f12377,f1962,f12412,f12413,f12414,f12415,f12587,f12627,f12628,f12629,f12630,f12666,f12667,f12668,f12669,f12704,f12705,f12706,f12707,f12742,f12743,f12744,f12745,f12780,f12781,f12782,f12783,f12943,f12949,f13008,f13012,f13016,f12945,f13050,f13051,f13052,f13053,f13089,f13090,f13091,f13092,f13128,f13129,f13130,f13131,f3704,f3702,f13288,f13291,f13295,f13300,f13365,f13367,f13289,f13380,f13296,f13384,f1054,f13301,f13415])).
fof(f13415,plain,(
  $false | ($spl0 | $spl36)),
  inference(equality_resolution,[],[f13357])).
fof(f13357,plain,(
  ( ! [X138] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X138)) ) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f13356,f1067])).
fof(f1067,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl36),
  inference(cnf_transformation,[],[f1067_D])).
fof(f1067_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f13356,plain,(
  ( ! [X138] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X138)) ) | $spl0),
  inference(forward_demodulation,[],[f13247,f842])).
fof(f842,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f384])).
fof(f384,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X11] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X11)))),
  file('thm.extreal.lt_|48||49|','thm.extreal.mul_lzero')).
fof(f13247,plain,(
  ( ! [X138] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X138) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X138))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl0),
  inference(superposition,[],[f797,f12787])).
fof(f12787,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0)))) ) | $spl0),
  inference(backward_demodulation,[],[f1034,f811])).
fof(f811,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11] : s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal',X11)))),
  file('thm.extreal.lt_|48||49|','thm.extreal.mul_lone')).
fof(f797,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(cnf_transformation,[],[f570])).
fof(f570,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) & (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(flattening,[],[f353])).
fof(f353,plain,(
  ! [X0] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X11] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X11) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_infty|split|5')).
fof(f13301,plain,(
  ( ! [X16] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16)))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X16)) ) | $spl302),
  inference(cnf_transformation,[],[f13301_D])).
fof(f13301_D,plain,(
  ( ! [X16] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16)))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X16)) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f1054,plain,(
  ( ! [X9] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl22),
  inference(cnf_transformation,[],[f1054_D])).
fof(f1054_D,plain,(
  ( ! [X9] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f13384,plain,(
  ( ! [X169,X170] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X170))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169)))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X169)) ) | $spl0),
  inference(forward_demodulation,[],[f13383,f842])).
fof(f13383,plain,(
  ( ! [X169,X170] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169)))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X169) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X170))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169)))) ) | $spl0),
  inference(forward_demodulation,[],[f13382,f842])).
fof(f13382,plain,(
  ( ! [X169,X170] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X169) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X170))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13381,f838])).
fof(f838,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f380])).
fof(f380,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X7] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X7)))),
  file('thm.extreal.lt_|48||49|','thm.extreal.num_not_infty|split|1')).
fof(f13381,plain,(
  ( ! [X169,X170] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X169) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X170))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169)))) ) | $spl0),
  inference(forward_demodulation,[],[f13266,f842])).
fof(f13266,plain,(
  ( ! [X169,X170] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X169) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X170))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X169)))) ) | $spl0),
  inference(superposition,[],[f853,f12787])).
fof(f853,plain,(
  ( ! [X2,X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43(X0)))) | s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f588])).
fof(f588,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43(X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK43])],[f587])).
fof(f587,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f449])).
fof(f449,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | ? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(flattening,[],[f448])).
fof(f448,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(ennf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) => (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f168])).
fof(f168,axiom,(
  ! [X17] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X17) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X17) | ? [X11] : s('type.extreal.extreal',X17) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) => (? [X11] : s('type.extreal.extreal',X17) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X17) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X17)))),
  file('thm.extreal.lt_|48||49|','thm.extreal.extreal_46_prove_ltb')).
fof(f13296,plain,(
  ( ! [X14] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X14)) ) | $spl300),
  inference(cnf_transformation,[],[f13296_D])).
fof(f13296_D,plain,(
  ( ! [X14] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X14)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f13380,plain,(
  ( ! [X167,X168] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X168))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X167)) ) | $spl0),
  inference(forward_demodulation,[],[f13379,f842])).
fof(f13379,plain,(
  ( ! [X167,X168] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X167) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X168))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167)))) ) | $spl0),
  inference(forward_demodulation,[],[f13378,f842])).
fof(f13378,plain,(
  ( ! [X167,X168] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X167) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X168))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13377,f837])).
fof(f837,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f379])).
fof(f379,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f378])).
fof(f378,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X7] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X7)))),
  file('thm.extreal.lt_|48||49|','thm.extreal.num_not_infty|split|0')).
fof(f13377,plain,(
  ( ! [X167,X168] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X167) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X168))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167)))) ) | $spl0),
  inference(forward_demodulation,[],[f13265,f842])).
fof(f13265,plain,(
  ( ! [X167,X168] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X167) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X168))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X167)))) ) | $spl0),
  inference(superposition,[],[f853,f12787])).
fof(f13289,plain,(
  ( ! [X11] : (s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl298),
  inference(cnf_transformation,[],[f13289_D])).
fof(f13289_D,plain,(
  ( ! [X11] : (s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f13367,plain,(
  ( ! [X149,X150] : (s('type.extreal.extreal',X149) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X150) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f13366,f842])).
fof(f13366,plain,(
  ( ! [X149,X150] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X150))) | s('type.extreal.extreal',X149) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X150) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f13255,f842])).
fof(f13255,plain,(
  ( ! [X149,X150] : (s('type.extreal.extreal',X149) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X150) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X149))),s('type.extreal.extreal',X150)))) ) | $spl0),
  inference(superposition,[],[f834,f12787])).
fof(f834,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) )),
  inference(cnf_transformation,[],[f583])).
fof(f583,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) & ((s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))),
  inference(flattening,[],[f582])).
fof(f582,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) | (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) & ((s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))),
  inference(nnf_transformation,[],[f377])).
fof(f377,plain,(
  ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) <=> (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X11,X12] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X12))) <=> (s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X12) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.entire')).
fof(f13365,plain,(
  ( ! [X147,X148] : (s('type.extreal.extreal',X147) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X148) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f13364,f841])).
fof(f841,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f383])).
fof(f383,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X11] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.mul_rzero')).
fof(f13364,plain,(
  ( ! [X147,X148] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X148),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.extreal.extreal',X147) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X148) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f13254,f842])).
fof(f13254,plain,(
  ( ! [X147,X148] : (s('type.extreal.extreal',X147) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X148) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X148),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X147)))))) ) | $spl0),
  inference(superposition,[],[f834,f12787])).
fof(f13300,plain,(
  ( ! [X17,X16] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16)))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X16)) ) | $spl0),
  inference(forward_demodulation,[],[f13299,f842])).
fof(f13299,plain,(
  ( ! [X17,X16] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16)))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X16) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16)))) ) | $spl0),
  inference(forward_demodulation,[],[f13298,f842])).
fof(f13298,plain,(
  ( ! [X17,X16] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X16) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13297,f838])).
fof(f13297,plain,(
  ( ! [X17,X16] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X16) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16)))) ) | $spl0),
  inference(forward_demodulation,[],[f13184,f842])).
fof(f13184,plain,(
  ( ! [X17,X16] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X16) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X16)))) ) | $spl0),
  inference(superposition,[],[f12787,f853])).
fof(f13295,plain,(
  ( ! [X14,X15] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X15))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X14)) ) | $spl0),
  inference(forward_demodulation,[],[f13294,f842])).
fof(f13294,plain,(
  ( ! [X14,X15] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14)))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X14) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X15))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14)))) ) | $spl0),
  inference(forward_demodulation,[],[f13293,f842])).
fof(f13293,plain,(
  ( ! [X14,X15] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X14) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X15))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14)))) ) | $spl0),
  inference(subsumption_resolution,[],[f13292,f837])).
fof(f13292,plain,(
  ( ! [X14,X15] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X14) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X15))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14)))) ) | $spl0),
  inference(forward_demodulation,[],[f13183,f842])).
fof(f13183,plain,(
  ( ! [X14,X15] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X14) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X15))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X14)))) ) | $spl0),
  inference(superposition,[],[f12787,f853])).
fof(f13291,plain,(
  ( ! [X12,X13] : (s('type.extreal.extreal',X12) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X13) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f13290,f842])).
fof(f13290,plain,(
  ( ! [X12,X13] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X13))) | s('type.extreal.extreal',X12) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X13) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f13182,f842])).
fof(f13182,plain,(
  ( ! [X12,X13] : (s('type.extreal.extreal',X12) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X13) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X12))),s('type.extreal.extreal',X13)))) ) | $spl0),
  inference(superposition,[],[f12787,f834])).
fof(f13288,plain,(
  ( ! [X10,X11] : (s('type.extreal.extreal',X10) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f13287,f841])).
fof(f13287,plain,(
  ( ! [X10,X11] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.extreal.extreal',X10) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f13181,f842])).
fof(f13181,plain,(
  ( ! [X10,X11] : (s('type.extreal.extreal',X10) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X10)))))) ) | $spl0),
  inference(superposition,[],[f12787,f834])).
fof(f3702,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl138),
  inference(cnf_transformation,[],[f3702_D])).
fof(f3702_D,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f3704,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl140),
  inference(cnf_transformation,[],[f3704_D])).
fof(f3704_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f13131,plain,(
  ( ! [X19,X20] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) | $spl296),
  inference(cnf_transformation,[],[f13131_D])).
fof(f13131_D,plain,(
  ( ! [X19,X20] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f13130,plain,(
  ( ! [X21,X19,X20] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f13116,f838])).
fof(f13116,plain,(
  ( ! [X21,X19,X20] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f8645,f853])).
fof(f8645,plain,(
  ( ! [X130,X131] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X130,X131)) ) | $spl32),
  inference(resolution,[],[f5475,f663])).
fof(f663,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f501])).
fof(f501,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f500])).
fof(f500,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f499])).
fof(f499,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f460])).
fof(f460,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f5475,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl32),
  inference(resolution,[],[f4110,f787])).
fof(f787,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f569])).
fof(f569,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) & (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f340])).
fof(f340,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X11,X12] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X12)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X11)))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.extreal_lt_def')).
fof(f4110,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl32),
  inference(subsumption_resolution,[],[f4088,f794])).
fof(f794,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(cnf_transformation,[],[f350])).
fof(f350,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(flattening,[],[f349])).
fof(f349,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X11] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_infty|split|3')).
fof(f4088,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl32),
  inference(superposition,[],[f1063,f798])).
fof(f798,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f571])).
fof(f571,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) & (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f356])).
fof(f356,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f355])).
fof(f355,plain,(
  ! [X0] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X11] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X11) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X11)))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_infty|split|6')).
fof(f1063,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl32),
  inference(cnf_transformation,[],[f1063_D])).
fof(f1063_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f13129,plain,(
  ( ! [X17,X16] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17)) ) | $spl294),
  inference(cnf_transformation,[],[f13129_D])).
fof(f13129_D,plain,(
  ( ! [X17,X16] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f13128,plain,(
  ( ! [X17,X18,X16] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f13115,f837])).
fof(f13115,plain,(
  ( ! [X17,X18,X16] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f8645,f853])).
fof(f13092,plain,(
  ( ! [X19,X20] : (sP7(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) | $spl292),
  inference(cnf_transformation,[],[f13092_D])).
fof(f13092_D,plain,(
  ( ! [X19,X20] : (sP7(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f13091,plain,(
  ( ! [X21,X19,X20] : (sP7(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f13077,f838])).
fof(f13077,plain,(
  ( ! [X21,X19,X20] : (sP7(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f8644,f853])).
fof(f8644,plain,(
  ( ! [X128,X129] : (sP7(X128,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X129)) ) | $spl32),
  inference(resolution,[],[f5475,f662])).
fof(f662,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f501])).
fof(f13090,plain,(
  ( ! [X17,X16] : (sP7(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17)) ) | $spl290),
  inference(cnf_transformation,[],[f13090_D])).
fof(f13090_D,plain,(
  ( ! [X17,X16] : (sP7(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17)) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f13089,plain,(
  ( ! [X17,X18,X16] : (sP7(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f13076,f837])).
fof(f13076,plain,(
  ( ! [X17,X18,X16] : (sP7(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f8644,f853])).
fof(f13053,plain,(
  ( ! [X19,X20] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) | $spl288),
  inference(cnf_transformation,[],[f13053_D])).
fof(f13053_D,plain,(
  ( ! [X19,X20] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f13052,plain,(
  ( ! [X21,X19,X20] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f13038,f838])).
fof(f13038,plain,(
  ( ! [X21,X19,X20] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f8642,f853])).
fof(f8642,plain,(
  ( ! [X125,X124] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X124,X125)) ) | $spl32),
  inference(resolution,[],[f5475,f659])).
fof(f659,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f498])).
fof(f498,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f497])).
fof(f497,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f496])).
fof(f496,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f13051,plain,(
  ( ! [X17,X16] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17)) ) | $spl286),
  inference(cnf_transformation,[],[f13051_D])).
fof(f13051_D,plain,(
  ( ! [X17,X16] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f13050,plain,(
  ( ! [X17,X18,X16] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f13037,f837])).
fof(f13037,plain,(
  ( ! [X17,X18,X16] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f8642,f853])).
fof(f12945,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl284),
  inference(cnf_transformation,[],[f12945_D])).
fof(f12945_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f13016,plain,(
  ( ! [X86] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X86))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f13015,f1034])).
fof(f13015,plain,(
  ( ! [X86] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X86))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(forward_demodulation,[],[f13014,f1034])).
fof(f13014,plain,(
  ( ! [X86] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X86))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f13013,f838])).
fof(f13013,plain,(
  ( ! [X86] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X86))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f12917,f837])).
fof(f12917,plain,(
  ( ! [X86] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X86))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(superposition,[],[f853,f1034])).
fof(f13012,plain,(
  ( ! [X85] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X85))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f13011,f1034])).
fof(f13011,plain,(
  ( ! [X85] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X85))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(forward_demodulation,[],[f13010,f1034])).
fof(f13010,plain,(
  ( ! [X85] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X85))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f13009,f837])).
fof(f13009,plain,(
  ( ! [X85] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X85))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f12916,f838])).
fof(f12916,plain,(
  ( ! [X85] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X85))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(superposition,[],[f853,f1034])).
fof(f13008,plain,(
  ( ! [X84] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X84)))) ) | $spl0),
  inference(forward_demodulation,[],[f13007,f1034])).
fof(f13007,plain,(
  ( ! [X84] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X84))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f13006,f837])).
fof(f13006,plain,(
  ( ! [X84] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X84))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f12915,f838])).
fof(f12915,plain,(
  ( ! [X84] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X84))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl0),
  inference(superposition,[],[f853,f1034])).
fof(f12949,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f12948,f1034])).
fof(f12948,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(forward_demodulation,[],[f12947,f1034])).
fof(f12947,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f12946,f838])).
fof(f12946,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f12839,f837])).
fof(f12839,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(superposition,[],[f1034,f853])).
fof(f12943,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f12942,f1034])).
fof(f12942,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(forward_demodulation,[],[f12941,f1034])).
fof(f12941,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f12940,f837])).
fof(f12940,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f12838,f838])).
fof(f12838,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl0),
  inference(superposition,[],[f1034,f853])).
fof(f12783,plain,(
  ( ! [X19,X20] : (sP8(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) | $spl282),
  inference(cnf_transformation,[],[f12783_D])).
fof(f12783_D,plain,(
  ( ! [X19,X20] : (sP8(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f12782,plain,(
  ( ! [X21,X19,X20] : (sP8(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12769,f838])).
fof(f12769,plain,(
  ( ! [X21,X19,X20] : (sP8(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f8641,f853])).
fof(f8641,plain,(
  ( ! [X123,X122] : (sP8(X122,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X123)) ) | $spl32),
  inference(resolution,[],[f5475,f658])).
fof(f658,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f498])).
fof(f12781,plain,(
  ( ! [X17,X16] : (sP8(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17)) ) | $spl280),
  inference(cnf_transformation,[],[f12781_D])).
fof(f12781_D,plain,(
  ( ! [X17,X16] : (sP8(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f12780,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12768,f837])).
fof(f12768,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f8641,f853])).
fof(f12745,plain,(
  ( ! [X19,X20] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) | $spl278),
  inference(cnf_transformation,[],[f12745_D])).
fof(f12745_D,plain,(
  ( ! [X19,X20] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f12744,plain,(
  ( ! [X21,X19,X20] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12732,f838])).
fof(f12732,plain,(
  ( ! [X21,X19,X20] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f8629,f853])).
fof(f8629,plain,(
  ( ! [X99,X98] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X98,X99)) ) | $spl32),
  inference(resolution,[],[f5475,f642])).
fof(f642,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f490])).
fof(f490,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f489])).
fof(f489,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f12743,plain,(
  ( ! [X17,X16] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17)) ) | $spl276),
  inference(cnf_transformation,[],[f12743_D])).
fof(f12743_D,plain,(
  ( ! [X17,X16] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f12742,plain,(
  ( ! [X17,X18,X16] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12731,f837])).
fof(f12731,plain,(
  ( ! [X17,X18,X16] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f8629,f853])).
fof(f12707,plain,(
  ( ! [X19,X20] : (sP4(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) | $spl274),
  inference(cnf_transformation,[],[f12707_D])).
fof(f12707_D,plain,(
  ( ! [X19,X20] : (sP4(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f12706,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12694,f838])).
fof(f12694,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f8628,f853])).
fof(f8628,plain,(
  ( ! [X97,X96] : (sP4(X96,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X97)) ) | $spl32),
  inference(resolution,[],[f5475,f641])).
fof(f641,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f490])).
fof(f12705,plain,(
  ( ! [X17,X16] : (sP4(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17)) ) | $spl272),
  inference(cnf_transformation,[],[f12705_D])).
fof(f12705_D,plain,(
  ( ! [X17,X16] : (sP4(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f12704,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12693,f837])).
fof(f12693,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f8628,f853])).
fof(f12669,plain,(
  ( ! [X19,X20] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) | $spl270),
  inference(cnf_transformation,[],[f12669_D])).
fof(f12669_D,plain,(
  ( ! [X19,X20] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f12668,plain,(
  ( ! [X21,X19,X20] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12655,f838])).
fof(f12655,plain,(
  ( ! [X21,X19,X20] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f8615,f853])).
fof(f8615,plain,(
  ( ! [X70,X71] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X70,X71)) ) | $spl32),
  inference(resolution,[],[f5475,f625])).
fof(f625,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f483])).
fof(f483,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f482])).
fof(f482,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f481])).
fof(f481,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f453])).
fof(f453,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f12667,plain,(
  ( ! [X17,X16] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17)) ) | $spl268),
  inference(cnf_transformation,[],[f12667_D])).
fof(f12667_D,plain,(
  ( ! [X17,X16] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f12666,plain,(
  ( ! [X17,X18,X16] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12654,f837])).
fof(f12654,plain,(
  ( ! [X17,X18,X16] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f8615,f853])).
fof(f12630,plain,(
  ( ! [X19,X20] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) | $spl266),
  inference(cnf_transformation,[],[f12630_D])).
fof(f12630_D,plain,(
  ( ! [X19,X20] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f12629,plain,(
  ( ! [X21,X19,X20] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12617,f838])).
fof(f12617,plain,(
  ( ! [X21,X19,X20] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f8599,f853])).
fof(f8599,plain,(
  ( ! [X39,X38] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X38,X39)) ) | $spl32),
  inference(resolution,[],[f5475,f608])).
fof(f608,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f475])).
fof(f475,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f474])).
fof(f474,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f473])).
fof(f473,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f450])).
fof(f450,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f12628,plain,(
  ( ! [X17,X16] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17)) ) | $spl264),
  inference(cnf_transformation,[],[f12628_D])).
fof(f12628_D,plain,(
  ( ! [X17,X16] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17)) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f12627,plain,(
  ( ! [X17,X18,X16] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12616,f837])).
fof(f12616,plain,(
  ( ! [X17,X18,X16] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f8599,f853])).
fof(f12587,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f12424,f814])).
fof(f814,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.le_|48||49|')).
fof(f12424,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl2),
  inference(resolution,[],[f1036,f787])).
fof(f12415,plain,(
  ( ! [X19,X20] : (sP0(X19,X20,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl262),
  inference(cnf_transformation,[],[f12415_D])).
fof(f12415_D,plain,(
  ( ! [X19,X20] : (sP0(X19,X20,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f12414,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,X20,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12402,f838])).
fof(f12402,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,X20,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f8598,f853])).
fof(f8598,plain,(
  ( ! [X37,X36] : (sP0(X36,X37,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f606])).
fof(f606,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f475])).
fof(f12413,plain,(
  ( ! [X17,X16] : (sP0(X16,X17,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl260),
  inference(cnf_transformation,[],[f12413_D])).
fof(f12413_D,plain,(
  ( ! [X17,X16] : (sP0(X16,X17,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f12412,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,X17,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f12401,f837])).
fof(f12401,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,X17,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f8598,f853])).
fof(f1962,plain,(
  ( ! [X16] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16)) ) | $spl54),
  inference(cnf_transformation,[],[f1962_D])).
fof(f1962_D,plain,(
  ( ! [X16] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f12377,plain,(
  ( ! [X10,X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl182),
  inference(subsumption_resolution,[],[f12367,f837])).
fof(f12367,plain,(
  ( ! [X10,X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl182),
  inference(superposition,[],[f12190,f853])).
fof(f12190,plain,(
  ( ! [X155] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X155)) ) | $spl182),
  inference(resolution,[],[f5796,f704])).
fof(f704,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f518])).
fof(f518,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f517])).
fof(f517,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f466])).
fof(f466,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f1930,plain,(
  ( ! [X16] : (sP11(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl50),
  inference(cnf_transformation,[],[f1930_D])).
fof(f1930_D,plain,(
  ( ! [X16] : (sP11(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f12346,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl182),
  inference(subsumption_resolution,[],[f12336,f837])).
fof(f12336,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl182),
  inference(superposition,[],[f12189,f853])).
fof(f12189,plain,(
  ( ! [X154] : (sP11(X154,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl182),
  inference(resolution,[],[f5796,f703])).
fof(f703,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f518])).
fof(f1070,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl38),
  inference(cnf_transformation,[],[f1070_D])).
fof(f1070_D,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1072,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl40),
  inference(cnf_transformation,[],[f1072_D])).
fof(f1072_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f12317,plain,(
  ( ! [X19,X20] : (sP7('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) | $spl258),
  inference(cnf_transformation,[],[f12317_D])).
fof(f12317_D,plain,(
  ( ! [X19,X20] : (sP7('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f12316,plain,(
  ( ! [X21,X19,X20] : (sP7('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f12308,f838])).
fof(f12308,plain,(
  ( ! [X21,X19,X20] : (sP7('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f929,f853])).
fof(f929,plain,(
  ( ! [X134,X133] : (sP7('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X133,X134)) )),
  inference(resolution,[],[f589,f663])).
fof(f589,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_|48||49|')).
fof(f12315,plain,(
  ( ! [X17,X16] : (sP7('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17)) ) | $spl256),
  inference(cnf_transformation,[],[f12315_D])).
fof(f12315_D,plain,(
  ( ! [X17,X16] : (sP7('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17)) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f12314,plain,(
  ( ! [X17,X18,X16] : (sP7('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f12307,f837])).
fof(f12307,plain,(
  ( ! [X17,X18,X16] : (sP7('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f929,f853])).
fof(f12242,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl182),
  inference(subsumption_resolution,[],[f12241,f837])).
fof(f12241,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl182),
  inference(subsumption_resolution,[],[f12220,f793])).
fof(f793,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(flattening,[],[f347])).
fof(f347,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X11] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X11))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_refl')).
fof(f12220,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl182),
  inference(superposition,[],[f5796,f853])).
fof(f5406,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl178),
  inference(cnf_transformation,[],[f5406_D])).
fof(f5406_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f1043,plain,(
  ( ! [X18] : (p(s(bool,X18))) ) | $spl16),
  inference(cnf_transformation,[],[f1043_D])).
fof(f1043_D,plain,(
  ( ! [X18] : (p(s(bool,X18))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f12201,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl182),
  inference(resolution,[],[f5796,f855])).
fof(f855,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f594])).
fof(f594,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f469])).
fof(f469,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f468])).
fof(f468,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f398])).
fof(f398,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f397])).
fof(f397,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X11,X38,X12,X39] : ((s(bool,X11) = s(bool,X38) & (p(s(bool,X38)) => s(bool,X12) = s(bool,X39))) => ((p(s(bool,X11)) => p(s(bool,X12))) <=> (p(s(bool,X38)) => p(s(bool,X39)))))),
  file('thm.extreal.lt_|48||49|','thm.bool.IMP_CONG')).
fof(f1041,plain,(
  ( ! [X11,X9] : (s(bool,X9) != s(bool,X11) | ~p(s(bool,X9))) ) | $spl12),
  inference(cnf_transformation,[],[f1041_D])).
fof(f1041_D,plain,(
  ( ! [X11,X9] : (s(bool,X9) != s(bool,X11) | ~p(s(bool,X9))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f12121,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X26) != s(bool,X27)) ) | $spl182),
  inference(resolution,[],[f5796,f597])).
fof(f597,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f469])).
fof(f1042,plain,(
  ( ! [X20] : (~p(s(bool,X20))) ) | $spl14),
  inference(cnf_transformation,[],[f1042_D])).
fof(f1042_D,plain,(
  ( ! [X20] : (~p(s(bool,X20))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f12118,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl182),
  inference(resolution,[],[f5796,f596])).
fof(f596,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f469])).
fof(f1038,plain,(
  ( ! [X5,X3] : (s(bool,X3) != s(bool,X5) | p(s(bool,X3))) ) | $spl6),
  inference(cnf_transformation,[],[f1038_D])).
fof(f1038_D,plain,(
  ( ! [X5,X3] : (s(bool,X3) != s(bool,X5) | p(s(bool,X3))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f12117,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl182),
  inference(resolution,[],[f5796,f595])).
fof(f595,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f469])).
fof(f12116,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X11) != s(bool,X12)) ) | $spl182),
  inference(resolution,[],[f5796,f593])).
fof(f593,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f469])).
fof(f12114,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl182),
  inference(resolution,[],[f5796,f591])).
fof(f591,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f469])).
fof(f12103,plain,(
  ( ! [X19,X20] : (sP7(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) | $spl254),
  inference(cnf_transformation,[],[f12103_D])).
fof(f12103_D,plain,(
  ( ! [X19,X20] : (sP7(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f12102,plain,(
  ( ! [X21,X19,X20] : (sP7(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f12094,f838])).
fof(f12094,plain,(
  ( ! [X21,X19,X20] : (sP7(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f928,f853])).
fof(f928,plain,(
  ( ! [X132,X131] : (sP7(X131,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X132)) )),
  inference(resolution,[],[f589,f662])).
fof(f12101,plain,(
  ( ! [X17,X16] : (sP7(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17)) ) | $spl252),
  inference(cnf_transformation,[],[f12101_D])).
fof(f12101_D,plain,(
  ( ! [X17,X16] : (sP7(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f12100,plain,(
  ( ! [X17,X18,X16] : (sP7(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f12093,f837])).
fof(f12093,plain,(
  ( ! [X17,X18,X16] : (sP7(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f928,f853])).
fof(f12032,plain,(
  ( ! [X19,X20] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) | $spl250),
  inference(cnf_transformation,[],[f12032_D])).
fof(f12032_D,plain,(
  ( ! [X19,X20] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f12031,plain,(
  ( ! [X21,X19,X20] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f12023,f838])).
fof(f12023,plain,(
  ( ! [X21,X19,X20] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f926,f853])).
fof(f926,plain,(
  ( ! [X127,X128] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X127,X128)) )),
  inference(resolution,[],[f589,f659])).
fof(f12030,plain,(
  ( ! [X17,X16] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17)) ) | $spl248),
  inference(cnf_transformation,[],[f12030_D])).
fof(f12030_D,plain,(
  ( ! [X17,X16] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f12029,plain,(
  ( ! [X17,X18,X16] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f12022,f837])).
fof(f12022,plain,(
  ( ! [X17,X18,X16] : (sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f926,f853])).
fof(f4522,plain,(
  ( ! [X28,X27] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X27,X28)) ) | $spl172),
  inference(cnf_transformation,[],[f4522_D])).
fof(f4522_D,plain,(
  ( ! [X28,X27] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X27,X28)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f11961,plain,(
  ( ! [X21,X19,X20] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11954,f838])).
fof(f11954,plain,(
  ( ! [X21,X19,X20] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f5542,f853])).
fof(f5542,plain,(
  ( ! [X140,X141] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X140,X141)) ) | $spl32),
  inference(resolution,[],[f4110,f667])).
fof(f667,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f503])).
fof(f503,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f502])).
fof(f502,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f459])).
fof(f459,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f3221,plain,(
  ( ! [X28,X27] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X27,X28)) ) | $spl120),
  inference(cnf_transformation,[],[f3221_D])).
fof(f3221_D,plain,(
  ( ! [X28,X27] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X27,X28)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f11960,plain,(
  ( ! [X17,X18,X16] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11953,f837])).
fof(f11953,plain,(
  ( ! [X17,X18,X16] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f5542,f853])).
fof(f5796,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl182),
  inference(cnf_transformation,[],[f5796_D])).
fof(f5796_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f11949,plain,(
  ( ! [X6,X7] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f5542,f796])).
fof(f796,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f570])).
fof(f4486,plain,(
  ( ! [X28,X27] : (sP6(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28)) ) | $spl170),
  inference(cnf_transformation,[],[f4486_D])).
fof(f4486_D,plain,(
  ( ! [X28,X27] : (sP6(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f11924,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11917,f838])).
fof(f11917,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f5541,f853])).
fof(f5541,plain,(
  ( ! [X138,X139] : (sP6(X138,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X139)) ) | $spl32),
  inference(resolution,[],[f4110,f666])).
fof(f666,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f503])).
fof(f3186,plain,(
  ( ! [X28,X27] : (sP6(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28)) ) | $spl116),
  inference(cnf_transformation,[],[f3186_D])).
fof(f3186_D,plain,(
  ( ! [X28,X27] : (sP6(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f11923,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11916,f837])).
fof(f11916,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f5541,f853])).
fof(f11912,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f5541,f796])).
fof(f4293,plain,(
  ( ! [X28,X27] : (sP6(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl168),
  inference(cnf_transformation,[],[f4293_D])).
fof(f4293_D,plain,(
  ( ! [X28,X27] : (sP6(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f11887,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11880,f838])).
fof(f11880,plain,(
  ( ! [X21,X19,X20] : (sP6(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f5540,f853])).
fof(f5540,plain,(
  ( ! [X136,X137] : (sP6(X136,X137,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f4110,f665])).
fof(f665,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f503])).
fof(f3151,plain,(
  ( ! [X28,X27] : (sP6(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl112),
  inference(cnf_transformation,[],[f3151_D])).
fof(f3151_D,plain,(
  ( ! [X28,X27] : (sP6(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f11886,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11879,f837])).
fof(f11879,plain,(
  ( ! [X17,X18,X16] : (sP6(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f5540,f853])).
fof(f11875,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f5540,f796])).
fof(f11850,plain,(
  ( ! [X19,X20] : (sP8(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) | $spl246),
  inference(cnf_transformation,[],[f11850_D])).
fof(f11850_D,plain,(
  ( ! [X19,X20] : (sP8(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f11849,plain,(
  ( ! [X21,X19,X20] : (sP8(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f11842,f838])).
fof(f11842,plain,(
  ( ! [X21,X19,X20] : (sP8(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f925,f853])).
fof(f925,plain,(
  ( ! [X125,X126] : (sP8(X125,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X126)) )),
  inference(resolution,[],[f589,f658])).
fof(f11848,plain,(
  ( ! [X17,X16] : (sP8(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17)) ) | $spl244),
  inference(cnf_transformation,[],[f11848_D])).
fof(f11848_D,plain,(
  ( ! [X17,X16] : (sP8(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f11847,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f11841,f837])).
fof(f11841,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f925,f853])).
fof(f4257,plain,(
  ( ! [X28,X27] : (sP7(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl166),
  inference(cnf_transformation,[],[f4257_D])).
fof(f4257_D,plain,(
  ( ! [X28,X27] : (sP7(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f11780,plain,(
  ( ! [X21,X19,X20] : (sP7(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11774,f838])).
fof(f11774,plain,(
  ( ! [X21,X19,X20] : (sP7(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f5539,f853])).
fof(f5539,plain,(
  ( ! [X134,X135] : (sP7(X134,X135,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f4110,f661])).
fof(f661,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f501])).
fof(f3116,plain,(
  ( ! [X28,X27] : (sP7(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl108),
  inference(cnf_transformation,[],[f3116_D])).
fof(f3116_D,plain,(
  ( ! [X28,X27] : (sP7(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f11779,plain,(
  ( ! [X17,X18,X16] : (sP7(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11773,f837])).
fof(f11773,plain,(
  ( ! [X17,X18,X16] : (sP7(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f5539,f853])).
fof(f11769,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f5539,f796])).
fof(f4222,plain,(
  ( ! [X28,X27] : (sP8(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl164),
  inference(cnf_transformation,[],[f4222_D])).
fof(f4222_D,plain,(
  ( ! [X28,X27] : (sP8(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f11744,plain,(
  ( ! [X21,X19,X20] : (sP8(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11738,f838])).
fof(f11738,plain,(
  ( ! [X21,X19,X20] : (sP8(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f5536,f853])).
fof(f5536,plain,(
  ( ! [X128,X129] : (sP8(X128,X129,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f4110,f657])).
fof(f657,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f498])).
fof(f3082,plain,(
  ( ! [X28,X27] : (sP8(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl104),
  inference(cnf_transformation,[],[f3082_D])).
fof(f3082_D,plain,(
  ( ! [X28,X27] : (sP8(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f11743,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11737,f837])).
fof(f11737,plain,(
  ( ! [X17,X18,X16] : (sP8(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f5536,f853])).
fof(f11733,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f5536,f796])).
fof(f4187,plain,(
  ( ! [X28,X27] : (sP4(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl162),
  inference(cnf_transformation,[],[f4187_D])).
fof(f4187_D,plain,(
  ( ! [X28,X27] : (sP4(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f11708,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11702,f838])).
fof(f11702,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f5523,f853])).
fof(f5523,plain,(
  ( ! [X103,X102] : (sP4(X102,X103,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f4110,f640])).
fof(f640,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f490])).
fof(f3048,plain,(
  ( ! [X28,X27] : (sP4(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl100),
  inference(cnf_transformation,[],[f3048_D])).
fof(f3048_D,plain,(
  ( ! [X28,X27] : (sP4(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f11707,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11701,f837])).
fof(f11701,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f5523,f853])).
fof(f11697,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f5523,f796])).
fof(f4151,plain,(
  ( ! [X28,X27] : (sP2(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28)) ) | $spl160),
  inference(cnf_transformation,[],[f4151_D])).
fof(f4151_D,plain,(
  ( ! [X28,X27] : (sP2(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f11671,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11666,f838])).
fof(f11666,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f5507,f853])).
fof(f5507,plain,(
  ( ! [X70,X71] : (sP2(X70,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X71)) ) | $spl32),
  inference(resolution,[],[f4110,f624])).
fof(f624,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f483])).
fof(f3013,plain,(
  ( ! [X28,X27] : (sP2(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28)) ) | $spl96),
  inference(cnf_transformation,[],[f3013_D])).
fof(f3013_D,plain,(
  ( ! [X28,X27] : (sP2(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f11670,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11665,f837])).
fof(f11665,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f5507,f853])).
fof(f11661,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f5507,f796])).
fof(f3809,plain,(
  ( ! [X28,X27] : (sP2(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl146),
  inference(cnf_transformation,[],[f3809_D])).
fof(f3809_D,plain,(
  ( ! [X28,X27] : (sP2(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f11635,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11630,f838])).
fof(f11630,plain,(
  ( ! [X21,X19,X20] : (sP2(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f5506,f853])).
fof(f5506,plain,(
  ( ! [X68,X69] : (sP2(X68,X69,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f4110,f623])).
fof(f623,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f483])).
fof(f2979,plain,(
  ( ! [X28,X27] : (sP2(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl92),
  inference(cnf_transformation,[],[f2979_D])).
fof(f2979_D,plain,(
  ( ! [X28,X27] : (sP2(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f11634,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11629,f837])).
fof(f11629,plain,(
  ( ! [X17,X18,X16] : (sP2(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f5506,f853])).
fof(f11625,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f5506,f796])).
fof(f11599,plain,(
  ( ! [X19,X20] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) | $spl242),
  inference(cnf_transformation,[],[f11599_D])).
fof(f11599_D,plain,(
  ( ! [X19,X20] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f11598,plain,(
  ( ! [X21,X19,X20] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f11592,f838])).
fof(f11592,plain,(
  ( ! [X21,X19,X20] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f913,f853])).
fof(f913,plain,(
  ( ! [X101,X102] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X101,X102)) )),
  inference(resolution,[],[f589,f642])).
fof(f11597,plain,(
  ( ! [X17,X16] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17)) ) | $spl240),
  inference(cnf_transformation,[],[f11597_D])).
fof(f11597_D,plain,(
  ( ! [X17,X16] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f11596,plain,(
  ( ! [X17,X18,X16] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f11591,f837])).
fof(f11591,plain,(
  ( ! [X17,X18,X16] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f913,f853])).
fof(f3774,plain,(
  ( ! [X28,X27] : (sP0(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28)) ) | $spl144),
  inference(cnf_transformation,[],[f3774_D])).
fof(f3774_D,plain,(
  ( ! [X28,X27] : (sP0(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f11529,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11524,f838])).
fof(f11524,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21)))) ) | $spl32),
  inference(superposition,[],[f5493,f853])).
fof(f5493,plain,(
  ( ! [X43,X42] : (sP0(X42,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X43)) ) | $spl32),
  inference(resolution,[],[f4110,f607])).
fof(f607,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f475])).
fof(f2945,plain,(
  ( ! [X28,X27] : (sP0(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28)) ) | $spl88),
  inference(cnf_transformation,[],[f2945_D])).
fof(f2945_D,plain,(
  ( ! [X28,X27] : (sP0(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f11528,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11523,f837])).
fof(f11523,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))) ) | $spl32),
  inference(superposition,[],[f5493,f853])).
fof(f11519,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f5493,f796])).
fof(f3364,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl124),
  inference(cnf_transformation,[],[f3364_D])).
fof(f3364_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f11491,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11475,f838])).
fof(f11475,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(superposition,[],[f5477,f853])).
fof(f5477,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl32),
  inference(resolution,[],[f4110,f790])).
fof(f790,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f447])).
fof(f447,plain,(
  ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))),
  inference(ennf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X11,X12] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X12)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X12)))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_imp_le')).
fof(f2796,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl76),
  inference(cnf_transformation,[],[f2796_D])).
fof(f2796_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f11490,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11474,f837])).
fof(f11474,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl32),
  inference(superposition,[],[f5477,f853])).
fof(f11470,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl32),
  inference(superposition,[],[f5477,f796])).
fof(f10010,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl232),
  inference(cnf_transformation,[],[f10010_D])).
fof(f10010_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f11447,plain,(
  ( ! [X165,X164] : (p(s(bool,X164)) | p(s(bool,X165)) | s(bool,X165) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f5477,f855])).
fof(f11367,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X24) != s(bool,X25)) ) | $spl32),
  inference(resolution,[],[f5477,f597])).
fof(f11364,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f5477,f596])).
fof(f11363,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f5477,f595])).
fof(f11362,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X9) != s(bool,X10)) ) | $spl32),
  inference(resolution,[],[f5477,f593])).
fof(f11360,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f5477,f591])).
fof(f3365,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl125),
  inference(cnf_transformation,[],[f3365_D])).
fof(f3365_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f11352,plain,(
  ( ! [X5] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11326,f838])).
fof(f11326,plain,(
  ( ! [X5] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(superposition,[],[f5475,f853])).
fof(f2660,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl71),
  inference(cnf_transformation,[],[f2660_D])).
fof(f2660_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f11351,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11325,f837])).
fof(f11325,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl32),
  inference(superposition,[],[f5475,f853])).
fof(f11321,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl32),
  inference(superposition,[],[f5475,f796])).
fof(f8717,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl216),
  inference(cnf_transformation,[],[f8717_D])).
fof(f8717_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f11306,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f855])).
fof(f11221,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl32),
  inference(resolution,[],[f5475,f597])).
fof(f11219,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f596])).
fof(f11217,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f595])).
fof(f11216,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl32),
  inference(resolution,[],[f5475,f593])).
fof(f11214,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f591])).
fof(f11209,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11208,f838])).
fof(f11208,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11192,f794])).
fof(f11192,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(superposition,[],[f1063,f853])).
fof(f4105,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl158),
  inference(cnf_transformation,[],[f4105_D])).
fof(f4105_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f11156,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f855])).
fof(f11076,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X26) != s(bool,X27)) ) | $spl32),
  inference(resolution,[],[f1063,f597])).
fof(f11073,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f596])).
fof(f11072,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f595])).
fof(f11071,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X11) != s(bool,X12)) ) | $spl32),
  inference(resolution,[],[f1063,f593])).
fof(f11069,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f591])).
fof(f3739,plain,(
  ( ! [X18] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X18)) ) | $spl142),
  inference(cnf_transformation,[],[f3739_D])).
fof(f3739_D,plain,(
  ( ! [X18] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X18)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f11059,plain,(
  ( ! [X12,X13] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X12) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11053,f838])).
fof(f11053,plain,(
  ( ! [X12,X13] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl32),
  inference(superposition,[],[f9958,f853])).
fof(f9958,plain,(
  ( ! [X153] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X153)) ) | $spl32),
  inference(resolution,[],[f5477,f704])).
fof(f2911,plain,(
  ( ! [X18] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X18)) ) | $spl84),
  inference(cnf_transformation,[],[f2911_D])).
fof(f2911_D,plain,(
  ( ! [X18] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X18)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f11058,plain,(
  ( ! [X10,X11] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11052,f837])).
fof(f11052,plain,(
  ( ! [X10,X11] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl32),
  inference(superposition,[],[f9958,f853])).
fof(f11048,plain,(
  ( ! [X4] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f9958,f796])).
fof(f10724,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl236),
  inference(cnf_transformation,[],[f10724_D])).
fof(f10724_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f11055,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,X1))) ) | $spl32),
  inference(subsumption_resolution,[],[f11027,f744])).
fof(f744,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f546])).
fof(f546,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f545])).
fof(f545,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f104])).
fof(f104,axiom,(
  ! [X6] : ((p(s(bool,'F_0')) | p(s(bool,X6))) <=> p(s(bool,X6)))),
  file('thm.extreal.lt_|48||49|','thm.bool.OR_CLAUSES|split|2')).
fof(f11027,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,X1))) ) | $spl32),
  inference(resolution,[],[f9958,f702])).
fof(f702,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f518])).
fof(f11054,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,X0))) ) | $spl32),
  inference(subsumption_resolution,[],[f11026,f744])).
fof(f11026,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,X0))) ) | $spl32),
  inference(resolution,[],[f9958,f706])).
fof(f706,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f520])).
fof(f520,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f519])).
fof(f519,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f407,f466])).
fof(f407,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f406])).
fof(f406,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X18,X10] : ((~(p(s(bool,X10)) | p(s(bool,X18))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X10)) => p(s(bool,'F_0'))) => (~p(s(bool,X18)) => p(s(bool,'F_0')))))),
  file('thm.extreal.lt_|48||49|','thm.sat.sat_9_prove_ltb')).
fof(f3539,plain,(
  ( ! [X18] : (sP11(X18,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl128),
  inference(cnf_transformation,[],[f3539_D])).
fof(f3539_D,plain,(
  ( ! [X18] : (sP11(X18,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f11025,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11019,f838])).
fof(f11019,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl32),
  inference(superposition,[],[f9957,f853])).
fof(f9957,plain,(
  ( ! [X152] : (sP11(X152,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f5477,f703])).
fof(f2830,plain,(
  ( ! [X18] : (sP11(X18,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl80),
  inference(cnf_transformation,[],[f2830_D])).
fof(f2830_D,plain,(
  ( ! [X18] : (sP11(X18,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f11024,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl32),
  inference(subsumption_resolution,[],[f11018,f837])).
fof(f11018,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl32),
  inference(superposition,[],[f9957,f853])).
fof(f11014,plain,(
  ( ! [X4] : (sP11(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f9957,f796])).
fof(f11021,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(subsumption_resolution,[],[f10993,f744])).
fof(f10993,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(resolution,[],[f9957,f702])).
fof(f11020,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(subsumption_resolution,[],[f10992,f744])).
fof(f10992,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(resolution,[],[f9957,f706])).
fof(f10981,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10969,f838])).
fof(f10969,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f10627,f853])).
fof(f10627,plain,(
  ( ! [X152] : (sP11(X152,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5279,f703])).
fof(f5279,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f790])).
fof(f3699,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f3669,f1067])).
fof(f3669,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl33),
  inference(superposition,[],[f1064,f796])).
fof(f2828,plain,(
  ( ! [X16] : (sP11(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl78),
  inference(cnf_transformation,[],[f2828_D])).
fof(f2828_D,plain,(
  ( ! [X16] : (sP11(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f10980,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10968,f837])).
fof(f10968,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f10627,f853])).
fof(f10964,plain,(
  ( ! [X4] : (sP11(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f10627,f796])).
fof(f10978,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl238),
  inference(cnf_transformation,[],[f10978_D])).
fof(f10978_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f10979,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10951,f744])).
fof(f10951,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f10627,f702])).
fof(f10976,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10950,f744])).
fof(f10950,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f10627,f706])).
fof(f10947,plain,(
  ( ! [X10,X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10937,f837])).
fof(f10937,plain,(
  ( ! [X10,X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5356,f853])).
fof(f5356,plain,(
  ( ! [X155] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X155)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f704])).
fof(f10933,plain,(
  ( ! [X4] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5356,f796])).
fof(f10946,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X1))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10920,f744])).
fof(f10920,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X1))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5356,f702])).
fof(f10945,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X0))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10919,f744])).
fof(f10919,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X0))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5356,f706])).
fof(f10916,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10906,f837])).
fof(f10906,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5355,f853])).
fof(f5355,plain,(
  ( ! [X154] : (sP11(X154,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f703])).
fof(f10902,plain,(
  ( ! [X4] : (sP11(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5355,f796])).
fof(f10915,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10889,f744])).
fof(f10889,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5355,f702])).
fof(f10914,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10888,f744])).
fof(f10888,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5355,f706])).
fof(f10885,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10884,f837])).
fof(f10884,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10863,f1067])).
fof(f10863,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl33),
  inference(superposition,[],[f1064,f853])).
fof(f3691,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl134),
  inference(cnf_transformation,[],[f3691_D])).
fof(f3691_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f10828,plain,(
  ( ! [X171,X172] : (p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f855])).
fof(f10743,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X24) != s(bool,X26)) ) | $spl33),
  inference(resolution,[],[f1064,f597])).
fof(f10741,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f596])).
fof(f10739,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f595])).
fof(f10738,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X9) != s(bool,X11)) ) | $spl33),
  inference(resolution,[],[f1064,f593])).
fof(f10736,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f591])).
fof(f10729,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl32),
  inference(subsumption_resolution,[],[f10721,f838])).
fof(f10721,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl32),
  inference(superposition,[],[f9957,f853])).
fof(f10728,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl32),
  inference(subsumption_resolution,[],[f10720,f837])).
fof(f10720,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | $spl32),
  inference(superposition,[],[f9957,f853])).
fof(f10716,plain,(
  ( ! [X4] : (sP11(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl32),
  inference(superposition,[],[f9957,f796])).
fof(f10725,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(subsumption_resolution,[],[f10695,f744])).
fof(f10695,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(resolution,[],[f9957,f702])).
fof(f10722,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(subsumption_resolution,[],[f10694,f744])).
fof(f10694,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(resolution,[],[f9957,f706])).
fof(f10681,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10659,f838])).
fof(f10659,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5279,f853])).
fof(f2656,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl68),
  inference(cnf_transformation,[],[f2656_D])).
fof(f2656_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f10680,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10658,f837])).
fof(f10658,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5279,f853])).
fof(f10678,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl234),
  inference(cnf_transformation,[],[f10678_D])).
fof(f10678_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f10639,plain,(
  ( ! [X165,X164] : (p(s(bool,X164)) | p(s(bool,X165)) | s(bool,X165) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5279,f855])).
fof(f10559,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X24) != s(bool,X25)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5279,f597])).
fof(f10556,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5279,f596])).
fof(f10555,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5279,f595])).
fof(f10554,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X9) != s(bool,X10)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5279,f593])).
fof(f10552,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5279,f591])).
fof(f10546,plain,(
  ( ! [X5] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10526,f838])).
fof(f10526,plain,(
  ( ! [X5] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5277,f853])).
fof(f5277,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f787])).
fof(f2657,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl69),
  inference(cnf_transformation,[],[f2657_D])).
fof(f2657_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f10545,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10525,f837])).
fof(f10525,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5277,f853])).
fof(f9431,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl226),
  inference(cnf_transformation,[],[f9431_D])).
fof(f9431_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f10500,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f855])).
fof(f10415,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X21) != s(bool,X23)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f597])).
fof(f10413,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f596])).
fof(f10411,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f595])).
fof(f10410,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X6) != s(bool,X8)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f593])).
fof(f10408,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f591])).
fof(f10401,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10400,f837])).
fof(f10400,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10380,f793])).
fof(f10380,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f3699,f853])).
fof(f10361,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f855])).
fof(f10281,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X26) != s(bool,X27)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f597])).
fof(f10278,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f596])).
fof(f10277,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f595])).
fof(f10276,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X11) != s(bool,X12)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f593])).
fof(f10274,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f591])).
fof(f10261,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10260,f837])).
fof(f10260,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10239,f1067])).
fof(f10239,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl33),
  inference(superposition,[],[f1064,f853])).
fof(f10204,plain,(
  ( ! [X171,X172] : (p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f855])).
fof(f10119,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X24) != s(bool,X26)) ) | $spl33),
  inference(resolution,[],[f1064,f597])).
fof(f10117,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f596])).
fof(f10115,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f595])).
fof(f10114,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X9) != s(bool,X11)) ) | $spl33),
  inference(resolution,[],[f1064,f593])).
fof(f10112,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f591])).
fof(f10103,plain,(
  ( ! [X10,X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10093,f837])).
fof(f10093,plain,(
  ( ! [X10,X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5356,f853])).
fof(f10089,plain,(
  ( ! [X4] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5356,f796])).
fof(f10102,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X1))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10076,f744])).
fof(f10076,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X1))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5356,f702])).
fof(f10101,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X0))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10075,f744])).
fof(f10075,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X0))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5356,f706])).
fof(f10072,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10062,f837])).
fof(f10062,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5355,f853])).
fof(f10058,plain,(
  ( ! [X4] : (sP11(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5355,f796])).
fof(f10071,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10045,f744])).
fof(f10045,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5355,f702])).
fof(f10070,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f10044,f744])).
fof(f10044,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5355,f706])).
fof(f10015,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f9997,f838])).
fof(f9997,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(superposition,[],[f5477,f853])).
fof(f10014,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl32),
  inference(subsumption_resolution,[],[f9996,f837])).
fof(f9996,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl32),
  inference(superposition,[],[f5477,f853])).
fof(f9992,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl32),
  inference(superposition,[],[f5477,f796])).
fof(f9969,plain,(
  ( ! [X165,X164] : (p(s(bool,X164)) | p(s(bool,X165)) | s(bool,X165) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f5477,f855])).
fof(f9889,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X24) != s(bool,X25)) ) | $spl32),
  inference(resolution,[],[f5477,f597])).
fof(f9886,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f5477,f596])).
fof(f9885,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f5477,f595])).
fof(f9884,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X9) != s(bool,X10)) ) | $spl32),
  inference(resolution,[],[f5477,f593])).
fof(f9882,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f5477,f591])).
fof(f9878,plain,(
  ( ! [X19,X20] : (sP4(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) | $spl230),
  inference(cnf_transformation,[],[f9878_D])).
fof(f9878_D,plain,(
  ( ! [X19,X20] : (sP4(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f9877,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f9871,f838])).
fof(f9871,plain,(
  ( ! [X21,X19,X20] : (sP4(X19,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f912,f853])).
fof(f912,plain,(
  ( ! [X99,X100] : (sP4(X99,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X100)) )),
  inference(resolution,[],[f589,f641])).
fof(f9876,plain,(
  ( ! [X17,X16] : (sP4(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17)) ) | $spl228),
  inference(cnf_transformation,[],[f9876_D])).
fof(f9876_D,plain,(
  ( ! [X17,X16] : (sP4(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f9875,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f9870,f837])).
fof(f9870,plain,(
  ( ! [X17,X18,X16] : (sP4(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f912,f853])).
fof(f9804,plain,(
  ( ! [X5] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f9778,f838])).
fof(f9778,plain,(
  ( ! [X5] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(superposition,[],[f5475,f853])).
fof(f9803,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl32),
  inference(subsumption_resolution,[],[f9777,f837])).
fof(f9777,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl32),
  inference(superposition,[],[f5475,f853])).
fof(f9773,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl32),
  inference(superposition,[],[f5475,f796])).
fof(f9758,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f855])).
fof(f9673,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl32),
  inference(resolution,[],[f5475,f597])).
fof(f9671,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f596])).
fof(f9669,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f595])).
fof(f9668,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl32),
  inference(resolution,[],[f5475,f593])).
fof(f9666,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f591])).
fof(f9623,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f9622,f838])).
fof(f9622,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f9606,f794])).
fof(f9606,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(superposition,[],[f1063,f853])).
fof(f9570,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f855])).
fof(f9490,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X26) != s(bool,X27)) ) | $spl32),
  inference(resolution,[],[f1063,f597])).
fof(f9487,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f596])).
fof(f9486,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f595])).
fof(f9485,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X11) != s(bool,X12)) ) | $spl32),
  inference(resolution,[],[f1063,f593])).
fof(f9483,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f591])).
fof(f9439,plain,(
  ( ! [X5] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f9418,f838])).
fof(f9418,plain,(
  ( ! [X5] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5277,f853])).
fof(f9438,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f9417,f837])).
fof(f9417,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5277,f853])).
fof(f9392,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f855])).
fof(f9307,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X21) != s(bool,X23)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f597])).
fof(f9305,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f596])).
fof(f9303,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f595])).
fof(f9302,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X6) != s(bool,X8)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f593])).
fof(f9300,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5277,f591])).
fof(f9297,plain,(
  ( ! [X19,X20] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) | $spl224),
  inference(cnf_transformation,[],[f9297_D])).
fof(f9297_D,plain,(
  ( ! [X19,X20] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f9296,plain,(
  ( ! [X21,X19,X20] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f9289,f838])).
fof(f9289,plain,(
  ( ! [X21,X19,X20] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f899,f853])).
fof(f899,plain,(
  ( ! [X74,X73] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X73,X74)) )),
  inference(resolution,[],[f589,f625])).
fof(f9295,plain,(
  ( ! [X17,X16] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17)) ) | $spl222),
  inference(cnf_transformation,[],[f9295_D])).
fof(f9295_D,plain,(
  ( ! [X17,X16] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f9294,plain,(
  ( ! [X17,X18,X16] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f9288,f837])).
fof(f9288,plain,(
  ( ! [X17,X18,X16] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f899,f853])).
fof(f9222,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f9221,f837])).
fof(f9221,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f9201,f793])).
fof(f9201,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f3699,f853])).
fof(f9182,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f855])).
fof(f9102,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X26) != s(bool,X27)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f597])).
fof(f9099,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f596])).
fof(f9098,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f595])).
fof(f9097,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X11) != s(bool,X12)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f593])).
fof(f9095,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f591])).
fof(f2371,plain,(
  ( ! [X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X11)) ) | $spl62),
  inference(cnf_transformation,[],[f2371_D])).
fof(f2371_D,plain,(
  ( ! [X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X11)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f9084,plain,(
  ( ! [X12,X13] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X12) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl136),
  inference(subsumption_resolution,[],[f9079,f838])).
fof(f9079,plain,(
  ( ! [X12,X13] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl136),
  inference(superposition,[],[f6101,f853])).
fof(f6101,plain,(
  ( ! [X155] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X155)) ) | $spl136),
  inference(resolution,[],[f3697,f704])).
fof(f9051,plain,(
  ( ! [X19,X20] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) | $spl220),
  inference(cnf_transformation,[],[f9051_D])).
fof(f9051_D,plain,(
  ( ! [X19,X20] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f9050,plain,(
  ( ! [X21,X19,X20] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f9044,f838])).
fof(f9044,plain,(
  ( ! [X21,X19,X20] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')),X19,X20) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f883,f853])).
fof(f883,plain,(
  ( ! [X41,X42] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X41,X42)) )),
  inference(resolution,[],[f589,f608])).
fof(f9049,plain,(
  ( ! [X17,X16] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17)) ) | $spl218),
  inference(cnf_transformation,[],[f9049_D])).
fof(f9049_D,plain,(
  ( ! [X17,X16] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f9048,plain,(
  ( ! [X17,X18,X16] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f9043,f837])).
fof(f9043,plain,(
  ( ! [X17,X18,X16] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X16,X17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f883,f853])).
fof(f2339,plain,(
  ( ! [X11] : (sP11(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl60),
  inference(cnf_transformation,[],[f2339_D])).
fof(f2339_D,plain,(
  ( ! [X11] : (sP11(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8981,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl136),
  inference(subsumption_resolution,[],[f8976,f838])).
fof(f8976,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl136),
  inference(superposition,[],[f6100,f853])).
fof(f6100,plain,(
  ( ! [X154] : (sP11(X154,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl136),
  inference(resolution,[],[f3697,f703])).
fof(f8946,plain,(
  ( ! [X10,X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f8936,f837])).
fof(f8936,plain,(
  ( ! [X10,X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X10) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5356,f853])).
fof(f8932,plain,(
  ( ! [X4] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5356,f796])).
fof(f8945,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X1))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f8919,f744])).
fof(f8919,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X1))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5356,f702])).
fof(f8944,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X0))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f8918,f744])).
fof(f8918,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,X0))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5356,f706])).
fof(f8915,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f8905,f837])).
fof(f8905,plain,(
  ( ! [X10,X11] : (sP11(X10,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5355,f853])).
fof(f8901,plain,(
  ( ! [X4] : (sP11(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5355,f796])).
fof(f8914,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f8888,f744])).
fof(f8888,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5355,f702])).
fof(f8913,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f8887,f744])).
fof(f8887,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5355,f706])).
fof(f8884,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f8883,f837])).
fof(f8883,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f8862,f1067])).
fof(f8862,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl33),
  inference(superposition,[],[f1064,f853])).
fof(f8827,plain,(
  ( ! [X171,X172] : (p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f855])).
fof(f8742,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X24) != s(bool,X26)) ) | $spl33),
  inference(resolution,[],[f1064,f597])).
fof(f8740,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f596])).
fof(f8738,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f595])).
fof(f8737,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X9) != s(bool,X11)) ) | $spl33),
  inference(resolution,[],[f1064,f593])).
fof(f8735,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f591])).
fof(f3859,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl149),
  inference(cnf_transformation,[],[f3859_D])).
fof(f3859_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f1064,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl33),
  inference(cnf_transformation,[],[f1064_D])).
fof(f1064_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f7269,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(superposition,[],[f1060,f806])).
fof(f806,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f139])).
fof(f139,axiom,(
  ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X7))),
  file('thm.extreal.lt_|48||49|','thm.numeral.numeral_distrib|split|17')).
fof(f7270,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(superposition,[],[f1060,f808])).
fof(f808,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f574])).
fof(f574,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f362])).
fof(f362,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X7))),
  file('thm.extreal.lt_|48||49|','thm.numeral.numeral_distrib|split|18')).
fof(f8723,plain,(
  ( ! [X5] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f8696,f838])).
fof(f8696,plain,(
  ( ! [X5] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(superposition,[],[f5475,f853])).
fof(f8722,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl32),
  inference(subsumption_resolution,[],[f8695,f837])).
fof(f8695,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl32),
  inference(superposition,[],[f5475,f853])).
fof(f8691,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl32),
  inference(superposition,[],[f5475,f796])).
fof(f8676,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f855])).
fof(f8591,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl32),
  inference(resolution,[],[f5475,f597])).
fof(f8589,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f596])).
fof(f8587,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f595])).
fof(f8586,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl32),
  inference(resolution,[],[f5475,f593])).
fof(f8584,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl32),
  inference(resolution,[],[f5475,f591])).
fof(f8581,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl32),
  inference(subsumption_resolution,[],[f5827,f838])).
fof(f5827,plain,(
  ( ! [X12,X13] : (sP11(X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl32),
  inference(superposition,[],[f5553,f853])).
fof(f5553,plain,(
  ( ! [X154] : (sP11(X154,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f4110,f703])).
fof(f8580,plain,(
  ( ! [X12,X13] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X12) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl32),
  inference(subsumption_resolution,[],[f5859,f838])).
fof(f5859,plain,(
  ( ! [X12,X13] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl32),
  inference(superposition,[],[f5554,f853])).
fof(f5554,plain,(
  ( ! [X155] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X155)) ) | $spl32),
  inference(resolution,[],[f4110,f704])).
fof(f8579,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f8578,f838])).
fof(f8578,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f5996,f794])).
fof(f5996,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(superposition,[],[f1063,f853])).
fof(f8577,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl136),
  inference(subsumption_resolution,[],[f8576,f838])).
fof(f8576,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl136),
  inference(subsumption_resolution,[],[f6140,f793])).
fof(f6140,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl136),
  inference(superposition,[],[f3697,f853])).
fof(f8567,plain,(
  $false | $spl138),
  inference(equality_resolution,[],[f5421])).
fof(f5421,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X1)))) ) | $spl138),
  inference(superposition,[],[f3702,f843])).
fof(f843,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f385])).
fof(f385,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X7))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7)))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.extreal_of_num_def')).
fof(f8415,plain,(
  ( ! [X19,X20] : (sP0(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl214),
  inference(cnf_transformation,[],[f8415_D])).
fof(f8415_D,plain,(
  ( ! [X19,X20] : (sP0(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f8414,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f8408,f838])).
fof(f8408,plain,(
  ( ! [X21,X19,X20] : (sP0(X19,X20,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X21))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f882,f853])).
fof(f882,plain,(
  ( ! [X39,X40] : (sP0(X39,X40,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(resolution,[],[f589,f606])).
fof(f8413,plain,(
  ( ! [X17,X16] : (sP0(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl212),
  inference(cnf_transformation,[],[f8413_D])).
fof(f8413_D,plain,(
  ( ! [X17,X16] : (sP0(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f8412,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f8407,f837])).
fof(f8407,plain,(
  ( ! [X17,X18,X16] : (sP0(X16,X17,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f882,f853])).
fof(f8182,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl36),
  inference(subsumption_resolution,[],[f8181,f837])).
fof(f8181,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl36),
  inference(subsumption_resolution,[],[f8147,f1235])).
fof(f1235,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl36),
  inference(resolution,[],[f1067,f787])).
fof(f8147,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f864,f853])).
fof(f864,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f589,f788])).
fof(f788,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f569])).
fof(f8177,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl210),
  inference(cnf_transformation,[],[f8177_D])).
fof(f8177_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f8104,plain,(
  ( ! [X165,X164] : (p(s(bool,X164)) | p(s(bool,X165)) | s(bool,X165) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(resolution,[],[f864,f855])).
fof(f8024,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X24) != s(bool,X25)) )),
  inference(resolution,[],[f864,f597])).
fof(f8021,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(resolution,[],[f864,f596])).
fof(f8020,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(resolution,[],[f864,f595])).
fof(f8019,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X9) != s(bool,X10)) )),
  inference(resolution,[],[f864,f593])).
fof(f8017,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(resolution,[],[f864,f591])).
fof(f7850,plain,(
  ( ! [X9] : (s('type.extreal.extreal',X9) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal',X9)))) ) | $spl208),
  inference(cnf_transformation,[],[f7850_D])).
fof(f7850_D,plain,(
  ( ! [X9] : (s('type.extreal.extreal',X9) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal',X9)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f7843,plain,(
  ( ! [X8,X9] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8))) | s('type.extreal.extreal',X9) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal',X9)))) ) | $spl38),
  inference(superposition,[],[f1070,f834])).
fof(f7849,plain,(
  ( ! [X7] : (s('type.extreal.extreal',X7) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl206),
  inference(cnf_transformation,[],[f7849_D])).
fof(f7849_D,plain,(
  ( ! [X7] : (s('type.extreal.extreal',X7) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f7842,plain,(
  ( ! [X6,X7] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))) | s('type.extreal.extreal',X7) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl38),
  inference(superposition,[],[f1070,f834])).
fof(f7848,plain,(
  ( ! [X12,X13] : (s('type.extreal.extreal',X13) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))),s('type.extreal.extreal',X13)))) ) | $spl204),
  inference(cnf_transformation,[],[f7848_D])).
fof(f7848_D,plain,(
  ( ! [X12,X13] : (s('type.extreal.extreal',X13) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))),s('type.extreal.extreal',X13)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f1033,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl1),
  inference(cnf_transformation,[],[f1033_D])).
fof(f1033_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f7801,plain,(
  ( ! [X12,X13] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal',X13) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))),s('type.extreal.extreal',X13)))) ) | $spl38),
  inference(superposition,[],[f1070,f834])).
fof(f7847,plain,(
  ( ! [X10,X11] : (s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))))) ) | $spl202),
  inference(cnf_transformation,[],[f7847_D])).
fof(f7847_D,plain,(
  ( ! [X10,X11] : (s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f7800,plain,(
  ( ! [X10,X11] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))))) ) | $spl38),
  inference(superposition,[],[f1070,f834])).
fof(f7846,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl200),
  inference(cnf_transformation,[],[f7846_D])).
fof(f7846_D,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f7793,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl38),
  inference(superposition,[],[f1070,f840])).
fof(f840,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f584])).
fof(f584,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(nnf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X11] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11))) <=> s('type.realax.real',X11) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.extreal_eq_zero')).
fof(f7622,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl34),
  inference(subsumption_resolution,[],[f7621,f837])).
fof(f7621,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl34),
  inference(subsumption_resolution,[],[f7599,f793])).
fof(f7599,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl34),
  inference(superposition,[],[f1066,f853])).
fof(f7619,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl198),
  inference(cnf_transformation,[],[f7619_D])).
fof(f7619_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f7556,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl34),
  inference(resolution,[],[f1066,f855])).
fof(f7476,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X26) != s(bool,X27)) ) | $spl34),
  inference(resolution,[],[f1066,f597])).
fof(f7473,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl34),
  inference(resolution,[],[f1066,f596])).
fof(f7472,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl34),
  inference(resolution,[],[f1066,f595])).
fof(f7471,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X11) != s(bool,X12)) ) | $spl34),
  inference(resolution,[],[f1066,f593])).
fof(f7469,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl34),
  inference(resolution,[],[f1066,f591])).
fof(f6714,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl192),
  inference(cnf_transformation,[],[f6714_D])).
fof(f6714_D,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f6716,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl194),
  inference(cnf_transformation,[],[f6716_D])).
fof(f6716_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f7303,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl29 | $spl36)),
  inference(subsumption_resolution,[],[f7302,f837])).
fof(f7302,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl29 | $spl36)),
  inference(subsumption_resolution,[],[f7281,f1067])).
fof(f7281,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl29),
  inference(superposition,[],[f1060,f853])).
fof(f6704,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl190),
  inference(cnf_transformation,[],[f6704_D])).
fof(f6704_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7234,plain,(
  ( ! [X171,X172] : (p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl29),
  inference(resolution,[],[f1060,f855])).
fof(f7149,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s(bool,X24) != s(bool,X26)) ) | $spl29),
  inference(resolution,[],[f1060,f597])).
fof(f7147,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl29),
  inference(resolution,[],[f1060,f596])).
fof(f7145,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl29),
  inference(resolution,[],[f1060,f595])).
fof(f7144,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s(bool,X9) != s(bool,X11)) ) | $spl29),
  inference(resolution,[],[f1060,f593])).
fof(f7142,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl29),
  inference(resolution,[],[f1060,f591])).
fof(f6701,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl186),
  inference(cnf_transformation,[],[f6701_D])).
fof(f6701_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f6703,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl188),
  inference(cnf_transformation,[],[f6703_D])).
fof(f6703_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f7137,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(resolution,[],[f1060,f791])).
fof(f791,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1)) )),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0,X1] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X11,X12] : (s('type.extreal.extreal',X11) = s('type.extreal.extreal',X12) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X12)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X11)))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_total')).
fof(f7136,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(resolution,[],[f1060,f791])).
fof(f7033,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f7028,f793])).
fof(f7028,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl28),
  inference(trivial_inequality_removal,[],[f6997])).
fof(f6997,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl28),
  inference(superposition,[],[f1059,f807])).
fof(f807,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f573])).
fof(f1059,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl28),
  inference(cnf_transformation,[],[f1059_D])).
fof(f1059_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f7032,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f7029,f793])).
fof(f7029,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl28),
  inference(trivial_inequality_removal,[],[f6996])).
fof(f6996,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl28),
  inference(superposition,[],[f1059,f809])).
fof(f809,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f574])).
fof(f7030,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl196),
  inference(cnf_transformation,[],[f7030_D])).
fof(f7030_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f6967,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl28),
  inference(resolution,[],[f1059,f855])).
fof(f6887,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s(bool,X26) != s(bool,X27)) ) | $spl28),
  inference(resolution,[],[f1059,f597])).
fof(f6884,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl28),
  inference(resolution,[],[f1059,f596])).
fof(f6883,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl28),
  inference(resolution,[],[f1059,f595])).
fof(f6882,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s(bool,X11) != s(bool,X12)) ) | $spl28),
  inference(resolution,[],[f1059,f593])).
fof(f6880,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl28),
  inference(resolution,[],[f1059,f591])).
fof(f6713,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl29 | $spl36)),
  inference(subsumption_resolution,[],[f6712,f837])).
fof(f6712,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | ($spl29 | $spl36)),
  inference(subsumption_resolution,[],[f6684,f1067])).
fof(f6684,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl29),
  inference(superposition,[],[f1060,f853])).
fof(f6637,plain,(
  ( ! [X171,X172] : (p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl29),
  inference(resolution,[],[f1060,f855])).
fof(f6552,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s(bool,X24) != s(bool,X26)) ) | $spl29),
  inference(resolution,[],[f1060,f597])).
fof(f6550,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl29),
  inference(resolution,[],[f1060,f596])).
fof(f6548,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl29),
  inference(resolution,[],[f1060,f595])).
fof(f6547,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s(bool,X9) != s(bool,X11)) ) | $spl29),
  inference(resolution,[],[f1060,f593])).
fof(f6545,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl29),
  inference(resolution,[],[f1060,f591])).
fof(f6540,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(resolution,[],[f1060,f791])).
fof(f6539,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl29),
  inference(resolution,[],[f1060,f791])).
fof(f6384,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f6383,f838])).
fof(f6383,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f6366,f793])).
fof(f6366,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl20),
  inference(superposition,[],[f1051,f853])).
fof(f1051,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl20),
  inference(cnf_transformation,[],[f1051_D])).
fof(f1051_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f6379,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl184),
  inference(cnf_transformation,[],[f6379_D])).
fof(f6379_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f6314,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl20),
  inference(resolution,[],[f1051,f855])).
fof(f6234,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,X26) != s(bool,X27)) ) | $spl20),
  inference(resolution,[],[f1051,f597])).
fof(f6231,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl20),
  inference(resolution,[],[f1051,f596])).
fof(f6230,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl20),
  inference(resolution,[],[f1051,f595])).
fof(f6229,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,X11) != s(bool,X12)) ) | $spl20),
  inference(resolution,[],[f1051,f593])).
fof(f6227,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl20),
  inference(resolution,[],[f1051,f591])).
fof(f5606,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl180),
  inference(cnf_transformation,[],[f5606_D])).
fof(f5606_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f6112,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl136),
  inference(resolution,[],[f3697,f855])).
fof(f6032,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X26) != s(bool,X27)) ) | $spl136),
  inference(resolution,[],[f3697,f597])).
fof(f6029,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl136),
  inference(resolution,[],[f3697,f596])).
fof(f6028,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl136),
  inference(resolution,[],[f3697,f595])).
fof(f6027,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X11) != s(bool,X12)) ) | $spl136),
  inference(resolution,[],[f3697,f593])).
fof(f6025,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl136),
  inference(resolution,[],[f3697,f591])).
fof(f5960,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f855])).
fof(f5880,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X26) != s(bool,X27)) ) | $spl32),
  inference(resolution,[],[f1063,f597])).
fof(f5877,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f596])).
fof(f5876,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f595])).
fof(f5875,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X11) != s(bool,X12)) ) | $spl32),
  inference(resolution,[],[f1063,f593])).
fof(f5873,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f591])).
fof(f3697,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl136),
  inference(cnf_transformation,[],[f3697_D])).
fof(f3697_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f5855,plain,(
  ( ! [X5] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X5) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(superposition,[],[f5554,f798])).
fof(f5861,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,X1))) ) | $spl32),
  inference(subsumption_resolution,[],[f5833,f744])).
fof(f5833,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,X1))) ) | $spl32),
  inference(resolution,[],[f5554,f702])).
fof(f5860,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,X0))) ) | $spl32),
  inference(subsumption_resolution,[],[f5832,f744])).
fof(f5832,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | p(s(bool,X0))) ) | $spl32),
  inference(resolution,[],[f5554,f706])).
fof(f5823,plain,(
  ( ! [X5] : (sP11(X5,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(superposition,[],[f5553,f798])).
fof(f5829,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(subsumption_resolution,[],[f5801,f744])).
fof(f5801,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(resolution,[],[f5553,f702])).
fof(f5828,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(subsumption_resolution,[],[f5800,f744])).
fof(f5800,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl32),
  inference(resolution,[],[f5553,f706])).
fof(f5782,plain,(
  ( ! [X4] : (sP11(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f5355,f796])).
fof(f5797,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f5769,f744])).
fof(f5769,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5355,f702])).
fof(f5794,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f5768,f744])).
fof(f5768,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f5355,f706])).
fof(f5710,plain,(
  ( ! [X171,X172] : (p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f855])).
fof(f5625,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X24) != s(bool,X26)) ) | $spl33),
  inference(resolution,[],[f1064,f597])).
fof(f5623,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f596])).
fof(f5621,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f595])).
fof(f5620,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X9) != s(bool,X11)) ) | $spl33),
  inference(resolution,[],[f1064,f593])).
fof(f5618,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f591])).
fof(f5565,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f4110,f855])).
fof(f5485,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X26) != s(bool,X27)) ) | $spl32),
  inference(resolution,[],[f4110,f597])).
fof(f5482,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f4110,f596])).
fof(f5481,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f4110,f595])).
fof(f5480,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X11) != s(bool,X12)) ) | $spl32),
  inference(resolution,[],[f4110,f593])).
fof(f5478,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f4110,f591])).
fof(f5409,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f5408,f837])).
fof(f5408,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f5386,f793])).
fof(f5386,plain,(
  ( ! [X4] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(superposition,[],[f3699,f853])).
fof(f5367,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f855])).
fof(f5287,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X26) != s(bool,X27)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f597])).
fof(f5284,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f596])).
fof(f5283,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f595])).
fof(f5282,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X11) != s(bool,X12)) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f593])).
fof(f5280,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | ($spl33 | $spl36)),
  inference(resolution,[],[f3699,f591])).
fof(f5267,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f5266,f837])).
fof(f5266,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f5245,f1067])).
fof(f5245,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl33),
  inference(superposition,[],[f1064,f853])).
fof(f5210,plain,(
  ( ! [X171,X172] : (p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f855])).
fof(f5125,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X24) != s(bool,X26)) ) | $spl33),
  inference(resolution,[],[f1064,f597])).
fof(f5123,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f596])).
fof(f5121,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f595])).
fof(f5120,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X9) != s(bool,X11)) ) | $spl33),
  inference(resolution,[],[f1064,f593])).
fof(f5118,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f591])).
fof(f5023,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl176),
  inference(cnf_transformation,[],[f5023_D])).
fof(f5023_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f5091,plain,(
  ( ! [X86] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X86))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl130),
  inference(forward_demodulation,[],[f5090,f3688])).
fof(f5090,plain,(
  ( ! [X86] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X86)))) ) | $spl130),
  inference(forward_demodulation,[],[f5089,f3688])).
fof(f5089,plain,(
  ( ! [X86] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X86)))) ) | $spl130),
  inference(subsumption_resolution,[],[f5088,f838])).
fof(f5088,plain,(
  ( ! [X86] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X86)))) ) | $spl130),
  inference(subsumption_resolution,[],[f4995,f837])).
fof(f4995,plain,(
  ( ! [X86] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X86)))) ) | $spl130),
  inference(superposition,[],[f853,f3688])).
fof(f5087,plain,(
  ( ! [X85] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X85))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl130),
  inference(forward_demodulation,[],[f5086,f3688])).
fof(f5086,plain,(
  ( ! [X85] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X85)))) ) | $spl130),
  inference(forward_demodulation,[],[f5085,f3688])).
fof(f5085,plain,(
  ( ! [X85] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X85)))) ) | $spl130),
  inference(subsumption_resolution,[],[f5084,f837])).
fof(f5084,plain,(
  ( ! [X85] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X85)))) ) | $spl130),
  inference(subsumption_resolution,[],[f4994,f838])).
fof(f4994,plain,(
  ( ! [X85] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X85)))) ) | $spl130),
  inference(superposition,[],[f853,f3688])).
fof(f5027,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl130),
  inference(forward_demodulation,[],[f5026,f3688])).
fof(f5026,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl130),
  inference(forward_demodulation,[],[f5025,f3688])).
fof(f5025,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl130),
  inference(subsumption_resolution,[],[f5024,f838])).
fof(f5024,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl130),
  inference(subsumption_resolution,[],[f4916,f837])).
fof(f4916,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3)))) ) | $spl130),
  inference(superposition,[],[f3688,f853])).
fof(f5021,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) ) | $spl130),
  inference(forward_demodulation,[],[f5020,f3688])).
fof(f5020,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl130),
  inference(forward_demodulation,[],[f5019,f3688])).
fof(f5019,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl130),
  inference(subsumption_resolution,[],[f5018,f837])).
fof(f5018,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl130),
  inference(subsumption_resolution,[],[f4915,f838])).
fof(f4915,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) ) | $spl130),
  inference(superposition,[],[f3688,f853])).
fof(f4886,plain,(
  $false | $spl21),
  inference(subsumption_resolution,[],[f4885,f837])).
fof(f4885,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl21),
  inference(subsumption_resolution,[],[f4720,f794])).
fof(f4720,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl21),
  inference(resolution,[],[f1052,f791])).
fof(f4884,plain,(
  $false | $spl21),
  inference(subsumption_resolution,[],[f4883,f837])).
fof(f4883,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl21),
  inference(subsumption_resolution,[],[f4719,f794])).
fof(f4719,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl21),
  inference(resolution,[],[f1052,f791])).
fof(f4882,plain,(
  $false | $spl21),
  inference(subsumption_resolution,[],[f4718,f837])).
fof(f4718,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl21),
  inference(resolution,[],[f1052,f798])).
fof(f4646,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | $spl174),
  inference(cnf_transformation,[],[f4646_D])).
fof(f4646_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f4717,plain,(
  ( ! [X89] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X89))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4716,f1056])).
fof(f4716,plain,(
  ( ! [X89] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X89))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4715,f1056])).
fof(f4715,plain,(
  ( ! [X89] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X89))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4714,f838])).
fof(f4714,plain,(
  ( ! [X89] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X89))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4713,f1056])).
fof(f4713,plain,(
  ( ! [X89] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X89))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4621,f837])).
fof(f4621,plain,(
  ( ! [X89] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X89))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(superposition,[],[f853,f1056])).
fof(f4712,plain,(
  ( ! [X88] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X88))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4711,f1056])).
fof(f4711,plain,(
  ( ! [X88] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X88))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4710,f1056])).
fof(f4710,plain,(
  ( ! [X88] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X88))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4709,f849])).
fof(f849,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f391])).
fof(f391,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f232])).
fof(f232,axiom,(
  ! [X17] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))),
  file('thm.extreal.lt_|48||49|','thm.extreal.extreal_distinct|split|1')).
fof(f4709,plain,(
  ( ! [X88] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X88))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4620,f838])).
fof(f4620,plain,(
  ( ! [X88] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X88))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(superposition,[],[f853,f1056])).
fof(f4708,plain,(
  ( ! [X87] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X87)))) ) | $spl24),
  inference(forward_demodulation,[],[f4707,f1056])).
fof(f4707,plain,(
  ( ! [X87] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X87))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4706,f838])).
fof(f4706,plain,(
  ( ! [X87] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X87))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4705,f1056])).
fof(f4705,plain,(
  ( ! [X87] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X87))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4619,f849])).
fof(f4619,plain,(
  ( ! [X87] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X87))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | $spl24),
  inference(superposition,[],[f853,f1056])).
fof(f4651,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4650,f1056])).
fof(f4650,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4649,f1056])).
fof(f4649,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4648,f838])).
fof(f4648,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4647,f1056])).
fof(f4647,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4535,f837])).
fof(f4535,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(superposition,[],[f1056,f853])).
fof(f4644,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4643,f1056])).
fof(f4643,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(forward_demodulation,[],[f4642,f1056])).
fof(f4642,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4641,f849])).
fof(f4641,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f4534,f838])).
fof(f4534,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl24),
  inference(superposition,[],[f1056,f853])).
fof(f1056,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl24),
  inference(cnf_transformation,[],[f1056_D])).
fof(f1056_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4521,plain,(
  ( ! [X28,X29,X27] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X27,X28) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4514,f838])).
fof(f4514,plain,(
  ( ! [X28,X29,X27] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X27,X28) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(superposition,[],[f2040,f853])).
fof(f2040,plain,(
  ( ! [X140,X141] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X140,X141)) ) | $spl18),
  inference(resolution,[],[f1050,f667])).
fof(f4520,plain,(
  ( ! [X26,X24,X25] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X24,X25) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4513,f837])).
fof(f4513,plain,(
  ( ! [X26,X24,X25] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X24,X25) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(superposition,[],[f2040,f853])).
fof(f4485,plain,(
  ( ! [X28,X29,X27] : (sP6(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4478,f838])).
fof(f4478,plain,(
  ( ! [X28,X29,X27] : (sP6(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(superposition,[],[f2039,f853])).
fof(f2039,plain,(
  ( ! [X138,X139] : (sP6(X138,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X139)) ) | $spl18),
  inference(resolution,[],[f1050,f666])).
fof(f4484,plain,(
  ( ! [X26,X24,X25] : (sP6(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4477,f837])).
fof(f4477,plain,(
  ( ! [X26,X24,X25] : (sP6(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(superposition,[],[f2039,f853])).
fof(f4448,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f4447,f837])).
fof(f4447,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f4427,f1067])).
fof(f4427,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl33),
  inference(superposition,[],[f1064,f853])).
fof(f3698,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl137),
  inference(cnf_transformation,[],[f3698_D])).
fof(f3698_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f4445,plain,(
  ( ! [X9] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl33),
  inference(subsumption_resolution,[],[f4410,f838])).
fof(f4410,plain,(
  ( ! [X9] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl33),
  inference(superposition,[],[f1064,f853])).
fof(f4392,plain,(
  ( ! [X171,X172] : (p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f855])).
fof(f4307,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X24) != s(bool,X26)) ) | $spl33),
  inference(resolution,[],[f1064,f597])).
fof(f4305,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f596])).
fof(f4303,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f595])).
fof(f4302,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X9) != s(bool,X11)) ) | $spl33),
  inference(resolution,[],[f1064,f593])).
fof(f4300,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f591])).
fof(f3688,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl130),
  inference(cnf_transformation,[],[f3688_D])).
fof(f3688_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f3690,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl132),
  inference(cnf_transformation,[],[f3690_D])).
fof(f3690_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f4295,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl33),
  inference(resolution,[],[f1064,f791])).
fof(f4294,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl33),
  inference(resolution,[],[f1064,f791])).
fof(f4292,plain,(
  ( ! [X28,X29,X27] : (sP6(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4285,f838])).
fof(f4285,plain,(
  ( ! [X28,X29,X27] : (sP6(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(superposition,[],[f2038,f853])).
fof(f2038,plain,(
  ( ! [X136,X137] : (sP6(X136,X137,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f665])).
fof(f4291,plain,(
  ( ! [X26,X24,X25] : (sP6(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4284,f837])).
fof(f4284,plain,(
  ( ! [X26,X24,X25] : (sP6(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(superposition,[],[f2038,f853])).
fof(f4256,plain,(
  ( ! [X28,X29,X27] : (sP7(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4250,f838])).
fof(f4250,plain,(
  ( ! [X28,X29,X27] : (sP7(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(superposition,[],[f2037,f853])).
fof(f2037,plain,(
  ( ! [X134,X135] : (sP7(X134,X135,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f661])).
fof(f4255,plain,(
  ( ! [X26,X24,X25] : (sP7(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4249,f837])).
fof(f4249,plain,(
  ( ! [X26,X24,X25] : (sP7(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(superposition,[],[f2037,f853])).
fof(f4221,plain,(
  ( ! [X28,X29,X27] : (sP8(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4215,f838])).
fof(f4215,plain,(
  ( ! [X28,X29,X27] : (sP8(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(superposition,[],[f2034,f853])).
fof(f2034,plain,(
  ( ! [X128,X129] : (sP8(X128,X129,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f657])).
fof(f4220,plain,(
  ( ! [X26,X24,X25] : (sP8(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4214,f837])).
fof(f4214,plain,(
  ( ! [X26,X24,X25] : (sP8(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(superposition,[],[f2034,f853])).
fof(f4186,plain,(
  ( ! [X28,X29,X27] : (sP4(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4180,f838])).
fof(f4180,plain,(
  ( ! [X28,X29,X27] : (sP4(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(superposition,[],[f2021,f853])).
fof(f2021,plain,(
  ( ! [X103,X102] : (sP4(X102,X103,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f640])).
fof(f4185,plain,(
  ( ! [X26,X24,X25] : (sP4(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4179,f837])).
fof(f4179,plain,(
  ( ! [X26,X24,X25] : (sP4(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(superposition,[],[f2021,f853])).
fof(f4150,plain,(
  ( ! [X28,X29,X27] : (sP2(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4145,f838])).
fof(f4145,plain,(
  ( ! [X28,X29,X27] : (sP2(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(superposition,[],[f2005,f853])).
fof(f2005,plain,(
  ( ! [X70,X71] : (sP2(X70,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X71)) ) | $spl18),
  inference(resolution,[],[f1050,f624])).
fof(f4149,plain,(
  ( ! [X26,X24,X25] : (sP2(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4144,f837])).
fof(f4144,plain,(
  ( ! [X26,X24,X25] : (sP2(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(superposition,[],[f2005,f853])).
fof(f4114,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f4113,f838])).
fof(f4113,plain,(
  ( ! [X5] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(subsumption_resolution,[],[f4092,f794])).
fof(f4092,plain,(
  ( ! [X5] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl32),
  inference(superposition,[],[f1063,f853])).
fof(f4109,plain,(
  ( ! [X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl32),
  inference(subsumption_resolution,[],[f4074,f838])).
fof(f4074,plain,(
  ( ! [X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl32),
  inference(superposition,[],[f1063,f853])).
fof(f4108,plain,(
  ( ! [X8] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl32),
  inference(subsumption_resolution,[],[f4107,f837])).
fof(f4107,plain,(
  ( ! [X8] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl32),
  inference(subsumption_resolution,[],[f4073,f795])).
fof(f795,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f352])).
fof(f352,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))),
  inference(flattening,[],[f351])).
fof(f351,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X11] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X11))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_infty|split|4')).
fof(f4073,plain,(
  ( ! [X8] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl32),
  inference(superposition,[],[f1063,f853])).
fof(f4056,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f855])).
fof(f3976,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X26) != s(bool,X27)) ) | $spl32),
  inference(resolution,[],[f1063,f597])).
fof(f3973,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f596])).
fof(f3972,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f595])).
fof(f3971,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X11) != s(bool,X12)) ) | $spl32),
  inference(resolution,[],[f1063,f593])).
fof(f3969,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl32),
  inference(resolution,[],[f1063,f591])).
fof(f3865,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl155),
  inference(cnf_transformation,[],[f3865_D])).
fof(f3865_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f3867,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl157),
  inference(cnf_transformation,[],[f3867_D])).
fof(f3867_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f3851,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl31),
  inference(superposition,[],[f1062,f808])).
fof(f3850,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl31),
  inference(superposition,[],[f1062,f806])).
fof(f3861,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl151),
  inference(cnf_transformation,[],[f3861_D])).
fof(f3861_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f3863,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl153),
  inference(cnf_transformation,[],[f3863_D])).
fof(f3863_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f3833,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl31),
  inference(superposition,[],[f1062,f807])).
fof(f3832,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl31),
  inference(superposition,[],[f1062,f809])).
fof(f1058,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl27),
  inference(cnf_transformation,[],[f1058_D])).
fof(f1058_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f3827,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl31),
  inference(superposition,[],[f1062,f808])).
fof(f3826,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl31),
  inference(superposition,[],[f1062,f806])).
fof(f3808,plain,(
  ( ! [X28,X29,X27] : (sP2(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3803,f838])).
fof(f3803,plain,(
  ( ! [X28,X29,X27] : (sP2(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(superposition,[],[f2004,f853])).
fof(f2004,plain,(
  ( ! [X68,X69] : (sP2(X68,X69,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f623])).
fof(f3807,plain,(
  ( ! [X26,X24,X25] : (sP2(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3802,f837])).
fof(f3802,plain,(
  ( ! [X26,X24,X25] : (sP2(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(superposition,[],[f2004,f853])).
fof(f3773,plain,(
  ( ! [X28,X29,X27] : (sP0(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3768,f838])).
fof(f3768,plain,(
  ( ! [X28,X29,X27] : (sP0(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X28) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl18),
  inference(superposition,[],[f1991,f853])).
fof(f1991,plain,(
  ( ! [X43,X42] : (sP0(X42,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X43)) ) | $spl18),
  inference(resolution,[],[f1050,f607])).
fof(f3772,plain,(
  ( ! [X26,X24,X25] : (sP0(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3767,f837])).
fof(f3767,plain,(
  ( ! [X26,X24,X25] : (sP0(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl18),
  inference(superposition,[],[f1991,f853])).
fof(f3738,plain,(
  ( ! [X19,X18] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X18) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3732,f838])).
fof(f3732,plain,(
  ( ! [X19,X18] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X18) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl18),
  inference(superposition,[],[f3449,f853])).
fof(f3449,plain,(
  ( ! [X153] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X153)) ) | $spl18),
  inference(resolution,[],[f1975,f704])).
fof(f1975,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl18),
  inference(resolution,[],[f1050,f790])).
fof(f3737,plain,(
  ( ! [X17,X16] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3731,f837])).
fof(f3731,plain,(
  ( ! [X17,X16] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl18),
  inference(superposition,[],[f3449,f853])).
fof(f3701,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f3700,f837])).
fof(f3700,plain,(
  ( ! [X4] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | ($spl33 | $spl36)),
  inference(subsumption_resolution,[],[f3673,f1067])).
fof(f3673,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4)))) ) | $spl33),
  inference(superposition,[],[f1064,f853])).
fof(f3696,plain,(
  ( ! [X9] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl33),
  inference(subsumption_resolution,[],[f3656,f838])).
fof(f3656,plain,(
  ( ! [X9] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl33),
  inference(superposition,[],[f1064,f853])).
fof(f3638,plain,(
  ( ! [X171,X172] : (p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f855])).
fof(f3553,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X24) != s(bool,X26)) ) | $spl33),
  inference(resolution,[],[f1064,f597])).
fof(f3551,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f596])).
fof(f3549,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f595])).
fof(f3548,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,X9) != s(bool,X11)) ) | $spl33),
  inference(resolution,[],[f1064,f593])).
fof(f3546,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl33),
  inference(resolution,[],[f1064,f591])).
fof(f3541,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl33),
  inference(resolution,[],[f1064,f791])).
fof(f3540,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl33),
  inference(resolution,[],[f1064,f791])).
fof(f3538,plain,(
  ( ! [X19,X18] : (sP11(X18,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3532,f838])).
fof(f3532,plain,(
  ( ! [X19,X18] : (sP11(X18,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl18),
  inference(superposition,[],[f3448,f853])).
fof(f3448,plain,(
  ( ! [X152] : (sP11(X152,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1975,f703])).
fof(f3537,plain,(
  ( ! [X17,X16] : (sP11(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3531,f837])).
fof(f3531,plain,(
  ( ! [X17,X16] : (sP11(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl18),
  inference(superposition,[],[f3448,f853])).
fof(f3504,plain,(
  ( ! [X9] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3486,f838])).
fof(f3486,plain,(
  ( ! [X9] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl18),
  inference(superposition,[],[f1975,f853])).
fof(f3503,plain,(
  ( ! [X8] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3485,f837])).
fof(f3485,plain,(
  ( ! [X8] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl18),
  inference(superposition,[],[f1975,f853])).
fof(f3499,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl126),
  inference(cnf_transformation,[],[f3499_D])).
fof(f3499_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f3460,plain,(
  ( ! [X165,X164] : (p(s(bool,X164)) | p(s(bool,X165)) | s(bool,X165) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1975,f855])).
fof(f3380,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X24) != s(bool,X25)) ) | $spl18),
  inference(resolution,[],[f1975,f597])).
fof(f3377,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1975,f596])).
fof(f3376,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1975,f595])).
fof(f3375,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X9) != s(bool,X10)) ) | $spl18),
  inference(resolution,[],[f1975,f593])).
fof(f3373,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1975,f591])).
fof(f3363,plain,(
  ( ! [X9] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3336,f838])).
fof(f3336,plain,(
  ( ! [X9] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl18),
  inference(superposition,[],[f1973,f853])).
fof(f1973,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl18),
  inference(resolution,[],[f1050,f787])).
fof(f3362,plain,(
  ( ! [X8] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl18),
  inference(subsumption_resolution,[],[f3335,f837])).
fof(f3335,plain,(
  ( ! [X8] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl18),
  inference(superposition,[],[f1973,f853])).
fof(f3357,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl122),
  inference(cnf_transformation,[],[f3357_D])).
fof(f3357_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f3318,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl18),
  inference(resolution,[],[f1973,f855])).
fof(f3233,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl18),
  inference(resolution,[],[f1973,f597])).
fof(f3231,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl18),
  inference(resolution,[],[f1973,f596])).
fof(f3229,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl18),
  inference(resolution,[],[f1973,f595])).
fof(f3228,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl18),
  inference(resolution,[],[f1973,f593])).
fof(f3226,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl18),
  inference(resolution,[],[f1973,f591])).
fof(f3220,plain,(
  ( ! [X28,X29,X27] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X27,X28) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3208,f838])).
fof(f3208,plain,(
  ( ! [X28,X29,X27] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X27,X28) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(superposition,[],[f1302,f853])).
fof(f1302,plain,(
  ( ! [X140,X141] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X140,X141)) ) | $spl36),
  inference(resolution,[],[f1067,f667])).
fof(f3219,plain,(
  ( ! [X24,X25] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X24,X25)) ) | $spl118),
  inference(cnf_transformation,[],[f3219_D])).
fof(f3219_D,plain,(
  ( ! [X24,X25] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X24,X25)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f3218,plain,(
  ( ! [X26,X24,X25] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X24,X25) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3207,f837])).
fof(f3207,plain,(
  ( ! [X26,X24,X25] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X24,X25) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(superposition,[],[f1302,f853])).
fof(f3185,plain,(
  ( ! [X28,X29,X27] : (sP6(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3173,f838])).
fof(f3173,plain,(
  ( ! [X28,X29,X27] : (sP6(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(superposition,[],[f1301,f853])).
fof(f1301,plain,(
  ( ! [X138,X139] : (sP6(X138,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X139)) ) | $spl36),
  inference(resolution,[],[f1067,f666])).
fof(f3184,plain,(
  ( ! [X24,X25] : (sP6(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25)) ) | $spl114),
  inference(cnf_transformation,[],[f3184_D])).
fof(f3184_D,plain,(
  ( ! [X24,X25] : (sP6(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f3183,plain,(
  ( ! [X26,X24,X25] : (sP6(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3172,f837])).
fof(f3172,plain,(
  ( ! [X26,X24,X25] : (sP6(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(superposition,[],[f1301,f853])).
fof(f3150,plain,(
  ( ! [X28,X29,X27] : (sP6(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3138,f838])).
fof(f3138,plain,(
  ( ! [X28,X29,X27] : (sP6(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(superposition,[],[f1300,f853])).
fof(f1300,plain,(
  ( ! [X136,X137] : (sP6(X136,X137,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f665])).
fof(f3149,plain,(
  ( ! [X24,X25] : (sP6(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl110),
  inference(cnf_transformation,[],[f3149_D])).
fof(f3149_D,plain,(
  ( ! [X24,X25] : (sP6(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f3148,plain,(
  ( ! [X26,X24,X25] : (sP6(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3137,f837])).
fof(f3137,plain,(
  ( ! [X26,X24,X25] : (sP6(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(superposition,[],[f1300,f853])).
fof(f3115,plain,(
  ( ! [X28,X29,X27] : (sP7(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3104,f838])).
fof(f3104,plain,(
  ( ! [X28,X29,X27] : (sP7(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(superposition,[],[f1299,f853])).
fof(f1299,plain,(
  ( ! [X134,X135] : (sP7(X134,X135,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f661])).
fof(f3114,plain,(
  ( ! [X24,X25] : (sP7(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl106),
  inference(cnf_transformation,[],[f3114_D])).
fof(f3114_D,plain,(
  ( ! [X24,X25] : (sP7(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f3113,plain,(
  ( ! [X26,X24,X25] : (sP7(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3103,f837])).
fof(f3103,plain,(
  ( ! [X26,X24,X25] : (sP7(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(superposition,[],[f1299,f853])).
fof(f3081,plain,(
  ( ! [X28,X29,X27] : (sP8(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3070,f838])).
fof(f3070,plain,(
  ( ! [X28,X29,X27] : (sP8(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(superposition,[],[f1296,f853])).
fof(f1296,plain,(
  ( ! [X128,X129] : (sP8(X128,X129,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f657])).
fof(f3080,plain,(
  ( ! [X24,X25] : (sP8(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl102),
  inference(cnf_transformation,[],[f3080_D])).
fof(f3080_D,plain,(
  ( ! [X24,X25] : (sP8(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f3079,plain,(
  ( ! [X26,X24,X25] : (sP8(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3069,f837])).
fof(f3069,plain,(
  ( ! [X26,X24,X25] : (sP8(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(superposition,[],[f1296,f853])).
fof(f3047,plain,(
  ( ! [X28,X29,X27] : (sP4(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3036,f838])).
fof(f3036,plain,(
  ( ! [X28,X29,X27] : (sP4(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(superposition,[],[f1283,f853])).
fof(f1283,plain,(
  ( ! [X103,X102] : (sP4(X102,X103,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f640])).
fof(f3046,plain,(
  ( ! [X24,X25] : (sP4(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl98),
  inference(cnf_transformation,[],[f3046_D])).
fof(f3046_D,plain,(
  ( ! [X24,X25] : (sP4(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f3045,plain,(
  ( ! [X26,X24,X25] : (sP4(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3035,f837])).
fof(f3035,plain,(
  ( ! [X26,X24,X25] : (sP4(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(superposition,[],[f1283,f853])).
fof(f3012,plain,(
  ( ! [X28,X29,X27] : (sP2(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3002,f838])).
fof(f3002,plain,(
  ( ! [X28,X29,X27] : (sP2(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(superposition,[],[f1267,f853])).
fof(f1267,plain,(
  ( ! [X70,X71] : (sP2(X70,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X71)) ) | $spl36),
  inference(resolution,[],[f1067,f624])).
fof(f3011,plain,(
  ( ! [X24,X25] : (sP2(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25)) ) | $spl94),
  inference(cnf_transformation,[],[f3011_D])).
fof(f3011_D,plain,(
  ( ! [X24,X25] : (sP2(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f3010,plain,(
  ( ! [X26,X24,X25] : (sP2(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(subsumption_resolution,[],[f3001,f837])).
fof(f3001,plain,(
  ( ! [X26,X24,X25] : (sP2(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(superposition,[],[f1267,f853])).
fof(f2978,plain,(
  ( ! [X28,X29,X27] : (sP2(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2968,f838])).
fof(f2968,plain,(
  ( ! [X28,X29,X27] : (sP2(X27,X28,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(superposition,[],[f1266,f853])).
fof(f1266,plain,(
  ( ! [X68,X69] : (sP2(X68,X69,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f623])).
fof(f2977,plain,(
  ( ! [X24,X25] : (sP2(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl90),
  inference(cnf_transformation,[],[f2977_D])).
fof(f2977_D,plain,(
  ( ! [X24,X25] : (sP2(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f2976,plain,(
  ( ! [X26,X24,X25] : (sP2(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2967,f837])).
fof(f2967,plain,(
  ( ! [X26,X24,X25] : (sP2(X24,X25,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(superposition,[],[f1266,f853])).
fof(f2944,plain,(
  ( ! [X28,X29,X27] : (sP0(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2934,f838])).
fof(f2934,plain,(
  ( ! [X28,X29,X27] : (sP0(X27,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X28) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))) ) | $spl36),
  inference(superposition,[],[f1253,f853])).
fof(f1253,plain,(
  ( ! [X43,X42] : (sP0(X42,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X43)) ) | $spl36),
  inference(resolution,[],[f1067,f607])).
fof(f2943,plain,(
  ( ! [X24,X25] : (sP0(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25)) ) | $spl86),
  inference(cnf_transformation,[],[f2943_D])).
fof(f2943_D,plain,(
  ( ! [X24,X25] : (sP0(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f2942,plain,(
  ( ! [X26,X24,X25] : (sP0(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2933,f837])).
fof(f2933,plain,(
  ( ! [X26,X24,X25] : (sP0(X24,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X25) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X26)))) ) | $spl36),
  inference(superposition,[],[f1253,f853])).
fof(f2910,plain,(
  ( ! [X19,X18] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X18) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2899,f838])).
fof(f2899,plain,(
  ( ! [X19,X18] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X18) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl36),
  inference(superposition,[],[f2742,f853])).
fof(f2742,plain,(
  ( ! [X153] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X153)) ) | $spl36),
  inference(resolution,[],[f1237,f704])).
fof(f1237,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl36),
  inference(resolution,[],[f1067,f790])).
fof(f2909,plain,(
  ( ! [X16] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16)) ) | $spl82),
  inference(cnf_transformation,[],[f2909_D])).
fof(f2909_D,plain,(
  ( ! [X16] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2908,plain,(
  ( ! [X17,X16] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2898,f837])).
fof(f2898,plain,(
  ( ! [X17,X16] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl36),
  inference(superposition,[],[f2742,f853])).
fof(f2829,plain,(
  ( ! [X19,X18] : (sP11(X18,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2818,f838])).
fof(f2818,plain,(
  ( ! [X19,X18] : (sP11(X18,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl36),
  inference(superposition,[],[f2741,f853])).
fof(f2741,plain,(
  ( ! [X152] : (sP11(X152,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1237,f703])).
fof(f2827,plain,(
  ( ! [X17,X16] : (sP11(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2817,f837])).
fof(f2817,plain,(
  ( ! [X17,X16] : (sP11(X16,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl36),
  inference(superposition,[],[f2741,f853])).
fof(f2794,plain,(
  ( ! [X9] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2771,f838])).
fof(f2771,plain,(
  ( ! [X9] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl36),
  inference(superposition,[],[f1237,f853])).
fof(f2793,plain,(
  ( ! [X8] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2770,f837])).
fof(f2770,plain,(
  ( ! [X8] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl36),
  inference(superposition,[],[f1237,f853])).
fof(f2790,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl72),
  inference(cnf_transformation,[],[f2790_D])).
fof(f2790_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2753,plain,(
  ( ! [X165,X164] : (p(s(bool,X164)) | p(s(bool,X165)) | s(bool,X165) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1237,f855])).
fof(f2673,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X24) != s(bool,X25)) ) | $spl36),
  inference(resolution,[],[f1237,f597])).
fof(f2791,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl74),
  inference(cnf_transformation,[],[f2791_D])).
fof(f2791_D,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f1040,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,X7) | p(s(bool,X7))) ) | $spl10),
  inference(cnf_transformation,[],[f1040_D])).
fof(f1040_D,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,X7) | p(s(bool,X7))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2672,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X21) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1237,f597])).
fof(f2670,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1237,f596])).
fof(f2669,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1237,f595])).
fof(f2668,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X9) != s(bool,X10)) ) | $spl36),
  inference(resolution,[],[f1237,f593])).
fof(f2667,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1237,f593])).
fof(f2666,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1237,f591])).
fof(f2658,plain,(
  ( ! [X9] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2633,f838])).
fof(f2633,plain,(
  ( ! [X9] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl36),
  inference(superposition,[],[f1235,f853])).
fof(f2655,plain,(
  ( ! [X8] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl36),
  inference(subsumption_resolution,[],[f2632,f837])).
fof(f2632,plain,(
  ( ! [X8] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl36),
  inference(superposition,[],[f1235,f853])).
fof(f2646,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl64),
  inference(cnf_transformation,[],[f2646_D])).
fof(f2646_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2609,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl36),
  inference(resolution,[],[f1235,f855])).
fof(f2524,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X21) != s(bool,X23)) ) | $spl36),
  inference(resolution,[],[f1235,f597])).
fof(f2522,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl36),
  inference(resolution,[],[f1235,f596])).
fof(f2647,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl66),
  inference(cnf_transformation,[],[f2647_D])).
fof(f2647_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2521,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X12) != s(bool,X13)) ) | $spl36),
  inference(resolution,[],[f1235,f595])).
fof(f2520,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl36),
  inference(resolution,[],[f1235,f595])).
fof(f2519,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X6) != s(bool,X8)) ) | $spl36),
  inference(resolution,[],[f1235,f593])).
fof(f2518,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X3) != s(bool,X4)) ) | $spl36),
  inference(resolution,[],[f1235,f591])).
fof(f2517,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl36),
  inference(resolution,[],[f1235,f591])).
fof(f2512,plain,(
  ( ! [X9] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl18),
  inference(subsumption_resolution,[],[f2511,f838])).
fof(f2511,plain,(
  ( ! [X9] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl18),
  inference(subsumption_resolution,[],[f2496,f793])).
fof(f2496,plain,(
  ( ! [X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl18),
  inference(superposition,[],[f1050,f853])).
fof(f2102,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl56),
  inference(cnf_transformation,[],[f2102_D])).
fof(f2102_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2470,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f855])).
fof(f2390,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X26) != s(bool,X27)) ) | $spl18),
  inference(resolution,[],[f1050,f597])).
fof(f2103,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl58),
  inference(cnf_transformation,[],[f2103_D])).
fof(f2103_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2389,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f597])).
fof(f2387,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f596])).
fof(f2386,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f595])).
fof(f2385,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X11) != s(bool,X12)) ) | $spl18),
  inference(resolution,[],[f1050,f593])).
fof(f2384,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f593])).
fof(f2383,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f591])).
fof(f2372,plain,(
  ( ! [X19,X18] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X18) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl18),
  inference(subsumption_resolution,[],[f2366,f838])).
fof(f2366,plain,(
  ( ! [X19,X18] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X18) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl18),
  inference(superposition,[],[f2052,f853])).
fof(f2052,plain,(
  ( ! [X155] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X155)) ) | $spl18),
  inference(resolution,[],[f1050,f704])).
fof(f2340,plain,(
  ( ! [X19,X18] : (sP11(X18,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl18),
  inference(subsumption_resolution,[],[f2334,f838])).
fof(f2334,plain,(
  ( ! [X19,X18] : (sP11(X18,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl18),
  inference(superposition,[],[f2051,f853])).
fof(f2051,plain,(
  ( ! [X154] : (sP11(X154,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f703])).
fof(f2297,plain,(
  $false | $spl19),
  inference(subsumption_resolution,[],[f2296,f837])).
fof(f2296,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl19),
  inference(subsumption_resolution,[],[f2157,f794])).
fof(f2157,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl19),
  inference(resolution,[],[f1049,f791])).
fof(f1049,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl19),
  inference(cnf_transformation,[],[f1049_D])).
fof(f1049_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f2295,plain,(
  $false | $spl19),
  inference(subsumption_resolution,[],[f2294,f837])).
fof(f2294,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl19),
  inference(subsumption_resolution,[],[f2156,f794])).
fof(f2156,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl19),
  inference(resolution,[],[f1049,f791])).
fof(f2293,plain,(
  $false | $spl19),
  inference(subsumption_resolution,[],[f2155,f837])).
fof(f2155,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl19),
  inference(resolution,[],[f1049,f798])).
fof(f2108,plain,(
  ( ! [X9] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl18),
  inference(subsumption_resolution,[],[f2107,f838])).
fof(f2107,plain,(
  ( ! [X9] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl18),
  inference(subsumption_resolution,[],[f2089,f793])).
fof(f2089,plain,(
  ( ! [X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) ) | $spl18),
  inference(superposition,[],[f1050,f853])).
fof(f2063,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f855])).
fof(f1983,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X26) != s(bool,X27)) ) | $spl18),
  inference(resolution,[],[f1050,f597])).
fof(f1982,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f597])).
fof(f1980,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f596])).
fof(f1979,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f595])).
fof(f1978,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X11) != s(bool,X12)) ) | $spl18),
  inference(resolution,[],[f1050,f593])).
fof(f1977,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f593])).
fof(f1976,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl18),
  inference(resolution,[],[f1050,f591])).
fof(f1960,plain,(
  ( ! [X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11)) ) | $spl52),
  inference(cnf_transformation,[],[f1960_D])).
fof(f1960_D,plain,(
  ( ! [X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1963,plain,(
  ( ! [X19,X18] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X18) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl36),
  inference(subsumption_resolution,[],[f1951,f838])).
fof(f1951,plain,(
  ( ! [X19,X18] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X18) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl36),
  inference(superposition,[],[f1314,f853])).
fof(f1314,plain,(
  ( ! [X155] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')),X155)) ) | $spl36),
  inference(resolution,[],[f1067,f704])).
fof(f1961,plain,(
  ( ! [X17,X16] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl36),
  inference(subsumption_resolution,[],[f1950,f837])).
fof(f1950,plain,(
  ( ! [X17,X16] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X16) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl36),
  inference(superposition,[],[f1314,f853])).
fof(f1050,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl18),
  inference(cnf_transformation,[],[f1050_D])).
fof(f1050_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1947,plain,(
  ( ! [X11] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl36),
  inference(superposition,[],[f1314,f798])).
fof(f1928,plain,(
  ( ! [X11] : (sP11(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl48),
  inference(cnf_transformation,[],[f1928_D])).
fof(f1928_D,plain,(
  ( ! [X11] : (sP11(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1931,plain,(
  ( ! [X19,X18] : (sP11(X18,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl36),
  inference(subsumption_resolution,[],[f1919,f838])).
fof(f1919,plain,(
  ( ! [X19,X18] : (sP11(X18,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19)))) ) | $spl36),
  inference(superposition,[],[f1313,f853])).
fof(f1313,plain,(
  ( ! [X154] : (sP11(X154,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f703])).
fof(f1929,plain,(
  ( ! [X17,X16] : (sP11(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl36),
  inference(subsumption_resolution,[],[f1918,f837])).
fof(f1918,plain,(
  ( ! [X17,X16] : (sP11(X16,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))) ) | $spl36),
  inference(superposition,[],[f1313,f853])).
fof(f1915,plain,(
  ( ! [X11] : (sP11(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl36),
  inference(superposition,[],[f1313,f798])).
fof(f1899,plain,(
  $false | $spl22),
  inference(equality_resolution,[],[f1881])).
fof(f1881,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f1868])).
fof(f1868,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl22),
  inference(superposition,[],[f1054,f840])).
fof(f1861,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1776,f1651])).
fof(f1651,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl12),
  inference(equality_resolution,[],[f1041])).
fof(f1776,plain,(
  ( ! [X296,X295] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X295),s('type.extreal.extreal',X296))))) ) | $spl12),
  inference(resolution,[],[f1651,f788])).
fof(f1860,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1765,f1651])).
fof(f1765,plain,(
  ( ! [X280,X281] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X280),s('type.extreal.extreal',X281))))) ) | $spl12),
  inference(resolution,[],[f1651,f788])).
fof(f1770,plain,(
  $false | ($spl12 | $spl36)),
  inference(resolution,[],[f1651,f1067])).
fof(f1771,plain,(
  $false | $spl12),
  inference(resolution,[],[f1651,f801])).
fof(f801,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X12] : p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_infty|split|1')).
fof(f1773,plain,(
  $false | $spl12),
  inference(resolution,[],[f1651,f800])).
fof(f800,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))) )),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X12] : p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_infty|split|0')).
fof(f1774,plain,(
  $false | $spl12),
  inference(resolution,[],[f1651,f802])).
fof(f802,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(cnf_transformation,[],[f54])).
fof(f54,axiom,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_infty|split|2')).
fof(f1777,plain,(
  $false | $spl12),
  inference(resolution,[],[f1651,f814])).
fof(f1799,plain,(
  $false | $spl12),
  inference(resolution,[],[f1651,f813])).
fof(f813,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(cnf_transformation,[],[f127])).
fof(f127,axiom,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.extreal.lt_|48||49|','thm.real.REAL_LT_|48||49|')).
fof(f1800,plain,(
  $false | $spl12),
  inference(resolution,[],[f1651,f815])).
fof(f815,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(cnf_transformation,[],[f72])).
fof(f72,axiom,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.extreal.lt_|48||49|','thm.real.REAL_LE_|48||49|')).
fof(f1617,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl46),
  inference(cnf_transformation,[],[f1617_D])).
fof(f1617_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1613,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1591])).
fof(f1591,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30)) | p(s(bool,X30))) ) | $spl6),
  inference(superposition,[],[f1038,f765])).
fof(f765,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f558])).
fof(f558,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : (s(bool,'F_0') = s(bool,X6) <=> ~p(s(bool,X6)))),
  file('thm.extreal.lt_|48||49|','thm.bool.EQ_CLAUSES|split|3')).
fof(f1614,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f1590])).
fof(f1590,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl6),
  inference(superposition,[],[f1038,f763])).
fof(f763,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f557])).
fof(f557,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f313])).
fof(f313,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : (s(bool,'F_0') = s(bool,X6) <=> ~p(s(bool,X6)))),
  file('thm.extreal.lt_|48||49|','thm.bool.EQ_CLAUSES|split|2')).
fof(f1581,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1492,f1043])).
fof(f1492,plain,(
  ( ! [X298,X299] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X298),s('type.extreal.extreal',X299))))) ) | $spl16),
  inference(resolution,[],[f1043,f787])).
fof(f1580,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1484,f1043])).
fof(f1484,plain,(
  ( ! [X290,X291] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X290),s('type.extreal.extreal',X291))))) ) | $spl16),
  inference(resolution,[],[f1043,f787])).
fof(f1579,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1483,f1043])).
fof(f1483,plain,(
  ( ! [X288,X289] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X288),s('type.extreal.extreal',X289))))) ) | $spl16),
  inference(resolution,[],[f1043,f786])).
fof(f786,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f445])).
fof(f445,plain,(
  ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(ennf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X11,X12] : ~(p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X12)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X11)))))),
  file('thm.extreal.lt_|48||49|','thm.extreal.lt_antisym')).
fof(f1578,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1482,f1043])).
fof(f1482,plain,(
  ( ! [X287,X286] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X286),s('type.extreal.extreal',X287))))) ) | $spl16),
  inference(resolution,[],[f1043,f786])).
fof(f1577,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1468,f1043])).
fof(f1468,plain,(
  ( ! [X264] : (~p(s(bool,X264))) ) | $spl16),
  inference(resolution,[],[f1043,f773])).
fof(f773,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f425])).
fof(f425,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f323])).
fof(f323,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X6] : ((p(s(bool,X6)) => p(s(bool,'F_0'))) => ~p(s(bool,X6)))),
  file('thm.extreal.lt_|48||49|','thm.bool.IMP_F')).
fof(f1576,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1463,f1043])).
fof(f1463,plain,(
  ( ! [X259] : (~p(s(bool,X259))) ) | $spl16),
  inference(resolution,[],[f1043,f742])).
fof(f742,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f544])).
fof(f544,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f543])).
fof(f543,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f419])).
fof(f419,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X6] : ((p(s(bool,X6)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X6)))),
  file('thm.extreal.lt_|48||49|','thm.bool.IMP_CLAUSES|split|4')).
fof(f1470,plain,(
  $false | $spl16),
  inference(resolution,[],[f1043,f782])).
fof(f782,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.extreal.lt_|48||49|',aHL_FALSITY)).
fof(f1477,plain,(
  $false | $spl16),
  inference(resolution,[],[f1043,f793])).
fof(f1488,plain,(
  $false | $spl16),
  inference(resolution,[],[f1043,f794])).
fof(f1489,plain,(
  $false | $spl16),
  inference(resolution,[],[f1043,f589])).
fof(f1490,plain,(
  $false | $spl16),
  inference(resolution,[],[f1043,f795])).
fof(f1366,plain,(
  ( ! [X8] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl36),
  inference(subsumption_resolution,[],[f1365,f837])).
fof(f1365,plain,(
  ( ! [X8] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl36),
  inference(subsumption_resolution,[],[f1342,f793])).
fof(f1342,plain,(
  ( ! [X8] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X8)))) ) | $spl36),
  inference(superposition,[],[f1067,f853])).
fof(f1362,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl42),
  inference(cnf_transformation,[],[f1362_D])).
fof(f1362_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1325,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X167) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f855])).
fof(f1245,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X26) != s(bool,X27)) ) | $spl36),
  inference(resolution,[],[f1067,f597])).
fof(f1363,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl44),
  inference(cnf_transformation,[],[f1363_D])).
fof(f1363_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1244,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f597])).
fof(f1242,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f596])).
fof(f1241,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f595])).
fof(f1240,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X11) != s(bool,X12)) ) | $spl36),
  inference(resolution,[],[f1067,f593])).
fof(f1239,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X9) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f593])).
fof(f1238,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(resolution,[],[f1067,f591])).
fof(f1215,plain,(
  $false | $spl37),
  inference(subsumption_resolution,[],[f1214,f838])).
fof(f1214,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl37),
  inference(subsumption_resolution,[],[f1077,f795])).
fof(f1077,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl37),
  inference(resolution,[],[f1068,f791])).
fof(f1213,plain,(
  $false | $spl37),
  inference(subsumption_resolution,[],[f1212,f838])).
fof(f1212,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl37),
  inference(subsumption_resolution,[],[f1076,f795])).
fof(f1076,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl37),
  inference(resolution,[],[f1068,f791])).
fof(f1211,plain,(
  $false | $spl37),
  inference(subsumption_resolution,[],[f1075,f838])).
fof(f1075,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl37),
  inference(resolution,[],[f1068,f796])).
fof(f1068,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl37),
  inference(cnf_transformation,[],[f1068_D])).
fof(f1068_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f1069,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(subsumption_resolution,[],[f1019,f837])).
fof(f1019,plain,(
  ( ! [X4] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X4))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f589,f853])).
fof(f1066,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl34),
  inference(cnf_transformation,[],[f1066_D])).
fof(f1066_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1015,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(superposition,[],[f589,f796])).
fof(f1062,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl31),
  inference(cnf_transformation,[],[f1062_D])).
fof(f1062_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f1008,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f589,f808])).
fof(f1007,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f589,f806])).
fof(f1060,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl29),
  inference(cnf_transformation,[],[f1060_D])).
fof(f1060_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f996,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f589,f808])).
fof(f995,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f589,f806])).
fof(f1052,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl21),
  inference(cnf_transformation,[],[f1052_D])).
fof(f1052_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1053,plain,(
  ( ! [X9] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) )),
  inference(subsumption_resolution,[],[f978,f838])).
fof(f978,plain,(
  ( ! [X9] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK43('const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))) )),
  inference(superposition,[],[f589,f853])).
fof(f974,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f589,f798])).
fof(f1037,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl4),
  inference(cnf_transformation,[],[f1037_D])).
fof(f1037_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f960,plain,(
  ( ! [X171,X172] : (p(s(bool,X171)) | ~p(s(bool,X172)) | s(bool,X172) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(resolution,[],[f589,f855])).
fof(f875,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f589,f597])).
fof(f873,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(resolution,[],[f589,f596])).
fof(f1039,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl8),
  inference(cnf_transformation,[],[f1039_D])).
fof(f1039_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f872,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f589,f595])).
fof(f871,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(resolution,[],[f589,f595])).
fof(f870,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f589,f593])).
fof(f869,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,X6) != s(bool,X7)) )),
  inference(resolution,[],[f589,f591])).
fof(f868,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(resolution,[],[f589,f591])).
fof(f1034,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl0),
  inference(cnf_transformation,[],[f1034_D])).
fof(f1034_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1036,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl2),
  inference(cnf_transformation,[],[f1036_D])).
fof(f1036_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f863,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(resolution,[],[f589,f791])).
fof(f862,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(resolution,[],[f589,f791])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 372
% Passive clauses: 7312
% Generated clauses: 13412
% Final active clauses: 305
% Final passive clauses: 3715
% Input formulas: 261
% Initial clauses: 266
% 
% Selected by SInE selection: 98
% SInE iterations: 1
% 
% Duplicate literals: 112
% Trivial inequalities: 88
% Fw subsumption resolutions: 1058
% Fw demodulations: 322
% Bw demodulations: 17
% 
% Simple tautologies: 199
% Equational tautologies: 1
% Forward subsumptions: 3261
% 
% Binary resolution: 4972
% Forward superposition: 5360
% Backward superposition: 372
% Self superposition: 1
% Equality resolution: 6
% 
% Split clauses: 537
% Split components: 1266
% SAT solver clauses: 577
% SAT solver unit clauses: 33
% SAT solver binary clauses: 26
% 
% Sat splits: 537
% Sat splitting refutations: 35
% 
% Memory used [KB]: 5756
% Time elapsed: 0.174 s
% ------------------------------
% ------------------------------
% Success in time 11.955 s
