Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (91962)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 128
% Passive clauses: 1301
% Generated clauses: 24308
% Final active clauses: 124
% Final passive clauses: 1090
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 1632
% Trivial inequalities: 35
% Fw subsumption resolutions: 2519
% Fw demodulations: 4173
% Bw demodulations: 2
% 
% Simple tautologies: 211
% Equational tautologies: 18
% Forward subsumptions: 13287
% 
% Binary resolution: 11
% Forward superposition: 9348
% Backward superposition: 6325
% Self superposition: 172
% Equality factoring: 34
% Equality resolution: 4
% 
% Split clauses: 660
% Split components: 1325
% SAT solver clauses: 1212
% SAT solver unit clauses: 1
% SAT solver binary clauses: 23
% 
% Sat splits: 660
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1212
% TWLsolver calls for satisfiability: 90
% 
% Memory used [KB]: 15095
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (91983)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 110
% Passive clauses: 899
% Generated clauses: 19223
% Final active clauses: 106
% Final passive clauses: 780
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 1072
% Trivial inequalities: 37
% Fw subsumption resolutions: 1134
% Fw demodulations: 2476
% Bw demodulations: 9
% Condensations: 1836
% 
% Simple tautologies: 162
% Equational tautologies: 21
% Forward subsumptions: 10021
% 
% Binary resolution: 9
% Forward superposition: 5939
% Backward superposition: 5576
% Self superposition: 163
% Equality factoring: 44
% Equality resolution: 3
% 
% Split clauses: 18
% Split components: 41
% SAT solver clauses: 1271
% SAT solver unit clauses: 38
% SAT solver binary clauses: 20
% 
% Sat splits: 18
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1271
% TWLsolver calls for satisfiability: 121
% 
% Memory used [KB]: 17910
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% (92003)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 286
% Passive clauses: 6072
% Generated clauses: 160521
% Final active clauses: 101
% Final passive clauses: 589
% Input formulas: 21
% Initial clauses: 22
% Discarded non-redundant clauses: 8591
% 
% Duplicate literals: 7003
% Trivial inequalities: 894
% Fw subsumption resolutions: 7685
% Fw demodulations: 17775
% Bw demodulations: 3138
% Condensations: 13968
% Global subsumptions: 63
% 
% Simple tautologies: 990
% Equational tautologies: 509
% Forward subsumptions: 102754
% Bw demodulations to eq. taut.: 389
% 
% Unit resulting resolution: 1177
% Forward superposition: 36128
% Backward superposition: 39120
% Self superposition: 867
% Equality factoring: 367
% Equality resolution: 33
% 
% Split clauses: 331
% Split components: 669
% SAT solver clauses: 6863
% SAT solver unit clauses: 50
% SAT solver binary clauses: 217
% 
% Sat splits: 331
% Sat splitting refutations: 64
% 
% Memory used [KB]: 74710
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 7
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 7
% Input formulas: 21
% Initial clauses: 9
% 
% Selected by SInE selection: 8
% SInE iterations: 1
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 1
% 
% Memory used [KB]: 767
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.ringNorm
% (92287)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 87
% Passive clauses: 1256
% Generated clauses: 13421
% Final active clauses: 48
% Final passive clauses: 243
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 241
% Trivial inequalities: 201
% Fw subsumption resolutions: 996
% Fw demodulations: 1684
% Bw demodulations: 439
% Global subsumptions: 5
% 
% Simple tautologies: 247
% Equational tautologies: 65
% Forward subsumptions: 8403
% Bw demodulations to eq. taut.: 146
% 
% Unit resulting resolution: 316
% Forward superposition: 4237
% Backward superposition: 4179
% Self superposition: 143
% Equality factoring: 42
% Equality resolution: 13
% 
% Split clauses: 13
% Split components: 26
% SAT solver clauses: 1321
% SAT solver unit clauses: 45
% SAT solver binary clauses: 149
% 
% Sat splits: 13
% Sat splitting refutations: 44
% 
% Memory used [KB]: 12792
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 20
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 20
% Input formulas: 21
% Initial clauses: 22
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 3
% SAT solver unit clauses: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 20
% Passive clauses: 1
% Generated clauses: 2
% Final active clauses: 20
% Input formulas: 21
% Initial clauses: 22
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 5
% SAT solver unit clauses: 2
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 20
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 20
% Input formulas: 21
% Initial clauses: 22
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 1
% SAT solver unit clauses: 1
% 
% TWLsolver clauses: 1
% TWLsolver calls for satisfiability: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 115
lrs+11_28_cond=on:gs=on:gsaa=from_current:gsem=on:gsssp=full:lwlo=on:nm=64:nwc=1.7:stl=30:sac=on:sdd=off:sfr=on:ssfp=40000:ssfq=1.1:smm=off:ssnc=none:urr=on:updr=off_109 on thm.ringNorm
% (92296)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 110
% Passive clauses: 923
% Generated clauses: 37183
% Final active clauses: 97
% Final passive clauses: 192
% Input formulas: 21
% Initial clauses: 22
% Discarded non-redundant clauses: 8658
% 
% Duplicate literals: 3744
% Trivial inequalities: 320
% Fw subsumption resolutions: 1994
% Fw demodulations: 1566
% Bw demodulations: 141
% Condensations: 7177
% Global subsumptions: 1
% 
% Simple tautologies: 450
% Equational tautologies: 988
% Forward subsumptions: 19103
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 31
% Unit resulting resolution: 2580
% Forward superposition: 10414
% Backward superposition: 8341
% Self superposition: 568
% Equality factoring: 422
% Equality resolution: 28
% 
% Split clauses: 13
% Split components: 27
% SAT solver clauses: 999
% SAT solver unit clauses: 43
% SAT solver binary clauses: 43
% SAT solver learnt clauses: 38
% SAT solver learnt literals: 188
% 
% Sat splits: 13
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 940
% TWLsolver calls for satisfiability: 47417
% 
% Memory used [KB]: 14711
% Time elapsed: 11.600 s
% ------------------------------
% ------------------------------
% remaining time: 59 next slice time: 59
ott+11_3:1_br=off:gs=on:gsem=on:nwc=1:sas=minisat:sos=all:spl=off:urr=on:updr=off_135 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 20
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 20
% Input formulas: 21
% Initial clauses: 22
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 2
% SAT solver unit clauses: 1
% 
% Memory used [KB]: 1023
% Time elapsed: 0.027 s
% ------------------------------
% ------------------------------
% remaining time: 59 next slice time: 59
lrs+3_5_bd=preordered:bsr=unit_only:gsp=input_only:gs=on:gsem=on:gsssp=full:lwlo=on:nm=64:nwc=1:sas=minisat:stl=90:spl=off:sp=occurrence:urr=ec_only:updr=off_815 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f27364,plain,(
  $false),
  inference(subsumption_resolution,[],[f65,f27363])).
fof(f27363,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),X5))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X5)))))) )),
  inference(forward_demodulation,[],[f27362,f16397])).
fof(f16397,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f16274,f78])).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X19,X20,X21] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X6),X19),s(X6,X20),s('type.canonical.canonical_sum'(X6),X21))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun(X6,fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X6),X19))),s(X6,X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|0','ARITY_1_3_const.ringNorm.r_canonical_sum_scalar')).
fof(f16274,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f79,f69])).
fof(f69,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ! [X1] : s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun(X6,fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_scalar_1'(s('type.ring.ring'(X6),X7))) = s(fun(X6,fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|0','thm.ringNorm.r_canonical_sum_scalar_def')).
fof(f79,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X19,X20,X21] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X6),X19),s(X6,X20),s('type.canonical.canonical_sum'(X6),X21))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun(X6,fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X6),X19))),s(X6,X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|0','ARITY_1_3_const.canonical.canonical_sum_scalar')).
fof(f27362,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),X5))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X5)))))) )),
  inference(forward_demodulation,[],[f27088,f16397])).
fof(f27088,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),X5))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s('type.canonical.canonical_sum'(X0),X5)))))) )),
  inference(superposition,[],[f85,f21573])).
fof(f21573,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s(X0,X3)))) )),
  inference(forward_demodulation,[],[f21543,f77])).
fof(f77,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X19,X20,X21] : s(X6,'const.ring.ring_RM_3'(s('type.ring.ring'(X6),X19),s(X6,X20),s(X6,X21))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X19))),s(X6,X20))),s(X6,X21)))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|0','ARITY_1_3_const.ring.ring_RM')).
fof(f21543,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(superposition,[],[f74,f20534])).
fof(f20534,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(superposition,[],[f83,f70])).
fof(f70,plain,(
  ( ! [X0,X1] : (s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1] : s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7))) = s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X7))),s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X7))),s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X7))),s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|0','thm.ring.semi_ring_of_def')).
fof(f83,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X15,X16,X17,X18] : s(fun(X6,fun(X6,X6)),X18) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X15),s(X6,X16),s(fun(X6,fun(X6,X6)),X17),s(fun(X6,fun(X6,X6)),X18)))))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|0','thm.semi_ring.semi_ring_SRM')).
fof(f74,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X19,X20,X21] : s(X6,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X6),X19),s(X6,X20),s(X6,X21))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X19))),s(X6,X20))),s(X6,X21)))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|0','ARITY_1_3_const.semi_ring.semi_ring_SRM')).
fof(f85,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),X5))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X5)))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),X5))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X5)))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),X5))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X5)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X14,X8,X9,X10,X11] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X6),X14),s(X6,X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X9),s('type.list.list'('type.quote.index'),X10),s('type.canonical.canonical_sum'(X6),X11))))) = s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X6),X14),s(X6,X8),s(X6,X9))),s('type.list.list'('type.quote.index'),X10),s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X6),X14),s(X6,X8),s('type.canonical.canonical_sum'(X6),X11)))))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|0','thm.canonical.canonical_sum_scalar_def|split|0')).
fof(f65,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK3),s('type.list.list'('type.quote.index'),sK4),s('type.canonical.canonical_sum'(sK0),sK5))))) != s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,'const.ring.ring_RM_3'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s(sK0,sK3))),s('type.list.list'('type.quote.index'),sK4),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s('type.canonical.canonical_sum'(sK0),sK5)))))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK3),s('type.list.list'('type.quote.index'),sK4),s('type.canonical.canonical_sum'(sK0),sK5))))) != s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,'const.ring.ring_RM_3'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s(sK0,sK3))),s('type.list.list'('type.quote.index'),sK4),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(sK0),sK1),s(sK0,sK2),s('type.canonical.canonical_sum'(sK0),sK5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f60])).
fof(f60,plain,(
  ? [X0,X1,X2,X3,X4,X5] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),X5))))) != s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X5)))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),X5))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X5)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),X5))))) = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))),s('type.list.list'('type.quote.index'),X4),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X5)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X6),X7),s(X6,X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X9),s('type.list.list'('type.quote.index'),X10),s('type.canonical.canonical_sum'(X6),X11))))) = s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,'const.ring.ring_RM_3'(s('type.ring.ring'(X6),X7),s(X6,X8),s(X6,X9))),s('type.list.list'('type.quote.index'),X10),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X6),X7),s(X6,X8),s('type.canonical.canonical_sum'(X6),X11)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X6),X7),s(X6,X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X9),s('type.list.list'('type.quote.index'),X10),s('type.canonical.canonical_sum'(X6),X11))))) = s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,'const.ring.ring_RM_3'(s('type.ring.ring'(X6),X7),s(X6,X8),s(X6,X9))),s('type.list.list'('type.quote.index'),X10),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar_3'(s('type.ring.ring'(X6),X7),s(X6,X8),s('type.canonical.canonical_sum'(X6),X11)))))),
  file('thm.ringNorm.canonical_sum_scalar_def|split|0','thm.ringNorm.canonical_sum_scalar_def|split|0')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 80
% Passive clauses: 865
% Generated clauses: 26490
% Final active clauses: 79
% Final passive clauses: 785
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 2004
% Trivial inequalities: 320
% Fw subsumption resolutions: 2992
% Bw subsumption resolutions: 1
% Fw demodulations: 317
% 
% Simple tautologies: 459
% Equational tautologies: 173
% Forward subsumptions: 19086
% 
% Binary resolution: 47
% Forward superposition: 9541
% Backward superposition: 10003
% Self superposition: 653
% Equality factoring: 573
% Equality resolution: 17
% 
% SAT solver clauses: 866
% SAT solver unit clauses: 31
% SAT solver binary clauses: 17
% 
% Memory used [KB]: 12537
% Time elapsed: 5.030 s
% ------------------------------
% ------------------------------
% Success in time 29.083 s
