Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f3670,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3669,f3594,f177,f275,f273,f176,f274,f272])).
fof(f272,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) | $spl0),
  inference(cnf_transformation,[],[f272_D])).
fof(f272_D,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f274,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl2),
  inference(cnf_transformation,[],[f274_D])).
fof(f274_D,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f176,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) = s('type.num.num',sK1)) & (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) != s('type.num.num',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f160])).
fof(f160,plain,(
  ? [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ? [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ~! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.numRing.num_rewrites|split|19','thm.numRing.num_rewrites|split|19')).
fof(f273,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) | $spl1),
  inference(cnf_transformation,[],[f273_D])).
fof(f273_D,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f275,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | $spl3),
  inference(cnf_transformation,[],[f275_D])).
fof(f275_D,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f177,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f161])).
fof(f3594,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3538,f275])).
fof(f3538,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(equality_resolution,[],[f3237])).
fof(f3237,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))) | s('type.num.num',sK1) = s('type.num.num',X3)) ) | $spl0),
  inference(superposition,[],[f268,f272])).
fof(f268,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X10,X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) <=> s('type.num.num',X7) = s('type.num.num',X10))),
  file('thm.numRing.num_rewrites|split|19','thm.numeral.numeral_distrib|split|19')).
fof(f3669,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f3668])).
fof(f3668,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f273,f274])).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 122
% Passive clauses: 377
% Generated clauses: 3457
% Final active clauses: 82
% Final passive clauses: 192
% Input formulas: 81
% Initial clauses: 95
% 
% Duplicate literals: 1
% Trivial inequalities: 22
% Fw subsumption resolutions: 36
% Fw demodulations: 616
% Bw demodulations: 28
% 
% Simple tautologies: 8
% Equational tautologies: 9
% Forward subsumptions: 2385
% Fw demodulations to eq. taut.: 10
% 
% Binary resolution: 7
% Forward superposition: 1261
% Backward superposition: 1325
% Self superposition: 26
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 32
% Split components: 64
% SAT solver clauses: 42
% SAT solver unit clauses: 2
% SAT solver binary clauses: 20
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 32
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 40
% TWLsolver calls for satisfiability: 14
% 
% Memory used [KB]: 2174
% Time elapsed: 0.059 s
% ------------------------------
% ------------------------------
% Success in time 0.06 s
